Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relation
% (56991)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 157
% Passive clauses: 3044
% Generated clauses: 16959
% Final active clauses: 132
% Final passive clauses: 2350
% Input formulas: 30
% Initial clauses: 70
% 
% Duplicate literals: 750
% Trivial inequalities: 390
% Fw subsumption resolutions: 2335
% Fw demodulations: 2226
% Bw demodulations: 17
% 
% Simple tautologies: 389
% Equational tautologies: 57
% Forward subsumptions: 6382
% Fw demodulations to eq. taut.: 18
% 
% Binary resolution: 419
% Factoring: 2
% Forward superposition: 4410
% Backward superposition: 5900
% Self superposition: 167
% Equality factoring: 130
% Equality resolution: 40
% 
% Split clauses: 552
% Split components: 1148
% SAT solver clauses: 1009
% SAT solver unit clauses: 29
% SAT solver binary clauses: 184
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 552
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 966
% TWLsolver calls for satisfiability: 185
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relation
% (57029)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 193
% Passive clauses: 1451
% Generated clauses: 15132
% Final active clauses: 121
% Final passive clauses: 935
% Input formulas: 30
% Initial clauses: 70
% 
% Duplicate literals: 229
% Trivial inequalities: 310
% Fw subsumption resolutions: 1135
% Fw demodulations: 2192
% Bw demodulations: 299
% Condensations: 426
% Global subsumptions: 5
% 
% Simple tautologies: 151
% Equational tautologies: 27
% Forward subsumptions: 6885
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 22
% 
% Binary resolution: 224
% Factoring: 1
% Forward superposition: 3610
% Backward superposition: 4022
% Self superposition: 112
% Equality factoring: 44
% Equality resolution: 38
% 
% Split clauses: 112
% Split components: 227
% SAT solver clauses: 2206
% SAT solver unit clauses: 76
% SAT solver binary clauses: 82
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 5
% 
% Sat splits: 112
% Sat splitting refutations: 47
% 
% TWLsolver clauses: 2182
% TWLsolver calls for satisfiability: 447
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relation
% SZS output start Proof for thm.clos_relation
fof(f13277,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f125,f213,f124,f215,f211,f123,f216,f210,f492,f495,f516,f531,f532,f534,f538,f871,f873,f979,f981,f983,f985,f991,f993,f995,f999,f1001,f1003,f1005,f1007,f1009,f1013,f1015,f1017,f1019,f1021,f1023,f1025,f1079,f1080,f1082,f1085,f1086,f1087,f1088,f1089,f1090,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1719,f1727,f1751,f1752,f4196,f4183,f4239,f4176,f4167,f4224,f4113,f4429,f4055,f4515,f4810,f4811,f4813,f4815,f4816,f4818,f4820,f4824,f4826,f4830,f4832,f4834,f4836,f4837,f4839,f4840,f4842,f4843,f4845,f4848,f4854,f4859,f4861,f4863,f4865,f4867,f4886,f4890,f4904,f4910,f4915,f4924,f4926,f6466,f6492,f6465,f6493,f6383,f6382,f6504,f6531,f6532,f6616,f6534,f6618,f6612,f6535,f6619,f6615,f6639,f6640,f6641,f6642,f6643,f6644,f6651,f6815,f6817,f7667,f7720,f8149,f8151,f8153,f8157,f8159,f8161,f8163,f8167,f8169,f8171,f8173,f8177,f8179,f8182,f8184,f8187,f8189,f8204,f8206,f8245,f8246,f8247,f8248,f8474,f8583,f8606,f8607,f8608,f8609,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8618,f8619,f8620,f8621,f8622,f8623,f8624,f8625,f8626,f8627,f8628,f8629,f8630,f8631,f8632,f8633,f8634,f8635,f10108,f10109,f10924,f11190,f874,f11738,f11741,f12461,f12462,f12464,f12504,f12505,f11739,f12540,f12541,f12542,f12543,f12701,f12702,f12704,f12707,f12708,f12710,f12711,f12713,f12714,f12715,f12718,f12722,f12724,f12726,f12732,f12734,f12735,f12736,f12737,f12739,f12741,f12745,f12748,f12751,f12752,f12787,f12788,f12791,f12793,f12825,f12944,f12982,f11953,f13004,f1725,f13005,f13008,f6505,f4514,f11687,f6614,f11684,f13014,f13049,f10925,f13051,f11191,f13096,f13013,f9894,f13097,f494,f9652,f9651,f13154,f13158,f13161,f13163,f7719,f13007,f13210,f13266])).
fof(f13266,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f13265,f128])).
fof(f128,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relation.clos_relation_87_prove_ltb',aHL_FALSITY)).
fof(f13265,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f13264,f13015])).
fof(f13015,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl3),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f213,f8645,f8665,f8908,f8946,f8985,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10362,f10367,f10579,f10580,f10585,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10628,f10924,f11190,f11256,f11373,f11451,f11736,f11517,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11850,f11851,f11856,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11740,f11972,f11971,f11762,f11768,f11769,f11829,f11835,f11836,f11858,f11863,f11865,f11866,f11867,f11868,f11869,f11871,f11872,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11886,f9520,f3468,f3492,f9519,f3467,f3491,f2459,f2627,f10357,f10622,f1630,f1669,f1670,f2343,f2380,f3247,f9870,f10905,f11165,f11859,f11862,f164,f165,f166,f849,f870,f1319,f2192,f2217,f2234,f2342,f2379,f2996,f2997,f3246,f9869,f10861,f10864,f11173,f11844,f1664,f11172,f11843,f11852,f11853,f11171,f1437,f1650,f1683,f10863,f1430,f1450,f10862,f853,f1474,f2449,f2465,f2561,f2562,f2617,f2633,f3326,f10121,f10124,f10871,f10876,f10878,f11129,f11134,f2616,f3325,f10870,f11128,f11181,f10123,f10122,f966,f1480,f10889,f11148,f11534,f11911,f11912,f11913,f11914,f11916,f11917,f11938,f11939,f11947,f11948,f12009,f12013,f12012,f12010,f12008,f11973,f11968,f11967,f11966,f11962,f11961,f12007,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11964,f11963,f11720,f10587,f10578,f9871,f9868,f9522,f9521,f9518,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f11737,f11751,f11749,f11529,f11746,f11744,f11742,f11725,f11722,f11718,f11468,f13009,f13010,f11463,f11462,f13011,f13012,f11454,f11687,f11684,f11425])).
fof(f11425,plain,(
  ( ! [X47,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool))) ) | $spl3),
  inference(superposition,[],[f406,f8908])).
fof(f11454,plain,(
  ( ! [X101,X100] : (s(bool,X100) != s(bool,X101) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X100)),bool))) ) | $spl3),
  inference(superposition,[],[f163,f8908])).
fof(f13012,plain,(
  ( ! [X103,X102] : (s(bool,'T_0') = s(bool,X102) | p(s(bool,happ(s(fun(bool,bool),X103),s(bool,X102)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X103,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f11455,f2072])).
fof(f11455,plain,(
  ( ! [X103,X102] : (p(s(bool,happ(s(fun(bool,bool),X103),s(bool,X102)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X103,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X102)),bool))) ) | $spl3),
  inference(superposition,[],[f164,f8908])).
fof(f13011,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') = s(bool,X104) | ~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X104)))) | s(bool,X106) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X106))))) ) | $spl3),
  inference(forward_demodulation,[],[f11456,f2072])).
fof(f11456,plain,(
  ( ! [X105,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X104)))) | s(bool,X106) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X106)))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X104)),bool))) ) | $spl3),
  inference(superposition,[],[f166,f8908])).
fof(f11462,plain,(
  ( ! [X125,X123,X126,X124] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,X124) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X125,bool),X126))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X123)),bool))) ) | $spl3),
  inference(superposition,[],[f409,f8908])).
fof(f11463,plain,(
  ( ! [X127,X130,X128,X129] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X127))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X128))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X129),'const.option.NONE_0'),s(bool,X128),s(fun(X129,bool),X130))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X127)),bool))) ) | $spl3),
  inference(superposition,[],[f409,f8908])).
fof(f13010,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'T_0') = s(bool,X131) | p(s(bool,happ(s(fun(bool,bool),X132),s(bool,X131)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X133),'const.option.NONE_0'),s(fun(bool,bool),X132),s(fun(X133,fun(bool,bool)),X134)),bool))) ) | $spl3),
  inference(forward_demodulation,[],[f11464,f2072])).
fof(f11464,plain,(
  ( ! [X134,X132,X133,X131] : (p(s(bool,happ(s(fun(bool,bool),X132),s(bool,X131)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X133),'const.option.NONE_0'),s(fun(bool,bool),X132),s(fun(X133,fun(bool,bool)),X134)),bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X131)),bool))) ) | $spl3),
  inference(superposition,[],[f849,f8908])).
fof(f13009,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,'T_0') = s(bool,X135) | p(s(bool,happ(s(fun(bool,bool),X136),s(bool,X135)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X137),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X137,bool),X138)),X136,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f11465,f2072])).
fof(f11465,plain,(
  ( ! [X138,X136,X135,X137] : (p(s(bool,happ(s(fun(bool,bool),X136),s(bool,X135)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X137),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X137,bool),X138)),X136,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X135)),bool))) ) | $spl3),
  inference(superposition,[],[f870,f8908])).
fof(f11468,plain,(
  ( ! [X144,X143] : (s('type.option.option'(bool),X144) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X143))) | s(bool,sK4(X144,bool)) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X144) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X143)),bool))) ) | $spl3),
  inference(superposition,[],[f956,f8908])).
fof(f11718,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11541,f2072])).
fof(f11541,plain,(
  ( ! [X5] : (p(s(bool,sK4('const.option.SOME_1'(s(bool,X5)),bool))) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f11473])).
fof(f11473,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK4('const.option.SOME_1'(s(bool,X5)),bool))) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f131,f8908])).
fof(f11722,plain,(
  ( ! [X14,X13] : (s(bool,X13) = s(bool,X14) | ~p(s(bool,X14)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11721,f2072])).
fof(f11721,plain,(
  ( ! [X14,X13] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,sK4('const.option.SOME_1'(s(bool,X13)),bool)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(subsumption_resolution,[],[f11479,f126])).
fof(f11479,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X14) = s(bool,sK4('const.option.SOME_1'(s(bool,X13)),bool)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f170,f8908])).
fof(f11725,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | p(s(bool,X16)) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11724,f2072])).
fof(f11724,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,sK4('const.option.SOME_1'(s(bool,X15)),bool)) | p(s(bool,X16)) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(subsumption_resolution,[],[f11480,f126])).
fof(f11480,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,sK4('const.option.SOME_1'(s(bool,X15)),bool)) | p(s(bool,X16)) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f204,f8908])).
fof(f11742,plain,(
  ( ! [X88,X89] : (s(bool,X88) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X88)),bool),X89,bool)) | p(s(bool,happ(s(fun(bool,bool),X89),s(bool,'T_0')))) | s(bool,X88) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11521,f2072])).
fof(f11521,plain,(
  ( ! [X88,X89] : (p(s(bool,happ(s(fun(bool,bool),X89),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X88)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X88)),bool),X89,bool)) | s(bool,X88) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f164,f8908])).
fof(f11744,plain,(
  ( ! [X92,X90,X91] : (s(bool,X90) != s(bool,X92) | ~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X92)))) | s(bool,X90) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11522,f2072])).
fof(f11522,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,'T_0')))) | s(bool,X92) != s(bool,sK4('const.option.SOME_1'(s(bool,X90)),bool)) | p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X92)))) | s(bool,X90) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f166,f8908])).
fof(f11746,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X110) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(bool,X109),s(fun(X111,bool),X112))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,X109) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11528,f2072])).
fof(f11528,plain,(
  ( ! [X111,X109,X112,X110] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,X110) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(bool,sK4('const.option.SOME_1'(s(bool,X109)),bool)),s(fun(X111,bool),X112))) | s(bool,X109) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f409,f8908])).
fof(f11529,plain,(
  ( ! [X116,X114,X115,X113] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X114))) | s(bool,sK4('const.option.SOME_1'(s(bool,X113)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X115),'const.option.NONE_0'),s(bool,X114),s(fun(X115,bool),X116))) | s(bool,X113) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f409,f8908])).
fof(f11749,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X117) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X117)),bool),'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.NONE_0'),s(fun(bool,bool),X118),s(fun(X119,fun(bool,bool)),X120)),bool)) | p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0')))) | s(bool,X117) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11530,f2072])).
fof(f11530,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X117)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X117)),bool),'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.NONE_0'),s(fun(bool,bool),X118),s(fun(X119,fun(bool,bool)),X120)),bool)) | s(bool,X117) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f849,f8908])).
fof(f11751,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,X121) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X123),'const.option.NONE_0'),s(bool,X121),s(fun(X123,bool),X124)),X122,bool)) | p(s(bool,happ(s(fun(bool,bool),X122),s(bool,'T_0')))) | s(bool,X121) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11531,f2072])).
fof(f11531,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,happ(s(fun(bool,bool),X122),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X121)),bool)) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X123),'const.option.NONE_0'),s(bool,sK4('const.option.SOME_1'(s(bool,X121)),bool)),s(fun(X123,bool),X124)),X122,bool)) | s(bool,X121) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f870,f8908])).
fof(f11737,plain,(
  ( ! [X85,X84] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X84))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) | s(bool,'T_0') != s(bool,X85) | s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11519,f2072])).
fof(f11519,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,X85) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK4('const.option.SOME_1'(s(bool,X84)),bool)))) | s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f163,f8908])).
fof(f135,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,X1),s(fun(X0,bool),X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f175,f129])).
fof(f398,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f217,f162])).
fof(f406,plain,(
  ( ! [X4,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f162,f129])).
fof(f414,plain,(
  ( ! [X8,X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f162,f129])).
fof(f441,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') != s(bool,X8) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f163,f129])).
fof(f445,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f163,f129])).
fof(f2157,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X3)),bool))) )),
  inference(superposition,[],[f2072,f129])).
fof(f3245,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X61),'const.option.NONE_0'),s(bool,X59),s(fun(X61,bool),X62))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X58),'const.option.NONE_0'),s(bool,X59),s(fun(X58,bool),X60)))) )),
  inference(superposition,[],[f303,f175])).
fof(f3281,plain,(
  ( ! [X92,X90,X93,X91] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X93))) | s(bool,X93) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X90),'const.option.NONE_0'),s(bool,X91),s(fun(X90,bool),X92))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f162,f303])).
fof(f3315,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f3261,f126])).
fof(f3261,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f135,f303])).
fof(f3317,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,X26) | ~p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24)) )),
  inference(forward_demodulation,[],[f3316,f175])).
fof(f3316,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X23),'const.option.NONE_0'),s(bool,X24),s(fun(X23,bool),X25))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3265,f126])).
fof(f3265,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X23),'const.option.NONE_0'),s(bool,X24),s(fun(X23,bool),X25))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f170,f303])).
fof(f3319,plain,(
  ( ! [X30,X28] : (s(bool,X28) != s(bool,X30) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f3318,f175])).
fof(f3318,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.NONE_0'),s(bool,X28),s(fun(X27,bool),X29))) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f3266,f126])).
fof(f3266,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X30) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.NONE_0'),s(bool,X28),s(fun(X27,bool),X29))) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f204,f303])).
fof(f3323,plain,(
  ( ! [X101,X99] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X99))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) | s(bool,'T_0') != s(bool,X101) | s(bool,'F_0') = s(bool,X99)) )),
  inference(forward_demodulation,[],[f3283,f175])).
fof(f3283,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,X101) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(bool,X99),s(fun(X98,bool),X100))))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f163,f303])).
fof(f3466,plain,(
  ( ! [X39,X41,X42,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X40))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),'const.option.NONE_0'),s(bool,X40),s(fun(X39,bool),X41)))) )),
  inference(superposition,[],[f406,f175])).
fof(f3469,plain,(
  ( ! [X54,X53] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X53))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X53)),bool))) )),
  inference(superposition,[],[f406,f2072])).
fof(f3470,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK4(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f406,f154])).
fof(f3490,plain,(
  ( ! [X43,X41,X44,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f406,f175])).
fof(f3493,plain,(
  ( ! [X56,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X55))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X55)),bool)) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f406,f2072])).
fof(f3494,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,sK4(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f406,f154])).
fof(f9518,plain,(
  ( ! [X39,X37,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) | s(bool,X40) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X37),'const.option.NONE_0'),s(bool,X38),s(fun(X37,bool),X39))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f414,f175])).
fof(f9521,plain,(
  ( ! [X52,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) | s(bool,X52) = s(bool,sK4('const.option.SOME_1'(s(bool,X51)),bool)) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f414,f2072])).
fof(f9522,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,sK4(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f414,f154])).
fof(f9868,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X40) != s(bool,X42) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),'const.option.NONE_0'),s(bool,X40),s(fun(X39,bool),X41)))) )),
  inference(superposition,[],[f445,f175])).
fof(f9871,plain,(
  ( ! [X54,X53] : (s(bool,X53) != s(bool,X54) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X53)),bool))) )),
  inference(superposition,[],[f445,f2072])).
fof(f10578,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.NONE_0'),s(bool,X28),s(fun(X30,bool),X31))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f409,f129])).
fof(f10587,plain,(
  ( ! [X70,X72,X71,X73] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X70))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) | s(bool,X71) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X72),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X72,bool),X73))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f409,f129])).
fof(f11720,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X7)),bool)) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(subsumption_resolution,[],[f11475,f126])).
fof(f11475,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X7)),bool)) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f135,f8908])).
fof(f11963,plain,(
  ( ! [X165,X163,X166,X164] : (~p(s(bool,happ(s(fun(X163,bool),X166),s(X163,sK4(X164,X163))))) | s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X163),X164),s(bool,X165),s(fun(X163,bool),X166))) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X164)) )),
  inference(superposition,[],[f135,f1387])).
fof(f11964,plain,(
  ( ! [X167,X169,X170,X168] : (p(s(bool,happ(s(fun(X167,bool),X170),s(X167,sK4(X168,X167))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X167),X168),s(bool,X169),s(fun(X167,bool),X170))) | s('type.option.option'(X167),'const.option.NONE_0') = s('type.option.option'(X167),X168)) )),
  inference(superposition,[],[f136,f1387])).
fof(f11974,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X220),'const.option.NONE_0'),s(bool,happ(s(fun(X216,bool),X219),s(X216,sK4(X217,X216)))),s(fun(X220,bool),X221))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X216),X217),s(bool,X218),s(fun(X216,bool),X219))) | s('type.option.option'(X216),'const.option.NONE_0') = s('type.option.option'(X216),X217)) )),
  inference(superposition,[],[f303,f1387])).
fof(f11975,plain,(
  ( ! [X222,X225,X223,X226,X224] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X226))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X222,bool),X225),s(X222,sK4(X223,X222)))))) | s(bool,'F_0') = s(bool,X226) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X222),X223),s(bool,X224),s(fun(X222,bool),X225))) | s('type.option.option'(X222),'const.option.NONE_0') = s('type.option.option'(X222),X223)) )),
  inference(superposition,[],[f406,f1387])).
fof(f11976,plain,(
  ( ! [X231,X229,X227,X230,X228] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X231))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X227,bool),X230),s(X227,sK4(X228,X227)))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X227),X228),s(bool,X229),s(fun(X227,bool),X230))) | s(bool,'T_0') = s(bool,X231) | s('type.option.option'(X227),'const.option.NONE_0') = s('type.option.option'(X227),X228)) )),
  inference(superposition,[],[f406,f1387])).
fof(f11977,plain,(
  ( ! [X235,X233,X236,X234,X232] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X232,bool),X235),s(X232,sK4(X233,X232)))))) | s(bool,X236) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X232),X233),s(bool,X234),s(fun(X232,bool),X235))) | s(bool,'T_0') = s(bool,X236) | s('type.option.option'(X232),'const.option.NONE_0') = s('type.option.option'(X232),X233)) )),
  inference(superposition,[],[f414,f1387])).
fof(f11978,plain,(
  ( ! [X241,X239,X237,X240,X238] : (s(bool,'F_0') != s(bool,happ(s(fun(X237,bool),X240),s(X237,sK4(X238,X237)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X241))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X237),X238),s(bool,X239),s(fun(X237,bool),X240))))) | s(bool,'T_0') = s(bool,X241) | s('type.option.option'(X237),'const.option.NONE_0') = s('type.option.option'(X237),X238)) )),
  inference(superposition,[],[f441,f1387])).
fof(f11979,plain,(
  ( ! [X243,X246,X244,X242,X245] : (s(bool,X246) != s(bool,happ(s(fun(X242,bool),X245),s(X242,sK4(X243,X242)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X246))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X242),X243),s(bool,X244),s(fun(X242,bool),X245))) | s('type.option.option'(X242),'const.option.NONE_0') = s('type.option.option'(X242),X243)) )),
  inference(superposition,[],[f445,f1387])).
fof(f11980,plain,(
  ( ! [X251,X249,X250,X248,X247] : (s(bool,X251) != s(bool,happ(s(fun(X247,bool),X250),s(X247,sK4(X248,X247)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X247),X248),s(bool,X249),s(fun(X247,bool),X250))))) | s(bool,'T_0') = s(bool,X251) | s('type.option.option'(X247),'const.option.NONE_0') = s('type.option.option'(X247),X248)) )),
  inference(superposition,[],[f445,f1387])).
fof(f11982,plain,(
  ( ! [X257,X259,X260,X256,X258] : (p(s(bool,happ(s(fun(bool,bool),X260),s(bool,happ(s(fun(X256,bool),X259),s(X256,sK4(X257,X256))))))) | s(bool,'F_0') = s(bool,sK5('F_0',X260,bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X256),X257),s(bool,X258),s(fun(X256,bool),X259))) | s('type.option.option'(X256),'const.option.NONE_0') = s('type.option.option'(X256),X257)) )),
  inference(superposition,[],[f853,f1387])).
fof(f12007,plain,(
  ( ! [X372,X374,X370,X373,X371] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X370,fun(bool,bool)),X373),s(X370,sK4(X371,X370)))),s(bool,X374)))) | s(bool,'F_0') = s(bool,sK5('F_0','const.option.option_CASE_3'(s('type.option.option'(X370),X371),s(fun(bool,bool),X372),s(fun(X370,fun(bool,bool)),X373)),bool)) | s(bool,'T_0') = s(bool,X374) | s('type.option.option'(X370),'const.option.NONE_0') = s('type.option.option'(X370),X371)) )),
  inference(superposition,[],[f853,f1387])).
fof(f11961,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,happ(s(fun(X155,bool),X158),s(X155,sK4(X156,X155)))) | p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X155),X156),s(bool,X157),s(fun(X155,bool),X158)))) | s('type.option.option'(X155),'const.option.NONE_0') = s('type.option.option'(X155),X156)) )),
  inference(superposition,[],[f131,f1387])).
fof(f11962,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,happ(s(fun(X159,bool),X162),s(X159,sK4(X160,X159))))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X159),X160),s(bool,X161),s(fun(X159,bool),X162))) | s('type.option.option'(X159),'const.option.NONE_0') = s('type.option.option'(X159),X160)) )),
  inference(superposition,[],[f132,f1387])).
fof(f11966,plain,(
  ( ! [X177,X180,X178,X176,X179] : (p(s(bool,happ(s(fun(X176,bool),X179),s(X176,sK4(X177,X176))))) | p(s(bool,X180)) | s(bool,X180) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X176),X177),s(bool,X178),s(fun(X176,bool),X179))) | s('type.option.option'(X176),'const.option.NONE_0') = s('type.option.option'(X176),X177)) )),
  inference(superposition,[],[f167,f1387])).
fof(f11967,plain,(
  ( ! [X185,X182,X184,X183,X181] : (~p(s(bool,happ(s(fun(X181,bool),X184),s(X181,sK4(X182,X181))))) | ~p(s(bool,X185)) | s(bool,X185) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X181),X182),s(bool,X183),s(fun(X181,bool),X184))) | s('type.option.option'(X181),'const.option.NONE_0') = s('type.option.option'(X181),X182)) )),
  inference(superposition,[],[f170,f1387])).
fof(f11968,plain,(
  ( ! [X189,X187,X190,X188,X186] : (~p(s(bool,happ(s(fun(X186,bool),X189),s(X186,sK4(X187,X186))))) | s(bool,X190) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X186),X187),s(bool,X188),s(fun(X186,bool),X189))) | p(s(bool,X190)) | s('type.option.option'(X186),'const.option.NONE_0') = s('type.option.option'(X186),X187)) )),
  inference(superposition,[],[f204,f1387])).
fof(f11973,plain,(
  ( ! [X212,X215,X213,X211,X214] : (p(s(bool,happ(s(fun(X211,bool),X214),s(X211,sK4(X212,X211))))) | s(bool,X215) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X211),X212),s(bool,X213),s(fun(X211,bool),X214))) | s(bool,'T_0') = s(bool,X215) | s('type.option.option'(X211),'const.option.NONE_0') = s('type.option.option'(X211),X212)) )),
  inference(superposition,[],[f272,f1387])).
fof(f12008,plain,(
  ( ! [X377,X379,X380,X376,X378,X375] : (p(s(bool,happ(s(fun(X375,bool),happ(s(fun(X376,fun(X375,bool)),X379),s(X376,sK4(X377,X376)))),s(X375,X380)))) | s(X375,X380) = s(X375,sK5(X380,'const.option.option_CASE_3'(s('type.option.option'(X376),X377),s(fun(X375,bool),X378),s(fun(X376,fun(X375,bool)),X379)),X375)) | s('type.option.option'(X376),'const.option.NONE_0') = s('type.option.option'(X376),X377)) )),
  inference(superposition,[],[f164,f1387])).
fof(f12010,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X392] : (~p(s(bool,happ(s(fun(X387,bool),happ(s(fun(X388,fun(X387,bool)),X391),s(X388,sK4(X389,X388)))),s(X387,X392)))) | s(X387,X392) != s(X387,X393) | p(s(bool,happ(s(fun(X387,bool),'const.option.option_CASE_3'(s('type.option.option'(X388),X389),s(fun(X387,bool),X390),s(fun(X388,fun(X387,bool)),X391))),s(X387,X393)))) | s('type.option.option'(X388),'const.option.NONE_0') = s('type.option.option'(X388),X389)) )),
  inference(superposition,[],[f166,f1387])).
fof(f12012,plain,(
  ( ! [X405,X407,X401,X403,X404,X406,X400,X402] : (p(s(bool,happ(s(fun(X400,bool),happ(s(fun(X401,fun(X400,bool)),X404),s(X401,sK4(X402,X401)))),s(X400,X405)))) | s(X400,X405) = s(X400,sK5(X405,'const.option.option_CASE_3'(s('type.option.option'(X406),'const.option.NONE_0'),s(fun(X400,bool),'const.option.option_CASE_3'(s('type.option.option'(X401),X402),s(fun(X400,bool),X403),s(fun(X401,fun(X400,bool)),X404))),s(fun(X406,fun(X400,bool)),X407)),X400)) | s('type.option.option'(X401),'const.option.NONE_0') = s('type.option.option'(X401),X402)) )),
  inference(superposition,[],[f849,f1387])).
fof(f12013,plain,(
  ( ! [X412,X414,X408,X410,X413,X415,X409,X411] : (p(s(bool,happ(s(fun(X408,bool),happ(s(fun(X409,fun(X408,bool)),X412),s(X409,sK4(X410,X409)))),s(X408,X413)))) | s(X408,X413) = s(X408,sK5('const.option.option_CASE_3'(s('type.option.option'(X414),'const.option.NONE_0'),s(X408,X413),s(fun(X414,X408),X415)),'const.option.option_CASE_3'(s('type.option.option'(X409),X410),s(fun(X408,bool),X411),s(fun(X409,fun(X408,bool)),X412)),X408)) | s('type.option.option'(X409),'const.option.NONE_0') = s('type.option.option'(X409),X410)) )),
  inference(superposition,[],[f870,f1387])).
fof(f12009,plain,(
  ( ! [X383,X385,X382,X384,X386,X381] : (~p(s(bool,happ(s(fun(X381,bool),happ(s(fun(X382,fun(X381,bool)),X385),s(X382,sK4(X383,X382)))),s(X381,sK5(X386,'const.option.option_CASE_3'(s('type.option.option'(X382),X383),s(fun(X381,bool),X384),s(fun(X382,fun(X381,bool)),X385)),X381))))) | p(s(bool,happ(s(fun(X381,bool),'const.option.option_CASE_3'(s('type.option.option'(X382),X383),s(fun(X381,bool),X384),s(fun(X382,fun(X381,bool)),X385))),s(X381,X386)))) | s('type.option.option'(X382),'const.option.NONE_0') = s('type.option.option'(X382),X383)) )),
  inference(superposition,[],[f165,f1387])).
fof(f11948,plain,(
  ( ! [X111,X118,X116,X114,X112,X117,X115,X113] : (p(s(bool,happ(s(fun(X111,bool),X116),s(X111,happ(s(fun(X112,X111),X115),s(X112,sK4(X113,X112))))))) | s(X111,'const.option.option_CASE_3'(s('type.option.option'(X112),X113),s(X111,X114),s(fun(X112,X111),X115))) = s(X111,sK5('const.option.option_CASE_3'(s('type.option.option'(X117),'const.option.NONE_0'),s(X111,'const.option.option_CASE_3'(s('type.option.option'(X112),X113),s(X111,X114),s(fun(X112,X111),X115))),s(fun(X117,X111),X118)),X116,X111)) | s('type.option.option'(X112),'const.option.NONE_0') = s('type.option.option'(X112),X113)) )),
  inference(superposition,[],[f870,f1387])).
fof(f11947,plain,(
  ( ! [X103,X109,X107,X105,X110,X108,X106,X104] : (p(s(bool,happ(s(fun(X103,bool),X108),s(X103,happ(s(fun(X104,X103),X107),s(X104,sK4(X105,X104))))))) | s(X103,'const.option.option_CASE_3'(s('type.option.option'(X104),X105),s(X103,X106),s(fun(X104,X103),X107))) = s(X103,sK5('const.option.option_CASE_3'(s('type.option.option'(X104),X105),s(X103,X106),s(fun(X104,X103),X107)),'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(fun(X103,bool),X108),s(fun(X109,fun(X103,bool)),X110)),X103)) | s('type.option.option'(X104),'const.option.NONE_0') = s('type.option.option'(X104),X105)) )),
  inference(superposition,[],[f849,f1387])).
fof(f11939,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X49),s(X44,happ(s(fun(X45,X44),X48),s(X45,sK4(X46,X45))))))) | s(X44,X50) != s(X44,'const.option.option_CASE_3'(s('type.option.option'(X45),X46),s(X44,X47),s(fun(X45,X44),X48))) | p(s(bool,happ(s(fun(X44,bool),X49),s(X44,X50)))) | s('type.option.option'(X45),'const.option.NONE_0') = s('type.option.option'(X45),X46)) )),
  inference(superposition,[],[f166,f1387])).
fof(f11938,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),X43),s(X38,happ(s(fun(X39,X38),X42),s(X39,sK4(X40,X39))))))) | s(X38,'const.option.option_CASE_3'(s('type.option.option'(X39),X40),s(X38,X41),s(fun(X39,X38),X42))) = s(X38,sK5('const.option.option_CASE_3'(s('type.option.option'(X39),X40),s(X38,X41),s(fun(X39,X38),X42)),X43,X38)) | s('type.option.option'(X39),'const.option.NONE_0') = s('type.option.option'(X39),X40)) )),
  inference(superposition,[],[f164,f1387])).
fof(f11917,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X439] : (p(s(bool,happ(s(fun(X438,bool),'const.option.option_CASE_3'(s('type.option.option'(X439),X441),s(fun(X438,bool),X442),s(fun(X439,fun(X438,bool)),X440))),s(X438,X443)))) | s(X438,X443) = s(X438,sK5('const.option.option_CASE_3'(s('type.option.option'(X444),'const.option.NONE_0'),s(X438,X443),s(fun(X444,X438),X445)),happ(s(fun(X439,fun(X438,bool)),X440),s(X439,sK4(X441,X439))),X438)) | s('type.option.option'(X439),'const.option.NONE_0') = s('type.option.option'(X439),X441)) )),
  inference(superposition,[],[f870,f1387])).
fof(f11916,plain,(
  ( ! [X430,X436,X432,X434,X431,X437,X433,X435] : (p(s(bool,happ(s(fun(X430,bool),'const.option.option_CASE_3'(s('type.option.option'(X431),X433),s(fun(X430,bool),X434),s(fun(X431,fun(X430,bool)),X432))),s(X430,X435)))) | s(X430,X435) = s(X430,sK5(X435,'const.option.option_CASE_3'(s('type.option.option'(X436),'const.option.NONE_0'),s(fun(X430,bool),happ(s(fun(X431,fun(X430,bool)),X432),s(X431,sK4(X433,X431)))),s(fun(X436,fun(X430,bool)),X437)),X430)) | s('type.option.option'(X431),'const.option.NONE_0') = s('type.option.option'(X431),X433)) )),
  inference(superposition,[],[f849,f1387])).
fof(f11914,plain,(
  ( ! [X420,X422,X418,X421,X423,X417,X419] : (~p(s(bool,happ(s(fun(X417,bool),'const.option.option_CASE_3'(s('type.option.option'(X418),X420),s(fun(X417,bool),X421),s(fun(X418,fun(X417,bool)),X419))),s(X417,X422)))) | s(X417,X422) != s(X417,X423) | p(s(bool,happ(s(fun(X417,bool),happ(s(fun(X418,fun(X417,bool)),X419),s(X418,sK4(X420,X418)))),s(X417,X423)))) | s('type.option.option'(X418),'const.option.NONE_0') = s('type.option.option'(X418),X420)) )),
  inference(superposition,[],[f166,f1387])).
fof(f11913,plain,(
  ( ! [X412,X414,X416,X413,X415,X411] : (~p(s(bool,happ(s(fun(X411,bool),'const.option.option_CASE_3'(s('type.option.option'(X412),X414),s(fun(X411,bool),X415),s(fun(X412,fun(X411,bool)),X413))),s(X411,sK5(X416,happ(s(fun(X412,fun(X411,bool)),X413),s(X412,sK4(X414,X412))),X411))))) | p(s(bool,happ(s(fun(X411,bool),happ(s(fun(X412,fun(X411,bool)),X413),s(X412,sK4(X414,X412)))),s(X411,X416)))) | s('type.option.option'(X412),'const.option.NONE_0') = s('type.option.option'(X412),X414)) )),
  inference(superposition,[],[f165,f1387])).
fof(f11912,plain,(
  ( ! [X408,X410,X405,X407,X409,X406] : (p(s(bool,happ(s(fun(X405,bool),'const.option.option_CASE_3'(s('type.option.option'(X406),X408),s(fun(X405,bool),X409),s(fun(X406,fun(X405,bool)),X407))),s(X405,X410)))) | s(X405,X410) = s(X405,sK5(X410,happ(s(fun(X406,fun(X405,bool)),X407),s(X406,sK4(X408,X406))),X405)) | s('type.option.option'(X406),'const.option.NONE_0') = s('type.option.option'(X406),X408)) )),
  inference(superposition,[],[f164,f1387])).
fof(f11911,plain,(
  ( ! [X401,X403,X404,X400,X402] : (p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X400),X402),s(fun(bool,bool),X403),s(fun(X400,fun(bool,bool)),X401))),s(bool,X404)))) | s(bool,'F_0') = s(bool,sK5('F_0',happ(s(fun(X400,fun(bool,bool)),X401),s(X400,sK4(X402,X400))),bool)) | s(bool,'T_0') = s(bool,X404) | s('type.option.option'(X400),'const.option.NONE_0') = s('type.option.option'(X400),X402)) )),
  inference(superposition,[],[f853,f1387])).
fof(f11534,plain,(
  ( ! [X130,X129] : (s('type.option.option'(bool),X130) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK4(X130,bool)) = s(bool,sK4('const.option.SOME_1'(s(bool,X129)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X130) | s(bool,X129) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f956,f8908])).
fof(f11148,plain,(
  ( ! [X125,X123,X121,X124,X122] : (p(s(bool,happ(s(fun('type.option.option'(X121),bool),X123),s('type.option.option'(X121),X122)))) | s('type.option.option'(X121),'const.option.SOME_1'(s(X121,sK4(X122,X121)))) = s('type.option.option'(X121),sK5('const.option.SOME_1'(s(X121,sK4(X122,X121))),'const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.NONE_0'),s(fun('type.option.option'(X121),bool),X123),s(fun(X124,fun('type.option.option'(X121),bool)),X125)),'type.option.option'(X121))) | s('type.option.option'(X121),'const.option.NONE_0') = s('type.option.option'(X121),X122)) )),
  inference(superposition,[],[f849,f154])).
fof(f10889,plain,(
  ( ! [X121,X118,X120,X119,X117] : (p(s(bool,happ(s(fun('type.option.option'(X117),bool),X119),s('type.option.option'(X117),X118)))) | s('type.option.option'(X117),'const.option.SOME_1'(s(X117,sK4(X118,X117)))) = s('type.option.option'(X117),sK5('const.option.option_CASE_3'(s('type.option.option'(X120),'const.option.NONE_0'),s('type.option.option'(X117),'const.option.SOME_1'(s(X117,sK4(X118,X117)))),s(fun(X120,'type.option.option'(X117)),X121)),X119,'type.option.option'(X117))) | s('type.option.option'(X117),'const.option.NONE_0') = s('type.option.option'(X117),X118)) )),
  inference(superposition,[],[f870,f154])).
fof(f1480,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun('type.option.option'(X37),bool),X39),s('type.option.option'(X37),X38)))) | s('type.option.option'(X37),X40) != s('type.option.option'(X37),'const.option.SOME_1'(s(X37,sK4(X38,X37)))) | p(s(bool,happ(s(fun('type.option.option'(X37),bool),X39),s('type.option.option'(X37),X40)))) | s('type.option.option'(X37),'const.option.NONE_0') = s('type.option.option'(X37),X38)) )),
  inference(superposition,[],[f166,f154])).
fof(f966,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun('type.option.option'(X24),bool),X26),s('type.option.option'(X24),X25)))) | s('type.option.option'(X24),'const.option.SOME_1'(s(X24,sK4(X25,X24)))) = s('type.option.option'(X24),sK5('const.option.SOME_1'(s(X24,sK4(X25,X24))),X26,'type.option.option'(X24))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X25)) )),
  inference(superposition,[],[f164,f154])).
fof(f10122,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,X8)))) | s(bool,'F_0') = s(bool,sK5('F_0','const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f853,f177])).
fof(f10123,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(fun(bool,bool),X11),s(fun(X9,fun(bool,bool)),X12))),s(bool,X13)))) | s(bool,'F_0') = s(bool,sK5('F_0',happ(s(fun(fun(X9,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X9),fun(fun(bool,bool),fun(fun(X9,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(fun(bool,bool),X11))),s(fun(X9,fun(bool,bool)),X12)),bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f853,f178])).
fof(f11181,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (s(bool,X92) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X91),'const.option.NONE_0'),s(bool,X92),s(fun(X91,bool),X93)),'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(fun(bool,bool),X94),s(fun(X95,fun(bool,bool)),X96)),bool)) | p(s(bool,happ(s(fun(bool,bool),X94),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X92)) )),
  inference(forward_demodulation,[],[f11141,f175])).
fof(f11141,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (p(s(bool,happ(s(fun(bool,bool),X94),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X91),'const.option.NONE_0'),s(bool,X92),s(fun(X91,bool),X93))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X91),'const.option.NONE_0'),s(bool,X92),s(fun(X91,bool),X93)),'const.option.option_CASE_3'(s('type.option.option'(X95),'const.option.NONE_0'),s(fun(bool,bool),X94),s(fun(X95,fun(bool,bool)),X96)),bool)) | s(bool,'F_0') = s(bool,X92)) )),
  inference(superposition,[],[f849,f303])).
fof(f11128,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | s(bool,X32) = s(bool,sK5(X32,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(fun(bool,bool),X33),s(fun(X34,fun(bool,bool)),X35)),bool)) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f849,f239])).
fof(f10870,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | s(bool,X32) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X32),s(fun(X34,bool),X35)),X33,bool)) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f870,f239])).
fof(f3325,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X107) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X106),'const.option.NONE_0'),s(bool,X107),s(fun(X106,bool),X108)),X109,bool)) | p(s(bool,happ(s(fun(bool,bool),X109),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X107)) )),
  inference(forward_demodulation,[],[f3285,f175])).
fof(f3285,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,happ(s(fun(bool,bool),X109),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X106),'const.option.NONE_0'),s(bool,X107),s(fun(X106,bool),X108))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X106),'const.option.NONE_0'),s(bool,X107),s(fun(X106,bool),X108)),X109,bool)) | s(bool,'F_0') = s(bool,X107)) )),
  inference(superposition,[],[f164,f303])).
fof(f2616,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,happ(s(fun(bool,bool),X122),s(bool,'T_0')))) | s(bool,X121) = s(bool,sK5(X121,X122,bool)) | s(bool,X121) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f164,f239])).
fof(f11134,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun(bool,bool),X61),s(bool,X60)))) | s(bool,'T_0') = s(bool,sK5('T_0','const.option.option_CASE_3'(s('type.option.option'(X62),'const.option.NONE_0'),s(fun(bool,bool),X61),s(fun(X62,fun(bool,bool)),X63)),bool)) | s(bool,X60) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f849,f239])).
fof(f11129,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,X37) = s(bool,sK5(X37,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(fun(bool,bool),X39),s(fun(X40,fun(bool,bool)),X41)),bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f849,f239])).
fof(f10878,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X70)))) | s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X72),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X72,bool),X73)),X71,bool)) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f870,f129])).
fof(f10876,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun(bool,bool),X61),s(bool,X60)))) | s(bool,'T_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X62),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X62,bool),X63)),X61,bool)) | s(bool,X60) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f870,f239])).
fof(f10871,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,X37) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X37),s(fun(X40,bool),X41)),X39,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f870,f239])).
fof(f10124,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'F_0') = s(bool,sK5('F_0',sK4('const.option.SOME_1'(s(fun(bool,bool),X14)),fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f853,f2072])).
fof(f10121,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | s(bool,'F_0') = s(bool,sK5('F_0','const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f853,f175])).
fof(f3326,plain,(
  ( ! [X111,X114,X113] : (s(bool,X111) != s(bool,X114) | ~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X113),s(bool,X114)))) | s(bool,'F_0') = s(bool,X111)) )),
  inference(forward_demodulation,[],[f3286,f175])).
fof(f3286,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'T_0')))) | s(bool,X114) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X110),'const.option.NONE_0'),s(bool,X111),s(fun(X110,bool),X112))) | p(s(bool,happ(s(fun(bool,bool),X113),s(bool,X114)))) | s(bool,'F_0') = s(bool,X111)) )),
  inference(superposition,[],[f166,f303])).
fof(f2633,plain,(
  ( ! [X189,X187,X188] : (~p(s(bool,happ(s(fun(bool,bool),X188),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X188),s(bool,X187)))) | s(bool,X189) = s(bool,sK5(X187,X188,bool)) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f165,f239])).
fof(f2617,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,happ(s(fun(bool,bool),X125),s(bool,'T_0')))) | s(bool,X124) != s(bool,X126) | p(s(bool,happ(s(fun(bool,bool),X125),s(bool,X126)))) | s(bool,X124) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f166,f239])).
fof(f2562,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,X74)))) | s(bool,'T_0') != s(bool,X76) | p(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76)))) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f166,f239])).
fof(f2561,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),X72),s(bool,X71)))) | s(bool,'T_0') = s(bool,sK5('T_0',X72,bool)) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f164,f239])).
fof(f2465,plain,(
  ( ! [X198,X196,X197] : (~p(s(bool,happ(s(fun(bool,bool),X197),s(bool,X198)))) | p(s(bool,happ(s(fun(bool,bool),X197),s(bool,X196)))) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,sK5(X196,X197,bool))) )),
  inference(superposition,[],[f165,f239])).
fof(f2449,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,happ(s(fun(bool,bool),X134),s(bool,X133)))) | s(bool,X132) != s(bool,X135) | p(s(bool,happ(s(fun(bool,bool),X134),s(bool,X135)))) | s(bool,'T_0') = s(bool,X133) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f166,f239])).
fof(f1474,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,'F_0') != s(bool,X21) | p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X21)))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f166,f129])).
fof(f853,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'F_0') = s(bool,sK5('F_0',X7,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f164,f129])).
fof(f10862,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X8,fun(X7,bool)),X11),s(X8,X9))),s(X7,X12)))) | s(X7,X12) = s(X7,sK5('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(X7,X12),s(fun(X13,X7),X14)),'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(fun(X7,bool),X10),s(fun(X8,fun(X7,bool)),X11)),X7))) )),
  inference(superposition,[],[f870,f177])).
fof(f1450,plain,(
  ( ! [X127,X125,X123,X128,X126,X124] : (p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X124,fun(X123,bool)),X127),s(X124,X125))),s(X123,X128)))) | ~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X124,fun(X123,bool)),X127),s(X124,X125))),s(X123,sK5(X128,'const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.SOME_1'(s(X124,X125))),s(fun(X123,bool),X126),s(fun(X124,fun(X123,bool)),X127)),X123)))))) )),
  inference(forward_demodulation,[],[f1428,f177])).
fof(f1428,plain,(
  ( ! [X127,X125,X123,X128,X126,X124] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X124,fun(X123,bool)),X127),s(X124,X125))),s(X123,sK5(X128,'const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.SOME_1'(s(X124,X125))),s(fun(X123,bool),X126),s(fun(X124,fun(X123,bool)),X127)),X123))))) | p(s(bool,happ(s(fun(X123,bool),'const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.SOME_1'(s(X124,X125))),s(fun(X123,bool),X126),s(fun(X124,fun(X123,bool)),X127))),s(X123,X128))))) )),
  inference(superposition,[],[f165,f177])).
fof(f1430,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(X135,bool),happ(s(fun(X136,fun(X135,bool)),X139),s(X136,X137))),s(X135,X140)))) | s(X135,X140) = s(X135,sK5(X140,'const.option.option_CASE_3'(s('type.option.option'(X136),'const.option.SOME_1'(s(X136,X137))),s(fun(X135,bool),X138),s(fun(X136,fun(X135,bool)),X139)),X135))) )),
  inference(superposition,[],[f164,f177])).
fof(f10863,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,happ(s(fun(X15,bool),'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(fun(X15,bool),X18),s(fun(X16,fun(X15,bool)),X19))),s(X15,X20)))) | s(X15,X20) = s(X15,sK5('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(X15,X20),s(fun(X21,X15),X22)),happ(s(fun(fun(X16,fun(X15,bool)),fun(X15,bool)),happ(s(fun(fun(X15,bool),fun(fun(X16,fun(X15,bool)),fun(X15,bool))),happ(s(fun('type.option.option'(X16),fun(fun(X15,bool),fun(fun(X16,fun(X15,bool)),fun(X15,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(fun(X15,bool),X18))),s(fun(X16,fun(X15,bool)),X19)),X15))) )),
  inference(superposition,[],[f870,f178])).
fof(f1683,plain,(
  ( ! [X210,X208,X206,X211,X209,X207] : (p(s(bool,happ(s(fun(X206,bool),'const.option.option_CASE_3'(s('type.option.option'(X207),X208),s(fun(X206,bool),X209),s(fun(X207,fun(X206,bool)),X210))),s(X206,X211)))) | ~p(s(bool,happ(s(fun(X206,bool),'const.option.option_CASE_3'(s('type.option.option'(X207),X208),s(fun(X206,bool),X209),s(fun(X207,fun(X206,bool)),X210))),s(X206,sK5(X211,happ(s(fun(fun(X207,fun(X206,bool)),fun(X206,bool)),happ(s(fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool))),happ(s(fun('type.option.option'(X207),fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X207),X208))),s(fun(X206,bool),X209))),s(fun(X207,fun(X206,bool)),X210)),X206)))))) )),
  inference(forward_demodulation,[],[f1651,f178])).
fof(f1651,plain,(
  ( ! [X210,X208,X206,X211,X209,X207] : (~p(s(bool,happ(s(fun(X206,bool),'const.option.option_CASE_3'(s('type.option.option'(X207),X208),s(fun(X206,bool),X209),s(fun(X207,fun(X206,bool)),X210))),s(X206,sK5(X211,happ(s(fun(fun(X207,fun(X206,bool)),fun(X206,bool)),happ(s(fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool))),happ(s(fun('type.option.option'(X207),fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X207),X208))),s(fun(X206,bool),X209))),s(fun(X207,fun(X206,bool)),X210)),X206))))) | p(s(bool,happ(s(fun(X206,bool),happ(s(fun(fun(X207,fun(X206,bool)),fun(X206,bool)),happ(s(fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool))),happ(s(fun('type.option.option'(X207),fun(fun(X206,bool),fun(fun(X207,fun(X206,bool)),fun(X206,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X207),X208))),s(fun(X206,bool),X209))),s(fun(X207,fun(X206,bool)),X210))),s(X206,X211))))) )),
  inference(superposition,[],[f165,f178])).
fof(f1650,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (p(s(bool,happ(s(fun(X200,bool),'const.option.option_CASE_3'(s('type.option.option'(X201),X202),s(fun(X200,bool),X203),s(fun(X201,fun(X200,bool)),X204))),s(X200,X205)))) | s(X200,X205) = s(X200,sK5(X205,happ(s(fun(fun(X201,fun(X200,bool)),fun(X200,bool)),happ(s(fun(fun(X200,bool),fun(fun(X201,fun(X200,bool)),fun(X200,bool))),happ(s(fun('type.option.option'(X201),fun(fun(X200,bool),fun(fun(X201,fun(X200,bool)),fun(X200,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X201),X202))),s(fun(X200,bool),X203))),s(fun(X201,fun(X200,bool)),X204)),X200))) )),
  inference(superposition,[],[f164,f178])).
fof(f1437,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(X29,happ(s(fun(X30,X29),X33),s(X30,X31))) = s(X29,sK5('const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33)),X34,X29)) | p(s(bool,happ(s(fun(X29,bool),X34),s(X29,happ(s(fun(X30,X29),X33),s(X30,X31))))))) )),
  inference(forward_demodulation,[],[f1409,f177])).
fof(f1409,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X29,bool),X34),s(X29,happ(s(fun(X30,X29),X33),s(X30,X31)))))) | s(X29,'const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33))) = s(X29,sK5('const.option.option_CASE_3'(s('type.option.option'(X30),'const.option.SOME_1'(s(X30,X31))),s(X29,X32),s(fun(X30,X29),X33)),X34,X29))) )),
  inference(superposition,[],[f164,f177])).
fof(f11171,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))) = s(X7,sK5('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(fun(X7,bool),X12),s(fun(X13,fun(X7,bool)),X14)),X7)) | p(s(bool,happ(s(fun(X7,bool),X12),s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))))))) )),
  inference(forward_demodulation,[],[f11124,f177])).
fof(f11124,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,happ(s(fun(X7,bool),X12),s(X7,happ(s(fun(X8,X7),X11),s(X8,X9)))))) | s(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11))) = s(X7,sK5('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(fun(X7,bool),X12),s(fun(X13,fun(X7,bool)),X14)),X7))) )),
  inference(superposition,[],[f849,f177])).
fof(f11853,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X115,X113] : (p(s(bool,happ(s(fun(X108,bool),X113),s(X108,'const.option.option_CASE_3'(s('type.option.option'(X109),X111),s(X108,X112),s(fun(X109,X108),X110)))))) | s(X108,happ(s(fun(X109,X108),X110),s(X109,sK4(X111,X109)))) = s(X108,sK5('const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s(X108,happ(s(fun(X109,X108),X110),s(X109,sK4(X111,X109)))),s(fun(X114,X108),X115)),X113,X108)) | s('type.option.option'(X109),'const.option.NONE_0') = s('type.option.option'(X109),X111)) )),
  inference(superposition,[],[f870,f1387])).
fof(f11852,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (p(s(bool,happ(s(fun(X100,bool),X105),s(X100,'const.option.option_CASE_3'(s('type.option.option'(X101),X103),s(X100,X104),s(fun(X101,X100),X102)))))) | s(X100,happ(s(fun(X101,X100),X102),s(X101,sK4(X103,X101)))) = s(X100,sK5(happ(s(fun(X101,X100),X102),s(X101,sK4(X103,X101))),'const.option.option_CASE_3'(s('type.option.option'(X106),'const.option.NONE_0'),s(fun(X100,bool),X105),s(fun(X106,fun(X100,bool)),X107)),X100)) | s('type.option.option'(X101),'const.option.NONE_0') = s('type.option.option'(X101),X103)) )),
  inference(superposition,[],[f849,f1387])).
fof(f11843,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,'const.option.option_CASE_3'(s('type.option.option'(X36),X38),s(X35,X39),s(fun(X36,X35),X37)))))) | s(X35,happ(s(fun(X36,X35),X37),s(X36,sK4(X38,X36)))) = s(X35,sK5(happ(s(fun(X36,X35),X37),s(X36,sK4(X38,X36))),X40,X35)) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),X38)) )),
  inference(superposition,[],[f164,f1387])).
fof(f11172,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))) = s(X15,sK5(happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(fun(X15,bool),X20),s(fun(X21,fun(X15,bool)),X22)),X15)) | p(s(bool,happ(s(fun(X15,bool),X20),s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))))))) )),
  inference(forward_demodulation,[],[f11125,f178])).
fof(f11125,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (p(s(bool,happ(s(fun(X15,bool),X20),s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19)))))) | s(X15,happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19))) = s(X15,sK5(happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(fun(X15,bool),X20),s(fun(X21,fun(X15,bool)),X22)),X15))) )),
  inference(superposition,[],[f849,f178])).
fof(f1664,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(X29,'const.option.option_CASE_3'(s('type.option.option'(X30),X31),s(X29,X32),s(fun(X30,X29),X33))) = s(X29,sK5(happ(s(fun(fun(X30,X29),X29),happ(s(fun(X29,fun(fun(X30,X29),X29)),happ(s(fun('type.option.option'(X30),fun(X29,fun(fun(X30,X29),X29))),'const.option.option_CASE_0'),s('type.option.option'(X30),X31))),s(X29,X32))),s(fun(X30,X29),X33)),X34,X29)) | p(s(bool,happ(s(fun(X29,bool),X34),s(X29,'const.option.option_CASE_3'(s('type.option.option'(X30),X31),s(X29,X32),s(fun(X30,X29),X33))))))) )),
  inference(forward_demodulation,[],[f1617,f178])).
fof(f1617,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,happ(s(fun(X29,bool),X34),s(X29,'const.option.option_CASE_3'(s('type.option.option'(X30),X31),s(X29,X32),s(fun(X30,X29),X33)))))) | s(X29,happ(s(fun(fun(X30,X29),X29),happ(s(fun(X29,fun(fun(X30,X29),X29)),happ(s(fun('type.option.option'(X30),fun(X29,fun(fun(X30,X29),X29))),'const.option.option_CASE_0'),s('type.option.option'(X30),X31))),s(X29,X32))),s(fun(X30,X29),X33))) = s(X29,sK5(happ(s(fun(fun(X30,X29),X29),happ(s(fun(X29,fun(fun(X30,X29),X29)),happ(s(fun('type.option.option'(X30),fun(X29,fun(fun(X30,X29),X29))),'const.option.option_CASE_0'),s('type.option.option'(X30),X31))),s(X29,X32))),s(fun(X30,X29),X33)),X34,X29))) )),
  inference(superposition,[],[f164,f178])).
fof(f11844,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(X41,bool),X46),s(X41,'const.option.option_CASE_3'(s('type.option.option'(X42),X44),s(X41,X45),s(fun(X42,X41),X43)))))) | s(X41,X47) != s(X41,happ(s(fun(X42,X41),X43),s(X42,sK4(X44,X42)))) | p(s(bool,happ(s(fun(X41,bool),X46),s(X41,X47)))) | s('type.option.option'(X42),'const.option.NONE_0') = s('type.option.option'(X42),X44)) )),
  inference(superposition,[],[f166,f1387])).
fof(f11173,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(X23,X24) = s(X23,sK5(sK4('const.option.SOME_1'(s(X23,X24)),X23),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(X23,bool),X25),s(fun(X26,fun(X23,bool)),X27)),X23)) | p(s(bool,happ(s(fun(X23,bool),X25),s(X23,X24))))) )),
  inference(forward_demodulation,[],[f11126,f2072])).
fof(f11126,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,happ(s(fun(X23,bool),X25),s(X23,X24)))) | s(X23,sK4('const.option.SOME_1'(s(X23,X24)),X23)) = s(X23,sK5(sK4('const.option.SOME_1'(s(X23,X24)),X23),'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(fun(X23,bool),X25),s(fun(X26,fun(X23,bool)),X27)),X23))) )),
  inference(superposition,[],[f849,f2072])).
fof(f10864,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) | s(X23,X25) = s(X23,sK5('const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(X23,X25),s(fun(X26,X23),X27)),sK4('const.option.SOME_1'(s(fun(X23,bool),X24)),fun(X23,bool)),X23))) )),
  inference(superposition,[],[f870,f2072])).
fof(f10861,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,X4) = s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(X0,X4),s(fun(X5,X0),X6)),'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) )),
  inference(superposition,[],[f870,f175])).
fof(f9869,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),X46),s(X43,X44))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(bool,X45),s(fun(X43,bool),X46)))) )),
  inference(superposition,[],[f445,f177])).
fof(f3246,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X67),'const.option.NONE_0'),s(bool,happ(s(fun(X63,bool),X66),s(X63,X64))),s(fun(X67,bool),X68))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.SOME_1'(s(X63,X64))),s(bool,X65),s(fun(X63,bool),X66)))) )),
  inference(superposition,[],[f303,f177])).
fof(f2997,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) | s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK5(X99,X98,X97)))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f2608,f126])).
fof(f2608,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) | s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK5(X99,X98,X97)))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f165,f239])).
fof(f2996,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(X77,X79) != s(X77,X80) | p(s(bool,happ(s(fun(X77,bool),X78),s(X77,X80)))) | s(bool,X81) = s(bool,happ(s(fun(X77,bool),X78),s(X77,X79))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f2604,f126])).
fof(f2604,plain,(
  ( ! [X80,X78,X81,X79,X77] : (~p(s(bool,'T_0')) | s(X77,X79) != s(X77,X80) | p(s(bool,happ(s(fun(X77,bool),X78),s(X77,X80)))) | s(bool,X81) = s(bool,happ(s(fun(X77,bool),X78),s(X77,X79))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f166,f239])).
fof(f2379,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),X31),s(X28,X29))) | s(bool,X32) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X28),'const.option.SOME_1'(s(X28,X29))),s(bool,X30),s(fun(X28,bool),X31))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f239,f177])).
fof(f2342,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X43) = s(bool,happ(s(fun(X39,bool),X42),s(X39,X40))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),'const.option.SOME_1'(s(X39,X40))),s(bool,X41),s(fun(X39,bool),X42)))) )),
  inference(superposition,[],[f239,f177])).
fof(f2234,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,happ(s(fun(X89,bool),X90),s(X89,X91)))) | ~p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK5(X91,sK4('const.option.SOME_1'(s(fun(X89,bool),X90)),fun(X89,bool)),X89)))))) )),
  inference(forward_demodulation,[],[f2193,f2072])).
fof(f2193,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK5(X91,sK4('const.option.SOME_1'(s(fun(X89,bool),X90)),fun(X89,bool)),X89))))) | p(s(bool,happ(s(fun(X89,bool),sK4('const.option.SOME_1'(s(fun(X89,bool),X90)),fun(X89,bool))),s(X89,X91))))) )),
  inference(superposition,[],[f165,f2072])).
fof(f2217,plain,(
  ( ! [X17,X18,X16] : (s(X16,X17) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16)) | p(s(bool,happ(s(fun(X16,bool),X18),s(X16,X17))))) )),
  inference(forward_demodulation,[],[f2163,f2072])).
fof(f2163,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),X18),s(X16,X17)))) | s(X16,sK4('const.option.SOME_1'(s(X16,X17)),X16)) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16))) )),
  inference(superposition,[],[f164,f2072])).
fof(f2192,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,happ(s(fun(X86,bool),X87),s(X86,X88)))) | s(X86,X88) = s(X86,sK5(X88,sK4('const.option.SOME_1'(s(fun(X86,bool),X87)),fun(X86,bool)),X86))) )),
  inference(superposition,[],[f164,f2072])).
fof(f1319,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1312,f175])).
fof(f1312,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3))),s(X0,X4))))) )),
  inference(superposition,[],[f165,f175])).
fof(f870,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,happ(s(fun(X14,bool),X18),s(X14,X16)))) | s(X14,X16) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) )),
  inference(forward_demodulation,[],[f858,f175])).
fof(f858,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,happ(s(fun(X14,bool),X18),s(X14,X16)))) | s(X14,'const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17))) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) )),
  inference(superposition,[],[f164,f175])).
fof(f849,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,X4) = s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) )),
  inference(superposition,[],[f164,f175])).
fof(f166,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X8,X19] : (! [X13] : (s(X6,X13) = s(X6,X19) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X13))))) <=> p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X19)))))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.bool.UNWIND_FORALL_THM2')).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f164,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f115])).
fof(f11862,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X160),X162),s(bool,X163),s(fun(X160,bool),X161)))) | s(X160,X164) != s(X160,sK4(X162,X160)) | p(s(bool,happ(s(fun(X160,bool),X161),s(X160,X164)))) | s('type.option.option'(X160),'const.option.NONE_0') = s('type.option.option'(X160),X162)) )),
  inference(superposition,[],[f166,f1387])).
fof(f11859,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X146),X148),s(bool,X149),s(fun(X146,bool),X147)))) | s(X146,sK4(X148,X146)) = s(X146,sK5('const.option.option_CASE_3'(s('type.option.option'(X150),'const.option.NONE_0'),s(X146,sK4(X148,X146)),s(fun(X150,X146),X151)),X147,X146)) | s('type.option.option'(X146),'const.option.NONE_0') = s('type.option.option'(X146),X148)) )),
  inference(superposition,[],[f870,f1387])).
fof(f11165,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X51),X52),s(bool,X53),s(fun(X51,bool),X54)))) | s(fun(X51,bool),X54) = s(fun(X51,bool),sK5(X54,'const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.NONE_0'),s(fun(fun(X51,bool),bool),happ(s(fun(bool,fun(fun(X51,bool),bool)),happ(s(fun('type.option.option'(X51),fun(bool,fun(fun(X51,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X51),X52))),s(bool,X53))),s(fun(X55,fun(fun(X51,bool),bool)),X56)),fun(X51,bool)))) )),
  inference(superposition,[],[f849,f178])).
fof(f10905,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X51),X52),s(bool,X53),s(fun(X51,bool),X54)))) | s(fun(X51,bool),X54) = s(fun(X51,bool),sK5('const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.NONE_0'),s(fun(X51,bool),X54),s(fun(X55,fun(X51,bool)),X56)),happ(s(fun(bool,fun(fun(X51,bool),bool)),happ(s(fun('type.option.option'(X51),fun(bool,fun(fun(X51,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X51),X52))),s(bool,X53)),fun(X51,bool)))) )),
  inference(superposition,[],[f870,f178])).
fof(f9870,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X52) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X48),X49),s(bool,X50),s(fun(X48,bool),X51))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X48,bool),bool),happ(s(fun(bool,fun(fun(X48,bool),bool)),happ(s(fun('type.option.option'(X48),fun(bool,fun(fun(X48,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X48),X49))),s(bool,X50))),s(fun(X48,bool),X51)))) )),
  inference(superposition,[],[f445,f178])).
fof(f3247,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X69),X70),s(bool,X71),s(fun(X69,bool),X72))),s(fun(X73,bool),X74))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X69,bool),bool),happ(s(fun(bool,fun(fun(X69,bool),bool)),happ(s(fun('type.option.option'(X69),fun(bool,fun(fun(X69,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X69),X70))),s(bool,X71))),s(fun(X69,bool),X72)))) )),
  inference(superposition,[],[f303,f178])).
fof(f2380,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s(bool,X35),s(fun(X33,bool),X36))) | s(bool,X37) = s(bool,happ(s(fun(fun(X33,bool),bool),happ(s(fun(bool,fun(fun(X33,bool),bool)),happ(s(fun('type.option.option'(X33),fun(bool,fun(fun(X33,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s(bool,X35))),s(fun(X33,bool),X36))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f239,f178])).
fof(f2343,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X48) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X44),X45),s(bool,X46),s(fun(X44,bool),X47))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X44,bool),bool),happ(s(fun(bool,fun(fun(X44,bool),bool)),happ(s(fun('type.option.option'(X44),fun(bool,fun(fun(X44,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X44),X45))),s(bool,X46))),s(fun(X44,bool),X47)))) )),
  inference(superposition,[],[f239,f178])).
fof(f1670,plain,(
  ( ! [X118,X116,X119,X117] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),X117),s(bool,X118),s(fun(X116,bool),X119)))) | ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),X117),s(bool,X118),s(fun(X116,bool),sK5(X119,happ(s(fun(bool,fun(fun(X116,bool),bool)),happ(s(fun('type.option.option'(X116),fun(bool,fun(fun(X116,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X116),X117))),s(bool,X118)),fun(X116,bool))))))) )),
  inference(forward_demodulation,[],[f1632,f178])).
fof(f1632,plain,(
  ( ! [X118,X116,X119,X117] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X116),X117),s(bool,X118),s(fun(X116,bool),sK5(X119,happ(s(fun(bool,fun(fun(X116,bool),bool)),happ(s(fun('type.option.option'(X116),fun(bool,fun(fun(X116,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X116),X117))),s(bool,X118)),fun(X116,bool)))))) | p(s(bool,happ(s(fun(fun(X116,bool),bool),happ(s(fun(bool,fun(fun(X116,bool),bool)),happ(s(fun('type.option.option'(X116),fun(bool,fun(fun(X116,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X116),X117))),s(bool,X118))),s(fun(X116,bool),X119))))) )),
  inference(superposition,[],[f165,f178])).
fof(f1669,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),X96),s(bool,X97),s(fun(X95,bool),X99)))) | ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),X96),s(bool,X97),s(fun(X95,bool),X98)))) | s(fun(X95,bool),X98) != s(fun(X95,bool),X99)) )),
  inference(forward_demodulation,[],[f1627,f178])).
fof(f1627,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X95),X96),s(bool,X97),s(fun(X95,bool),X98)))) | s(fun(X95,bool),X98) != s(fun(X95,bool),X99) | p(s(bool,happ(s(fun(fun(X95,bool),bool),happ(s(fun(bool,fun(fun(X95,bool),bool)),happ(s(fun('type.option.option'(X95),fun(bool,fun(fun(X95,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X95),X96))),s(bool,X97))),s(fun(X95,bool),X99))))) )),
  inference(superposition,[],[f166,f178])).
fof(f1630,plain,(
  ( ! [X111,X109,X110,X108] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X108),X109),s(bool,X110),s(fun(X108,bool),X111)))) | s(fun(X108,bool),X111) = s(fun(X108,bool),sK5(X111,happ(s(fun(bool,fun(fun(X108,bool),bool)),happ(s(fun('type.option.option'(X108),fun(bool,fun(fun(X108,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X108),X109))),s(bool,X110)),fun(X108,bool)))) )),
  inference(superposition,[],[f164,f178])).
fof(f10622,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X88))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X85),'const.option.NONE_0'),s(bool,X86),s(fun(X85,bool),X87))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X89),'const.option.NONE_0'),s(bool,X88),s(fun(X89,bool),X90))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f409,f303])).
fof(f10357,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.option.option'(bool),X54) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK4(X54,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X51),'const.option.NONE_0'),s(bool,X52),s(fun(X51,bool),X53))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X54) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f956,f303])).
fof(f2627,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,'T_0') != s(bool,f(s(X165,sK6(X166,X167,bool,X165)))) | s(fun(X165,bool),X166) = s(fun(X165,bool),X167) | s(bool,X168) = s(bool,g(s(X165,sK6(X166,X167,bool,X165)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f176,f239])).
fof(f2459,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,X177) != s(bool,f(s(X174,sK6(X175,X176,bool,X174)))) | s(fun(X174,bool),X175) = s(fun(X174,bool),X176) | s(bool,'T_0') = s(bool,X177) | s(bool,'T_0') = s(bool,g(s(X174,sK6(X175,X176,bool,X174))))) )),
  inference(superposition,[],[f176,f239])).
fof(f3491,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X45,bool),X48),s(X45,X46))))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f406,f177])).
fof(f3467,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X43,bool),X46),s(X43,X44))))) | s(bool,'F_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(bool,X45),s(fun(X43,bool),X46)))) )),
  inference(superposition,[],[f406,f177])).
fof(f9519,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X41,bool),X44),s(X41,X42))))) | s(bool,X45) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.SOME_1'(s(X41,X42))),s(bool,X43),s(fun(X41,bool),X44))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f414,f177])).
fof(f3492,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f406,f178])).
fof(f3468,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X48),X49),s(bool,X50),s(fun(X48,bool),X51))))) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X48,bool),bool),happ(s(fun(bool,fun(fun(X48,bool),bool)),happ(s(fun('type.option.option'(X48),fun(bool,fun(fun(X48,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X48),X49))),s(bool,X50))),s(fun(X48,bool),X51)))) )),
  inference(superposition,[],[f406,f178])).
fof(f9520,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X46),X47),s(bool,X48),s(fun(X46,bool),X49))))) | s(bool,X50) = s(bool,happ(s(fun(fun(X46,bool),bool),happ(s(fun(bool,fun(fun(X46,bool),bool)),happ(s(fun('type.option.option'(X46),fun(bool,fun(fun(X46,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X46),X47))),s(bool,X48))),s(fun(X46,bool),X49))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f414,f178])).
fof(f11886,plain,(
  ( ! [X280,X277,X279,X281,X278] : (p(s(bool,happ(s(fun(bool,bool),X281),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X277),X279),s(bool,X280),s(fun(X277,bool),X278)))))) | s(bool,'F_0') = s(bool,sK5('F_0',X281,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X277,bool),X278),s(X277,sK4(X279,X277)))) | s('type.option.option'(X277),'const.option.NONE_0') = s('type.option.option'(X277),X279)) )),
  inference(superposition,[],[f853,f1387])).
fof(f11884,plain,(
  ( ! [X269,X271,X268,X270,X272] : (s(bool,X272) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X268),X270),s(bool,X271),s(fun(X268,bool),X269))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X268,bool),X269),s(X268,sK4(X270,X268)))))) | s(bool,'T_0') = s(bool,X272) | s('type.option.option'(X268),'const.option.NONE_0') = s('type.option.option'(X268),X270)) )),
  inference(superposition,[],[f445,f1387])).
fof(f11883,plain,(
  ( ! [X263,X265,X267,X264,X266] : (s(bool,X267) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X263),X265),s(bool,X266),s(fun(X263,bool),X264))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X267))) | s(bool,'T_0') = s(bool,happ(s(fun(X263,bool),X264),s(X263,sK4(X265,X263)))) | s('type.option.option'(X263),'const.option.NONE_0') = s('type.option.option'(X263),X265)) )),
  inference(superposition,[],[f445,f1387])).
fof(f11882,plain,(
  ( ! [X261,X259,X260,X262,X258] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X258),X260),s(bool,X261),s(fun(X258,bool),X259))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X262))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X258,bool),X259),s(X258,sK4(X260,X258)))))) | s(bool,'T_0') = s(bool,X262) | s('type.option.option'(X258),'const.option.NONE_0') = s('type.option.option'(X258),X260)) )),
  inference(superposition,[],[f441,f1387])).
fof(f11881,plain,(
  ( ! [X255,X253,X257,X254,X256] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X253),X255),s(bool,X256),s(fun(X253,bool),X254))))) | s(bool,X257) = s(bool,happ(s(fun(X253,bool),X254),s(X253,sK4(X255,X253)))) | s(bool,'T_0') = s(bool,X257) | s('type.option.option'(X253),'const.option.NONE_0') = s('type.option.option'(X253),X255)) )),
  inference(superposition,[],[f414,f1387])).
fof(f11880,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X252))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X248),X250),s(bool,X251),s(fun(X248,bool),X249))))) | s(bool,'F_0') = s(bool,happ(s(fun(X248,bool),X249),s(X248,sK4(X250,X248)))) | s(bool,'T_0') = s(bool,X252) | s('type.option.option'(X248),'const.option.NONE_0') = s('type.option.option'(X248),X250)) )),
  inference(superposition,[],[f406,f1387])).
fof(f11879,plain,(
  ( ! [X243,X246,X244,X247,X245] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X247))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X243),X245),s(bool,X246),s(fun(X243,bool),X244))))) | s(bool,'F_0') = s(bool,X247) | s(bool,'T_0') = s(bool,happ(s(fun(X243,bool),X244),s(X243,sK4(X245,X243)))) | s('type.option.option'(X243),'const.option.NONE_0') = s('type.option.option'(X243),X245)) )),
  inference(superposition,[],[f406,f1387])).
fof(f11878,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X241),'const.option.NONE_0'),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X237),X239),s(bool,X240),s(fun(X237,bool),X238))),s(fun(X241,bool),X242))) | s(bool,'F_0') = s(bool,happ(s(fun(X237,bool),X238),s(X237,sK4(X239,X237)))) | s('type.option.option'(X237),'const.option.NONE_0') = s('type.option.option'(X237),X239)) )),
  inference(superposition,[],[f303,f1387])).
fof(f11877,plain,(
  ( ! [X235,X233,X236,X234,X232] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X232),X234),s(bool,X235),s(fun(X232,bool),X233)))) | s(bool,X236) = s(bool,happ(s(fun(X232,bool),X233),s(X232,sK4(X234,X232)))) | s(bool,'T_0') = s(bool,X236) | s('type.option.option'(X232),'const.option.NONE_0') = s('type.option.option'(X232),X234)) )),
  inference(superposition,[],[f272,f1387])).
fof(f11872,plain,(
  ( ! [X210,X208,X211,X209,X207] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X207),X209),s(bool,X210),s(fun(X207,bool),X208)))) | s(bool,X211) != s(bool,happ(s(fun(X207,bool),X208),s(X207,sK4(X209,X207)))) | p(s(bool,X211)) | s('type.option.option'(X207),'const.option.NONE_0') = s('type.option.option'(X207),X209)) )),
  inference(superposition,[],[f204,f1387])).
fof(f11871,plain,(
  ( ! [X206,X204,X202,X205,X203] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X202),X204),s(bool,X205),s(fun(X202,bool),X203)))) | ~p(s(bool,X206)) | s(bool,X206) = s(bool,happ(s(fun(X202,bool),X203),s(X202,sK4(X204,X202)))) | s('type.option.option'(X202),'const.option.NONE_0') = s('type.option.option'(X202),X204)) )),
  inference(superposition,[],[f170,f1387])).
fof(f11869,plain,(
  ( ! [X196,X194,X192,X195,X193] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X192),X194),s(bool,X195),s(fun(X192,bool),X193)))) | p(s(bool,X196)) | s(bool,X196) = s(bool,happ(s(fun(X192,bool),X193),s(X192,sK4(X194,X192)))) | s('type.option.option'(X192),'const.option.NONE_0') = s('type.option.option'(X192),X194)) )),
  inference(superposition,[],[f167,f1387])).
fof(f11868,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X188),X190),s(bool,X191),s(fun(X188,bool),X189)))) | s(bool,'F_0') = s(bool,happ(s(fun(X188,bool),X189),s(X188,sK4(X190,X188)))) | s('type.option.option'(X188),'const.option.NONE_0') = s('type.option.option'(X188),X190)) )),
  inference(superposition,[],[f136,f1387])).
fof(f11867,plain,(
  ( ! [X187,X185,X186,X184] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X184),X186),s(bool,X187),s(fun(X184,bool),X185)))) | s(bool,'F_0') != s(bool,happ(s(fun(X184,bool),X185),s(X184,sK4(X186,X184)))) | s('type.option.option'(X184),'const.option.NONE_0') = s('type.option.option'(X184),X186)) )),
  inference(superposition,[],[f135,f1387])).
fof(f11866,plain,(
  ( ! [X182,X180,X183,X181] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X180),X182),s(bool,X183),s(fun(X180,bool),X181)))) | s(bool,'T_0') = s(bool,happ(s(fun(X180,bool),X181),s(X180,sK4(X182,X180)))) | s('type.option.option'(X180),'const.option.NONE_0') = s('type.option.option'(X180),X182)) )),
  inference(superposition,[],[f132,f1387])).
fof(f11865,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X176),X178),s(bool,X179),s(fun(X176,bool),X177))) | p(s(bool,happ(s(fun(X176,bool),X177),s(X176,sK4(X178,X176))))) | s('type.option.option'(X176),'const.option.NONE_0') = s('type.option.option'(X176),X178)) )),
  inference(superposition,[],[f131,f1387])).
fof(f11863,plain,(
  ( ! [X167,X165,X166,X168] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X165),X167),s(bool,X168),s(fun(X165,bool),X166)))) | s(X165,sK4(X167,X165)) = s(X165,sK5(sK4(X167,X165),X166,X165)) | s('type.option.option'(X165),'const.option.NONE_0') = s('type.option.option'(X165),X167)) )),
  inference(superposition,[],[f164,f1387])).
fof(f11858,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X140),X142),s(bool,X143),s(fun(X140,bool),X141)))) | s(X140,sK4(X142,X140)) = s(X140,sK5(sK4(X142,X140),'const.option.option_CASE_3'(s('type.option.option'(X144),'const.option.NONE_0'),s(fun(X140,bool),X141),s(fun(X144,fun(X140,bool)),X145)),X140)) | s('type.option.option'(X140),'const.option.NONE_0') = s('type.option.option'(X140),X142)) )),
  inference(superposition,[],[f849,f1387])).
fof(f11836,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,happ(s(fun(X54,bool),X57),s(X54,sK4(X55,X54)))) | s('type.option.option'(X54),'const.option.NONE_0') = s('type.option.option'(X54),X55) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X54),X55),s(bool,X56),s(fun(X54,bool),X57)))),bool))) ) | $spl3),
  inference(superposition,[],[f1387,f8908])).
fof(f11835,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,happ(s(fun(X50,bool),X53),s(X50,sK4(X51,X50)))) | s('type.option.option'(X50),'const.option.NONE_0') = s('type.option.option'(X50),X51) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53)))) ) | $spl3),
  inference(superposition,[],[f1387,f8665])).
fof(f11829,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(X22,bool),X25),s(X22,sK4(X23,X22)))) | s('type.option.option'(X22),'const.option.NONE_0') = s('type.option.option'(X22),X23) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),X23),s(bool,X24),s(fun(X22,bool),X25)))) )),
  inference(superposition,[],[f1387,f129])).
fof(f11769,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X49),X51),s(bool,X52),s(fun(X49,bool),X50))) | s('type.option.option'(X49),'const.option.NONE_0') = s('type.option.option'(X49),X51) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X49,bool),X50),s(X49,sK4(X51,X49))))),bool))) ) | $spl3),
  inference(superposition,[],[f1387,f8908])).
fof(f11768,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),X47),s(bool,X48),s(fun(X45,bool),X46))) | s('type.option.option'(X45),'const.option.NONE_0') = s('type.option.option'(X45),X47) | s(bool,'T_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK4(X47,X45))))) ) | $spl3),
  inference(superposition,[],[f1387,f8665])).
fof(f11762,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X17),X19),s(bool,X20),s(fun(X17,bool),X18))) | s('type.option.option'(X17),'const.option.NONE_0') = s('type.option.option'(X17),X19) | s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),X18),s(X17,sK4(X19,X17))))) )),
  inference(superposition,[],[f1387,f129])).
fof(f11971,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X205) = s(bool,happ(s(fun(X201,bool),X204),s(X201,sK4(X202,X201)))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X201),X202),s(bool,X203),s(fun(X201,bool),X204))) | s(bool,'T_0') = s(bool,X205) | s('type.option.option'(X201),'const.option.NONE_0') = s('type.option.option'(X201),X202)) )),
  inference(superposition,[],[f239,f1387])).
fof(f11972,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,'T_0') = s(bool,happ(s(fun(X206,bool),X209),s(X206,sK4(X207,X206)))) | s(bool,X210) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X206),X207),s(bool,X208),s(fun(X206,bool),X209))) | s(bool,'T_0') = s(bool,X210) | s('type.option.option'(X206),'const.option.NONE_0') = s('type.option.option'(X206),X207)) )),
  inference(superposition,[],[f239,f1387])).
fof(f11740,plain,(
  ( ! [X87,X86] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X86))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X87))) | s(bool,'T_0') != s(bool,X87) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(forward_demodulation,[],[f11520,f2072])).
fof(f11520,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,X87) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X87))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK4('const.option.SOME_1'(s(bool,X86)),bool)))) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f163,f8908])).
fof(f12022,plain,(
  ( ! [X472,X469,X471,X467,X468,X470] : (s('type.option.option'(X467),'const.option.SOME_1'(s(X467,X472))) != s('type.option.option'(X467),happ(s(fun(X468,'type.option.option'(X467)),X471),s(X468,sK4(X469,X468)))) | s(X467,X472) = s(X467,sK4('const.option.option_CASE_3'(s('type.option.option'(X468),X469),s('type.option.option'(X467),X470),s(fun(X468,'type.option.option'(X467)),X471)),X467)) | s('type.option.option'(X467),'const.option.NONE_0') = s('type.option.option'(X467),'const.option.option_CASE_3'(s('type.option.option'(X468),X469),s('type.option.option'(X467),X470),s(fun(X468,'type.option.option'(X467)),X471))) | s('type.option.option'(X468),'const.option.NONE_0') = s('type.option.option'(X468),X469)) )),
  inference(superposition,[],[f956,f1387])).
fof(f11951,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s('type.option.option'(X131),X136) != s('type.option.option'(X131),'const.option.SOME_1'(s(X131,happ(s(fun(X132,X131),X135),s(X132,sK4(X133,X132)))))) | s(X131,sK4(X136,X131)) = s(X131,'const.option.option_CASE_3'(s('type.option.option'(X132),X133),s(X131,X134),s(fun(X132,X131),X135))) | s('type.option.option'(X131),'const.option.NONE_0') = s('type.option.option'(X131),X136) | s('type.option.option'(X132),'const.option.NONE_0') = s('type.option.option'(X132),X133)) )),
  inference(superposition,[],[f956,f1387])).
fof(f11946,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (s('type.option.option'(X95),'const.option.SOME_1'(s(X95,X100))) != s('type.option.option'(X95),'const.option.SOME_1'(s(X95,happ(s(fun(X96,X95),X99),s(X96,sK4(X97,X96)))))) | s(X95,'const.option.option_CASE_3'(s('type.option.option'(X101),'const.option.NONE_0'),s(X95,X100),s(fun(X101,X95),X102))) = s(X95,'const.option.option_CASE_3'(s('type.option.option'(X96),X97),s(X95,X98),s(fun(X96,X95),X99))) | s('type.option.option'(X96),'const.option.NONE_0') = s('type.option.option'(X96),X97)) )),
  inference(superposition,[],[f409,f1387])).
fof(f11945,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (s('type.option.option'(X87),'const.option.SOME_1'(s(X87,X92))) != s('type.option.option'(X87),'const.option.SOME_1'(s(X87,happ(s(fun(X88,X87),X91),s(X88,sK4(X89,X88)))))) | s(X87,X92) = s(X87,'const.option.option_CASE_3'(s('type.option.option'(X93),'const.option.NONE_0'),s(X87,'const.option.option_CASE_3'(s('type.option.option'(X88),X89),s(X87,X90),s(fun(X88,X87),X91))),s(fun(X93,X87),X94))) | s('type.option.option'(X88),'const.option.NONE_0') = s('type.option.option'(X88),X89)) )),
  inference(superposition,[],[f409,f1387])).
fof(f11936,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X26,X31) != s(X26,happ(s(fun(X27,X26),X30),s(X27,sK4(X28,X27)))) | s('type.option.option'(X26),'const.option.SOME_1'(s(X26,X31))) = s('type.option.option'(X26),'const.option.SOME_1'(s(X26,'const.option.option_CASE_3'(s('type.option.option'(X27),X28),s(X26,X29),s(fun(X27,X26),X30))))) | s('type.option.option'(X27),'const.option.NONE_0') = s('type.option.option'(X27),X28)) )),
  inference(superposition,[],[f163,f1387])).
fof(f11934,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s('type.option.option'(X14),'const.option.SOME_1'(s(X14,X19))) != s('type.option.option'(X14),'const.option.SOME_1'(s(X14,happ(s(fun(X15,X14),X18),s(X15,sK4(X16,X15)))))) | s(X14,X19) = s(X14,'const.option.option_CASE_3'(s('type.option.option'(X15),X16),s(X14,X17),s(fun(X15,X14),X18))) | s('type.option.option'(X15),'const.option.NONE_0') = s('type.option.option'(X15),X16)) )),
  inference(superposition,[],[f162,f1387])).
fof(f11926,plain,(
  ( ! [X499,X500,X502,X498,X501,X497] : (s('type.option.option'(X497),'const.option.SOME_1'(s(X497,X502))) != s('type.option.option'(X497),'const.option.option_CASE_3'(s('type.option.option'(X498),X500),s('type.option.option'(X497),X501),s(fun(X498,'type.option.option'(X497)),X499))) | s(X497,X502) = s(X497,sK4(happ(s(fun(X498,'type.option.option'(X497)),X499),s(X498,sK4(X500,X498))),X497)) | s('type.option.option'(X497),'const.option.NONE_0') = s('type.option.option'(X497),happ(s(fun(X498,'type.option.option'(X497)),X499),s(X498,sK4(X500,X498)))) | s('type.option.option'(X498),'const.option.NONE_0') = s('type.option.option'(X498),X500)) )),
  inference(superposition,[],[f956,f1387])).
fof(f11856,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s('type.option.option'(X128),X133) != s('type.option.option'(X128),'const.option.SOME_1'(s(X128,'const.option.option_CASE_3'(s('type.option.option'(X129),X131),s(X128,X132),s(fun(X129,X128),X130))))) | s(X128,sK4(X133,X128)) = s(X128,happ(s(fun(X129,X128),X130),s(X129,sK4(X131,X129)))) | s('type.option.option'(X128),'const.option.NONE_0') = s('type.option.option'(X128),X133) | s('type.option.option'(X129),'const.option.NONE_0') = s('type.option.option'(X129),X131)) )),
  inference(superposition,[],[f956,f1387])).
fof(f11851,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (s('type.option.option'(X92),'const.option.SOME_1'(s(X92,X97))) != s('type.option.option'(X92),'const.option.SOME_1'(s(X92,'const.option.option_CASE_3'(s('type.option.option'(X93),X95),s(X92,X96),s(fun(X93,X92),X94))))) | s(X92,happ(s(fun(X93,X92),X94),s(X93,sK4(X95,X93)))) = s(X92,'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(X92,X97),s(fun(X98,X92),X99))) | s('type.option.option'(X93),'const.option.NONE_0') = s('type.option.option'(X93),X95)) )),
  inference(superposition,[],[f409,f1387])).
fof(f11850,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86,X84] : (s('type.option.option'(X84),'const.option.SOME_1'(s(X84,X89))) != s('type.option.option'(X84),'const.option.SOME_1'(s(X84,'const.option.option_CASE_3'(s('type.option.option'(X85),X87),s(X84,X88),s(fun(X85,X84),X86))))) | s(X84,X89) = s(X84,'const.option.option_CASE_3'(s('type.option.option'(X90),'const.option.NONE_0'),s(X84,happ(s(fun(X85,X84),X86),s(X85,sK4(X87,X85)))),s(fun(X90,X84),X91))) | s('type.option.option'(X85),'const.option.NONE_0') = s('type.option.option'(X85),X87)) )),
  inference(superposition,[],[f409,f1387])).
fof(f11841,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(X23,X28) != s(X23,'const.option.option_CASE_3'(s('type.option.option'(X24),X26),s(X23,X27),s(fun(X24,X23),X25))) | s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X28))) = s('type.option.option'(X23),'const.option.SOME_1'(s(X23,happ(s(fun(X24,X23),X25),s(X24,sK4(X26,X24)))))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X26)) )),
  inference(superposition,[],[f163,f1387])).
fof(f11839,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.option.option'(X11),'const.option.SOME_1'(s(X11,X16))) != s('type.option.option'(X11),'const.option.SOME_1'(s(X11,'const.option.option_CASE_3'(s('type.option.option'(X12),X14),s(X11,X15),s(fun(X12,X11),X13))))) | s(X11,X16) = s(X11,happ(s(fun(X12,X11),X13),s(X12,sK4(X14,X12)))) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),X14)) )),
  inference(superposition,[],[f162,f1387])).
fof(f12051,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s('type.option.option'(X52),'const.option.NONE_0') = s('type.option.option'(X52),X53) | s(X54,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(X54,X56),s(fun(X52,X54),X55))) = s(X54,happ(s(fun(X52,X54),X55),s(X52,sK4(sK4('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)),X52))))) )),
  inference(forward_demodulation,[],[f11778,f2072])).
fof(f11778,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(X54,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(X54,X56),s(fun(X52,X54),X55))) = s(X54,happ(s(fun(X52,X54),X55),s(X52,sK4(sK4('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)),X52)))) | s('type.option.option'(X52),'const.option.NONE_0') = s('type.option.option'(X52),sK4('const.option.SOME_1'(s('type.option.option'(X52),X53)),'type.option.option'(X52)))) )),
  inference(superposition,[],[f1387,f2072])).
fof(f11777,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (s(X49,happ(s(fun(X44,X49),X50),s(X44,sK4(happ(s(fun(X45,'type.option.option'(X44)),X46),s(X45,sK4(X47,X45))),X44)))) = s(X49,'const.option.option_CASE_3'(s('type.option.option'(X44),'const.option.option_CASE_3'(s('type.option.option'(X45),X47),s('type.option.option'(X44),X48),s(fun(X45,'type.option.option'(X44)),X46))),s(X49,X51),s(fun(X44,X49),X50))) | s('type.option.option'(X44),'const.option.NONE_0') = s('type.option.option'(X44),happ(s(fun(X45,'type.option.option'(X44)),X46),s(X45,sK4(X47,X45)))) | s('type.option.option'(X45),'const.option.NONE_0') = s('type.option.option'(X45),X47)) )),
  inference(superposition,[],[f1387,f1387])).
fof(f12050,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))) | s(X41,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))),s(X41,X43),s(fun(X36,X41),X42))) = s(X41,happ(s(fun(X36,X41),X42),s(X36,sK4(happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)),X36))))) )),
  inference(forward_demodulation,[],[f11776,f178])).
fof(f11776,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(X41,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.option_CASE_3'(s('type.option.option'(X37),X38),s('type.option.option'(X36),X39),s(fun(X37,'type.option.option'(X36)),X40))),s(X41,X43),s(fun(X36,X41),X42))) = s(X41,happ(s(fun(X36,X41),X42),s(X36,sK4(happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)),X36)))) | s('type.option.option'(X36),'const.option.NONE_0') = s('type.option.option'(X36),happ(s(fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)),happ(s(fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36))),happ(s(fun('type.option.option'(X37),fun('type.option.option'(X36),fun(fun(X37,'type.option.option'(X36)),'type.option.option'(X36)))),'const.option.option_CASE_0'),s('type.option.option'(X37),X38))),s('type.option.option'(X36),X39))),s(fun(X37,'type.option.option'(X36)),X40)))) )),
  inference(superposition,[],[f1387,f178])).
fof(f11775,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X33,'const.option.option_CASE_3'(s('type.option.option'(X28),happ(s(fun(X29,'type.option.option'(X28)),X32),s(X29,sK4(X30,X29)))),s(X33,X35),s(fun(X28,X33),X34))) = s(X33,happ(s(fun(X28,X33),X34),s(X28,sK4('const.option.option_CASE_3'(s('type.option.option'(X29),X30),s('type.option.option'(X28),X31),s(fun(X29,'type.option.option'(X28)),X32)),X28)))) | s('type.option.option'(X28),'const.option.NONE_0') = s('type.option.option'(X28),'const.option.option_CASE_3'(s('type.option.option'(X29),X30),s('type.option.option'(X28),X31),s(fun(X29,'type.option.option'(X28)),X32))) | s('type.option.option'(X29),'const.option.NONE_0') = s('type.option.option'(X29),X30)) )),
  inference(superposition,[],[f1387,f1387])).
fof(f12049,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s('type.option.option'(X20),'const.option.NONE_0') = s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))) | s(X25,'const.option.option_CASE_3'(s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))),s(X25,X27),s(fun(X20,X25),X26))) = s(X25,happ(s(fun(X20,X25),X26),s(X20,sK4('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)),X20))))) )),
  inference(forward_demodulation,[],[f11774,f177])).
fof(f11774,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X25,'const.option.option_CASE_3'(s('type.option.option'(X20),happ(s(fun(X21,'type.option.option'(X20)),X24),s(X21,X22))),s(X25,X27),s(fun(X20,X25),X26))) = s(X25,happ(s(fun(X20,X25),X26),s(X20,sK4('const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)),X20)))) | s('type.option.option'(X20),'const.option.NONE_0') = s('type.option.option'(X20),'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X22))),s('type.option.option'(X20),X23),s(fun(X21,'type.option.option'(X20)),X24)))) )),
  inference(superposition,[],[f1387,f177])).
fof(f12048,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s('type.option.option'(X13),'const.option.NONE_0') = s('type.option.option'(X13),X15) | s(X17,'const.option.option_CASE_3'(s('type.option.option'(X13),X15),s(X17,X19),s(fun(X13,X17),X18))) = s(X17,happ(s(fun(X13,X17),X18),s(X13,sK4('const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)),X13))))) )),
  inference(forward_demodulation,[],[f11773,f175])).
fof(f11773,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(X17,'const.option.option_CASE_3'(s('type.option.option'(X13),X15),s(X17,X19),s(fun(X13,X17),X18))) = s(X17,happ(s(fun(X13,X17),X18),s(X13,sK4('const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)),X13)))) | s('type.option.option'(X13),'const.option.NONE_0') = s('type.option.option'(X13),'const.option.option_CASE_3'(s('type.option.option'(X14),'const.option.NONE_0'),s('type.option.option'(X13),X15),s(fun(X14,'type.option.option'(X13)),X16)))) )),
  inference(superposition,[],[f1387,f175])).
fof(f12031,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X10),s(fun(X7,X6),X8))) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X11),s(fun(X7,X6),X8))) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9)) )),
  inference(duplicate_literal_removal,[],[f11760])).
fof(f11760,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X10),s(fun(X7,X6),X8))) = s(X6,'const.option.option_CASE_3'(s('type.option.option'(X7),X9),s(X6,X11),s(fun(X7,X6),X8))) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9)) )),
  inference(superposition,[],[f1387,f1387])).
fof(f11759,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),X2),s(X0,X3),s(fun(X1,X0),sK4(X4,fun(X1,X0))))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(fun(X1,X0)),X4),s(X0,X5),s(fun(fun(X1,X0),X0),happ(s(fun(X0,fun(fun(X1,X0),X0)),happ(s(fun('type.option.option'(X1),fun(X0,fun(fun(X1,X0),X0))),'const.option.option_CASE_0'),s('type.option.option'(X1),X2))),s(X0,X3))))) | s('type.option.option'(fun(X1,X0)),'const.option.NONE_0') = s('type.option.option'(fun(X1,X0)),X4)) )),
  inference(superposition,[],[f1387,f178])).
fof(f11517,plain,(
  ( ! [X80,X81] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X81))) | s(bool,X81) = s(bool,sK4('const.option.SOME_1'(s(bool,X80)),bool)) | s(bool,X80) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f162,f8908])).
fof(f11736,plain,(
  ( ! [X78] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) | s(bool,X78) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(subsumption_resolution,[],[f11515,f159])).
fof(f11515,plain,(
  ( ! [X78] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X78))) | s(bool,X78) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f154,f8908])).
fof(f11451,plain,(
  ( ! [X94,X95] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X94))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95))) | s(bool,X95) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X94)),bool))) ) | $spl3),
  inference(superposition,[],[f162,f8908])).
fof(f11373,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X125))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X124)),bool))) ) | $spl3),
  inference(superposition,[],[f163,f8908])).
fof(f11256,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl3),
  inference(superposition,[],[f8908,f8908])).
fof(f10628,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s('type.option.option'('type.option.option'(X111)),'const.option.SOME_1'(s('type.option.option'(X111),X112))) != s('type.option.option'('type.option.option'(X111)),'const.option.SOME_1'(s('type.option.option'(X111),X113))) | s('type.option.option'(X111),'const.option.SOME_1'(s(X111,sK4(X112,X111)))) = s('type.option.option'(X111),'const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s('type.option.option'(X111),X113),s(fun(X114,'type.option.option'(X111)),X115))) | s('type.option.option'(X111),'const.option.NONE_0') = s('type.option.option'(X111),X112)) )),
  inference(superposition,[],[f409,f154])).
fof(f10616,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X60))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X62),'const.option.NONE_0'),s(bool,X61),s(fun(X62,bool),X63))) | s(bool,X60) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f409,f239])).
fof(f10611,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X37) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X39),s(fun(X40,bool),X41))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f409,f239])).
fof(f10610,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X32) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X33),s(fun(X34,bool),X35))) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f409,f239])).
fof(f10608,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X24))) != s('type.option.option'(X23),'const.option.SOME_1'(s(X23,X25))) | s(X23,sK4('const.option.SOME_1'(s(X23,X24)),X23)) = s(X23,'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(X23,X25),s(fun(X26,X23),X27)))) )),
  inference(superposition,[],[f409,f2072])).
fof(f10607,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.option.option'(X15),'const.option.SOME_1'(s(X15,X20))) != s('type.option.option'(X15),'const.option.SOME_1'(s(X15,'const.option.option_CASE_3'(s('type.option.option'(X16),X17),s(X15,X18),s(fun(X16,X15),X19))))) | s(X15,'const.option.option_CASE_3'(s('type.option.option'(X21),'const.option.NONE_0'),s(X15,X20),s(fun(X21,X15),X22))) = s(X15,happ(s(fun(fun(X16,X15),X15),happ(s(fun(X15,fun(fun(X16,X15),X15)),happ(s(fun('type.option.option'(X16),fun(X15,fun(fun(X16,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X16),X17))),s(X15,X18))),s(fun(X16,X15),X19)))) )),
  inference(superposition,[],[f409,f178])).
fof(f10606,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X12))) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,happ(s(fun(X8,X7),X11),s(X8,X9))))) | s(X7,'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.NONE_0'),s(X7,X12),s(fun(X13,X7),X14))) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.SOME_1'(s(X8,X9))),s(X7,X10),s(fun(X8,X7),X11)))) )),
  inference(superposition,[],[f409,f177])).
fof(f10605,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X4))) | s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(X0,X4),s(fun(X5,X0),X6)))) )),
  inference(superposition,[],[f409,f175])).
fof(f10602,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(X0,sK4(X1,X0)),s(fun(X3,X0),X4))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f409,f154])).
fof(f10597,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s('type.option.option'('type.option.option'(X111)),'const.option.SOME_1'(s('type.option.option'(X111),X112))) != s('type.option.option'('type.option.option'(X111)),'const.option.SOME_1'(s('type.option.option'(X111),X113))) | s('type.option.option'(X111),X113) = s('type.option.option'(X111),'const.option.option_CASE_3'(s('type.option.option'(X114),'const.option.NONE_0'),s('type.option.option'(X111),'const.option.SOME_1'(s(X111,sK4(X112,X111)))),s(fun(X114,'type.option.option'(X111)),X115))) | s('type.option.option'(X111),'const.option.NONE_0') = s('type.option.option'(X111),X112)) )),
  inference(superposition,[],[f409,f154])).
fof(f10585,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X60))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X61))) | s(bool,X61) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X62),'const.option.NONE_0'),s(bool,'T_0'),s(fun(X62,bool),X63))) | s(bool,X60) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f409,f239])).
fof(f10580,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X39) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),'const.option.NONE_0'),s(bool,X37),s(fun(X40,bool),X41))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f409,f239])).
fof(f10579,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X33) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X34),'const.option.NONE_0'),s(bool,X32),s(fun(X34,bool),X35))) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f409,f239])).
fof(f10367,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),X2) | s(X0,sK4(X1,X0)) = s(X0,sK4(X2,X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X2) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f956,f154])).
fof(f10362,plain,(
  ( ! [X64,X65,X63] : (s('type.option.option'('type.option.option'(X63)),X65) != s('type.option.option'('type.option.option'(X63)),'const.option.SOME_1'(s('type.option.option'(X63),X64))) | s('type.option.option'(X63),sK4(X65,'type.option.option'(X63))) = s('type.option.option'(X63),'const.option.SOME_1'(s(X63,sK4(X64,X63)))) | s('type.option.option'('type.option.option'(X63)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(X63)),X65) | s('type.option.option'(X63),'const.option.NONE_0') = s('type.option.option'(X63),X64)) )),
  inference(superposition,[],[f956,f154])).
fof(f10351,plain,(
  ( ! [X39,X38,X40] : (s('type.option.option'(bool),X39) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38))) | s(bool,'T_0') = s(bool,sK4(X39,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f956,f239])).
fof(f10346,plain,(
  ( ! [X26,X27,X25] : (s('type.option.option'(bool),X27) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(bool,X25) = s(bool,sK4(X27,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f956,f239])).
fof(f10345,plain,(
  ( ! [X24,X23,X22] : (s('type.option.option'(bool),X23) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X22) = s(bool,sK4(X23,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f956,f239])).
fof(f10343,plain,(
  ( ! [X19,X17,X18] : (s('type.option.option'(X17),X19) != s('type.option.option'(X17),'const.option.SOME_1'(s(X17,X18))) | s(X17,sK4(X19,X17)) = s(X17,sK4('const.option.SOME_1'(s(X17,X18)),X17)) | s('type.option.option'(X17),'const.option.NONE_0') = s('type.option.option'(X17),X19)) )),
  inference(superposition,[],[f956,f2072])).
fof(f10342,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.option.option'(X11),X16) != s('type.option.option'(X11),'const.option.SOME_1'(s(X11,'const.option.option_CASE_3'(s('type.option.option'(X12),X13),s(X11,X14),s(fun(X12,X11),X15))))) | s(X11,sK4(X16,X11)) = s(X11,happ(s(fun(fun(X12,X11),X11),happ(s(fun(X11,fun(fun(X12,X11),X11)),happ(s(fun('type.option.option'(X12),fun(X11,fun(fun(X12,X11),X11))),'const.option.option_CASE_0'),s('type.option.option'(X12),X13))),s(X11,X14))),s(fun(X12,X11),X15))) | s('type.option.option'(X11),'const.option.NONE_0') = s('type.option.option'(X11),X16)) )),
  inference(superposition,[],[f956,f178])).
fof(f10341,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),X10) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,happ(s(fun(X6,X5),X9),s(X6,X7))))) | s(X5,sK4(X10,X5)) = s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s(X5,X8),s(fun(X6,X5),X9))) | s('type.option.option'(X5),'const.option.NONE_0') = s('type.option.option'(X5),X10)) )),
  inference(superposition,[],[f956,f177])).
fof(f10340,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'(X0),X4) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,sK4(X4,X0)) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X4)) )),
  inference(superposition,[],[f956,f175])).
fof(f10377,plain,(
  ( ! [X26,X24,X25] : (s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X25) | s('type.option.option'(X24),X25) != s('type.option.option'(X24),'const.option.SOME_1'(s(X24,X26))) | s(X24,X26) = s(X24,sK4(sK4('const.option.SOME_1'(s('type.option.option'(X24),X25)),'type.option.option'(X24)),X24))) )),
  inference(forward_demodulation,[],[f10339,f2072])).
fof(f10339,plain,(
  ( ! [X26,X24,X25] : (s('type.option.option'(X24),X25) != s('type.option.option'(X24),'const.option.SOME_1'(s(X24,X26))) | s(X24,X26) = s(X24,sK4(sK4('const.option.SOME_1'(s('type.option.option'(X24),X25)),'type.option.option'(X24)),X24)) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),sK4('const.option.SOME_1'(s('type.option.option'(X24),X25)),'type.option.option'(X24)))) )),
  inference(superposition,[],[f956,f2072])).
fof(f10376,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'(X18),'const.option.NONE_0') = s('type.option.option'(X18),'const.option.option_CASE_3'(s('type.option.option'(X19),X20),s('type.option.option'(X18),X21),s(fun(X19,'type.option.option'(X18)),X22))) | s('type.option.option'(X18),'const.option.SOME_1'(s(X18,X23))) != s('type.option.option'(X18),'const.option.option_CASE_3'(s('type.option.option'(X19),X20),s('type.option.option'(X18),X21),s(fun(X19,'type.option.option'(X18)),X22))) | s(X18,X23) = s(X18,sK4(happ(s(fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)),happ(s(fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18))),happ(s(fun('type.option.option'(X19),fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)))),'const.option.option_CASE_0'),s('type.option.option'(X19),X20))),s('type.option.option'(X18),X21))),s(fun(X19,'type.option.option'(X18)),X22)),X18))) )),
  inference(forward_demodulation,[],[f10338,f178])).
fof(f10338,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'(X18),'const.option.SOME_1'(s(X18,X23))) != s('type.option.option'(X18),'const.option.option_CASE_3'(s('type.option.option'(X19),X20),s('type.option.option'(X18),X21),s(fun(X19,'type.option.option'(X18)),X22))) | s(X18,X23) = s(X18,sK4(happ(s(fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)),happ(s(fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18))),happ(s(fun('type.option.option'(X19),fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)))),'const.option.option_CASE_0'),s('type.option.option'(X19),X20))),s('type.option.option'(X18),X21))),s(fun(X19,'type.option.option'(X18)),X22)),X18)) | s('type.option.option'(X18),'const.option.NONE_0') = s('type.option.option'(X18),happ(s(fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)),happ(s(fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18))),happ(s(fun('type.option.option'(X19),fun('type.option.option'(X18),fun(fun(X19,'type.option.option'(X18)),'type.option.option'(X18)))),'const.option.option_CASE_0'),s('type.option.option'(X19),X20))),s('type.option.option'(X18),X21))),s(fun(X19,'type.option.option'(X18)),X22)))) )),
  inference(superposition,[],[f956,f178])).
fof(f10375,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),happ(s(fun(X13,'type.option.option'(X12)),X16),s(X13,X14))) | s('type.option.option'(X12),'const.option.SOME_1'(s(X12,X17))) != s('type.option.option'(X12),happ(s(fun(X13,'type.option.option'(X12)),X16),s(X13,X14))) | s(X12,X17) = s(X12,sK4('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(X12),X15),s(fun(X13,'type.option.option'(X12)),X16)),X12))) )),
  inference(forward_demodulation,[],[f10337,f177])).
fof(f10337,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'(X12),'const.option.SOME_1'(s(X12,X17))) != s('type.option.option'(X12),happ(s(fun(X13,'type.option.option'(X12)),X16),s(X13,X14))) | s(X12,X17) = s(X12,sK4('const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(X12),X15),s(fun(X13,'type.option.option'(X12)),X16)),X12)) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),'const.option.option_CASE_3'(s('type.option.option'(X13),'const.option.SOME_1'(s(X13,X14))),s('type.option.option'(X12),X15),s(fun(X13,'type.option.option'(X12)),X16)))) )),
  inference(superposition,[],[f956,f177])).
fof(f10374,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),X9) | s('type.option.option'(X7),X9) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X11))) | s(X7,X11) = s(X7,sK4('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s('type.option.option'(X7),X9),s(fun(X8,'type.option.option'(X7)),X10)),X7))) )),
  inference(forward_demodulation,[],[f10336,f175])).
fof(f10336,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.option.option'(X7),X9) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X11))) | s(X7,X11) = s(X7,sK4('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s('type.option.option'(X7),X9),s(fun(X8,'type.option.option'(X7)),X10)),X7)) | s('type.option.option'(X7),'const.option.NONE_0') = s('type.option.option'(X7),'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s('type.option.option'(X7),X9),s(fun(X8,'type.option.option'(X7)),X10)))) )),
  inference(superposition,[],[f956,f175])).
fof(f10371,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK4('const.option.SOME_1'(s(X0,sK4(X1,X0))),X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(subsumption_resolution,[],[f10333,f159])).
fof(f10333,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK4('const.option.SOME_1'(s(X0,sK4(X1,X0))),X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f956,f154])).
fof(f8985,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,X33) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X32)) ) | $spl3),
  inference(superposition,[],[f162,f8665])).
fof(f8946,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'T_0') = s(bool,X42)) ) | $spl3),
  inference(superposition,[],[f163,f8665])).
fof(f8908,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X27)),bool))) ) | $spl3),
  inference(superposition,[],[f8665,f2072])).
fof(f8665,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f8653,f126])).
fof(f8653,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f213,f239])).
fof(f8645,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f126,f213,f204])).
fof(f3120,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.option.option'(X0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.NONE_0'))))) != s('type.option.option'('type.option.option'('type.option.option'(X0))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))))) )),
  inference(unit_resulting_resolution,[],[f403,f162])).
fof(f2634,plain,(
  ( ! [X191,X192,X190,X193] : (s(X192,f(s(bool,'T_0'))) != s(X192,g(s(bool,'T_0'))) | s(fun(bool,X192),X190) = s(fun(bool,X192),X191) | s(bool,X193) = s(bool,sK6(X190,X191,X192,bool)) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f176,f239])).
fof(f2631,plain,(
  ( ! [X184,X183] : (s('type.option.option'(bool),X183) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X183) | s(bool,X184) = s(bool,sK4(X183,bool)) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f154,f239])).
fof(f2614,plain,(
  ( ! [X116,X117,X115] : (s(bool,'T_0') != s(bool,X116) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X115))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X116))) | s(bool,X115) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f163,f239])).
fof(f2612,plain,(
  ( ! [X111,X109,X110] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,X109) = s(bool,X110) | s(bool,X109) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f162,f239])).
fof(f2990,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2579,f126])).
fof(f2579,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f170,f239])).
fof(f2989,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2575,f126])).
fof(f2575,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f135,f239])).
fof(f2559,plain,(
  ( ! [X66,X67,X65] : (s(bool,X65) != s(bool,X66) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X66))) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f163,f239])).
fof(f2536,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f131,f239])).
fof(f2466,plain,(
  ( ! [X202,X200,X199,X201] : (s(X201,f(s(bool,X202))) != s(X201,g(s(bool,X202))) | s(fun(bool,X201),X199) = s(fun(bool,X201),X200) | s(bool,'T_0') = s(bool,X202) | s(bool,'T_0') = s(bool,sK6(X199,X200,X201,bool))) )),
  inference(superposition,[],[f176,f239])).
fof(f2463,plain,(
  ( ! [X192,X193] : (s('type.option.option'(bool),X192) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X193))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X192) | s(bool,'T_0') = s(bool,X193) | s(bool,'T_0') = s(bool,sK4(X192,bool))) )),
  inference(superposition,[],[f154,f239])).
fof(f2446,plain,(
  ( ! [X125,X123,X124] : (s(bool,X124) != s(bool,X125) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X125))) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f163,f239])).
fof(f2444,plain,(
  ( ! [X118,X119,X117] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X118))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f162,f239])).
fof(f2381,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,sK4('const.option.SOME_1'(s(bool,X38)),bool)) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f239,f2072])).
fof(f2378,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X24),'const.option.NONE_0'),s(bool,X25),s(fun(X24,bool),X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f239,f175])).
fof(f2365,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f239,f239])).
fof(f2344,plain,(
  ( ! [X50,X49] : (s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X49)),bool))) )),
  inference(superposition,[],[f239,f2072])).
fof(f2341,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X35),'const.option.NONE_0'),s(bool,X36),s(fun(X35,bool),X37)))) )),
  inference(superposition,[],[f239,f175])).
fof(f2092,plain,(
  ( ! [X8,X7,X9] : (s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X8))) != s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X9))) | s(X7,X9) = s(X7,sK4('const.option.SOME_1'(s(X7,X8)),X7))) )),
  inference(superposition,[],[f162,f951])).
fof(f2072,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK4('const.option.SOME_1'(s(X0,X1)),X0))) )),
  inference(unit_resulting_resolution,[],[f951,f162])).
fof(f1613,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X10))) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),X7),s(X5,X8),s(fun(X6,X5),X9))))) | s(X5,X10) = s(X5,happ(s(fun(fun(X6,X5),X5),happ(s(fun(X5,fun(fun(X6,X5),X5)),happ(s(fun('type.option.option'(X6),fun(X5,fun(fun(X6,X5),X5))),'const.option.option_CASE_0'),s('type.option.option'(X6),X7))),s(X5,X8))),s(fun(X6,X5),X9)))) )),
  inference(superposition,[],[f162,f178])).
fof(f1405,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X10))) != s('type.option.option'(X5),'const.option.SOME_1'(s(X5,happ(s(fun(X6,X5),X9),s(X6,X7))))) | s(X5,X10) = s(X5,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X7))),s(X5,X8),s(fun(X6,X5),X9)))) )),
  inference(superposition,[],[f162,f177])).
fof(f1387,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK4(X1,X0)))) = s(X2,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(X2,X4),s(fun(X0,X2),X3))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f177,f154])).
fof(f982,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0)) | $spl3),
  inference(subsumption_resolution,[],[f974,f128])).
fof(f974,plain,(
  p(s(bool,'F_0')) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0)) | $spl3),
  inference(superposition,[],[f892,f510])).
fof(f972,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK3)))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | $spl3),
  inference(superposition,[],[f892,f175])).
fof(f964,plain,(
  ( ! [X19,X20,X18] : (s('type.option.option'(X18),X19) != s('type.option.option'(X18),X20) | s('type.option.option'('type.option.option'(X18)),'const.option.SOME_1'(s('type.option.option'(X18),X20))) = s('type.option.option'('type.option.option'(X18)),'const.option.SOME_1'(s('type.option.option'(X18),'const.option.SOME_1'(s(X18,sK4(X19,X18)))))) | s('type.option.option'(X18),'const.option.NONE_0') = s('type.option.option'(X18),X19)) )),
  inference(superposition,[],[f163,f154])).
fof(f962,plain,(
  ( ! [X14,X12,X13] : (s('type.option.option'('type.option.option'(X12)),'const.option.SOME_1'(s('type.option.option'(X12),X13))) != s('type.option.option'('type.option.option'(X12)),'const.option.SOME_1'(s('type.option.option'(X12),X14))) | s('type.option.option'(X12),X14) = s('type.option.option'(X12),'const.option.SOME_1'(s(X12,sK4(X13,X12)))) | s('type.option.option'(X12),'const.option.NONE_0') = s('type.option.option'(X12),X13)) )),
  inference(superposition,[],[f162,f154])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK4(X1,X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f162,f154])).
fof(f951,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4('const.option.SOME_1'(s(X0,X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f159,f154])).
fof(f892,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK3)))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X4,sK0))) ) | $spl3),
  inference(forward_demodulation,[],[f886,f875])).
fof(f886,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK3)))) | s(sK0,sK5(sK3,sK2,sK0)) = s(sK0,sK5(sK5(sK3,sK2,sK0),X4,sK0))) ) | $spl3),
  inference(superposition,[],[f164,f875])).
fof(f875,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0)) | $spl3),
  inference(subsumption_resolution,[],[f860,f128])).
fof(f860,plain,(
  p(s(bool,'F_0')) | s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0)) | $spl3),
  inference(superposition,[],[f164,f510])).
fof(f869,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X13,bool)) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f857,f510])).
fof(f857,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X13,bool))) ) | $spl3),
  inference(superposition,[],[f164,f510])).
fof(f868,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X12,bool)) | p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(forward_demodulation,[],[f856,f510])).
fof(f856,plain,(
  ( ! [X12,X11] : (p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X12,bool)) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(superposition,[],[f164,f586])).
fof(f717,plain,(
  ( ! [X21,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) | s(bool,X22) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X21)) ) | $spl3),
  inference(superposition,[],[f162,f586])).
fof(f696,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X15)) ) | $spl3),
  inference(superposition,[],[f163,f586])).
fof(f577,plain,(
  ( ! [X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl3),
  inference(superposition,[],[f162,f510])).
fof(f586,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f574,f128])).
fof(f574,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f272,f510])).
fof(f510,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f213,f136])).
fof(f511,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f213,f131])).
fof(f409,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.option.option'(X9),'const.option.SOME_1'(s(X9,X11))) != s('type.option.option'(X9),'const.option.SOME_1'(s(X9,X13))) | s(X9,X13) = s(X9,'const.option.option_CASE_3'(s('type.option.option'(X10),'const.option.NONE_0'),s(X9,X11),s(fun(X10,X9),X12)))) )),
  inference(superposition,[],[f162,f175])).
fof(f403,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.NONE_0'))) != s('type.option.option'('type.option.option'(X0)),'const.option.SOME_1'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f159,f162])).
fof(f272,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f266,f128])).
fof(f266,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f167,f129])).
fof(f239,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f129,f129])).
fof(f217,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f128,f131])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f203])).
fof(f203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f192])).
fof(f192,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13,X17,X14,X18] : ((s(bool,X13) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X14) = s(bool,X18))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.bool.IMP_CONG')).
fof(f178,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X15] : ! [X20,X21,X22] : s(X15,'const.option.option_CASE_3'(s('type.option.option'(X6),X20),s(X15,X21),s(fun(X6,X15),X22))) = s(X15,happ(s(fun(fun(X6,X15),X15),happ(s(fun(X15,fun(fun(X6,X15),X15)),happ(s(fun('type.option.option'(X6),fun(X15,fun(fun(X6,X15),X15))),'const.option.option_CASE_0'),s('type.option.option'(X6),X20))),s(X15,X21))),s(fun(X6,X15),X22)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','ARITY_0_3_const.option.option_CASE')).
fof(f177,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(X0,X3),s(fun(X1,X0),X4)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X15,X6] : ! [X13,X19,X8] : s(X15,happ(s(fun(X6,X15),X8),s(X6,X13))) = s(X15,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13))),s(X15,X19),s(fun(X6,X15),X8)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.option.option_case_def|split|1')).
fof(f176,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.clos_relation.clos_relation_87_prove_ltb',aHL_EXT)).
fof(f175,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X15] : ! [X19,X8] : s(X15,X19) = s(X15,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.NONE_0'),s(X15,X19),s(fun(X6,X15),X8)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.option.option_case_def|split|0')).
fof(f167,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f170,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f162,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X13,X14] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X14))) <=> s(X6,X13) = s(X6,X14))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.option.SOME_|49||49|')).
fof(f163,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f159,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : ~s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X13] : ~s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.option.NOT_NONE_SOME')).
fof(f154,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK4(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X7] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X7) | ? [X13] : s('type.option.option'(X6),X7) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X13))))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.option.option_nchotomy')).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relation.clos_relation_87_prove_ltb',aHL_BOOL_CASES)).
fof(f126,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relation.clos_relation_87_prove_ltb',aHL_TRUTH)).
fof(f13264,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f13241,f210])).
fof(f13241,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun(sK0,X2),X3),s(sK0,sK3))) = s(X2,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(X2,X4),s(fun(sK0,X2),X3)))) ) | $spl4),
  inference(superposition,[],[f177,f215])).
fof(f13210,plain,(
  $false | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f13209,f128])).
fof(f13209,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl32)),
  inference(forward_demodulation,[],[f210,f9641])).
fof(f9641,plain,(
  s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl32),
  inference(equality_resolution,[],[f6492])).
fof(f13007,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl77),
  inference(cnf_transformation,[],[f13007_D])).
fof(f13007_D,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f7719,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,X39)) ) | $spl58),
  inference(cnf_transformation,[],[f7719_D])).
fof(f7719_D,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,X39)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13163,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | $spl32),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10362,f10367,f10579,f10580,f10585,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10628,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11850,f11851,f11856,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11972,f11971,f11762,f11829,f11858,f11863,f11865,f11866,f11867,f11868,f11869,f11871,f11872,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11886,f9520,f3468,f3492,f9519,f3467,f3491,f11893,f2459,f2627,f10357,f10622,f1630,f1669,f1670,f2343,f2380,f3247,f9870,f10905,f11165,f11859,f11862,f164,f165,f166,f849,f870,f1319,f2192,f2217,f2234,f2342,f2379,f2996,f2997,f3246,f9869,f10861,f10864,f11173,f11844,f1664,f11172,f11843,f11852,f11853,f11171,f1437,f1650,f1683,f10863,f1430,f1450,f10862,f853,f1474,f2449,f2465,f2561,f2562,f2617,f2633,f3326,f10121,f10124,f10871,f10876,f10878,f11129,f11134,f2616,f3325,f10870,f11128,f11181,f10123,f10122,f966,f1480,f10889,f11148,f11911,f11912,f11913,f11914,f11916,f11917,f11938,f11939,f11947,f11948,f11953,f12009,f12013,f12012,f12010,f12008,f11973,f11968,f11967,f11966,f11962,f11961,f12007,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11964,f11963,f10587,f10578,f9871,f9868,f9522,f9521,f9518,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f9626,f9625,f9622,f9607,f6492,f11989,f9893,f13142,f9641,f13151,f13152,f9652,f9637,f13155,f9651,f13156,f13157,f13158,f13159,f13160,f13161,f13162])).
fof(f13162,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(subsumption_resolution,[],[f9610,f13155])).
fof(f9610,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f13160,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) ) | $spl32),
  inference(subsumption_resolution,[],[f9613,f13155])).
fof(f9613,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f13159,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X13),s(fun(sK0,bool),sK2))) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl32),
  inference(subsumption_resolution,[],[f9617,f217])).
fof(f9617,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X13),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f13157,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'F_0') != s(bool,happ(s(fun(X27,bool),X30),s(X27,X28))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s(bool,X29),s(fun(X27,bool),X30)))) ) | $spl32),
  inference(forward_demodulation,[],[f9623,f9641])).
fof(f9623,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,happ(s(fun(X27,bool),X30),s(X27,X28))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s(bool,X29),s(fun(X27,bool),X30)))) ) | $spl32),
  inference(superposition,[],[f6492,f177])).
fof(f13156,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X31),X32),s(bool,X33),s(fun(X31,bool),X34))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X31,bool),bool),happ(s(fun(bool,fun(fun(X31,bool),bool)),happ(s(fun('type.option.option'(X31),fun(bool,fun(fun(X31,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X31),X32))),s(bool,X33))),s(fun(X31,bool),X34)))) ) | $spl32),
  inference(forward_demodulation,[],[f9624,f9641])).
fof(f9624,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X31),X32),s(bool,X33),s(fun(X31,bool),X34))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X31,bool),bool),happ(s(fun(bool,fun(fun(X31,bool),bool)),happ(s(fun('type.option.option'(X31),fun(bool,fun(fun(X31,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X31),X32))),s(bool,X33))),s(fun(X31,bool),X34)))) ) | $spl32),
  inference(superposition,[],[f6492,f178])).
fof(f13155,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'F_0') = s(bool,X9)) ) | $spl32),
  inference(global_subsumption,[],[f9637,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10362,f10367,f10579,f10580,f10585,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10628,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11850,f11851,f11856,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11972,f11971,f11762,f11829,f11858,f11863,f11865,f11866,f11867,f11868,f11869,f11871,f11872,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11886,f9520,f3468,f3492,f9519,f3467,f3491,f2459,f2627,f10357,f10622,f1630,f1669,f1670,f2343,f2380,f3247,f9870,f10905,f11165,f11859,f11862,f164,f165,f166,f849,f870,f1319,f2192,f2217,f2234,f2342,f2379,f2996,f2997,f3246,f9869,f10861,f10864,f11173,f11844,f1664,f11172,f11843,f11852,f11853,f11171,f1437,f1650,f1683,f10863,f1430,f1450,f10862,f853,f1474,f2449,f2465,f2561,f2562,f2617,f2633,f3326,f10121,f10124,f10871,f10876,f10878,f11129,f11134,f2616,f3325,f10870,f11128,f11181,f10123,f10122,f966,f1480,f10889,f11148,f11911,f11912,f11913,f11914,f11916,f11917,f11938,f11939,f11947,f11948,f12009,f12013,f12012,f12010,f12008,f11973,f11968,f11967,f11966,f11962,f11961,f12007,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11964,f11963,f10587,f10578,f9871,f9868,f9522,f9521,f9518,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f9893,f13142])).
fof(f9637,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X8)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f13152,plain,(
  ( ! [X310,X312,X309,X311] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X309),X311),s(bool,X312),s(fun(X309,bool),X310))) | s(bool,'F_0') = s(bool,happ(s(fun(X309,bool),X310),s(X309,sK4(X311,X309)))) | s('type.option.option'(X309),'const.option.NONE_0') = s('type.option.option'(X309),X311)) ) | $spl32),
  inference(backward_demodulation,[],[f9641,f11893])).
fof(f13151,plain,(
  ( ! [X288,X290,X289,X291] : (s(bool,'F_0') != s(bool,happ(s(fun(X288,bool),X291),s(X288,sK4(X289,X288)))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X288),X289),s(bool,X290),s(fun(X288,bool),X291))) | s('type.option.option'(X288),'const.option.NONE_0') = s('type.option.option'(X288),X289)) ) | $spl32),
  inference(backward_demodulation,[],[f9641,f11989])).
fof(f13142,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f9880,f398])).
fof(f9880,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f445,f239])).
fof(f9893,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_resolution,[],[f445])).
fof(f11989,plain,(
  ( ! [X288,X290,X289,X291] : (s(bool,happ(s(fun(X288,bool),X291),s(X288,sK4(X289,X288)))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X288),X289),s(bool,X290),s(fun(X288,bool),X291))) | s('type.option.option'(X288),'const.option.NONE_0') = s('type.option.option'(X288),X289)) ) | $spl32),
  inference(superposition,[],[f6492,f1387])).
fof(f9607,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(sK0,bool),bool),happ(s(fun(bool,fun(fun(sK0,bool),bool)),happ(s(fun('type.option.option'(sK0),fun(bool,fun(fun(sK0,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(sK0),sK1))),s(bool,'T_0'))),s(fun(sK0,bool),sK2))) | $spl32),
  inference(unit_resulting_resolution,[],[f178,f6492])).
fof(f9622,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X24),'const.option.NONE_0'),s(bool,X25),s(fun(X24,bool),X26)))) ) | $spl32),
  inference(superposition,[],[f6492,f175])).
fof(f9625,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool))) ) | $spl32),
  inference(superposition,[],[f6492,f2072])).
fof(f9626,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X0),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f11893,plain,(
  ( ! [X310,X312,X309,X311] : (s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X309),X311),s(bool,X312),s(fun(X309,bool),X310))) | s(bool,'F_0') = s(bool,happ(s(fun(X309,bool),X310),s(X309,sK4(X311,X309)))) | s('type.option.option'(X309),'const.option.NONE_0') = s('type.option.option'(X309),X311)) ) | $spl32),
  inference(superposition,[],[f6492,f1387])).
fof(f13161,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X7)) ) | $spl32),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10362,f10367,f10579,f10580,f10585,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10628,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11850,f11851,f11856,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11972,f11971,f11762,f11829,f11858,f11863,f11865,f11866,f11867,f11868,f11869,f11871,f11872,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11886,f9520,f3468,f3492,f9519,f3467,f3491,f11893,f2459,f2627,f10357,f10622,f1630,f1669,f1670,f2343,f2380,f3247,f9870,f10905,f11165,f11859,f11862,f164,f165,f166,f849,f870,f1319,f2192,f2217,f2234,f2342,f2379,f2996,f2997,f3246,f9869,f10861,f10864,f11173,f11844,f1664,f11172,f11843,f11852,f11853,f11171,f1437,f1650,f1683,f10863,f1430,f1450,f10862,f853,f1474,f2449,f2465,f2561,f2562,f2617,f2633,f3326,f10121,f10124,f10871,f10876,f10878,f11129,f11134,f2616,f3325,f10870,f11128,f11181,f10123,f10122,f966,f1480,f10889,f11148,f11911,f11912,f11913,f11914,f11916,f11917,f11938,f11939,f11947,f11948,f11953,f12009,f12013,f12012,f12010,f12008,f11973,f11968,f11967,f11966,f11962,f11961,f12007,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11964,f11963,f10587,f10578,f9871,f9868,f9522,f9521,f9518,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f9626,f9625,f9622,f9607,f6492,f11989,f9893,f13142,f9641,f13151,f13152,f9652,f9637,f13155,f9651,f13156,f13157,f13158,f13159,f13160])).
fof(f13158,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X18)) ) | $spl32),
  inference(subsumption_resolution,[],[f9620,f3315])).
fof(f9620,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X17),'const.option.NONE_0'),s(bool,X18),s(fun(X17,bool),X19))) | s(bool,'F_0') = s(bool,X18)) ) | $spl32),
  inference(superposition,[],[f6492,f303])).
fof(f13154,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X11)) ) | $spl84),
  inference(cnf_transformation,[],[f13154_D])).
fof(f13154_D,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X11)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9651,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl32),
  inference(subsumption_resolution,[],[f9635,f129])).
fof(f9635,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f9652,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X11)) ) | $spl32),
  inference(subsumption_resolution,[],[f9638,f129])).
fof(f9638,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X11)) ) | $spl32),
  inference(superposition,[],[f6492,f239])).
fof(f494,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13097,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl82),
  inference(cnf_transformation,[],[f13097_D])).
fof(f13097_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9894,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f9857])).
fof(f9857,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f445,f8665])).
fof(f13013,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl78),
  inference(cnf_transformation,[],[f13013_D])).
fof(f13013_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f13096,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X34)),bool))) ) | $spl3),
  inference(subsumption_resolution,[],[f11322,f13047])).
fof(f13047,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X19)),bool))) ) | $spl3),
  inference(subsumption_resolution,[],[f13046,f128])).
fof(f13046,plain,(
  ( ! [X19] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X19) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X19)),bool))) ) | $spl3),
  inference(forward_demodulation,[],[f11410,f13015])).
fof(f11410,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X19)),bool))) ) | $spl3),
  inference(superposition,[],[f131,f8908])).
fof(f11322,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X34)),bool))) ) | $spl3),
  inference(superposition,[],[f445,f8908])).
fof(f11191,plain,(
  ( ! [X4,X3] : (s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X3,fun(sK0,bool)),X4)),sK0))) ) | $spl68),
  inference(cnf_transformation,[],[f11191_D])).
fof(f11191_D,plain,(
  ( ! [X4,X3] : (s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X3,fun(sK0,bool)),X4)),sK0))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13051,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,X7) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X8,fun(sK0,bool)),X9)),sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f13050,f132])).
fof(f13050,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,X7) | p(s(bool,X7)) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X8,fun(sK0,bool)),X9)),sK0))) ) | $spl3),
  inference(forward_demodulation,[],[f11403,f2072])).
fof(f11403,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X8,fun(sK0,bool)),X9)),sK0)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X7)),bool))) ) | $spl3),
  inference(superposition,[],[f849,f8908])).
fof(f10925,plain,(
  ( ! [X4,X3] : (s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(sK0,sK3),s(fun(X3,sK0),X4)),sK2,sK0))) ) | $spl66),
  inference(cnf_transformation,[],[f10925_D])).
fof(f10925_D,plain,(
  ( ! [X4,X3] : (s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(sK0,sK3),s(fun(X3,sK0),X4)),sK2,sK0))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f13049,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,X10) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X11),'const.option.NONE_0'),s(sK0,sK3),s(fun(X11,sK0),X12)),sK2,sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f13048,f132])).
fof(f13048,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,X10) | p(s(bool,X10)) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X11),'const.option.NONE_0'),s(sK0,sK3),s(fun(X11,sK0),X12)),sK2,sK0))) ) | $spl3),
  inference(forward_demodulation,[],[f11404,f2072])).
fof(f11404,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X11),'const.option.NONE_0'),s(sK0,sK3),s(fun(X11,sK0),X12)),sK2,sK0)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X10)),bool))) ) | $spl3),
  inference(superposition,[],[f870,f8908])).
fof(f13014,plain,(
  ( ! [X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl80),
  inference(cnf_transformation,[],[f13014_D])).
fof(f13014_D,plain,(
  ( ! [X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X51)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11684,plain,(
  ( ! [X50,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X50)),bool))) ) | $spl3),
  inference(subsumption_resolution,[],[f11427,f2157])).
fof(f11427,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,X50) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X50)),bool))) ) | $spl3),
  inference(superposition,[],[f441,f8908])).
fof(f6614,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | $spl42),
  inference(cnf_transformation,[],[f6614_D])).
fof(f6614_D,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f11687,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))))) ) | $spl3),
  inference(condensation,[],[f11686])).
fof(f11686,plain,(
  ( ! [X54,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X54)),bool))) ) | $spl3),
  inference(subsumption_resolution,[],[f11429,f2344])).
fof(f11429,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X55) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X54)),bool))) ) | $spl3),
  inference(superposition,[],[f445,f8908])).
fof(f4514,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl28),
  inference(cnf_transformation,[],[f4514_D])).
fof(f4514_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6505,plain,(
  ( ! [X55] : (~p(s(bool,X55))) ) | $spl38),
  inference(cnf_transformation,[],[f6505_D])).
fof(f6505_D,plain,(
  ( ! [X55] : (~p(s(bool,X55))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f13008,plain,(
  ( ! [X14,X13] : (~p(s(bool,X14)) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl3 | $spl34)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f213,f8645,f8665,f8908,f8946,f8985,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10362,f10367,f10579,f10580,f10585,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10628,f10924,f11190,f11256,f11373,f11451,f11736,f11517,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11850,f11851,f11856,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11740,f11972,f11971,f11762,f11768,f11769,f11829,f11835,f11836,f11858,f11863,f11865,f11866,f11867,f11868,f11869,f11871,f11872,f11877,f11878,f11879,f11880,f11881,f11882,f11883,f11884,f11886,f9520,f3468,f3492,f9519,f3467,f3491,f2459,f2627,f10357,f10622,f1630,f1669,f1670,f2343,f2380,f3247,f9870,f10905,f11165,f11859,f11862,f164,f165,f166,f849,f870,f1319,f2192,f2217,f2234,f2342,f2379,f2996,f2997,f3246,f9869,f10861,f10864,f11173,f11844,f1664,f11172,f11843,f11852,f11853,f11171,f1437,f1650,f1683,f10863,f1430,f1450,f10862,f853,f1474,f2449,f2465,f2561,f2562,f2617,f2633,f3326,f10121,f10124,f10871,f10876,f10878,f11129,f11134,f2616,f3325,f10870,f11128,f11181,f10123,f10122,f966,f1480,f10889,f11148,f11534,f11911,f11912,f11913,f11914,f11916,f11917,f11938,f11939,f11947,f11948,f12009,f12013,f12012,f12010,f12008,f11973,f11968,f11967,f11966,f11962,f11961,f12007,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11964,f11963,f11720,f10587,f10578,f9871,f9868,f9522,f9521,f9518,f6493,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f11737,f11751,f11749,f11529,f11746,f11744,f11742,f11725,f11722])).
fof(f13005,plain,(
  ( ! [X52] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,X52) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl3 | $spl32)),
  inference(subsumption_resolution,[],[f11499,f11720])).
fof(f11499,plain,(
  ( ! [X52] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X52)),bool)) | s(bool,X52) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl3 | $spl32)),
  inference(superposition,[],[f6492,f8908])).
fof(f1725,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | $spl18),
  inference(cnf_transformation,[],[f1725_D])).
fof(f1725_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f13004,plain,(
  ( ! [X142] : (s(bool,X142) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'F_0') = s(bool,X142)) ) | $spl74),
  inference(cnf_transformation,[],[f13004_D])).
fof(f13004_D,plain,(
  ( ! [X142] : (s(bool,X142) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'F_0') = s(bool,X142)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f11953,plain,(
  ( ! [X142] : (s(bool,X142) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'F_0') = s(bool,X142) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1)) ) | $spl32),
  inference(superposition,[],[f6492,f1387])).
fof(f12982,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750,f12790,f929,f12943,f12981])).
fof(f12981,plain,(
  ( ! [X294,X296,X298,X295] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X298))) | s(bool,X298) = s(bool,happ(s(fun(X294,bool),X295),s(X294,sK4(X296,X294)))) | s('type.option.option'(X294),'const.option.NONE_0') = s('type.option.option'(X294),X296)) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12980,f12545])).
fof(f12980,plain,(
  ( ! [X294,X296,X298,X295] : (s(bool,X298) = s(bool,happ(s(fun(X294,bool),X295),s(X294,sK4(X296,X294)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X298))) | s('type.option.option'(X294),'const.option.NONE_0') = s('type.option.option'(X294),X296)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12979,f126])).
fof(f12979,plain,(
  ( ! [X294,X296,X298,X295] : (~p(s(bool,'T_0')) | s(bool,X298) = s(bool,happ(s(fun(X294,bool),X295),s(X294,sK4(X296,X294)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X298))) | s('type.option.option'(X294),'const.option.NONE_0') = s('type.option.option'(X294),X296)) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f11890,f12545])).
fof(f11890,plain,(
  ( ! [X294,X296,X298,X295,X297] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X294),X296),s(bool,X297),s(fun(X294,bool),X295)))) | s(bool,X298) = s(bool,happ(s(fun(X294,bool),X295),s(X294,sK4(X296,X294)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X298))) | s('type.option.option'(X294),'const.option.NONE_0') = s('type.option.option'(X294),X296)) ) | $spl1),
  inference(superposition,[],[f2304,f1387])).
fof(f12943,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) | s(X59,X61) = s(X59,sK5(X61,X60,X59))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12942,f12545])).
fof(f12942,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))))) | s(X59,X61) = s(X59,sK5(X61,X60,X59))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f931,f12545])).
fof(f929,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) | s(bool,X54) = s(bool,happ(s(fun(X51,bool),X52),s(X51,X53))) | s(X51,X53) = s(X51,sK5(X53,X52,X51))) ) | $spl1),
  inference(superposition,[],[f162,f877])).
fof(f12790,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f3188,f12545])).
fof(f12750,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (s(X422,X427) = s(X422,sK5(X427,'const.option.option_CASE_3'(s('type.option.option'(X423),X424),s(fun(X422,bool),X425),s(fun(X423,fun(X422,bool)),X426)),X422)) | s('type.option.option'(X423),'const.option.NONE_0') = s('type.option.option'(X423),X424)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12676,f12014])).
fof(f12676,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (s(bool,'T_0') != s(bool,happ(s(fun(X422,bool),happ(s(fun(X423,fun(X422,bool)),X426),s(X423,sK4(X424,X423)))),s(X422,X427))) | s(X422,X427) = s(X422,sK5(X427,'const.option.option_CASE_3'(s('type.option.option'(X423),X424),s(fun(X422,bool),X425),s(fun(X423,fun(X422,bool)),X426)),X422)) | s('type.option.option'(X423),'const.option.NONE_0') = s('type.option.option'(X423),X424)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f12015])).
fof(f12747,plain,(
  ( ! [X277,X273,X275,X276,X274] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X277))) | s(bool,X277) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X273),X274),s(bool,X275),s(fun(X273,bool),X276))) | s('type.option.option'(X273),'const.option.NONE_0') = s('type.option.option'(X273),X274)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12746,f126])).
fof(f12746,plain,(
  ( ! [X277,X273,X275,X276,X274] : (~p(s(bool,'T_0')) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X277))) | s(bool,X277) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X273),X274),s(bool,X275),s(fun(X273,bool),X276))) | s('type.option.option'(X273),'const.option.NONE_0') = s('type.option.option'(X273),X274)) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12662,f12545])).
fof(f12662,plain,(
  ( ! [X277,X273,X275,X276,X274] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X277))) | ~p(s(bool,happ(s(fun(X273,bool),X276),s(X273,sK4(X274,X273))))) | s(bool,X277) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X273),X274),s(bool,X275),s(fun(X273,bool),X276))) | s('type.option.option'(X273),'const.option.NONE_0') = s('type.option.option'(X273),X274)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f11986])).
fof(f12744,plain,(
  ( ! [X163,X164] : (s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X164)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12743,f126])).
fof(f12743,plain,(
  ( ! [X163,X164] : (~p(s(bool,'T_0')) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X164)) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12742,f12545])).
fof(f12742,plain,(
  ( ! [X163,X166,X164] : (~p(s(bool,happ(s(fun(X163,bool),X166),s(X163,sK4(X164,X163))))) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X164)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12649,f12545])).
fof(f12649,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X163),X164),s(bool,X165),s(fun(X163,bool),X166))) | ~p(s(bool,happ(s(fun(X163,bool),X166),s(X163,sK4(X164,X163))))) | s('type.option.option'(X163),'const.option.NONE_0') = s('type.option.option'(X163),X164)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f11963])).
fof(f12638,plain,(
  ( ! [X109,X107,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X107))) | s(bool,X108) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,X107),s(fun(X109,bool),X110))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108)))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f11115])).
fof(f12725,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12586,f12545])).
fof(f12586,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11))) | s(bool,'F_0') = s(bool,X10)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f3315])).
fof(f12723,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12583,f12545])).
fof(f12583,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f3189])).
fof(f12721,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12582,f1798])).
fof(f12582,plain,(
  ( ! [X30] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f3187])).
fof(f12717,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl1 | $spl20)),
  inference(condensation,[],[f12716])).
fof(f12716,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33)))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12569,f12545])).
fof(f12569,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') != s(bool,X33)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2310])).
fof(f12700,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | ($spl1 | $spl20)),
  inference(duplicate_literal_removal,[],[f12568])).
fof(f12568,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2309])).
fof(f12567,plain,(
  ( ! [X12,X11] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11))) | p(s(bool,X12)) | s(bool,X11) != s(bool,X12)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2306])).
fof(f12566,plain,(
  ( ! [X10,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X9))) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2304])).
fof(f12564,plain,(
  ( ! [X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | p(s(bool,X1))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2301])).
fof(f12709,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12563,f163])).
fof(f12563,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2292])).
fof(f12562,plain,(
  ( ! [X28,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,sK4('const.option.SOME_1'(s(bool,X28)),bool))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2274])).
fof(f12706,plain,(
  ( ! [X47,X45,X46] : (s(X45,X46) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12561,f12545])).
fof(f12561,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,happ(s(fun(X45,bool),X47),s(X45,X46))) | s(X45,X46) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f2220])).
fof(f12703,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12547,f12545])).
fof(f12547,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f135])).
fof(f12545,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl20)),
  inference(condensation,[],[f12544])).
fof(f12544,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12527,f175])).
fof(f12527,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,X8),s(fun(sK0,bool),sK2))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f2538])).
fof(f12525,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X2)),bool))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f2156])).
fof(f12524,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X13,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f1304])).
fof(f12523,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f942])).
fof(f12522,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f797])).
fof(f12521,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f775])).
fof(f12520,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.NONE_0'),s(bool,X7),s(fun(X6,bool),X8)))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f756])).
fof(f12519,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f640])).
fof(f12518,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7)))) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f630])).
fof(f12514,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f11956,f1762])).
fof(f11956,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | $spl1),
  inference(superposition,[],[f560,f1387])).
fof(f11983,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X261,bool),X264),s(X261,sK4(X262,X261))))),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X261),X262),s(bool,X263),s(fun(X261,bool),X264))))) | s('type.option.option'(X261),'const.option.NONE_0') = s('type.option.option'(X261),X262)) ) | $spl1),
  inference(superposition,[],[f2116,f1387])).
fof(f11984,plain,(
  ( ! [X265,X267,X268,X266] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X265,bool),X268),s(X265,sK4(X266,X265)))))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X265),X266),s(bool,X267),s(fun(X265,bool),X268))) | s('type.option.option'(X265),'const.option.NONE_0') = s('type.option.option'(X265),X266)) ) | $spl1),
  inference(superposition,[],[f2292,f1387])).
fof(f11985,plain,(
  ( ! [X269,X271,X270,X272] : (p(s(bool,happ(s(fun(X269,bool),X272),s(X269,sK4(X270,X269))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X269),X270),s(bool,X271),s(fun(X269,bool),X272))))) | s('type.option.option'(X269),'const.option.NONE_0') = s('type.option.option'(X269),X270)) ) | $spl1),
  inference(superposition,[],[f2301,f1387])).
fof(f11986,plain,(
  ( ! [X277,X273,X275,X276,X274] : (~p(s(bool,happ(s(fun(X273,bool),X276),s(X273,sK4(X274,X273))))) | s(bool,X277) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X273),X274),s(bool,X275),s(fun(X273,bool),X276))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X277))) | s('type.option.option'(X273),'const.option.NONE_0') = s('type.option.option'(X273),X274)) ) | $spl1),
  inference(superposition,[],[f2304,f1387])).
fof(f11987,plain,(
  ( ! [X280,X282,X279,X281,X278] : (p(s(bool,happ(s(fun(X278,bool),X281),s(X278,sK4(X279,X278))))) | s(bool,X282) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X278),X279),s(bool,X280),s(fun(X278,bool),X281))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X282))) | s('type.option.option'(X278),'const.option.NONE_0') = s('type.option.option'(X278),X279)) ) | $spl1),
  inference(superposition,[],[f2306,f1387])).
fof(f11988,plain,(
  ( ! [X285,X287,X283,X284,X286] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X283,bool),X286),s(X283,sK4(X284,X283))))),bool)) | s(bool,X287) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X283),X284),s(bool,X285),s(fun(X283,bool),X286))) | s(bool,'T_0') = s(bool,X287) | s('type.option.option'(X283),'const.option.NONE_0') = s('type.option.option'(X283),X284)) ) | $spl1),
  inference(superposition,[],[f3366,f1387])).
fof(f12014,plain,(
  ( ! [X420,X416,X418,X421,X417,X419] : (s(bool,'T_0') = s(bool,happ(s(fun(X416,bool),happ(s(fun(X417,fun(X416,bool)),X420),s(X417,sK4(X418,X417)))),s(X416,X421))) | s(X416,X421) = s(X416,sK5(X421,'const.option.option_CASE_3'(s('type.option.option'(X417),X418),s(fun(X416,bool),X419),s(fun(X417,fun(X416,bool)),X420)),X416)) | s('type.option.option'(X417),'const.option.NONE_0') = s('type.option.option'(X417),X418)) ) | $spl1),
  inference(superposition,[],[f877,f1387])).
fof(f12015,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (s(bool,'F_0') != s(bool,happ(s(fun(X422,bool),happ(s(fun(X423,fun(X422,bool)),X426),s(X423,sK4(X424,X423)))),s(X422,X427))) | s(X422,X427) = s(X422,sK5(X427,'const.option.option_CASE_3'(s('type.option.option'(X423),X424),s(fun(X422,bool),X425),s(fun(X423,fun(X422,bool)),X426)),X422)) | s('type.option.option'(X423),'const.option.NONE_0') = s('type.option.option'(X423),X424)) ) | $spl1),
  inference(superposition,[],[f948,f1387])).
fof(f12025,plain,(
  ( ! [X486,X488,X485,X487] : (s(sK0,sK4(sK1,sK0)) != s(sK0,happ(s(fun(X485,sK0),X488),s(X485,sK4(X486,X485)))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,'const.option.option_CASE_3'(s('type.option.option'(X485),X486),s(sK0,X487),s(fun(X485,sK0),X488)))))) | s('type.option.option'(X485),'const.option.NONE_0') = s('type.option.option'(X485),X486)) ) | $spl20),
  inference(superposition,[],[f1816,f1387])).
fof(f11150,plain,(
  ( ! [X132,X130,X133,X131] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X131),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X130))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,X130)),'const.option.option_CASE_3'(s('type.option.option'(X132),'const.option.NONE_0'),s(fun('type.option.option'(bool),bool),X131),s(fun(X132,fun('type.option.option'(bool),bool)),X133)),'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X130)) ) | $spl1),
  inference(superposition,[],[f849,f2292])).
fof(f11149,plain,(
  ( ! [X127,X128,X126,X129] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X127),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X126)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),'const.option.option_CASE_3'(s('type.option.option'(X128),'const.option.NONE_0'),s(fun('type.option.option'(bool),bool),X127),s(fun(X128,fun('type.option.option'(bool),bool)),X129)),'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X126)) ) | $spl1),
  inference(superposition,[],[f849,f2292])).
fof(f11185,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,X113) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X113)),bool),'const.option.option_CASE_3'(s('type.option.option'(X115),'const.option.NONE_0'),s(fun(bool,bool),X114),s(fun(X115,fun(bool,bool)),X116)),bool)) | p(s(bool,happ(s(fun(bool,bool),X114),s(bool,'T_0')))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X113)))) ) | $spl1),
  inference(forward_demodulation,[],[f11146,f2072])).
fof(f11146,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,happ(s(fun(bool,bool),X114),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X113)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X113)),bool),'const.option.option_CASE_3'(s('type.option.option'(X115),'const.option.NONE_0'),s(fun(bool,bool),X114),s(fun(X115,fun(bool,bool)),X116)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X113)))) ) | $spl1),
  inference(superposition,[],[f849,f2116])).
fof(f11145,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0')))) | s(bool,happ(s(fun(X107,bool),X108),s(X107,X109))) = s(bool,sK5(happ(s(fun(X107,bool),X108),s(X107,X109)),'const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(fun(bool,bool),X110),s(fun(X111,fun(bool,bool)),X112)),bool)) | s(X107,X109) = s(X107,sK5(X109,X108,X107))) ) | $spl1),
  inference(superposition,[],[f849,f877])).
fof(f11182,plain,(
  ( ! [X99,X97,X98] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(fun(bool,bool),X97),s(fun(X98,fun(bool,bool)),X99)),bool)) | p(s(bool,happ(s(fun(bool,bool),X97),s(bool,'T_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11142,f1762])).
fof(f11142,plain,(
  ( ! [X99,X97,X98] : (p(s(bool,happ(s(fun(bool,bool),X97),s(bool,'T_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(fun(bool,bool),X97),s(fun(X98,fun(bool,bool)),X99)),bool))) ) | $spl20),
  inference(superposition,[],[f849,f1762])).
fof(f11177,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X79),'const.option.NONE_0'),s(fun(bool,bool),X78),s(fun(X79,fun(bool,bool)),X80)),bool)) | p(s(bool,happ(s(fun(bool,bool),X78),s(bool,X77)))) | s(bool,'T_0') = s(bool,X77)) ) | $spl1),
  inference(forward_demodulation,[],[f11138,f558])).
fof(f11138,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(bool,bool),X78),s(bool,X77)))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X79),'const.option.NONE_0'),s(fun(bool,bool),X78),s(fun(X79,fun(bool,bool)),X80)),bool)) | s(bool,'T_0') = s(bool,X77)) ) | $spl1),
  inference(superposition,[],[f849,f640])).
fof(f11176,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X75),'const.option.NONE_0'),s(fun(bool,bool),X74),s(fun(X75,fun(bool,bool)),X76)),bool)) | p(s(bool,happ(s(fun(bool,bool),X74),s(bool,'F_0'))))) ) | $spl1),
  inference(forward_demodulation,[],[f11137,f558])).
fof(f11137,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,happ(s(fun(bool,bool),X74),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X75),'const.option.NONE_0'),s(fun(bool,bool),X74),s(fun(X75,fun(bool,bool)),X76)),bool))) ) | $spl1),
  inference(superposition,[],[f849,f558])).
fof(f11062,plain,(
  ( ! [X123,X124] : (s('type.option.option'(bool),X124) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK4(X124,bool)) = s(bool,sK4('const.option.SOME_1'(s(bool,X123)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X124) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123)))) ) | $spl1),
  inference(superposition,[],[f956,f2116])).
fof(f11116,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,X115) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X117),'const.option.NONE_0'),s(bool,X115),s(fun(X117,bool),X118)),X116,bool)) | p(s(bool,happ(s(fun(bool,bool),X116),s(bool,'T_0')))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X115)))) ) | $spl1),
  inference(forward_demodulation,[],[f11059,f2072])).
fof(f11059,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,happ(s(fun(bool,bool),X116),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X115)),bool)) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X117),'const.option.NONE_0'),s(bool,sK4('const.option.SOME_1'(s(bool,X115)),bool)),s(fun(X117,bool),X118)),X116,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X115)))) ) | $spl1),
  inference(superposition,[],[f870,f2116])).
fof(f11058,plain,(
  ( ! [X111,X114,X112,X113] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X112))) | s(bool,sK4('const.option.SOME_1'(s(bool,X111)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X113),'const.option.NONE_0'),s(bool,X112),s(fun(X113,bool),X114))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X111)))) ) | $spl1),
  inference(superposition,[],[f409,f2116])).
fof(f11115,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,X108) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,X107),s(fun(X109,bool),X110))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X107)))) ) | $spl1),
  inference(forward_demodulation,[],[f11057,f2072])).
fof(f11057,plain,(
  ( ! [X109,X107,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108))) | s(bool,X108) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,sK4('const.option.SOME_1'(s(bool,X107)),bool)),s(fun(X109,bool),X110))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X107)))) ) | $spl1),
  inference(superposition,[],[f409,f2116])).
fof(f10891,plain,(
  ( ! [X127,X128,X126,X129] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X127),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X126))) = s('type.option.option'(bool),sK5('const.option.option_CASE_3'(s('type.option.option'(X128),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X126))),s(fun(X128,'type.option.option'(bool)),X129)),X127,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X126)) ) | $spl1),
  inference(superposition,[],[f870,f2292])).
fof(f10890,plain,(
  ( ! [X125,X123,X124,X122] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X123),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X122)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.option_CASE_3'(s('type.option.option'(X124),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X124,'type.option.option'(bool)),X125)),X123,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X122)) ) | $spl1),
  inference(superposition,[],[f870,f2292])).
fof(f10887,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0')))) | s(bool,happ(s(fun(X107,bool),X108),s(X107,X109))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X111),'const.option.NONE_0'),s(bool,happ(s(fun(X107,bool),X108),s(X107,X109))),s(fun(X111,bool),X112)),X110,bool)) | s(X107,X109) = s(X107,sK5(X109,X108,X107))) ) | $spl1),
  inference(superposition,[],[f870,f877])).
fof(f10630,plain,(
  ( ! [X123,X121,X122,X120] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X121))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X122),'const.option.NONE_0'),s('type.option.option'(bool),X121),s(fun(X122,'type.option.option'(bool)),X123))) | s(bool,'T_0') = s(bool,X120)) ) | $spl1),
  inference(superposition,[],[f409,f2292])).
fof(f10629,plain,(
  ( ! [X118,X116,X119,X117] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X117))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X116))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X118),'const.option.NONE_0'),s('type.option.option'(bool),X117),s(fun(X118,'type.option.option'(bool)),X119))) | s(bool,'T_0') = s(bool,X116)) ) | $spl1),
  inference(superposition,[],[f409,f2292])).
fof(f10626,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X104))) | s(bool,happ(s(fun(X101,bool),X102),s(X101,X103))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X105),'const.option.NONE_0'),s(bool,X104),s(fun(X105,bool),X106))) | s(X101,X103) = s(X101,sK5(X103,X102,X101))) ) | $spl1),
  inference(superposition,[],[f409,f877])).
fof(f10599,plain,(
  ( ! [X123,X121,X122,X120] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X121))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X121) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X122),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X120))),s(fun(X122,'type.option.option'(bool)),X123))) | s(bool,'T_0') = s(bool,X120)) ) | $spl1),
  inference(superposition,[],[f409,f2292])).
fof(f10598,plain,(
  ( ! [X118,X116,X119,X117] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X117))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X116))))) | s('type.option.option'(bool),X117) = s('type.option.option'(bool),'const.option.option_CASE_3'(s('type.option.option'(X118),'const.option.NONE_0'),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(fun(X118,'type.option.option'(bool)),X119))) | s(bool,'T_0') = s(bool,X116)) ) | $spl1),
  inference(superposition,[],[f409,f2292])).
fof(f10595,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X104))) | s(bool,X104) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X105),'const.option.NONE_0'),s(bool,happ(s(fun(X101,bool),X102),s(X101,X103))),s(fun(X105,bool),X106))) | s(X101,X103) = s(X101,sK5(X103,X102,X101))) ) | $spl1),
  inference(superposition,[],[f409,f877])).
fof(f10364,plain,(
  ( ! [X68,X69] : (s('type.option.option'('type.option.option'(bool)),X69) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X68))) = s('type.option.option'(bool),sK4(X69,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X69) | s(bool,'T_0') = s(bool,X68)) ) | $spl1),
  inference(superposition,[],[f956,f2292])).
fof(f10363,plain,(
  ( ! [X66,X67] : (s('type.option.option'('type.option.option'(bool)),X67) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X66))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK4(X67,'type.option.option'(bool))) | s('type.option.option'('type.option.option'(bool)),'const.option.NONE_0') = s('type.option.option'('type.option.option'(bool)),X67) | s(bool,'T_0') = s(bool,X66)) ) | $spl1),
  inference(superposition,[],[f956,f2292])).
fof(f10361,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.option.option'(bool),X62) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,sK4(X62,bool)) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X62) | s(X59,X61) = s(X59,sK5(X61,X60,X59))) ) | $spl1),
  inference(superposition,[],[f956,f877])).
fof(f10373,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X6) = s(bool,sK4('const.option.SOME_1'(s(bool,X5)),bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(subsumption_resolution,[],[f10335,f159])).
fof(f10335,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X6) = s(bool,sK4('const.option.SOME_1'(s(bool,X5)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f956,f2292])).
fof(f10372,plain,(
  ( ! [X4,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,X4) = s(bool,sK4('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl1),
  inference(subsumption_resolution,[],[f10334,f159])).
fof(f10334,plain,(
  ( ! [X4,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,X4) = s(bool,sK4('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl1),
  inference(superposition,[],[f956,f2292])).
fof(f2537,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X6),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl1),
  inference(superposition,[],[f558,f239])).
fof(f558,plain,(
  s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f128,f211,f167])).
fof(f1539,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f1536,f558])).
fof(f1536,plain,(
  ( ! [X2,X0,X1] : (s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0))))) ) | $spl1),
  inference(superposition,[],[f176,f640])).
fof(f3320,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,happ(s(fun(X55,bool),X56),s(X55,X57))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),'const.option.NONE_0'),s(bool,X53),s(fun(X52,bool),X54))) | s(X55,X57) = s(X55,sK5(X57,X56,X55)) | s(bool,'F_0') = s(bool,X53)) ) | $spl1),
  inference(subsumption_resolution,[],[f3271,f126])).
fof(f3271,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,happ(s(fun(X55,bool),X56),s(X55,X57))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),'const.option.NONE_0'),s(bool,X53),s(fun(X52,bool),X54))) | s(X55,X57) = s(X55,sK5(X57,X56,X55)) | s(bool,'F_0') = s(bool,X53)) ) | $spl1),
  inference(superposition,[],[f949,f303])).
fof(f1609,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X5),X6),s(bool,X7),s(fun(X5,bool),X8))) | s(fun(X5,bool),X8) = s(fun(X5,bool),sK5(X8,happ(s(fun(bool,fun(fun(X5,bool),bool)),happ(s(fun('type.option.option'(X5),fun(bool,fun(fun(X5,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X5),X6))),s(bool,X7)),fun(X5,bool)))) ) | $spl1),
  inference(superposition,[],[f178,f877])).
fof(f1626,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X90),X91),s(bool,X92),s(fun(X90,bool),X93)))) | s(fun(X90,bool),X93) != s(fun(X90,bool),X94) | s(fun(X90,bool),X94) = s(fun(X90,bool),sK5(X94,happ(s(fun(bool,fun(fun(X90,bool),bool)),happ(s(fun('type.option.option'(X90),fun(bool,fun(fun(X90,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X90),X91))),s(bool,X92)),fun(X90,bool)))) ) | $spl1),
  inference(superposition,[],[f1498,f178])).
fof(f1628,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X100),X101),s(bool,X102),s(fun(X100,bool),X103)))) | s(fun(X100,bool),sK5(X103,happ(s(fun(bool,fun(fun(X100,bool),bool)),happ(s(fun('type.option.option'(X100),fun(bool,fun(fun(X100,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X100),X101))),s(bool,X102)),fun(X100,bool))) = s(fun(X100,bool),sK5(sK5(X103,happ(s(fun(bool,fun(fun(X100,bool),bool)),happ(s(fun('type.option.option'(X100),fun(bool,fun(fun(X100,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X100),X101))),s(bool,X102)),fun(X100,bool)),happ(s(fun(bool,fun(fun(X100,bool),bool)),happ(s(fun('type.option.option'(X100),fun(bool,fun(fun(X100,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X100),X101))),s(bool,X102)),fun(X100,bool)))) ) | $spl1),
  inference(superposition,[],[f1321,f178])).
fof(f1629,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X104),X105),s(bool,X106),s(fun(X104,bool),X107))) | s(fun(X104,bool),X107) = s(fun(X104,bool),sK5(X107,happ(s(fun(bool,fun(fun(X104,bool),bool)),happ(s(fun('type.option.option'(X104),fun(bool,fun(fun(X104,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X104),X105))),s(bool,X106)),fun(X104,bool)))) ) | $spl1),
  inference(superposition,[],[f948,f178])).
fof(f1661,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(bool,X11),s(fun(X9,bool),X12))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X9,bool),bool),happ(s(fun(bool,fun(fun(X9,bool),bool)),happ(s(fun('type.option.option'(X9),fun(bool,fun(fun(X9,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(bool,X11))),s(fun(X9,bool),X12)))) ) | $spl1),
  inference(forward_demodulation,[],[f1610,f558])).
fof(f1610,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),X10),s(bool,X11),s(fun(X9,bool),X12))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X9,bool),bool),happ(s(fun(bool,fun(fun(X9,bool),bool)),happ(s(fun('type.option.option'(X9),fun(bool,fun(fun(X9,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X9),X10))),s(bool,X11))),s(fun(X9,bool),X12)))) ) | $spl1),
  inference(superposition,[],[f178,f640])).
fof(f3062,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X26),X27),s(bool,X28),s(fun(X26,bool),X29))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | $spl1),
  inference(superposition,[],[f2292,f178])).
fof(f3442,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53))) | s(bool,X54) = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | $spl1),
  inference(forward_demodulation,[],[f3408,f2072])).
fof(f3408,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53)))),bool)) | s(bool,X54) = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | $spl1),
  inference(superposition,[],[f3366,f178])).
fof(f894,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))) | s(X0,X4) = s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | $spl1),
  inference(superposition,[],[f877,f175])).
fof(f903,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))) = s(bool,sK5(happ(s(fun(X14,bool),X15),s(X14,X16)),X17,bool)) | s(X14,X16) = s(X14,sK5(X16,X15,X14))) ) | $spl1),
  inference(superposition,[],[f877,f877])).
fof(f931,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'T_0') != s(bool,X62) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))))) | s(X59,X61) = s(X59,sK5(X61,X60,X59))) ) | $spl1),
  inference(superposition,[],[f163,f877])).
fof(f933,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),X70),s(bool,'T_0')))) | s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) = s(bool,sK5(happ(s(fun(X67,bool),X68),s(X67,X69)),X70,bool)) | s(X67,X69) = s(X67,sK5(X69,X68,X67))) ) | $spl1),
  inference(superposition,[],[f164,f877])).
fof(f946,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,X20) = s(X18,sK5('const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(X18,X20),s(fun(X19,X18),X21)),X22,X18)) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X22),s(X18,X20)))) ) | $spl1),
  inference(forward_demodulation,[],[f904,f175])).
fof(f904,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X22),s(X18,X20))) | s(X18,'const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(X18,X20),s(fun(X19,X18),X21))) = s(X18,sK5('const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(X18,X20),s(fun(X19,X18),X21)),X22,X18))) ) | $spl1),
  inference(superposition,[],[f877,f175])).
fof(f949,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X30)) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),X28),s(X27,X29))) | s(X27,X29) = s(X27,sK5(X29,X28,X27))) ) | $spl1),
  inference(subsumption_resolution,[],[f922,f126])).
fof(f922,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),X28),s(X27,X29))) | s(X27,X29) = s(X27,sK5(X29,X28,X27))) ) | $spl1),
  inference(superposition,[],[f170,f877])).
fof(f950,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,X34)) | s(bool,X34) != s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | $spl1),
  inference(subsumption_resolution,[],[f923,f126])).
fof(f923,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | p(s(bool,X34)) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | $spl1),
  inference(superposition,[],[f204,f877])).
fof(f1284,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))) | s(X0,X4) = s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | $spl1),
  inference(superposition,[],[f948,f175])).
fof(f1292,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) = s(bool,sK5(happ(s(fun(X15,bool),X16),s(X15,X17)),X18,bool)) | s(X15,X17) = s(X15,sK5(X17,X16,X15))) ) | $spl1),
  inference(superposition,[],[f948,f877])).
fof(f1301,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0)) | s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),X4),s(X0,X2)))) ) | $spl1),
  inference(forward_demodulation,[],[f1285,f175])).
fof(f1285,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),X4),s(X0,X2))) | s(X0,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3))) = s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0))) ) | $spl1),
  inference(superposition,[],[f948,f175])).
fof(f1321,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,sK5(X2,X1,X0)) = s(X0,sK5(sK5(X2,X1,X0),X1,X0))) ) | $spl1),
  inference(subsumption_resolution,[],[f1316,f126])).
fof(f1316,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,sK5(X2,X1,X0)) = s(X0,sK5(sK5(X2,X1,X0),X1,X0))) ) | $spl1),
  inference(superposition,[],[f165,f877])).
fof(f1353,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0)) = s(X0,sK5(sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0),'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | $spl1),
  inference(superposition,[],[f1321,f175])).
fof(f1354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X2)))) | s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0)) = s(X0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0),X4,X0))) ) | $spl1),
  inference(superposition,[],[f1321,f175])).
fof(f1434,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X5,bool),X8),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X6))),s(bool,X7),s(fun(X5,bool),X8)))) ) | $spl1),
  inference(forward_demodulation,[],[f1402,f558])).
fof(f1402,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,happ(s(fun(X5,bool),X8),s(X5,X6))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.SOME_1'(s(X5,X6))),s(bool,X7),s(fun(X5,bool),X8)))) ) | $spl1),
  inference(superposition,[],[f177,f640])).
fof(f1479,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,X36) != s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))) | p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X36)))) | s(X32,X34) = s(X32,sK5(X34,X33,X32))) ) | $spl1),
  inference(superposition,[],[f166,f877])).
fof(f1498,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X6)))) | s(X3,X5) != s(X3,X6) | s(X3,X5) = s(X3,sK5(X5,X4,X3))) ) | $spl1),
  inference(subsumption_resolution,[],[f1484,f126])).
fof(f1484,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X5) != s(X3,X6) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X6)))) | s(X3,X5) = s(X3,sK5(X5,X4,X3))) ) | $spl1),
  inference(superposition,[],[f166,f877])).
fof(f1531,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(X29,X31) = s(X29,sK5(X31,X30,X29)) | s(X26,X28) = s(X26,sK5(X28,X27,X26))) ) | $spl1),
  inference(subsumption_resolution,[],[f1523,f126])).
fof(f1523,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'T_0')) | s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(X29,X31) = s(X29,sK5(X31,X30,X29)) | s(X26,X28) = s(X26,sK5(X28,X27,X26))) ) | $spl1),
  inference(superposition,[],[f949,f877])).
fof(f1559,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | s(X0,X4) != s(X0,X5) | s(X0,X5) = s(X0,sK5(X5,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | $spl1),
  inference(superposition,[],[f1498,f175])).
fof(f1570,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,X36) != s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))) | s(bool,X36) = s(bool,sK5(X36,X35,bool)) | s(X32,X34) = s(X32,sK5(X34,X33,X32))) ) | $spl1),
  inference(superposition,[],[f1498,f877])).
fof(f2171,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X49)))) | s(X48,sK5(sK4('const.option.SOME_1'(s(X48,X49)),X48),X50,X48)) = s(X48,sK5(sK5(sK4('const.option.SOME_1'(s(X48,X49)),X48),X50,X48),X50,X48))) ) | $spl1),
  inference(superposition,[],[f1321,f2072])).
fof(f2195,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') = s(bool,happ(s(fun(X96,bool),X97),s(X96,X98))) | s(X96,X98) = s(X96,sK5(X98,sK4('const.option.SOME_1'(s(fun(X96,bool),X97)),fun(X96,bool)),X96))) ) | $spl1),
  inference(superposition,[],[f877,f2072])).
fof(f2196,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') != s(bool,happ(s(fun(X99,bool),X100),s(X99,X101))) | s(X99,X101) = s(X99,sK5(X101,sK4('const.option.SOME_1'(s(fun(X99,bool),X100)),fun(X99,bool)),X99))) ) | $spl1),
  inference(superposition,[],[f948,f2072])).
fof(f2197,plain,(
  ( ! [X103,X102,X104] : (p(s(bool,happ(s(fun(X102,bool),X103),s(X102,X104)))) | s(X102,sK5(X104,sK4('const.option.SOME_1'(s(fun(X102,bool),X103)),fun(X102,bool)),X102)) = s(X102,sK5(sK5(X104,sK4('const.option.SOME_1'(s(fun(X102,bool),X103)),fun(X102,bool)),X102),sK4('const.option.SOME_1'(s(fun(X102,bool),X103)),fun(X102,bool)),X102))) ) | $spl1),
  inference(superposition,[],[f1321,f2072])).
fof(f2198,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X105,bool),X106),s(X105,X107)))) | s(X105,X107) != s(X105,X108) | s(X105,X108) = s(X105,sK5(X108,sK4('const.option.SOME_1'(s(fun(X105,bool),X106)),fun(X105,bool)),X105))) ) | $spl1),
  inference(superposition,[],[f1498,f2072])).
fof(f2307,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,sK4('const.option.SOME_1'(s(bool,X16)),bool)) = s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))) | s(X17,X19) = s(X17,sK5(X19,X18,X17)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2268,f126])).
fof(f2268,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'T_0')) | s(bool,sK4('const.option.SOME_1'(s(bool,X16)),bool)) = s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))) | s(X17,X19) = s(X17,sK5(X19,X18,X17)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16)))) ) | $spl1),
  inference(superposition,[],[f949,f2116])).
fof(f3061,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X22,bool),X25),s(X22,X23))))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.SOME_1'(s(X22,X23))),s(bool,X24),s(fun(X22,bool),X25)))) ) | $spl1),
  inference(superposition,[],[f2292,f177])).
fof(f3441,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(X45,bool),X48),s(X45,X46))) | s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48))) | s(bool,'T_0') = s(bool,X49)) ) | $spl1),
  inference(forward_demodulation,[],[f3407,f2072])).
fof(f3407,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X45,bool),X48),s(X45,X46)))),bool)) | s(bool,X49) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48))) | s(bool,'T_0') = s(bool,X49)) ) | $spl1),
  inference(superposition,[],[f3366,f177])).
fof(f1624,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (p(s(bool,happ(s(fun(X77,bool),X82),s(X77,'const.option.option_CASE_3'(s('type.option.option'(X78),X79),s(X77,X80),s(fun(X78,X77),X81)))))) | s(X77,sK5(happ(s(fun(fun(X78,X77),X77),happ(s(fun(X77,fun(fun(X78,X77),X77)),happ(s(fun('type.option.option'(X78),fun(X77,fun(fun(X78,X77),X77))),'const.option.option_CASE_0'),s('type.option.option'(X78),X79))),s(X77,X80))),s(fun(X78,X77),X81)),X82,X77)) = s(X77,sK5(sK5(happ(s(fun(fun(X78,X77),X77),happ(s(fun(X77,fun(fun(X78,X77),X77)),happ(s(fun('type.option.option'(X78),fun(X77,fun(fun(X78,X77),X77))),'const.option.option_CASE_0'),s('type.option.option'(X78),X79))),s(X77,X80))),s(fun(X78,X77),X81)),X82,X77),X82,X77))) ) | $spl1),
  inference(superposition,[],[f1321,f178])).
fof(f1667,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(X71,'const.option.option_CASE_3'(s('type.option.option'(X72),X73),s(X71,X74),s(fun(X72,X71),X75))) = s(X71,sK5(happ(s(fun(fun(X72,X71),X71),happ(s(fun(X71,fun(fun(X72,X71),X71)),happ(s(fun('type.option.option'(X72),fun(X71,fun(fun(X72,X71),X71))),'const.option.option_CASE_0'),s('type.option.option'(X72),X73))),s(X71,X74))),s(fun(X72,X71),X75)),X76,X71)) | s(bool,'F_0') != s(bool,happ(s(fun(X71,bool),X76),s(X71,'const.option.option_CASE_3'(s('type.option.option'(X72),X73),s(X71,X74),s(fun(X72,X71),X75)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1623,f178])).
fof(f1623,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') != s(bool,happ(s(fun(X71,bool),X76),s(X71,'const.option.option_CASE_3'(s('type.option.option'(X72),X73),s(X71,X74),s(fun(X72,X71),X75))))) | s(X71,happ(s(fun(fun(X72,X71),X71),happ(s(fun(X71,fun(fun(X72,X71),X71)),happ(s(fun('type.option.option'(X72),fun(X71,fun(fun(X72,X71),X71))),'const.option.option_CASE_0'),s('type.option.option'(X72),X73))),s(X71,X74))),s(fun(X72,X71),X75))) = s(X71,sK5(happ(s(fun(fun(X72,X71),X71),happ(s(fun(X71,fun(fun(X72,X71),X71)),happ(s(fun('type.option.option'(X72),fun(X71,fun(fun(X72,X71),X71))),'const.option.option_CASE_0'),s('type.option.option'(X72),X73))),s(X71,X74))),s(fun(X72,X71),X75)),X76,X71))) ) | $spl1),
  inference(superposition,[],[f948,f178])).
fof(f1666,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(X65,'const.option.option_CASE_3'(s('type.option.option'(X66),X67),s(X65,X68),s(fun(X66,X65),X69))) = s(X65,sK5(happ(s(fun(fun(X66,X65),X65),happ(s(fun(X65,fun(fun(X66,X65),X65)),happ(s(fun('type.option.option'(X66),fun(X65,fun(fun(X66,X65),X65))),'const.option.option_CASE_0'),s('type.option.option'(X66),X67))),s(X65,X68))),s(fun(X66,X65),X69)),X70,X65)) | s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X70),s(X65,'const.option.option_CASE_3'(s('type.option.option'(X66),X67),s(X65,X68),s(fun(X66,X65),X69)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1622,f178])).
fof(f1622,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(X65,bool),X70),s(X65,'const.option.option_CASE_3'(s('type.option.option'(X66),X67),s(X65,X68),s(fun(X66,X65),X69))))) | s(X65,happ(s(fun(fun(X66,X65),X65),happ(s(fun(X65,fun(fun(X66,X65),X65)),happ(s(fun('type.option.option'(X66),fun(X65,fun(fun(X66,X65),X65))),'const.option.option_CASE_0'),s('type.option.option'(X66),X67))),s(X65,X68))),s(fun(X66,X65),X69))) = s(X65,sK5(happ(s(fun(fun(X66,X65),X65),happ(s(fun(X65,fun(fun(X66,X65),X65)),happ(s(fun('type.option.option'(X66),fun(X65,fun(fun(X66,X65),X65))),'const.option.option_CASE_0'),s('type.option.option'(X66),X67))),s(X65,X68))),s(fun(X66,X65),X69)),X70,X65))) ) | $spl1),
  inference(superposition,[],[f877,f178])).
fof(f1413,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,happ(s(fun(X54,bool),X59),s(X54,happ(s(fun(X55,X54),X58),s(X55,X56)))))) | s(X54,sK5('const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.SOME_1'(s(X55,X56))),s(X54,X57),s(fun(X55,X54),X58)),X59,X54)) = s(X54,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X55),'const.option.SOME_1'(s(X55,X56))),s(X54,X57),s(fun(X55,X54),X58)),X59,X54),X59,X54))) ) | $spl1),
  inference(superposition,[],[f1321,f177])).
fof(f1439,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X48,happ(s(fun(X49,X48),X52),s(X49,X50))) = s(X48,sK5('const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.SOME_1'(s(X49,X50))),s(X48,X51),s(fun(X49,X48),X52)),X53,X48)) | s(bool,'F_0') != s(bool,happ(s(fun(X48,bool),X53),s(X48,happ(s(fun(X49,X48),X52),s(X49,X50)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1412,f177])).
fof(f1412,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(X48,bool),X53),s(X48,happ(s(fun(X49,X48),X52),s(X49,X50))))) | s(X48,'const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.SOME_1'(s(X49,X50))),s(X48,X51),s(fun(X49,X48),X52))) = s(X48,sK5('const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.SOME_1'(s(X49,X50))),s(X48,X51),s(fun(X49,X48),X52)),X53,X48))) ) | $spl1),
  inference(superposition,[],[f948,f177])).
fof(f1438,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X42,happ(s(fun(X43,X42),X46),s(X43,X44))) = s(X42,sK5('const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(X42,X45),s(fun(X43,X42),X46)),X47,X42)) | s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X44)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1411,f177])).
fof(f1411,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X44))))) | s(X42,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(X42,X45),s(fun(X43,X42),X46))) = s(X42,sK5('const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(X42,X45),s(fun(X43,X42),X46)),X47,X42))) ) | $spl1),
  inference(superposition,[],[f877,f177])).
fof(f1560,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X8))),s(X6,X11)))) | s(X6,X11) != s(X6,X12) | s(X6,X12) = s(X6,sK5(X12,'const.option.option_CASE_3'(s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X8))),s(fun(X6,bool),X9),s(fun(X7,fun(X6,bool)),X10)),X6))) ) | $spl1),
  inference(superposition,[],[f1498,f177])).
fof(f1431,plain,(
  ( ! [X146,X144,X142,X145,X143,X141] : (s(bool,'T_0') = s(bool,happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),X145),s(X142,X143))),s(X141,X146))) | s(X141,X146) = s(X141,sK5(X146,'const.option.option_CASE_3'(s('type.option.option'(X142),'const.option.SOME_1'(s(X142,X143))),s(fun(X141,bool),X144),s(fun(X142,fun(X141,bool)),X145)),X141))) ) | $spl1),
  inference(superposition,[],[f877,f177])).
fof(f1429,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (s(bool,'F_0') != s(bool,happ(s(fun(X129,bool),happ(s(fun(X130,fun(X129,bool)),X133),s(X130,X131))),s(X129,X134))) | s(X129,X134) = s(X129,sK5(X134,'const.option.option_CASE_3'(s('type.option.option'(X130),'const.option.SOME_1'(s(X130,X131))),s(fun(X129,bool),X132),s(fun(X130,fun(X129,bool)),X133)),X129))) ) | $spl1),
  inference(superposition,[],[f948,f177])).
fof(f1427,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),X121),s(X118,X119))),s(X117,X122)))) | s(X117,sK5(X122,'const.option.option_CASE_3'(s('type.option.option'(X118),'const.option.SOME_1'(s(X118,X119))),s(fun(X117,bool),X120),s(fun(X118,fun(X117,bool)),X121)),X117)) = s(X117,sK5(sK5(X122,'const.option.option_CASE_3'(s('type.option.option'(X118),'const.option.SOME_1'(s(X118,X119))),s(fun(X117,bool),X120),s(fun(X118,fun(X117,bool)),X121)),X117),'const.option.option_CASE_3'(s('type.option.option'(X118),'const.option.SOME_1'(s(X118,X119))),s(fun(X117,bool),X120),s(fun(X118,fun(X117,bool)),X121)),X117))) ) | $spl1),
  inference(superposition,[],[f1321,f177])).
fof(f1656,plain,(
  ( ! [X243,X241,X239,X237,X242,X240,X238] : (p(s(bool,happ(s(fun(X237,bool),'const.option.option_CASE_3'(s('type.option.option'(X238),X239),s(fun(X237,bool),X240),s(fun(X238,fun(X237,bool)),X241))),s(X237,X242)))) | s(X237,X242) != s(X237,X243) | s(X237,X243) = s(X237,sK5(X243,happ(s(fun(fun(X238,fun(X237,bool)),fun(X237,bool)),happ(s(fun(fun(X237,bool),fun(fun(X238,fun(X237,bool)),fun(X237,bool))),happ(s(fun('type.option.option'(X238),fun(fun(X237,bool),fun(fun(X238,fun(X237,bool)),fun(X237,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X238),X239))),s(fun(X237,bool),X240))),s(fun(X238,fun(X237,bool)),X241)),X237))) ) | $spl1),
  inference(superposition,[],[f1498,f178])).
fof(f1655,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),'const.option.option_CASE_3'(s('type.option.option'(X232),X233),s(fun(X231,bool),X234),s(fun(X232,fun(X231,bool)),X235))),s(X231,X236)))) | s(X231,sK5(X236,happ(s(fun(fun(X232,fun(X231,bool)),fun(X231,bool)),happ(s(fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool))),happ(s(fun('type.option.option'(X232),fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X232),X233))),s(fun(X231,bool),X234))),s(fun(X232,fun(X231,bool)),X235)),X231)) = s(X231,sK5(sK5(X236,happ(s(fun(fun(X232,fun(X231,bool)),fun(X231,bool)),happ(s(fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool))),happ(s(fun('type.option.option'(X232),fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X232),X233))),s(fun(X231,bool),X234))),s(fun(X232,fun(X231,bool)),X235)),X231),happ(s(fun(fun(X232,fun(X231,bool)),fun(X231,bool)),happ(s(fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool))),happ(s(fun('type.option.option'(X232),fun(fun(X231,bool),fun(fun(X232,fun(X231,bool)),fun(X231,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X232),X233))),s(fun(X231,bool),X234))),s(fun(X232,fun(X231,bool)),X235)),X231))) ) | $spl1),
  inference(superposition,[],[f1321,f178])).
fof(f1654,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,'F_0') != s(bool,happ(s(fun(X225,bool),'const.option.option_CASE_3'(s('type.option.option'(X226),X227),s(fun(X225,bool),X228),s(fun(X226,fun(X225,bool)),X229))),s(X225,X230))) | s(X225,X230) = s(X225,sK5(X230,happ(s(fun(fun(X226,fun(X225,bool)),fun(X225,bool)),happ(s(fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool))),happ(s(fun('type.option.option'(X226),fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X226),X227))),s(fun(X225,bool),X228))),s(fun(X226,fun(X225,bool)),X229)),X225))) ) | $spl1),
  inference(superposition,[],[f948,f178])).
fof(f1653,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (s(bool,'T_0') = s(bool,happ(s(fun(X219,bool),'const.option.option_CASE_3'(s('type.option.option'(X220),X221),s(fun(X219,bool),X222),s(fun(X220,fun(X219,bool)),X223))),s(X219,X224))) | s(X219,X224) = s(X219,sK5(X224,happ(s(fun(fun(X220,fun(X219,bool)),fun(X219,bool)),happ(s(fun(fun(X219,bool),fun(fun(X220,fun(X219,bool)),fun(X219,bool))),happ(s(fun('type.option.option'(X220),fun(fun(X219,bool),fun(fun(X220,fun(X219,bool)),fun(X219,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X220),X221))),s(fun(X219,bool),X222))),s(fun(X220,fun(X219,bool)),X223)),X219))) ) | $spl1),
  inference(superposition,[],[f877,f178])).
fof(f866,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X9,bool)) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X8)))) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(forward_demodulation,[],[f854,f558])).
fof(f854,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X8)))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X9,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f164,f640])).
fof(f1320,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f1313,f558])).
fof(f1313,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) ) | $spl1),
  inference(superposition,[],[f165,f640])).
fof(f1357,plain,(
  ( ! [X10,X9] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,X9)))) | s(bool,sK5('F_0',X10,bool)) = s(bool,sK5(sK5('F_0',X10,bool),X10,bool)) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f1321,f129])).
fof(f1359,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X13,bool)) = s(bool,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X13,bool),X13,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(superposition,[],[f1321,f640])).
fof(f1463,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK5(X2,X1,bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f864,f129])).
fof(f1491,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0')))) | s(bool,X13) != s(bool,X15) | p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl1),
  inference(forward_demodulation,[],[f1472,f558])).
fof(f1472,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))))) | s(bool,X13) != s(bool,X15) | p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl1),
  inference(superposition,[],[f166,f640])).
fof(f1565,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,'F_0') != s(bool,X21) | s(bool,X21) = s(bool,sK5(X21,X20,bool)) | s(bool,'T_0') = s(bool,X19)) ) | $spl1),
  inference(superposition,[],[f1498,f129])).
fof(f2313,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,X40) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38)))) ) | $spl1),
  inference(forward_demodulation,[],[f2279,f2072])).
fof(f2279,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | s(bool,X40) != s(bool,sK4('const.option.SOME_1'(s(bool,X38)),bool)) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X38)))) ) | $spl1),
  inference(superposition,[],[f166,f2116])).
fof(f2456,plain,(
  ( ! [X163,X161,X162] : (p(s(bool,happ(s(fun(bool,bool),X163),s(bool,X162)))) | s(bool,sK5(X161,X163,bool)) = s(bool,sK5(sK5(X161,X163,bool),X163,bool)) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,X161)) ) | $spl1),
  inference(superposition,[],[f1321,f239])).
fof(f2457,plain,(
  ( ! [X167,X165,X166,X164] : (p(s(bool,happ(s(fun(bool,bool),X166),s(bool,X165)))) | s(bool,X164) != s(bool,X167) | s(bool,X167) = s(bool,sK5(X167,X166,bool)) | s(bool,'T_0') = s(bool,X165) | s(bool,'T_0') = s(bool,X164)) ) | $spl1),
  inference(superposition,[],[f1498,f239])).
fof(f2569,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,happ(s(fun(bool,bool),X104),s(bool,X103)))) | s(bool,sK5('T_0',X104,bool)) = s(bool,sK5(sK5('T_0',X104,bool),X104,bool)) | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) ) | $spl1),
  inference(superposition,[],[f1321,f239])).
fof(f2570,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,happ(s(fun(bool,bool),X107),s(bool,X106)))) | s(bool,'T_0') != s(bool,X108) | s(bool,X108) = s(bool,sK5(X108,X107,bool)) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) ) | $spl1),
  inference(superposition,[],[f1498,f239])).
fof(f1362,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,sK5(happ(s(fun(X17,bool),X18),s(X17,X19)),X20,bool)) = s(bool,sK5(sK5(happ(s(fun(X17,bool),X18),s(X17,X19)),X20,bool),X20,bool)) | s(X17,X19) = s(X17,sK5(X19,X18,X17))) ) | $spl1),
  inference(superposition,[],[f1321,f877])).
fof(f1810,plain,(
  ( ! [X33] : (p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X33,bool)) = s(bool,sK5(sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X33,bool),X33,bool))) ) | ($spl1 | $spl20)),
  inference(superposition,[],[f1321,f1762])).
fof(f1829,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X17,bool)) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f1802,f1762])).
fof(f1802,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X17,bool))) ) | $spl20),
  inference(superposition,[],[f164,f1762])).
fof(f1831,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X31,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0')))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f1808,f1762])).
fof(f1808,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X31,bool))) ) | ($spl1 | $spl20)),
  inference(superposition,[],[f877,f1762])).
fof(f1832,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X32,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) ) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f1809,f1762])).
fof(f1809,plain,(
  ( ! [X32] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X32,bool))) ) | ($spl1 | $spl20)),
  inference(superposition,[],[f948,f1762])).
fof(f2286,plain,(
  ( ! [X61,X60] : (p(s(bool,happ(s(fun(bool,bool),X61),s(bool,'T_0')))) | s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X60)),bool),X61,bool)) = s(bool,sK5(sK5(sK4('const.option.SOME_1'(s(bool,X60)),bool),X61,bool),X61,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X60)))) ) | $spl1),
  inference(superposition,[],[f1321,f2116])).
fof(f2312,plain,(
  ( ! [X37,X36] : (s(bool,X36) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X36)),bool),X37,bool)) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0')))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | $spl1),
  inference(forward_demodulation,[],[f2278,f2072])).
fof(f2278,plain,(
  ( ! [X37,X36] : (p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X36)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X36)),bool),X37,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | $spl1),
  inference(superposition,[],[f164,f2116])).
fof(f2314,plain,(
  ( ! [X57,X56] : (s(bool,X56) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X56)),bool),X57,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X57),s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56)))) ) | $spl1),
  inference(forward_demodulation,[],[f2284,f2072])).
fof(f2284,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X57),s(bool,'T_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X56)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X56)),bool),X57,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X56)))) ) | $spl1),
  inference(superposition,[],[f877,f2116])).
fof(f2315,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X58)),bool),X59,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) | $spl1),
  inference(forward_demodulation,[],[f2285,f2072])).
fof(f2285,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X58)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X58)),bool),X59,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) | $spl1),
  inference(superposition,[],[f948,f2116])).
fof(f2316,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,X64) | p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0')))) | s(bool,X64) = s(bool,sK5(X64,X63,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62)))) ) | $spl1),
  inference(forward_demodulation,[],[f2287,f2072])).
fof(f2287,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0')))) | s(bool,X64) != s(bool,sK4('const.option.SOME_1'(s(bool,X62)),bool)) | s(bool,X64) = s(bool,sK5(X64,X63,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X62)))) ) | $spl1),
  inference(superposition,[],[f1498,f2116])).
fof(f2624,plain,(
  ( ! [X154,X155,X153] : (p(s(bool,happ(s(fun(bool,bool),X154),s(bool,'T_0')))) | s(bool,sK5(X153,X154,bool)) = s(bool,sK5(sK5(X153,X154,bool),X154,bool)) | s(bool,X153) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) ) | $spl1),
  inference(superposition,[],[f1321,f239])).
fof(f2625,plain,(
  ( ! [X158,X156,X159,X157] : (p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0')))) | s(bool,X156) != s(bool,X158) | s(bool,X158) = s(bool,sK5(X158,X157,bool)) | s(bool,X156) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) ) | $spl1),
  inference(superposition,[],[f1498,f239])).
fof(f3294,plain,(
  ( ! [X152,X151,X149,X150] : (p(s(bool,happ(s(fun(bool,bool),X152),s(bool,'T_0')))) | s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X149),'const.option.NONE_0'),s(bool,X150),s(fun(X149,bool),X151)),X152,bool)) = s(bool,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X149),'const.option.NONE_0'),s(bool,X150),s(fun(X149,bool),X151)),X152,bool),X152,bool)) | s(bool,'F_0') = s(bool,X150)) ) | $spl1),
  inference(superposition,[],[f1321,f303])).
fof(f3327,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X142) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X141),'const.option.NONE_0'),s(bool,X142),s(fun(X141,bool),X143)),X144,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X144),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X142)) ) | $spl1),
  inference(forward_demodulation,[],[f3292,f175])).
fof(f3292,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X144),s(bool,'T_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X141),'const.option.NONE_0'),s(bool,X142),s(fun(X141,bool),X143))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X141),'const.option.NONE_0'),s(bool,X142),s(fun(X141,bool),X143)),X144,bool)) | s(bool,'F_0') = s(bool,X142)) ) | $spl1),
  inference(superposition,[],[f877,f303])).
fof(f3328,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X146) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X145),'const.option.NONE_0'),s(bool,X146),s(fun(X145,bool),X147)),X148,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X148),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X146)) ) | $spl1),
  inference(forward_demodulation,[],[f3293,f175])).
fof(f3293,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X148),s(bool,'T_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X145),'const.option.NONE_0'),s(bool,X146),s(fun(X145,bool),X147))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X145),'const.option.NONE_0'),s(bool,X146),s(fun(X145,bool),X147)),X148,bool)) | s(bool,'F_0') = s(bool,X146)) ) | $spl1),
  inference(superposition,[],[f948,f303])).
fof(f3329,plain,(
  ( ! [X156,X154,X157] : (s(bool,X154) != s(bool,X157) | p(s(bool,happ(s(fun(bool,bool),X156),s(bool,'T_0')))) | s(bool,X157) = s(bool,sK5(X157,X156,bool)) | s(bool,'F_0') = s(bool,X154)) ) | $spl1),
  inference(forward_demodulation,[],[f3295,f175])).
fof(f3295,plain,(
  ( ! [X156,X154,X157,X155,X153] : (p(s(bool,happ(s(fun(bool,bool),X156),s(bool,'T_0')))) | s(bool,X157) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X153),'const.option.NONE_0'),s(bool,X154),s(fun(X153,bool),X155))) | s(bool,X157) = s(bool,sK5(X157,X156,bool)) | s(bool,'F_0') = s(bool,X154)) ) | $spl1),
  inference(superposition,[],[f1498,f303])).
fof(f864,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,X0) = s(bool,sK5(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f850,f558])).
fof(f850,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))))) | s(bool,X0) = s(bool,sK5(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f164,f640])).
fof(f867,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X10,bool)) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'F_0'))))) ) | $spl1),
  inference(forward_demodulation,[],[f855,f558])).
fof(f855,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X10,bool))) ) | $spl1),
  inference(superposition,[],[f164,f558])).
fof(f943,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X10,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,'F_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f900,f558])).
fof(f900,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X10,bool))) ) | $spl1),
  inference(superposition,[],[f877,f558])).
fof(f1303,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X11,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0')))) ) | $spl1),
  inference(forward_demodulation,[],[f1289,f558])).
fof(f1289,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X11,bool))) ) | $spl1),
  inference(superposition,[],[f948,f558])).
fof(f1358,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0')))) | s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X11,bool)) = s(bool,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X11,bool),X11,bool))) ) | $spl1),
  inference(superposition,[],[f1321,f558])).
fof(f1370,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | s(bool,sK5(X5,X6,bool)) = s(bool,sK5(sK5(X5,X6,bool),X6,bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(forward_demodulation,[],[f1355,f558])).
fof(f1355,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))))) | s(bool,sK5(X5,X6,bool)) = s(bool,sK5(sK5(X5,X6,bool),X6,bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f1321,f640])).
fof(f1461,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,X3) = s(bool,sK5(X3,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl1),
  inference(superposition,[],[f864,f175])).
fof(f1580,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0')))) | s(bool,X13) != s(bool,X15) | s(bool,X15) = s(bool,sK5(X15,X14,bool)) | s(bool,'T_0') = s(bool,X13)) ) | $spl1),
  inference(forward_demodulation,[],[f1563,f558])).
fof(f1563,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))))) | s(bool,X13) != s(bool,X15) | s(bool,X15) = s(bool,sK5(X15,X14,bool)) | s(bool,'T_0') = s(bool,X13)) ) | $spl1),
  inference(superposition,[],[f1498,f640])).
fof(f2191,plain,(
  ( ! [X85,X84] : (p(s(bool,happ(s(fun(bool,bool),X84),s(bool,'F_0')))) | s(bool,X85) = s(bool,sK5(X85,sK4('const.option.SOME_1'(s(fun(bool,bool),X84)),fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X85)) ) | $spl1),
  inference(superposition,[],[f864,f2072])).
fof(f1462,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X4,fun(bool,bool)),X7),s(X4,X5))),s(bool,'F_0')))) | s(bool,X8) = s(bool,sK5(X8,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.SOME_1'(s(X4,X5))),s(fun(bool,bool),X6),s(fun(X4,fun(bool,bool)),X7)),bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f864,f177])).
fof(f1649,plain,(
  ( ! [X198,X196,X199,X197,X195] : (p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X195),X196),s(fun(bool,bool),X197),s(fun(X195,fun(bool,bool)),X198))),s(bool,'F_0')))) | s(bool,X199) = s(bool,sK5(X199,happ(s(fun(fun(X195,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X195,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X195),fun(fun(bool,bool),fun(fun(X195,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X195),X196))),s(fun(bool,bool),X197))),s(fun(X195,fun(bool,bool)),X198)),bool)) | s(bool,'T_0') = s(bool,X199)) ) | $spl1),
  inference(superposition,[],[f864,f178])).
fof(f3108,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X51),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),X52) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X50))) | s('type.option.option'(bool),X52) = s('type.option.option'(bool),sK5(X52,X51,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X50)) ) | $spl1),
  inference(superposition,[],[f1498,f2292])).
fof(f3107,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X49),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,X48)),X49,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,X48)),X49,'type.option.option'(bool)),X49,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X48)) ) | $spl1),
  inference(superposition,[],[f1321,f2292])).
fof(f3100,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X27),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),X28) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X27),s('type.option.option'(bool),X28)))) | s(bool,'T_0') = s(bool,X26)) ) | $spl1),
  inference(superposition,[],[f166,f2292])).
fof(f3099,plain,(
  ( ! [X24,X25] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X25),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X24))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,X24)),X25,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X24)) ) | $spl1),
  inference(superposition,[],[f164,f2292])).
fof(f3085,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X50),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49)))))) | s('type.option.option'(bool),X51) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X51) = s('type.option.option'(bool),sK5(X51,X50,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X49)) ) | $spl1),
  inference(superposition,[],[f1498,f2292])).
fof(f3084,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X48),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47)))))) | s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool)),X48,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X47)) ) | $spl1),
  inference(superposition,[],[f1321,f2292])).
fof(f3077,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun('type.option.option'(bool),bool),X26),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X25)))))) | s('type.option.option'(bool),X27) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | p(s(bool,happ(s(fun('type.option.option'(bool),bool),X26),s('type.option.option'(bool),X27)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl1),
  inference(superposition,[],[f166,f2292])).
fof(f3076,plain,(
  ( ! [X24,X23] : (p(s(bool,happ(s(fun('type.option.option'(bool),bool),X24),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X24,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X23)) ) | $spl1),
  inference(superposition,[],[f164,f2292])).
fof(f968,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(X31),bool),X33),s('type.option.option'(X31),X32))) | s('type.option.option'(X31),'const.option.SOME_1'(s(X31,sK4(X32,X31)))) = s('type.option.option'(X31),sK5('const.option.SOME_1'(s(X31,sK4(X32,X31))),X33,'type.option.option'(X31))) | s('type.option.option'(X31),'const.option.NONE_0') = s('type.option.option'(X31),X32)) ) | $spl1),
  inference(superposition,[],[f877,f154])).
fof(f1293,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,happ(s(fun('type.option.option'(X19),bool),X21),s('type.option.option'(X19),X20))) | s('type.option.option'(X19),'const.option.SOME_1'(s(X19,sK4(X20,X19)))) = s('type.option.option'(X19),sK5('const.option.SOME_1'(s(X19,sK4(X20,X19))),X21,'type.option.option'(X19))) | s('type.option.option'(X19),'const.option.NONE_0') = s('type.option.option'(X19),X20)) ) | $spl1),
  inference(superposition,[],[f948,f154])).
fof(f1363,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,happ(s(fun('type.option.option'(X21),bool),X23),s('type.option.option'(X21),X22)))) | s('type.option.option'(X21),sK5('const.option.SOME_1'(s(X21,sK4(X22,X21))),X23,'type.option.option'(X21))) = s('type.option.option'(X21),sK5(sK5('const.option.SOME_1'(s(X21,sK4(X22,X21))),X23,'type.option.option'(X21)),X23,'type.option.option'(X21))) | s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X22)) ) | $spl1),
  inference(superposition,[],[f1321,f154])).
fof(f1571,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,happ(s(fun('type.option.option'(X37),bool),X39),s('type.option.option'(X37),X38)))) | s('type.option.option'(X37),X40) != s('type.option.option'(X37),'const.option.SOME_1'(s(X37,sK4(X38,X37)))) | s('type.option.option'(X37),X40) = s('type.option.option'(X37),sK5(X40,X39,'type.option.option'(X37))) | s('type.option.option'(X37),'const.option.NONE_0') = s('type.option.option'(X37),X38)) ) | $spl1),
  inference(superposition,[],[f1498,f154])).
fof(f1816,plain,(
  ( ! [X1] : (s(sK0,X1) != s(sK0,sK4(sK1,sK0)) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X1))))) ) | $spl20),
  inference(subsumption_resolution,[],[f1780,f126])).
fof(f1780,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(sK0,X1) != s(sK0,sK4(sK1,sK0)) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X1))))) ) | $spl20),
  inference(superposition,[],[f166,f1762])).
fof(f2568,plain,(
  ( ! [X101,X102,X100] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X101),s(bool,X100))) | s(bool,'T_0') = s(bool,sK5('T_0',X101,bool)) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) ) | $spl1),
  inference(superposition,[],[f948,f239])).
fof(f2567,plain,(
  ( ! [X99,X97,X98] : (s(bool,X97) = s(bool,happ(s(fun(bool,bool),X98),s(bool,X97))) | s(bool,'T_0') = s(bool,sK5('T_0',X98,bool)) | s(bool,X97) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) ) | $spl1),
  inference(superposition,[],[f877,f239])).
fof(f898,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,X6))) | s(bool,'F_0') = s(bool,sK5('F_0',X7,bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl1),
  inference(superposition,[],[f877,f129])).
fof(f1288,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,happ(s(fun(bool,bool),X10),s(bool,X9))) | s(bool,'F_0') = s(bool,sK5('F_0',X10,bool)) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f948,f129])).
fof(f877,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK5(X3,X2,X1))) ) | $spl1),
  inference(subsumption_resolution,[],[f876,f128])).
fof(f876,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(X1,X3) = s(X1,sK5(X3,X2,X1)) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) ) | $spl1),
  inference(forward_demodulation,[],[f861,f558])).
fof(f861,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | s(X1,X3) = s(X1,sK5(X3,X2,X1)) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) ) | $spl1),
  inference(superposition,[],[f164,f640])).
fof(f948,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X13,bool),X14),s(X13,X15))) | s(X13,X15) = s(X13,sK5(X15,X14,X13))) ) | $spl1),
  inference(subsumption_resolution,[],[f918,f126])).
fof(f918,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X13,bool),X14),s(X13,X15))) | s(X13,X15) = s(X13,sK5(X15,X14,X13))) ) | $spl1),
  inference(superposition,[],[f135,f877])).
fof(f1372,plain,(
  ( ! [X6,X4,X5] : (s(X4,sK5(X6,X5,X4)) = s(X4,sK5(sK5(X6,X5,X4),X5,X4)) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl1),
  inference(subsumption_resolution,[],[f1371,f128])).
fof(f1371,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(X4,sK5(X6,X5,X4)) = s(X4,sK5(sK5(X6,X5,X4),X5,X4)) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl1),
  inference(forward_demodulation,[],[f1368,f558])).
fof(f1368,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | s(X4,sK5(X6,X5,X4)) = s(X4,sK5(sK5(X6,X5,X4),X5,X4)) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) ) | $spl1),
  inference(superposition,[],[f1321,f640])).
fof(f1551,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK5(X3,X2,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f1550,f128])).
fof(f1550,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK5(X3,X2,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f1540,f558])).
fof(f1540,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | s(bool,X0) != s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK5(X3,X2,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f950,f640])).
fof(f2991,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X27) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X28,X30) = s(X28,sK5(X30,X29,X28)) | s(bool,X27) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl1),
  inference(subsumption_resolution,[],[f2583,f126])).
fof(f2583,plain,(
  ( ! [X30,X28,X31,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X27) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X28,X30) = s(X28,sK5(X30,X29,X28)) | s(bool,X27) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl1),
  inference(superposition,[],[f949,f239])).
fof(f3188,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | $spl1),
  inference(subsumption_resolution,[],[f3168,f126])).
fof(f3168,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | $spl1),
  inference(superposition,[],[f2304,f877])).
fof(f3082,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(bool),bool),X44),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X44,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X43)) ) | $spl1),
  inference(superposition,[],[f877,f2292])).
fof(f3083,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,happ(s(fun('type.option.option'(bool),bool),X46),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X45))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X46,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X45)) ) | $spl1),
  inference(superposition,[],[f948,f2292])).
fof(f630,plain,(
  ( ! [X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | $spl1),
  inference(superposition,[],[f162,f558])).
fof(f940,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,X0) = s(bool,sK5(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f895,f558])).
fof(f895,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(bool,X0) = s(bool,sK5(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f877,f640])).
fof(f942,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(forward_demodulation,[],[f899,f558])).
fof(f899,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,X8))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X9,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f877,f640])).
fof(f969,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0) | s(bool,'T_0') = s(bool,sK4(X0,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f953,f558])).
fof(f953,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0) | s(bool,'T_0') = s(bool,sK4(X0,bool))) ) | $spl1),
  inference(superposition,[],[f154,f640])).
fof(f1302,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,X5) = s(bool,sK5(X5,X6,bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(forward_demodulation,[],[f1286,f558])).
fof(f1286,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X6),s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(bool,X5) = s(bool,sK5(X5,X6,bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f948,f640])).
fof(f1304,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X13,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(forward_demodulation,[],[f1290,f558])).
fof(f1290,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,X12))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)),X13,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(superposition,[],[f948,f640])).
fof(f1538,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK6(X0,X1,X2,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f1534,f558])).
fof(f1534,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) != s(X2,g(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK6(X0,X1,X2,bool))) ) | $spl1),
  inference(superposition,[],[f176,f640])).
fof(f2116,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f2089,f558])).
fof(f2089,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool))) ) | $spl1),
  inference(superposition,[],[f951,f640])).
fof(f2220,plain,(
  ( ! [X47,X45,X46] : (s(X45,X46) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45)) | s(bool,'F_0') != s(bool,happ(s(fun(X45,bool),X47),s(X45,X46)))) ) | $spl1),
  inference(forward_demodulation,[],[f2170,f2072])).
fof(f2170,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') != s(bool,happ(s(fun(X45,bool),X47),s(X45,X46))) | s(X45,sK4('const.option.SOME_1'(s(X45,X46)),X45)) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45))) ) | $spl1),
  inference(superposition,[],[f948,f2072])).
fof(f2274,plain,(
  ( ! [X28,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,sK4('const.option.SOME_1'(s(bool,X28)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28)))) ) | $spl1),
  inference(superposition,[],[f162,f2116])).
fof(f2292,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f1284,f1301,f1302,f1288,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291])).
fof(f2301,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl1),
  inference(forward_demodulation,[],[f2290,f2072])).
fof(f2290,plain,(
  ( ! [X1] : (p(s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl1),
  inference(trivial_inequality_removal,[],[f2258])).
fof(f2258,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1)))) ) | $spl1),
  inference(superposition,[],[f131,f2116])).
fof(f2302,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X3)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2260,f126])).
fof(f2260,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X3)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3)))) ) | $spl1),
  inference(superposition,[],[f135,f2116])).
fof(f2304,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X9)))) ) | $spl1),
  inference(forward_demodulation,[],[f2303,f2072])).
fof(f2303,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK4('const.option.SOME_1'(s(bool,X9)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X9)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2264,f126])).
fof(f2264,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK4('const.option.SOME_1'(s(bool,X9)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X9)))) ) | $spl1),
  inference(superposition,[],[f170,f2116])).
fof(f2306,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11)))) ) | $spl1),
  inference(forward_demodulation,[],[f2305,f2072])).
fof(f2305,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,sK4('const.option.SOME_1'(s(bool,X11)),bool)) | p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2265,f126])).
fof(f2265,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X12) != s(bool,sK4('const.option.SOME_1'(s(bool,X11)),bool)) | p(s(bool,X12)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11)))) ) | $spl1),
  inference(superposition,[],[f204,f2116])).
fof(f2309,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | $spl1),
  inference(subsumption_resolution,[],[f2272,f159])).
fof(f2272,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | $spl1),
  inference(superposition,[],[f154,f2116])).
fof(f2310,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') != s(bool,X33) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32)))) ) | $spl1),
  inference(forward_demodulation,[],[f2276,f2072])).
fof(f2276,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK4('const.option.SOME_1'(s(bool,X32)),bool)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32)))) ) | $spl1),
  inference(superposition,[],[f163,f2116])).
fof(f2455,plain,(
  ( ! [X158,X159,X160] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X160),s(bool,X159))) | s(bool,X158) = s(bool,sK5(X158,X160,bool)) | s(bool,'T_0') = s(bool,X159) | s(bool,'T_0') = s(bool,X158)) ) | $spl1),
  inference(superposition,[],[f948,f239])).
fof(f2623,plain,(
  ( ! [X152,X151,X150] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X151),s(bool,'T_0'))) | s(bool,X150) = s(bool,sK5(X150,X151,bool)) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) ) | $spl1),
  inference(superposition,[],[f948,f239])).
fof(f3060,plain,(
  ( ! [X21,X19,X20] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X19),'const.option.NONE_0'),s(bool,X20),s(fun(X19,bool),X21)))) ) | $spl1),
  inference(superposition,[],[f2292,f175])).
fof(f3066,plain,(
  ( ! [X4,X2,X3,X1] : (s(X2,happ(s(fun(bool,X2),X3),s(bool,'F_0'))) = s(X2,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))),s(X2,X4),s(fun(bool,X2),X3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f177,f2292])).
fof(f3072,plain,(
  ( ! [X15,X16] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X16))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X15))))) | s('type.option.option'(bool),X16) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl1),
  inference(superposition,[],[f162,f2292])).
fof(f3074,plain,(
  ( ! [X19,X20] : (s('type.option.option'(bool),X20) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X19))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X20))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X19)) ) | $spl1),
  inference(superposition,[],[f163,f2292])).
fof(f3088,plain,(
  ( ! [X4,X2,X3,X1] : (s(X2,happ(s(fun(bool,X2),X3),s(bool,X1))) = s(X2,'const.option.option_CASE_3'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))),s(X2,X4),s(fun(bool,X2),X3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f177,f2292])).
fof(f3095,plain,(
  ( ! [X17,X16] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X17))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),X17) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | $spl1),
  inference(superposition,[],[f162,f2292])).
fof(f3097,plain,(
  ( ! [X21,X20] : (s('type.option.option'(bool),X21) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),X21))) = s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))))) | s(bool,'T_0') = s(bool,X20)) ) | $spl1),
  inference(superposition,[],[f163,f2292])).
fof(f3105,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(bool),bool),X45),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,X44)),X45,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X44)) ) | $spl1),
  inference(superposition,[],[f877,f2292])).
fof(f3106,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,happ(s(fun('type.option.option'(bool),bool),X47),s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X46))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,X46)),X47,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X46)) ) | $spl1),
  inference(superposition,[],[f948,f2292])).
fof(f3174,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(subsumption_resolution,[],[f3154,f126])).
fof(f3154,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(superposition,[],[f2304,f239])).
fof(f3187,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30)))) ) | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f3167,f126])).
fof(f3167,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30)))) ) | ($spl1 | $spl20)),
  inference(superposition,[],[f2304,f1762])).
fof(f3189,plain,(
  ( ! [X35,X36] : (s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl1),
  inference(subsumption_resolution,[],[f3169,f126])).
fof(f3169,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | $spl1),
  inference(superposition,[],[f2304,f2116])).
fof(f3322,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X79) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X76),'const.option.NONE_0'),s(bool,X77),s(fun(X76,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X79))) | s(bool,'F_0') = s(bool,X77)) ) | $spl1),
  inference(subsumption_resolution,[],[f3277,f126])).
fof(f3277,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | s(bool,X79) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X76),'const.option.NONE_0'),s(bool,X77),s(fun(X76,bool),X78))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X79))) | s(bool,'F_0') = s(bool,X77)) ) | $spl1),
  inference(superposition,[],[f2304,f303])).
fof(f3366,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X10)),bool)) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | $spl1),
  inference(subsumption_resolution,[],[f3337,f398])).
fof(f3337,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X10)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | $spl1),
  inference(superposition,[],[f2302,f239])).
fof(f3440,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,X44) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43))) | s(bool,'T_0') = s(bool,X44)) ) | $spl1),
  inference(forward_demodulation,[],[f3406,f2072])).
fof(f3406,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X42)),bool)) | s(bool,X44) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43))) | s(bool,'T_0') = s(bool,X44)) ) | $spl1),
  inference(superposition,[],[f3366,f175])).
fof(f3041,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2292,f129])).
fof(f2622,plain,(
  ( ! [X149,X147,X148] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X148),s(bool,'T_0'))) | s(bool,X147) = s(bool,sK5(X147,X148,bool)) | s(bool,X147) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) ) | $spl1),
  inference(superposition,[],[f877,f239])).
fof(f2539,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X10),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(superposition,[],[f640,f239])).
fof(f2538,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X8),s(fun(sK0,bool),sK2))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f560,f239])).
fof(f2454,plain,(
  ( ! [X156,X157,X155] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X157),s(bool,X156))) | s(bool,X155) = s(bool,sK5(X155,X157,bool)) | s(bool,'T_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X155)) ) | $spl1),
  inference(superposition,[],[f877,f239])).
fof(f2291,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f2240,f558])).
fof(f2240,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f2116,f640])).
fof(f2219,plain,(
  ( ! [X43,X44,X42] : (s(X42,X43) = s(X42,sK5(sK4('const.option.SOME_1'(s(X42,X43)),X42),X44,X42)) | s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),X44),s(X42,X43)))) ) | $spl1),
  inference(forward_demodulation,[],[f2169,f2072])).
fof(f2169,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),X44),s(X42,X43))) | s(X42,sK4('const.option.SOME_1'(s(X42,X43)),X42)) = s(X42,sK5(sK4('const.option.SOME_1'(s(X42,X43)),X42),X44,X42))) ) | $spl1),
  inference(superposition,[],[f877,f2072])).
fof(f2156,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X2)),bool))) ) | $spl1),
  inference(superposition,[],[f2072,f640])).
fof(f1798,plain,(
  ( ! [X13] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) ) | $spl20),
  inference(superposition,[],[f162,f1762])).
fof(f1762,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | $spl20),
  inference(unit_resulting_resolution,[],[f1727,f132])).
fof(f560,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f211,f131])).
fof(f640,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f627,f128])).
fof(f627,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f272,f558])).
fof(f756,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.NONE_0'),s(bool,X7),s(fun(X6,bool),X8)))) ) | $spl1),
  inference(superposition,[],[f640,f175])).
fof(f775,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X16))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | $spl1),
  inference(superposition,[],[f163,f640])).
fof(f797,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s(bool,X23) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X22)) ) | $spl1),
  inference(superposition,[],[f162,f640])).
fof(f1586,plain,(
  ( ! [X10,X8,X7,X9] : (s(X7,X9) != s(X7,X10) | s(X7,X10) = s(X7,sK5(X10,X8,X7)) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl1),
  inference(subsumption_resolution,[],[f1585,f128])).
fof(f1585,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'F_0')) | s(X7,X9) != s(X7,X10) | s(X7,X10) = s(X7,sK5(X10,X8,X7)) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl1),
  inference(forward_demodulation,[],[f1576,f558])).
fof(f1576,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | s(X7,X9) != s(X7,X10) | s(X7,X10) = s(X7,sK5(X10,X8,X7)) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) ) | $spl1),
  inference(superposition,[],[f1498,f640])).
fof(f500,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | $spl1),
  inference(condensation,[],[f499])).
fof(f499,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(subsumption_resolution,[],[f487,f337])).
fof(f487,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f450,f337])).
fof(f483,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) != s(bool,X12) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X12))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11)))) ) | $spl1),
  inference(superposition,[],[f450,f175])).
fof(f465,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X12) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11))) | s(bool,'T_0') = s(bool,X12)) ) | $spl1),
  inference(superposition,[],[f419,f175])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f419,f129])).
fof(f440,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f163,f337])).
fof(f450,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(forward_demodulation,[],[f437,f225])).
fof(f437,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f163,f337])).
fof(f427,plain,(
  ( ! [X0] : (s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f422,f129])).
fof(f426,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f422,f162])).
fof(f422,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f417,f162])).
fof(f419,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(forward_demodulation,[],[f407,f225])).
fof(f407,plain,(
  ( ! [X6,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f162,f337])).
fof(f405,plain,(
  ( ! [X2] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,X2) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | $spl1),
  inference(superposition,[],[f162,f225])).
fof(f404,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,X1) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f162,f337])).
fof(f417,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl1),
  inference(forward_demodulation,[],[f402,f225])).
fof(f402,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f218,f162])).
fof(f344,plain,(
  ( ! [X6,X4,X5] : (s(bool,X5) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X4),'const.option.NONE_0'),s(bool,X5),s(fun(X4,bool),X6)))) ) | $spl1),
  inference(superposition,[],[f337,f175])).
fof(f337,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f333,f128])).
fof(f333,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f272,f225])).
fof(f225,plain,(
  s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f211,f136])).
fof(f218,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f211,f131])).
fof(f12944,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750,f12790,f929,f12943])).
fof(f12825,plain,(
  $false | ($spl1 | $spl12 | $spl20)),
  inference(global_subsumption,[],[f12824,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750,f12790,f929])).
fof(f12824,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X25,sK0))) ) | ($spl1 | $spl12 | $spl20)),
  inference(forward_demodulation,[],[f12823,f12545])).
fof(f12823,plain,(
  ( ! [X26,X25] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK3))))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X25,sK0))) ) | ($spl1 | $spl12 | $spl20)),
  inference(subsumption_resolution,[],[f1343,f12545])).
fof(f1343,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X26) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),X25),s(sK0,sK3))))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X25,sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f163,f1261])).
fof(f1261,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f1235,f873])).
fof(f1235,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))) | s(sK0,sK5(sK3,sK2,sK0)) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f877,f873])).
fof(f12793,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750,f12790,f929])).
fof(f12791,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750,f12790])).
fof(f12788,plain,(
  $false | ($spl1 | $spl12 | $spl20)),
  inference(global_subsumption,[],[f1341,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750])).
fof(f1341,plain,(
  ( ! [X21,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) | s(bool,X22) = s(bool,happ(s(fun(sK0,bool),X21),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X21,sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f162,f1261])).
fof(f12787,plain,(
  $false | ($spl1 | $spl12 | $spl20)),
  inference(global_subsumption,[],[f12786,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750])).
fof(f12786,plain,(
  ( ! [X28,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X28,sK0))) ) | ($spl1 | $spl12 | $spl20)),
  inference(forward_demodulation,[],[f3186,f12545])).
fof(f3186,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X28,sK0))) ) | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f3166,f126])).
fof(f3166,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X28,sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f2304,f1261])).
fof(f12752,plain,(
  $false | ($spl1 | $spl10 | $spl20)),
  inference(global_subsumption,[],[f12681,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750])).
fof(f12681,plain,(
  ( ! [X6] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl10 | $spl20)),
  inference(backward_demodulation,[],[f12545,f12515])).
fof(f12515,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6)))) ) | ($spl1 | $spl10 | $spl20)),
  inference(backward_demodulation,[],[f12514,f495])).
fof(f12751,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747,f12750])).
fof(f12748,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744,f12747])).
fof(f12745,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638,f12744])).
fof(f12741,plain,(
  $false | ($spl1 | $spl3 | $spl20)),
  inference(global_subsumption,[],[f12740,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638])).
fof(f12740,plain,(
  ( ! [X47,X46] : (s(bool,X47) = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool)) | s(bool,X46) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(subsumption_resolution,[],[f12648,f11517])).
fof(f12648,plain,(
  ( ! [X47,X46] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,X47) = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool)) | s(bool,X46) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(backward_demodulation,[],[f12545,f11728])).
fof(f11728,plain,(
  ( ! [X47,X46] : (s(bool,X47) = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,X46) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f11496,f126])).
fof(f11496,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | s(bool,X47) = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X47))) | s(bool,X46) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f2304,f8908])).
fof(f12739,plain,(
  $false | ($spl1 | $spl3 | $spl20)),
  inference(global_subsumption,[],[f12738,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638])).
fof(f12738,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(subsumption_resolution,[],[f12647,f8908])).
fof(f12647,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X7)),bool)) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(backward_demodulation,[],[f12545,f11720])).
fof(f12737,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725,f12638])).
fof(f12736,plain,(
  $false | ($spl1 | $spl3 | $spl20 | $spl32)),
  inference(global_subsumption,[],[f12619,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725])).
fof(f12619,plain,(
  ( ! [X43,X44] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44))) | s(bool,X44) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X43),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl3 | $spl20 | $spl32)),
  inference(backward_demodulation,[],[f12545,f10037])).
fof(f10037,plain,(
  ( ! [X43,X44] : (s(bool,X44) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X43),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f9983,f128])).
fof(f9983,plain,(
  ( ! [X43,X44] : (p(s(bool,'F_0')) | s(bool,X44) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X43),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X44)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f2306,f9726])).
fof(f9726,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f177,f8645,f9633])).
fof(f9633,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f6492,f640])).
fof(f12735,plain,(
  $false | ($spl1 | $spl3 | $spl20 | $spl32)),
  inference(global_subsumption,[],[f12618,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725])).
fof(f12618,plain,(
  ( ! [X57,X56] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,X57) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X56),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl3 | $spl20 | $spl32)),
  inference(backward_demodulation,[],[f12545,f9992])).
fof(f9992,plain,(
  ( ! [X57,X56] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X57))) | s(bool,X57) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X56),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f162,f9726])).
fof(f12734,plain,(
  $false | ($spl1 | $spl3 | $spl20)),
  inference(global_subsumption,[],[f12733,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725])).
fof(f12733,plain,(
  ( ! [X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(subsumption_resolution,[],[f12604,f11736])).
fof(f12604,plain,(
  ( ! [X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X36) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(backward_demodulation,[],[f12545,f9487])).
fof(f9487,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f9473,f128])).
fof(f9473,plain,(
  ( ! [X36] : (p(s(bool,'F_0')) | s(bool,X36) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36)))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f2306,f8657])).
fof(f8657,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f8650,f558])).
fof(f8650,plain,(
  s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f560,f213,f272])).
fof(f12732,plain,(
  $false | ($spl1 | $spl3 | $spl20)),
  inference(global_subsumption,[],[f12731,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725])).
fof(f12731,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(subsumption_resolution,[],[f12601,f11736])).
fof(f12601,plain,(
  ( ! [X10] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl20)),
  inference(backward_demodulation,[],[f12545,f8806])).
fof(f8806,plain,(
  ( ! [X10] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f162,f8657])).
fof(f12726,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723,f12725])).
fof(f12724,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721,f12723])).
fof(f12722,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717,f12721])).
fof(f12718,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700,f12717])).
fof(f12715,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567,f12700])).
fof(f12714,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566,f12567])).
fof(f12713,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564,f12566])).
fof(f12711,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709,f12564])).
fof(f12710,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562,f12709])).
fof(f12708,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706,f12562])).
fof(f12707,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703,f12706])).
fof(f12704,plain,(
  $false | ($spl1 | $spl20)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f1727,f1762,f1798,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3187,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1816,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1832,f1831,f1829,f1810,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10628,f10629,f10630,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10905,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11181,f11182,f11145,f11185,f11148,f11149,f11150,f11165,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f12025,f12015,f12014,f12013,f12012,f12010,f12009,f12008,f12007,f11988,f11987,f11986,f11985,f11984,f11983,f11982,f11980,f11979,f11978,f11977,f11976,f11975,f11974,f11973,f11972,f11971,f11968,f11967,f11966,f11964,f11963,f11962,f11961,f12514,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12545,f12703])).
fof(f12702,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12546,f126])).
fof(f12546,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f128])).
fof(f12701,plain,(
  $false | ($spl1 | $spl20)),
  inference(trivial_inequality_removal,[],[f12550])).
fof(f12550,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12545,f398])).
fof(f12543,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,X6),s(fun(sK0,bool),sK2)))) ) | $spl72),
  inference(cnf_transformation,[],[f12543_D])).
fof(f12543_D,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,X6),s(fun(sK0,bool),sK2)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f12542,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,X6),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl20 | $spl32)),
  inference(subsumption_resolution,[],[f12526,f9773])).
fof(f9773,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f175])).
fof(f12526,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,X6),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f2537])).
fof(f12541,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12517,f175])).
fof(f12517,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f560])).
fof(f12540,plain,(
  $false | ($spl1 | $spl20)),
  inference(subsumption_resolution,[],[f12539,f217])).
fof(f12539,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f12516,f175])).
fof(f12516,plain,(
  s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl20)),
  inference(backward_demodulation,[],[f12514,f558])).
fof(f11739,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,X85)) ) | $spl70),
  inference(cnf_transformation,[],[f11739_D])).
fof(f11739_D,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,X85)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f12505,plain,(
  ( ! [X85,X84] : (s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') != s(bool,X85)) ) | ($spl1 | $spl3 | $spl22 | $spl32 | $spl34)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f511,f211,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f944,f945,f947,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4224,f6493,f6492,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f213,f558,f2537,f8651,f8657,f8645,f8665,f8806,f8827,f8908,f8946,f8985,f9058,f9173,f9174,f9322,f9323,f9487,f9518,f9519,f9520,f9521,f9522,f9607,f9622,f9649,f9650,f9625,f9626,f9633,f9653,f9726,f9752,f9753,f9754,f9755,f9797,f9773,f9774,f9775,f9776,f9836,f9820,f9868,f9869,f9870,f9871,f9940,f10024,f10037,f9992,f10043,f10047,f10049,f10108,f10115,f10116,f10119,f10093,f10094,f10095,f10096,f10121,f10122,f10123,f10124,f10310,f10290,f10328,f10279,f10280,f10281,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10504,f10510,f10512,f10513,f10475,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10627,f10628,f10629,f10630,f10703,f10747,f10748,f10848,f10856,f10817,f10859,f10860,f10820,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10924,f10926,f10905,f10939,f10960,f10965,f10966,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11179,f11180,f11181,f11183,f11184,f11145,f11185,f11186,f11148,f11149,f11150,f11189,f11190,f11192,f11165,f11205,f11226,f11230,f11232,f11233,f11256,f11587,f11589,f11607,f11636,f11641,f11373,f11649,f11654,f11655,f11656,f11657,f11664,f11669,f11451,f11707,f11708,f11717,f11720,f11728,f11736,f11517,f11743,f11748,f11750,f11754,f11756,f11534,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11740,f12504,f11737])).
fof(f11756,plain,(
  ( ! [X127,X128] : (s(bool,'F_0') = s(bool,X127) | s(bool,X127) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X127)),bool),X128,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X128),s(bool,'T_0')))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11755,f8657])).
fof(f11755,plain,(
  ( ! [X127,X128] : (s(bool,X127) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X127)),bool),X128,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X128),s(bool,'T_0'))) | s(bool,X127) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11533,f2072])).
fof(f11533,plain,(
  ( ! [X127,X128] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X128),s(bool,'T_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X127)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X127)),bool),X128,bool)) | s(bool,X127) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f8908])).
fof(f11754,plain,(
  ( ! [X125,X126] : (s(bool,'F_0') = s(bool,X125) | s(bool,X125) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X125)),bool),X126,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X126),s(bool,'T_0')))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11753,f8657])).
fof(f11753,plain,(
  ( ! [X125,X126] : (s(bool,X125) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X125)),bool),X126,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X126),s(bool,'T_0'))) | s(bool,X125) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11532,f2072])).
fof(f11532,plain,(
  ( ! [X125,X126] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X126),s(bool,'T_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X125)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X125)),bool),X126,bool)) | s(bool,X125) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f8908])).
fof(f11750,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'F_0') = s(bool,X117) | s(bool,X117) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X117)),bool),'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.NONE_0'),s(fun(bool,bool),X118),s(fun(X119,fun(bool,bool)),X120)),bool)) | p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0'))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11749,f8657])).
fof(f11748,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') = s(bool,X113) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X114))) | s(bool,sK4('const.option.SOME_1'(s(bool,X113)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X115),'const.option.NONE_0'),s(bool,X114),s(fun(X115,bool),X116)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11529,f8657])).
fof(f11743,plain,(
  ( ! [X88,X89] : (s(bool,'F_0') = s(bool,X88) | s(bool,X88) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X88)),bool),X89,bool)) | p(s(bool,happ(s(fun(bool,bool),X89),s(bool,'T_0'))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11742,f8657])).
fof(f11717,plain,(
  ( ! [X144,X143] : (s(bool,'F_0') = s(bool,sK4(X144,bool)) | s('type.option.option'(bool),X144) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X143))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X144) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X143)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11468,f8657])).
fof(f11708,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X129),'const.option.NONE_0'),s(bool,X128),s(fun(X129,bool),X130))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X127))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X128))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X127)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11463,f8657])).
fof(f11707,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X124) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X125),'const.option.NONE_0'),s(bool,'F_0'),s(fun(X125,bool),X126))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X123))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X124))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X123)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11462,f8657])).
fof(f11669,plain,(
  ( ! [X194,X195,X193] : (s(X195,f(s(bool,'F_0'))) != s(X195,g(s(bool,'F_0'))) | s(fun(bool,X195),X193) = s(fun(bool,X195),X194) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,sK6(X193,X194,X195,bool))),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11397,f8657])).
fof(f11397,plain,(
  ( ! [X194,X195,X193] : (s(X195,f(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) != s(X195,g(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(fun(bool,X195),X193) = s(fun(bool,X195),X194) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,sK6(X193,X194,X195,bool))),bool))) ) | $spl3),
  inference(superposition,[],[f176,f8908])).
fof(f11664,plain,(
  ( ! [X189] : (s('type.option.option'(bool),X189) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X189) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,sK4(X189,bool))),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11394,f8657])).
fof(f11394,plain,(
  ( ! [X189] : (s('type.option.option'(bool),X189) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X189) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,sK4(X189,bool))),bool))) ) | $spl3),
  inference(superposition,[],[f154,f8908])).
fof(f11657,plain,(
  ( ! [X177,X175,X176] : (s(bool,'F_0') != s(bool,f(s(X175,sK6(X176,X177,bool,X175)))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,g(s(X175,sK6(X176,X177,bool,X175))))),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11390,f8657])).
fof(f11390,plain,(
  ( ! [X177,X175,X176] : (s(bool,f(s(X175,sK6(X176,X177,bool,X175)))) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(fun(X175,bool),X176) = s(fun(X175,bool),X177) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,g(s(X175,sK6(X176,X177,bool,X175))))),bool))) ) | $spl3),
  inference(superposition,[],[f176,f8908])).
fof(f11656,plain,(
  ( ! [X169,X170] : (s('type.option.option'(bool),X170) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X169) = s(bool,sK4(X170,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X170) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X169)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11388,f8657])).
fof(f11388,plain,(
  ( ! [X169,X170] : (s('type.option.option'(bool),X170) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,X169) = s(bool,sK4(X170,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X170) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X169)),bool))) ) | $spl3),
  inference(superposition,[],[f956,f8908])).
fof(f11655,plain,(
  ( ! [X167,X168] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X168),s(bool,'F_0'))) | s(bool,X167) = s(bool,sK5(X167,X168,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X167)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11387,f8657])).
fof(f11387,plain,(
  ( ! [X167,X168] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X168),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,X167) = s(bool,sK5(X167,X168,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X167)),bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f8908])).
fof(f11654,plain,(
  ( ! [X165,X166] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X166),s(bool,'F_0'))) | s(bool,X165) = s(bool,sK5(X165,X166,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X165)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11386,f8657])).
fof(f11386,plain,(
  ( ! [X165,X166] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X166),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,X165) = s(bool,sK5(X165,X166,bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X165)),bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f8908])).
fof(f11649,plain,(
  ( ! [X156,X154,X155,X153] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X154))) | s(bool,X153) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X155),'const.option.NONE_0'),s(bool,X154),s(fun(X155,bool),X156))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X153)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11383,f8657])).
fof(f11383,plain,(
  ( ! [X156,X154,X155,X153] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X154))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,X153) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X155),'const.option.NONE_0'),s(bool,X154),s(fun(X155,bool),X156))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X153)),bool))) ) | $spl3),
  inference(superposition,[],[f409,f8908])).
fof(f11641,plain,(
  ( ! [X121,X120] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X121))) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X120)),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11371,f8657])).
fof(f11371,plain,(
  ( ! [X121,X120] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X121))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X120)),bool))) ) | $spl3),
  inference(superposition,[],[f162,f8908])).
fof(f11636,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,sK5('F_0',X113,bool)) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X113),s(bool,X114)))),bool))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f11635,f128])).
fof(f11635,plain,(
  ( ! [X114,X113] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK5('F_0',X113,bool)) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X113),s(bool,X114)))),bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11366,f8657])).
fof(f11366,plain,(
  ( ! [X114,X113] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,'F_0') = s(bool,sK5('F_0',X113,bool)) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(bool,bool),X113),s(bool,X114)))),bool))) ) | $spl3),
  inference(superposition,[],[f853,f8908])).
fof(f11607,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X32)),bool))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f11321,f8657])).
fof(f11321,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X32)),bool))) ) | $spl3),
  inference(superposition,[],[f441,f8908])).
fof(f11589,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X40,bool),bool),happ(s(fun(bool,fun(fun(X40,bool),bool)),happ(s(fun('type.option.option'(X40),fun(bool,fun(fun(X40,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X40),X41))),s(bool,X42))),s(fun(X40,bool),X43))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),X41),s(bool,X42),s(fun(X40,bool),X43)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11588,f8657])).
fof(f11588,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),X41),s(bool,X42),s(fun(X40,bool),X43))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,happ(s(fun(fun(X40,bool),bool),happ(s(fun(bool,fun(fun(X40,bool),bool)),happ(s(fun('type.option.option'(X40),fun(bool,fun(fun(X40,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X40),X41))),s(bool,X42))),s(fun(X40,bool),X43)))) ) | $spl3),
  inference(forward_demodulation,[],[f11293,f2072])).
fof(f11293,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X40),X41),s(bool,X42),s(fun(X40,bool),X43)))),bool)) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,happ(s(fun(fun(X40,bool),bool),happ(s(fun(bool,fun(fun(X40,bool),bool)),happ(s(fun('type.option.option'(X40),fun(bool,fun(fun(X40,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X40),X41))),s(bool,X42))),s(fun(X40,bool),X43)))) ) | $spl3),
  inference(superposition,[],[f8908,f178])).
fof(f11587,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.SOME_1'(s(X36,X37))),s(bool,X38),s(fun(X36,bool),X39))) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X39),s(X36,X37)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11586,f8657])).
fof(f11586,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X39),s(X36,X37))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.SOME_1'(s(X36,X37))),s(bool,X38),s(fun(X36,bool),X39)))) ) | $spl3),
  inference(forward_demodulation,[],[f11292,f2072])).
fof(f11292,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X36,bool),X39),s(X36,X37)))),bool)) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X36),'const.option.SOME_1'(s(X36,X37))),s(bool,X38),s(fun(X36,bool),X39)))) ) | $spl3),
  inference(superposition,[],[f8908,f177])).
fof(f11233,plain,(
  ( ! [X72,X71,X73] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X71),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X71,fun(sK0,bool)),X72)),sK0),X73,sK0)) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X73),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11221,f11192])).
fof(f11221,plain,(
  ( ! [X72,X71,X73] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X73),s(sK0,sK3))) | s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X71),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X71,fun(sK0,bool)),X72)),sK0)) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X71),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X71,fun(sK0,bool)),X72)),sK0),X73,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f11192])).
fof(f11232,plain,(
  ( ! [X70,X68,X69] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X68,fun(sK0,bool)),X69)),sK0),X70,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11220,f11192])).
fof(f11220,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK3))) | s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X68,fun(sK0,bool)),X69)),sK0)) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X68),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X68,fun(sK0,bool)),X69)),sK0),X70,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f11192])).
fof(f11230,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X58),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X58,fun(sK0,bool)),X59)),sK0),'const.option.option_CASE_3'(s('type.option.option'(X61),'const.option.NONE_0'),s(fun(sK0,bool),X60),s(fun(X61,fun(sK0,bool)),X62)),sK0)) | p(s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK3))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11218,f11192])).
fof(f11218,plain,(
  ( ! [X61,X59,X62,X60,X58] : (p(s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK3)))) | s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X58),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X58,fun(sK0,bool)),X59)),sK0)) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X58),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X58,fun(sK0,bool)),X59)),sK0),'const.option.option_CASE_3'(s('type.option.option'(X61),'const.option.NONE_0'),s(fun(sK0,bool),X60),s(fun(X61,fun(sK0,bool)),X62)),sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f849,f11192])).
fof(f11226,plain,(
  ( ! [X21,X22,X20] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X20,fun(sK0,bool)),X21)),sK0),X22,sK0)) | p(s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11209,f11192])).
fof(f11209,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3)))) | s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X20,fun(sK0,bool)),X21)),sK0)) = s(sK0,sK5(sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X20,fun(sK0,bool)),X21)),sK0),X22,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f164,f11192])).
fof(f11205,plain,(
  ( ! [X10,X8,X9] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X10))) | s(sK0,X10) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X8,fun(sK0,bool)),X9)),sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f162,f11192])).
fof(f11192,plain,(
  ( ! [X6,X5] : (s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X5,fun(sK0,bool)),X6)),sK0))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f11156,f128])).
fof(f11156,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X5,fun(sK0,bool)),X6)),sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f849,f8657])).
fof(f11189,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(sK0,sK3) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X140),'const.option.NONE_0'),s(sK0,sK3),s(fun(X140,sK0),X141)),sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X143),'const.option.NONE_0'),s(fun(sK0,bool),X142),s(fun(X143,fun(sK0,bool)),X144)),sK0)) | p(s(bool,happ(s(fun(sK0,bool),X142),s(sK0,sK3))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11153,f10926])).
fof(f11153,plain,(
  ( ! [X144,X142,X140,X143,X141] : (p(s(bool,happ(s(fun(sK0,bool),X142),s(sK0,sK3)))) | s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X140),'const.option.NONE_0'),s(sK0,sK3),s(fun(X140,sK0),X141)),sK2,sK0)) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X140),'const.option.NONE_0'),s(sK0,sK3),s(fun(X140,sK0),X141)),sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X143),'const.option.NONE_0'),s(fun(sK0,bool),X142),s(fun(X143,fun(sK0,bool)),X144)),sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f849,f10926])).
fof(f11186,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X117) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X117)),bool),'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.NONE_0'),s(fun(bool,bool),X118),s(fun(X119,fun(bool,bool)),X120)),bool)) | p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X117)) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f11147,f2072])).
fof(f11147,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'F_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X117)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X117)),bool),'const.option.option_CASE_3'(s('type.option.option'(X119),'const.option.NONE_0'),s(fun(bool,bool),X118),s(fun(X119,fun(bool,bool)),X120)),bool)) | s(bool,'T_0') = s(bool,X117)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f849,f10310])).
fof(f11184,plain,(
  ( ! [X105,X106,X104] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X105),'const.option.NONE_0'),s(fun(bool,bool),X104),s(fun(X105,fun(bool,bool)),X106)),bool)) | p(s(bool,happ(s(fun(bool,bool),X104),s(bool,'F_0'))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11144,f8657])).
fof(f11144,plain,(
  ( ! [X105,X106,X104] : (p(s(bool,happ(s(fun(bool,bool),X104),s(bool,'F_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X105),'const.option.NONE_0'),s(fun(bool,bool),X104),s(fun(X105,fun(bool,bool)),X106)),bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f849,f8657])).
fof(f11183,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X102),'const.option.NONE_0'),s(fun(bool,bool),X101),s(fun(X102,fun(bool,bool)),X103)),bool)) | p(s(bool,happ(s(fun(bool,bool),X101),s(bool,X100)))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f11143,f8657])).
fof(f11143,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,happ(s(fun(bool,bool),X101),s(bool,X100)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X102),'const.option.NONE_0'),s(fun(bool,bool),X101),s(fun(X102,fun(bool,bool)),X103)),bool)) | s(bool,'T_0') = s(bool,X100)) ) | $spl3),
  inference(superposition,[],[f849,f8665])).
fof(f11180,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(bool,X86) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X85),'const.option.NONE_0'),s(bool,X86),s(fun(X85,bool),X87)),'const.option.option_CASE_3'(s('type.option.option'(X89),'const.option.NONE_0'),s(fun(bool,bool),X88),s(fun(X89,fun(bool,bool)),X90)),bool)) | p(s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X86)) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f11140,f175])).
fof(f11140,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X85),'const.option.NONE_0'),s(bool,X86),s(fun(X85,bool),X87))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X85),'const.option.NONE_0'),s(bool,X86),s(fun(X85,bool),X87)),'const.option.option_CASE_3'(s('type.option.option'(X89),'const.option.NONE_0'),s(fun(bool,bool),X88),s(fun(X89,fun(bool,bool)),X90)),bool)) | s(bool,'T_0') = s(bool,X86)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f849,f10703])).
fof(f11179,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X81),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X83),'const.option.NONE_0'),s(fun(bool,bool),X82),s(fun(X83,fun(bool,bool)),X84)),bool)) | p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'F_0'))))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f11178,f8657])).
fof(f11178,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X81),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X83),'const.option.NONE_0'),s(fun(bool,bool),X82),s(fun(X83,fun(bool,bool)),X84)),bool)) | p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'F_0'))))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f11139,f177])).
fof(f11139,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,happ(s(fun(bool,bool),X82),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X81),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X81),s(fun(sK0,bool),sK2)),'const.option.option_CASE_3'(s('type.option.option'(X83),'const.option.NONE_0'),s(fun(bool,bool),X82),s(fun(X83,fun(bool,bool)),X84)),bool))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f849,f9726])).
fof(f10966,plain,(
  ( ! [X68,X66,X67] : (s(sK0,sK3) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(sK0,sK3),s(fun(X66,sK0),X67)),sK2,sK0),X68,sK0)) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f10954,f10926])).
fof(f10954,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK3))) | s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(sK0,sK3),s(fun(X66,sK0),X67)),sK2,sK0)) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X66),'const.option.NONE_0'),s(sK0,sK3),s(fun(X66,sK0),X67)),sK2,sK0),X68,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f10926])).
fof(f10965,plain,(
  ( ! [X64,X65,X63] : (s(sK0,sK3) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(sK0,sK3),s(fun(X63,sK0),X64)),sK2,sK0),X65,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f10953,f10926])).
fof(f10953,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK3))) | s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(sK0,sK3),s(fun(X63,sK0),X64)),sK2,sK0)) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X63),'const.option.NONE_0'),s(sK0,sK3),s(fun(X63,sK0),X64)),sK2,sK0),X65,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f10926])).
fof(f10960,plain,(
  ( ! [X21,X22,X20] : (s(sK0,sK3) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(sK0,sK3),s(fun(X20,sK0),X21)),sK2,sK0),X22,sK0)) | p(s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f10943,f10926])).
fof(f10943,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3)))) | s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(sK0,sK3),s(fun(X20,sK0),X21)),sK2,sK0)) = s(sK0,sK5(sK5('const.option.option_CASE_3'(s('type.option.option'(X20),'const.option.NONE_0'),s(sK0,sK3),s(fun(X20,sK0),X21)),sK2,sK0),X22,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f164,f10926])).
fof(f10939,plain,(
  ( ! [X10,X8,X9] : (s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X10))) | s(sK0,X10) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X8),'const.option.NONE_0'),s(sK0,sK3),s(fun(X8,sK0),X9)),sK2,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f162,f10926])).
fof(f10926,plain,(
  ( ! [X6,X5] : (s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(sK0,sK3),s(fun(X5,sK0),X6)),sK2,sK0))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f10896,f128])).
fof(f10896,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X5),'const.option.NONE_0'),s(sK0,sK3),s(fun(X5,sK0),X6)),sK2,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f870,f8657])).
fof(f10820,plain,(
  ( ! [X237,X235,X238,X236] : (s('type.option.option'(bool),X238) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,sK4(X238,bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X235),'const.option.NONE_0'),s(bool,X236),s(fun(X235,bool),X237))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X238) | s(bool,'T_0') = s(bool,X236)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f956,f10703])).
fof(f10860,plain,(
  ( ! [X231,X233,X234,X232] : (s(bool,X232) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X231),'const.option.NONE_0'),s(bool,X232),s(fun(X231,bool),X233)),X234,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X234),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X232)) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10819,f175])).
fof(f10819,plain,(
  ( ! [X231,X233,X234,X232] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X234),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X231),'const.option.NONE_0'),s(bool,X232),s(fun(X231,bool),X233))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X231),'const.option.NONE_0'),s(bool,X232),s(fun(X231,bool),X233)),X234,bool)) | s(bool,'T_0') = s(bool,X232)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f948,f10703])).
fof(f10859,plain,(
  ( ! [X229,X227,X230,X228] : (s(bool,X228) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X227),'const.option.NONE_0'),s(bool,X228),s(fun(X227,bool),X229)),X230,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X230),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X228)) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10818,f175])).
fof(f10818,plain,(
  ( ! [X229,X227,X230,X228] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X230),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X227),'const.option.NONE_0'),s(bool,X228),s(fun(X227,bool),X229))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X227),'const.option.NONE_0'),s(bool,X228),s(fun(X227,bool),X229)),X230,bool)) | s(bool,'T_0') = s(bool,X228)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f877,f10703])).
fof(f10817,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X224))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X221),'const.option.NONE_0'),s(bool,X222),s(fun(X221,bool),X223))) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X225),'const.option.NONE_0'),s(bool,X224),s(fun(X225,bool),X226))) | s(bool,'T_0') = s(bool,X222)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f409,f10703])).
fof(f10856,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,X181) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X180),'const.option.NONE_0'),s(bool,X181),s(fun(X180,bool),X182)),X183,bool)) | p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X181)) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10809,f175])).
fof(f10809,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(X180),'const.option.NONE_0'),s(bool,X181),s(fun(X180,bool),X182))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(X180),'const.option.NONE_0'),s(bool,X181),s(fun(X180,bool),X182)),X183,bool)) | s(bool,'T_0') = s(bool,X181)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f164,f10703])).
fof(f10848,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X103) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X100),'const.option.NONE_0'),s(bool,X101),s(fun(X100,bool),X102))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,X101)) ) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f10787,f128])).
fof(f10787,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,'F_0')) | s(bool,X103) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X100),'const.option.NONE_0'),s(bool,X101),s(fun(X100,bool),X102))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,X101)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f2306,f10703])).
fof(f10748,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X82),'const.option.NONE_0'),s(bool,'const.option.option_CASE_3'(s('type.option.option'(X78),X79),s(bool,X80),s(fun(X78,bool),X81))),s(fun(X82,bool),X83))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X78,bool),bool),happ(s(fun(bool,fun(fun(X78,bool),bool)),happ(s(fun('type.option.option'(X78),fun(bool,fun(fun(X78,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X78),X79))),s(bool,X80))),s(fun(X78,bool),X81)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f10703,f178])).
fof(f10747,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X76),'const.option.NONE_0'),s(bool,happ(s(fun(X72,bool),X75),s(X72,X73))),s(fun(X76,bool),X77))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X72),'const.option.SOME_1'(s(X72,X73))),s(bool,X74),s(fun(X72,bool),X75)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f10703,f177])).
fof(f10703,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(bool,X0),s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl3 | $spl32)),
  inference(trivial_inequality_removal,[],[f10676])).
fof(f10676,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(bool,X0),s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f9836,f129])).
fof(f10627,plain,(
  ( ! [X109,X107,X110,X108] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X108))) | s(bool,sK4('const.option.SOME_1'(s(bool,X107)),bool)) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X109),'const.option.NONE_0'),s(bool,X108),s(fun(X109,bool),X110))) | s(bool,'T_0') = s(bool,X107)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f409,f10310])).
fof(f10475,plain,(
  ( ! [X105,X104] : (s('type.option.option'(bool),X105) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,sK4(X105,bool)) = s(bool,sK4('const.option.SOME_1'(s(bool,X104)),bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X105) | s(bool,'T_0') = s(bool,X104)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f956,f10310])).
fof(f10513,plain,(
  ( ! [X103,X102] : (s(bool,X102) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X102)),bool),X103,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X103),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f10474,f2072])).
fof(f10474,plain,(
  ( ! [X103,X102] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X103),s(bool,'F_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X102)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X102)),bool),X103,bool)) | s(bool,'T_0') = s(bool,X102)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f948,f10310])).
fof(f10512,plain,(
  ( ! [X101,X100] : (s(bool,X100) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X100)),bool),X101,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X101),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f10473,f2072])).
fof(f10473,plain,(
  ( ! [X101,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X101),s(bool,'F_0'))) | s(bool,sK4('const.option.SOME_1'(s(bool,X100)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X100)),bool),X101,bool)) | s(bool,'T_0') = s(bool,X100)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f877,f10310])).
fof(f10510,plain,(
  ( ! [X80,X79] : (s(bool,X79) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X79)),bool),X80,bool)) | p(s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f10466,f2072])).
fof(f10466,plain,(
  ( ! [X80,X79] : (p(s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0')))) | s(bool,sK4('const.option.SOME_1'(s(bool,X79)),bool)) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X79)),bool),X80,bool)) | s(bool,'T_0') = s(bool,X79)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f164,f10310])).
fof(f10504,plain,(
  ( ! [X47,X48] : (s(bool,X48) != s(bool,sK4('const.option.SOME_1'(s(bool,X47)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f10446,f128])).
fof(f10446,plain,(
  ( ! [X47,X48] : (p(s(bool,'F_0')) | s(bool,X48) != s(bool,sK4('const.option.SOME_1'(s(bool,X47)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f2306,f10310])).
fof(f10281,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X54)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9776,f178])).
fof(f10280,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X49) != s(bool,happ(s(fun(X45,bool),X48),s(X45,X46))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X49)),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9776,f177])).
fof(f10279,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X44)),bool)) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X41),'const.option.NONE_0'),s(bool,X42),s(fun(X41,bool),X43)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9776,f175])).
fof(f10328,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool))) ) | ($spl1 | $spl32)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f6492,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f558,f2537,f9518,f9519,f9520,f9521,f9522,f9607,f9622,f9649,f9650,f9625,f9626,f9633,f9752,f9753,f9754,f9755,f9797,f9773,f9774,f9775,f9776,f9868,f9869,f9870,f9871,f10108,f10115,f10116,f10119,f10093,f10094,f10095,f10096,f10121,f10122,f10123,f10124,f10310,f10290])).
fof(f10290,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl32)),
  inference(duplicate_literal_removal,[],[f10261])).
fof(f10261,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9776,f129])).
fof(f10310,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl32)),
  inference(condensation,[],[f10309])).
fof(f10309,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f10239,f129])).
fof(f10239,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9776,f129])).
fof(f10096,plain,(
  ( ! [X57,X58] : (s(bool,X57) != s(bool,X58) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X57)),bool)) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X58)),bool))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f2072])).
fof(f10095,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,X56) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X52),X53),s(bool,X54),s(fun(X52,bool),X55))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X52,bool),bool),happ(s(fun(bool,fun(fun(X52,bool),bool)),happ(s(fun('type.option.option'(X52),fun(bool,fun(fun(X52,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X52),X53))),s(bool,X54))),s(fun(X52,bool),X55))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X56)),bool))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f178])).
fof(f10094,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,X51) != s(bool,happ(s(fun(X47,bool),X50),s(X47,X48))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X47),'const.option.SOME_1'(s(X47,X48))),s(bool,X49),s(fun(X47,bool),X50))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X51)),bool))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f177])).
fof(f10093,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X44) != s(bool,X46) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.NONE_0'),s(bool,X44),s(fun(X43,bool),X45))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X46)),bool))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f175])).
fof(f10119,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X19)),bool)) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f10082,f217])).
fof(f10082,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X19)),bool)) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f239])).
fof(f10116,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool))) ) | ($spl1 | $spl32)),
  inference(global_subsumption,[],[f10115,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871])).
fof(f10115,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') = s(bool,X1)) ) | ($spl1 | $spl32)),
  inference(condensation,[],[f10114])).
fof(f10114,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X3)),bool)) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f10076,f9633])).
fof(f10076,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X3)),bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f239])).
fof(f10049,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X87),s(fun(sK0,bool),sK2)),X88,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0')))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10048,f8657])).
fof(f10048,plain,(
  ( ! [X88,X87] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X87),s(fun(sK0,bool),sK2)),X88,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0')))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10004,f177])).
fof(f10004,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X88),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X87),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X87),s(fun(sK0,bool),sK2)),X88,bool))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f948,f9726])).
fof(f10047,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X85),s(fun(sK0,bool),sK2)),X86,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X86),s(bool,'F_0')))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10046,f8657])).
fof(f10046,plain,(
  ( ! [X85,X86] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X85),s(fun(sK0,bool),sK2)),X86,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X86),s(bool,'F_0')))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10003,f177])).
fof(f10003,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X86),s(bool,'F_0'))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X85),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X85),s(fun(sK0,bool),sK2)),X86,bool))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f877,f9726])).
fof(f10043,plain,(
  ( ! [X64,X65] : (s(bool,'F_0') = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X64),s(fun(sK0,bool),sK2)),X65,bool)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'F_0'))))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f10042,f8657])).
fof(f10042,plain,(
  ( ! [X64,X65] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X64),s(fun(sK0,bool),sK2)),X65,bool)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'F_0'))))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f9996,f177])).
fof(f9996,plain,(
  ( ! [X64,X65] : (p(s(bool,happ(s(fun(bool,bool),X65),s(bool,'F_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X64),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X64),s(fun(sK0,bool),sK2)),X65,bool))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f164,f9726])).
fof(f10024,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f9953,f217])).
fof(f9953,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f9726,f239])).
fof(f9940,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X15),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f9726,f129])).
fof(f9820,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X33)),bool))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f9653,f2072])).
fof(f9836,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s(bool,X23),s(fun(X22,bool),X24)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f9817,f8657])).
fof(f9817,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X22),'const.option.NONE_0'),s(bool,X23),s(fun(X22,bool),X24)))) ) | ($spl1 | $spl3 | $spl32)),
  inference(superposition,[],[f9653,f175])).
fof(f9776,plain,(
  ( ! [X56,X55] : (s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X55)),bool)) | s(bool,'T_0') = s(bool,X56)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f2072])).
fof(f9775,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X50),X51),s(bool,X52),s(fun(X50,bool),X53))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f178])).
fof(f9774,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X49) != s(bool,happ(s(fun(X45,bool),X48),s(X45,X46))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X46))),s(bool,X47),s(fun(X45,bool),X48))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f177])).
fof(f9797,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f9763,f217])).
fof(f9763,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f239])).
fof(f9755,plain,(
  ( ! [X54,X53] : (s(bool,X53) != s(bool,X54) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X53)),bool))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f2072])).
fof(f9754,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X52) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X48),X49),s(bool,X50),s(fun(X48,bool),X51))) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X48,bool),bool),happ(s(fun(bool,fun(fun(X48,bool),bool)),happ(s(fun('type.option.option'(X48),fun(bool,fun(fun(X48,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X48),X49))),s(bool,X50))),s(fun(X48,bool),X51)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f178])).
fof(f9753,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X43,bool),X46),s(X43,X44))) | s(bool,'F_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X43),'const.option.SOME_1'(s(X43,X44))),s(bool,X45),s(fun(X43,bool),X46)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f177])).
fof(f9752,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X40) != s(bool,X42) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X39),'const.option.NONE_0'),s(bool,X40),s(fun(X39,bool),X41)))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9633,f175])).
fof(f9653,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f9640,f560])).
fof(f9640,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | ($spl3 | $spl32)),
  inference(superposition,[],[f6492,f8665])).
fof(f9650,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X31),X32),s(bool,X33),s(fun(X31,bool),X34))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X31,bool),bool),happ(s(fun(bool,fun(fun(X31,bool),bool)),happ(s(fun('type.option.option'(X31),fun(bool,fun(fun(X31,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X31),X32))),s(bool,X33))),s(fun(X31,bool),X34)))) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f9624,f558])).
fof(f9649,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'F_0') != s(bool,happ(s(fun(X27,bool),X30),s(X27,X28))) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X27),'const.option.SOME_1'(s(X27,X28))),s(bool,X29),s(fun(X27,bool),X30)))) ) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f9623,f558])).
fof(f9323,plain,(
  ( ! [X56] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X56,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0')))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f9290,f8657])).
fof(f9290,plain,(
  ( ! [X56] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X56),s(bool,'F_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X56,bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f8657])).
fof(f9322,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X55,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f9289,f8657])).
fof(f9289,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X55,bool)) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f948,f8665])).
fof(f9174,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X52,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0')))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f9100,f8657])).
fof(f9100,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X52,bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f8657])).
fof(f9173,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X51,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X51),s(bool,X50))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f9099,f8657])).
fof(f9099,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X51),s(bool,X50))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X51,bool)) | s(bool,'T_0') = s(bool,X50)) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f8665])).
fof(f9058,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X41,bool)) | p(s(bool,happ(s(fun(bool,bool),X41),s(bool,X40)))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f8989,f8657])).
fof(f8989,plain,(
  ( ! [X41,X40] : (p(s(bool,happ(s(fun(bool,bool),X41),s(bool,X40)))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X41,bool)) | s(bool,'T_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f164,f8665])).
fof(f8827,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool)) | p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0'))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f8810,f8657])).
fof(f8810,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'F_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f164,f8657])).
fof(f8651,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f213,f2301])).
fof(f947,plain,(
  ( ! [X23] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X23,sK0)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK3)))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f905,f875])).
fof(f905,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK3))) | s(sK0,sK5(sK3,sK2,sK0)) = s(sK0,sK5(sK5(sK3,sK2,sK0),X23,sK0))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f875])).
fof(f945,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X13,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f902,f510])).
fof(f902,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X13,bool))) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f510])).
fof(f944,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X12,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X11))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f901,f510])).
fof(f901,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X11))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X12,bool)) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl3)),
  inference(superposition,[],[f877,f586])).
fof(f12504,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,X87) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl3 | $spl22 | $spl32 | $spl34)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f511,f211,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f944,f945,f947,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1586,f1613,f797,f775,f756,f640,f560,f2072,f2092,f2156,f2219,f2291,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2454,f2463,f2466,f2536,f2538,f2539,f2559,f2989,f2990,f2612,f2614,f2622,f2631,f2634,f3041,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3440,f3366,f3323,f3322,f3319,f3317,f3315,f3281,f3245,f3189,f3174,f3106,f3105,f3097,f3095,f3088,f3074,f3072,f3066,f3060,f2623,f2455,f2310,f2309,f2306,f2304,f2302,f2301,f2292,f2274,f2220,f2157,f2116,f1538,f1304,f1302,f969,f942,f940,f630,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4224,f6493,f6492,f3083,f3082,f3188,f2991,f1551,f1372,f948,f929,f877,f1288,f898,f2567,f2568,f1571,f1480,f1363,f1293,f968,f966,f3076,f3077,f3084,f3085,f3099,f3100,f3107,f3108,f1649,f1462,f2191,f1580,f1461,f1370,f1358,f1303,f943,f867,f864,f3329,f3328,f3327,f3325,f3294,f2625,f2624,f2616,f2316,f2315,f2314,f2312,f2286,f1362,f3326,f2633,f2617,f2570,f2569,f2562,f2561,f2465,f2457,f2456,f2449,f2313,f1565,f1491,f1474,f1463,f1359,f1357,f1320,f866,f853,f1650,f1683,f1653,f1654,f1655,f1656,f1427,f1450,f1429,f1430,f1431,f1560,f1437,f1438,f1439,f1413,f1664,f1666,f1667,f1624,f3491,f3467,f3441,f3246,f3061,f2997,f2996,f2379,f2342,f2307,f2234,f2217,f2198,f2197,f2196,f2195,f2192,f2171,f1570,f1559,f1531,f1498,f1479,f1434,f1354,f1353,f1321,f1319,f1301,f1292,f1284,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f3492,f3468,f3442,f3247,f3062,f2380,f2343,f1670,f1669,f1661,f1630,f1629,f1628,f1626,f1609,f3320,f2627,f2459,f1539,f213,f558,f2537,f8651,f8657,f8645,f8665,f8806,f8827,f8908,f8946,f8985,f9058,f9173,f9174,f9322,f9323,f9487,f9518,f9519,f9520,f9521,f9522,f9607,f9622,f9649,f9650,f9625,f9626,f9633,f9653,f9726,f9752,f9753,f9754,f9755,f9797,f9773,f9774,f9775,f9776,f9836,f9820,f9868,f9869,f9870,f9871,f9940,f10024,f10037,f9992,f10043,f10047,f10049,f10108,f10115,f10116,f10119,f10093,f10094,f10095,f10096,f10121,f10122,f10123,f10124,f10310,f10290,f10328,f10279,f10280,f10281,f10371,f10372,f10373,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10361,f10362,f10363,f10364,f10367,f10504,f10510,f10512,f10513,f10475,f10578,f10579,f10580,f10585,f10587,f10595,f10597,f10598,f10599,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10626,f10627,f10628,f10629,f10630,f10703,f10747,f10748,f10848,f10856,f10817,f10859,f10860,f10820,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10887,f10889,f10890,f10891,f10924,f10926,f10905,f10939,f10960,f10965,f10966,f11115,f11058,f11116,f11062,f11171,f11172,f11173,f11128,f11129,f11134,f11176,f11177,f11179,f11180,f11181,f11183,f11184,f11145,f11185,f11186,f11148,f11149,f11150,f11189,f11190,f11192,f11165,f11205,f11226,f11230,f11232,f11233,f11256,f11587,f11589,f11607,f11636,f11641,f11373,f11649,f11654,f11655,f11656,f11657,f11664,f11669,f11451,f11707,f11708,f11717,f11720,f11728,f11736,f11517,f11737,f11743,f11748,f11750,f11754,f11756,f11534,f11759,f12031,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f11926,f11934,f11936,f11945,f11946,f11951,f12022,f11740])).
fof(f12464,plain,(
  $false | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f12463,f11739])).
fof(f12463,plain,(
  ( ! [X99,X97,X98] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(fun(bool,bool),X97),s(fun(X98,fun(bool,bool)),X99)),bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f12257,f128])).
fof(f12257,plain,(
  ( ! [X99,X97,X98] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),'const.option.option_CASE_3'(s('type.option.option'(X98),'const.option.NONE_0'),s(fun(bool,bool),X97),s(fun(X98,fun(bool,bool)),X99)),bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(backward_demodulation,[],[f12073,f11182])).
fof(f12073,plain,(
  ( ! [X352] : (s(bool,'F_0') = s(bool,X352)) ) | ($spl1 | $spl32 | $spl70)),
  inference(global_subsumption,[],[f12072,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10362,f10367,f10578,f10579,f10580,f10585,f10587,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10628,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10889,f10905,f11171,f11172,f11173,f11128,f11129,f11134,f11181,f11148,f11165,f11739,f11759,f12031,f12034,f12036,f12048,f12049,f11775,f12050,f11777,f12051,f11839,f11841,f11843,f11844,f11850,f11851,f11852,f11853,f11856,f11859,f11862,f12054,f12056])).
fof(f12056,plain,(
  ( ! [X259,X260,X262,X258] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X262))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X258,bool),X259),s(X258,sK4(X260,X258)))))) | s('type.option.option'(X258),'const.option.NONE_0') = s('type.option.option'(X258),X260)) ) | $spl70),
  inference(subsumption_resolution,[],[f12055,f11739])).
fof(f12055,plain,(
  ( ! [X259,X260,X262,X258] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X262))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X258,bool),X259),s(X258,sK4(X260,X258)))))) | s(bool,'T_0') = s(bool,X262) | s('type.option.option'(X258),'const.option.NONE_0') = s('type.option.option'(X258),X260)) ) | $spl70),
  inference(subsumption_resolution,[],[f11882,f12034])).
fof(f12054,plain,(
  ( ! [X187,X185,X186,X184] : (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X184),X186),s(bool,X187),s(fun(X184,bool),X185)))) | s('type.option.option'(X184),'const.option.NONE_0') = s('type.option.option'(X184),X186)) ) | $spl70),
  inference(subsumption_resolution,[],[f11867,f12036])).
fof(f12036,plain,(
  ( ! [X23,X21,X25,X22] : (s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X23) | s(bool,X25) = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK4(X23,X21))))) ) | $spl70),
  inference(subsumption_resolution,[],[f12035,f11739])).
fof(f12035,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X21),X23),s(bool,X24),s(fun(X21,bool),X22))) | s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X23) | s(bool,X25) = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK4(X23,X21))))) ) | $spl70),
  inference(subsumption_resolution,[],[f11763,f11739])).
fof(f11763,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X21),X23),s(bool,X24),s(fun(X21,bool),X22))) | s('type.option.option'(X21),'const.option.NONE_0') = s('type.option.option'(X21),X23) | s(bool,X25) = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK4(X23,X21)))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1387,f239])).
fof(f12034,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X17),X19),s(bool,X20),s(fun(X17,bool),X18))) | s('type.option.option'(X17),'const.option.NONE_0') = s('type.option.option'(X17),X19)) ) | $spl70),
  inference(subsumption_resolution,[],[f11762,f11739])).
fof(f12072,plain,(
  ( ! [X350,X352,X349,X351,X348] : (s(bool,X352) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X348),X350),s(bool,X351),s(fun(X348,bool),X349))) | s(bool,'F_0') = s(bool,X352) | s('type.option.option'(X348),'const.option.NONE_0') = s('type.option.option'(X348),X350)) ) | ($spl1 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f11902,f11739])).
fof(f11902,plain,(
  ( ! [X350,X352,X349,X351,X348] : (s(bool,X352) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X348),X350),s(bool,X351),s(fun(X348,bool),X349))) | s(bool,'F_0') = s(bool,X352) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X348,bool),X349),s(X348,sK4(X350,X348))))),bool)) | s('type.option.option'(X348),'const.option.NONE_0') = s('type.option.option'(X348),X350)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f1387])).
fof(f12462,plain,(
  $false | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f12374,f11739])).
fof(f12374,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X32,bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(trivial_inequality_removal,[],[f12239])).
fof(f12239,plain,(
  ( ! [X32] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X32,bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(backward_demodulation,[],[f12073,f1832])).
fof(f12461,plain,(
  $false | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f12460,f11739])).
fof(f12460,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X17,bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(subsumption_resolution,[],[f12237,f128])).
fof(f12237,plain,(
  ( ! [X17] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))),X17,bool))) ) | ($spl1 | $spl20 | $spl32 | $spl70)),
  inference(backward_demodulation,[],[f12073,f1829])).
fof(f11741,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,X87) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl3 | $spl26)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f4429,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f213,f8645,f8665,f8908,f8946,f8985,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10362,f10367,f10578,f10579,f10580,f10585,f10587,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10628,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10889,f10924,f10905,f11171,f11172,f11173,f11128,f11129,f11134,f11181,f11148,f11190,f11165,f11256,f11373,f11451,f11720,f11736,f11517,f11737,f11738,f11740])).
fof(f11738,plain,(
  ( ! [X85,X84] : (s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') != s(bool,X85)) ) | ($spl3 | $spl26)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f4429,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f213,f8645,f8665,f8908,f8946,f8985,f9518,f9519,f9520,f9521,f9522,f9868,f9869,f9870,f9871,f10121,f10122,f10123,f10124,f10371,f10374,f10375,f10376,f10377,f10340,f10341,f10342,f10343,f10345,f10346,f10351,f10357,f10362,f10367,f10578,f10579,f10580,f10585,f10587,f10597,f10602,f10605,f10606,f10607,f10608,f10610,f10611,f10616,f10622,f10628,f10861,f10862,f10863,f10864,f10870,f10871,f10876,f10878,f10889,f10924,f10905,f11171,f11172,f11173,f11128,f11129,f11134,f11181,f11148,f11190,f11165,f11256,f11373,f11451,f11720,f11736,f11517,f11737])).
fof(f874,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f874_D])).
fof(f874_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f11190,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X3,fun(sK0,bool)),X4)),sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f11155,f131])).
fof(f11155,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | s(sK0,sK3) = s(sK0,sK5(sK3,'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(fun(sK0,bool),sK2),s(fun(X3,fun(sK0,bool)),X4)),sK0)) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f849,f8665])).
fof(f10924,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(sK0,sK3),s(fun(X3,sK0),X4)),sK2,sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f10895,f131])).
fof(f10895,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | s(sK0,sK3) = s(sK0,sK5('const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.NONE_0'),s(sK0,sK3),s(fun(X3,sK0),X4)),sK2,sK0)) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f870,f8665])).
fof(f10109,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))),bool)) | s(X39,X41) = s(X39,sK5(X41,X40,X39))) ) | $spl64),
  inference(cnf_transformation,[],[f10109_D])).
fof(f10109_D,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))),bool)) | s(X39,X41) = s(X39,sK5(X41,X40,X39))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f10108,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))),bool)) | s(X39,X41) = s(X39,sK5(X41,X40,X39))) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f10070,f129])).
fof(f10070,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') != s(bool,X42) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))),bool)) | s(X39,X41) = s(X39,sK5(X41,X40,X39))) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f9755,f877])).
fof(f8635,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093,f7094,f213])).
fof(f7094,plain,(
  s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))),bool)) | $spl30),
  inference(unit_resulting_resolution,[],[f2072,f4515])).
fof(f7093,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X0,bool),X1)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f175,f4515])).
fof(f7089,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(unit_resulting_resolution,[],[f217,f4515])).
fof(f7136,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f7095])).
fof(f7095,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f4515,f129])).
fof(f7111,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X26),'const.option.NONE_0'),s(bool,X27),s(fun(X26,bool),X28)))) ) | $spl30),
  inference(superposition,[],[f4515,f175])).
fof(f7112,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) != s(bool,happ(s(fun(X29,bool),X32),s(X29,X30))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X29),'const.option.SOME_1'(s(X29,X30))),s(bool,X31),s(fun(X29,bool),X32)))) ) | $spl30),
  inference(superposition,[],[f4515,f177])).
fof(f7113,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X33),X34),s(bool,X35),s(fun(X33,bool),X36))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X33,bool),bool),happ(s(fun(bool,fun(fun(X33,bool),bool)),happ(s(fun('type.option.option'(X33),fun(bool,fun(fun(X33,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X33),X34))),s(bool,X35))),s(fun(X33,bool),X36)))) ) | $spl30),
  inference(superposition,[],[f4515,f178])).
fof(f7157,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(condensation,[],[f7156])).
fof(f7156,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl30),
  inference(subsumption_resolution,[],[f7117,f129])).
fof(f7117,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl30),
  inference(superposition,[],[f4515,f129])).
fof(f7160,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(condensation,[],[f7159])).
fof(f7159,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl30),
  inference(condensation,[],[f7158])).
fof(f7158,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl30),
  inference(subsumption_resolution,[],[f7119,f239])).
fof(f7119,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl30),
  inference(superposition,[],[f4515,f239])).
fof(f7163,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(condensation,[],[f7162])).
fof(f7162,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(condensation,[],[f7161])).
fof(f7161,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X8)) ) | $spl30),
  inference(subsumption_resolution,[],[f7120,f239])).
fof(f7120,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X8)) ) | $spl30),
  inference(superposition,[],[f4515,f239])).
fof(f7127,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(equality_resolution,[],[f4515])).
fof(f7399,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | $spl30),
  inference(forward_demodulation,[],[f7303,f177])).
fof(f7303,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2)))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f7092,f163])).
fof(f7400,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | $spl30),
  inference(forward_demodulation,[],[f7302,f177])).
fof(f7302,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2)))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f7092,f163])).
fof(f7401,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | $spl30),
  inference(forward_demodulation,[],[f7301,f177])).
fof(f7301,plain,(
  ( ! [X0] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f126,f7092,f204])).
fof(f7403,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | $spl30),
  inference(forward_demodulation,[],[f7299,f177])).
fof(f7299,plain,(
  ( ! [X0] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f7092,f131])).
fof(f7339,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(superposition,[],[f7092,f177])).
fof(f7406,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | $spl30),
  inference(forward_demodulation,[],[f7382,f177])).
fof(f7382,plain,(
  ( ! [X1] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X1),s(fun(sK0,bool),sK2))))) ) | $spl30),
  inference(trivial_inequality_removal,[],[f7341])).
fof(f7341,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X1),s(fun(sK0,bool),sK2))))) ) | $spl30),
  inference(superposition,[],[f131,f7092])).
fof(f7411,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | ~p(s(bool,X10))) ) | $spl30),
  inference(forward_demodulation,[],[f7410,f177])).
fof(f7410,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X9),s(fun(sK0,bool),sK2)))) ) | $spl30),
  inference(subsumption_resolution,[],[f7347,f126])).
fof(f7347,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X9),s(fun(sK0,bool),sK2)))) ) | $spl30),
  inference(superposition,[],[f170,f7092])).
fof(f7414,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | p(s(bool,X12))) ) | $spl30),
  inference(forward_demodulation,[],[f7413,f177])).
fof(f7413,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X11),s(fun(sK0,bool),sK2))) | p(s(bool,X12))) ) | $spl30),
  inference(subsumption_resolution,[],[f7348,f126])).
fof(f7348,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X12) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X11),s(fun(sK0,bool),sK2))) | p(s(bool,X12))) ) | $spl30),
  inference(superposition,[],[f204,f7092])).
fof(f7364,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl30),
  inference(superposition,[],[f177,f7092])).
fof(f7419,plain,(
  ( ! [X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') != s(bool,X49)) ) | $spl30),
  inference(forward_demodulation,[],[f7369,f177])).
fof(f7369,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X49))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X48),s(fun(sK0,bool),sK2)))))) ) | $spl30),
  inference(superposition,[],[f163,f7092])).
fof(f7485,plain,(
  ( ! [X90,X91] : (s(bool,X91) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | ~p(s(bool,happ(s(fun(bool,bool),X90),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X90),s(bool,X91))))) ) | $spl30),
  inference(forward_demodulation,[],[f7449,f177])).
fof(f7449,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X90),s(bool,'T_0')))) | s(bool,X91) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X89),s(fun(sK0,bool),sK2))) | p(s(bool,happ(s(fun(bool,bool),X90),s(bool,X91))))) ) | $spl30),
  inference(superposition,[],[f166,f7092])).
fof(f8454,plain,(
  ( ! [X70,X69] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X69),s(fun(sK0,bool),sK2)),X70,bool)) | p(s(bool,happ(s(fun(bool,bool),X70),s(bool,'T_0'))))) ) | $spl30),
  inference(forward_demodulation,[],[f8418,f177])).
fof(f8418,plain,(
  ( ! [X70,X69] : (p(s(bool,happ(s(fun(bool,bool),X70),s(bool,'T_0')))) | s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X69),s(fun(sK0,bool),sK2))) = s(bool,sK5('const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X69),s(fun(sK0,bool),sK2)),X70,bool))) ) | $spl30),
  inference(superposition,[],[f164,f7092])).
fof(f7366,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,X43) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X42),s(fun(sK0,bool),sK2)))) ) | $spl30),
  inference(superposition,[],[f162,f7092])).
fof(f7315,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X19),s(fun(sK0,bool),sK2))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl30),
  inference(superposition,[],[f7092,f239])).
fof(f7114,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X37)),bool))) ) | $spl30),
  inference(superposition,[],[f4515,f2072])).
fof(f7092,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))),s(bool,X0),s(fun(sK0,bool),sK2)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f177,f4515])).
fof(f8634,plain,(
  $false | ($spl1 | $spl3 | $spl30)),
  inference(global_subsumption,[],[f558,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093,f7094])).
fof(f8633,plain,(
  $false | ($spl1 | $spl3 | $spl30)),
  inference(global_subsumption,[],[f2537,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093,f7094])).
fof(f8632,plain,(
  $false | ($spl3 | $spl30 | $spl36)),
  inference(global_subsumption,[],[f6504,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093,f7094])).
fof(f8631,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093,f7094])).
fof(f8630,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089,f7093])).
fof(f8629,plain,(
  $false | ($spl1 | $spl3 | $spl30)),
  inference(global_subsumption,[],[f7139,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089])).
fof(f7139,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f7090,f558])).
fof(f7090,plain,(
  s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl30)),
  inference(unit_resulting_resolution,[],[f560,f4515])).
fof(f8628,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136,f7089])).
fof(f8627,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111,f7136])).
fof(f8626,plain,(
  $false | ($spl1 | $spl3 | $spl30)),
  inference(global_subsumption,[],[f7143,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111])).
fof(f7143,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl30)),
  inference(forward_demodulation,[],[f7131,f558])).
fof(f7131,plain,(
  ( ! [X9] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl30)),
  inference(duplicate_literal_removal,[],[f7100])).
fof(f7100,plain,(
  ( ! [X9] : (s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl30)),
  inference(superposition,[],[f4515,f640])).
fof(f8625,plain,(
  $false | ($spl1 | $spl3 | $spl30)),
  inference(global_subsumption,[],[f7151,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111])).
fof(f7151,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | ($spl1 | $spl30)),
  inference(subsumption_resolution,[],[f7106,f560])).
fof(f7106,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl30)),
  inference(superposition,[],[f4515,f558])).
fof(f8624,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112,f7111])).
fof(f8623,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113,f7112])).
fof(f8622,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157,f7113])).
fof(f8621,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160,f7157])).
fof(f8620,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163,f7160])).
fof(f8619,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127,f7163])).
fof(f8618,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399,f7127])).
fof(f8616,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400,f7399])).
fof(f8615,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401,f7400])).
fof(f8614,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403,f7401])).
fof(f8613,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339,f7403])).
fof(f8612,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406,f7339])).
fof(f8611,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411,f7406])).
fof(f8610,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414,f7411])).
fof(f8609,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364,f7414])).
fof(f8608,plain,(
  $false | ($spl3 | $spl30)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f4515,f7092,f7114,f7315,f7366,f1480,f966,f3325,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f1650,f1683,f1450,f1430,f1437,f1664,f3491,f3467,f3246,f2997,f2996,f2379,f2342,f2234,f2217,f2192,f1319,f870,f849,f166,f165,f164,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f2627,f2459,f8454,f7485,f7419,f7364])).
fof(f8607,plain,(
  $false | ($spl1 | $spl62)),
  inference(subsumption_resolution,[],[f8605,f217])).
fof(f8605,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl62)),
  inference(duplicate_literal_removal,[],[f8604])).
fof(f8604,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl62)),
  inference(condensation,[],[f8603])).
fof(f8603,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl62)),
  inference(backward_demodulation,[],[f8583,f2537])).
fof(f8606,plain,(
  $false | ($spl1 | $spl62)),
  inference(subsumption_resolution,[],[f8602,f217])).
fof(f8602,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl62)),
  inference(backward_demodulation,[],[f8583,f558])).
fof(f8583,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X0),s(fun(sK0,bool),sK2)))) ) | $spl62),
  inference(cnf_transformation,[],[f8583_D])).
fof(f8583_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X0),s(fun(sK0,bool),sK2)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8474,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,X0),s(fun(sK0,bool),sK2))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1)) ) | $spl60),
  inference(superposition,[],[f8315,f154])).
fof(f8315,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK4(sK1,sK0)))),s(bool,X0),s(fun(sK0,bool),sK2)))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f177,f7720])).
fof(f8248,plain,(
  $false | $spl58),
  inference(subsumption_resolution,[],[f8160,f8146])).
fof(f8146,plain,(
  ( ! [X54] : (~p(s(bool,X54))) ) | $spl58),
  inference(subsumption_resolution,[],[f8107,f8142])).
fof(f8142,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,sK5(X5,X6,bool))) ) | $spl58),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f3491,f3467,f3246,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2627,f2459,f3325,f1480,f966,f1450,f1430,f1650,f1683,f165,f2997,f1319,f2234,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f7719,f8017,f8028,f8033,f8037,f8038,f8081,f8141])).
fof(f8141,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X7) | s(bool,X7) = s(bool,sK5(X5,X6,bool))) ) | $spl58),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f3491,f3467,f3246,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2627,f2459,f3325,f1480,f966,f1450,f1430,f1650,f1683,f165,f2997,f1319,f2234,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f7719,f8017,f8028,f8033,f8037,f8038,f8081])).
fof(f8081,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X7) = s(bool,sK5(X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) ) | $spl58),
  inference(superposition,[],[f8017,f239])).
fof(f8038,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,X99)))) ) | $spl58),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f3494,f3493,f3490,f3470,f3469,f3466,f3323,f3319,f3317,f3315,f3281,f3245,f2157,f445,f441,f414,f406,f398,f303,f136,f135,f128,f2616,f3326,f2633,f2617,f2562,f2561,f2465,f2449,f1474,f853,f3491,f3467,f3246,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2627,f2459,f3325,f1480,f966,f1450,f1430,f1650,f1683,f165,f2997,f1319,f2234,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f7719,f8017,f8028,f8033,f8037])).
fof(f8037,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,X99))) | s(bool,'T_0') = s(bool,X100)) ) | $spl58),
  inference(forward_demodulation,[],[f8036,f8017])).
fof(f8036,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK5(X99,X98,X97)))) | s(bool,'T_0') = s(bool,X100)) ) | $spl58),
  inference(subsumption_resolution,[],[f7818,f128])).
fof(f7818,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'F_0')) | s(bool,X100) = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK5(X99,X98,X97)))) | s(bool,'T_0') = s(bool,X100)) ) | $spl58),
  inference(backward_demodulation,[],[f7719,f2997])).
fof(f8033,plain,(
  ( ! [X17,X18,X16] : (s(X16,X17) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16))) ) | $spl58),
  inference(subsumption_resolution,[],[f7812,f128])).
fof(f7812,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(X16,X17) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16))) ) | $spl58),
  inference(backward_demodulation,[],[f7719,f2217])).
fof(f8028,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(X14,X16) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) ) | $spl58),
  inference(subsumption_resolution,[],[f7785,f128])).
fof(f7785,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'F_0')) | s(X14,X16) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) ) | $spl58),
  inference(backward_demodulation,[],[f7719,f870])).
fof(f8017,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | $spl58),
  inference(subsumption_resolution,[],[f7781,f128])).
fof(f7781,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | $spl58),
  inference(backward_demodulation,[],[f7719,f164])).
fof(f8107,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | s(bool,'F_0') != s(bool,sK5(X54,X55,bool))) ) | $spl58),
  inference(superposition,[],[f135,f8017])).
fof(f8160,plain,(
  p(s(bool,'F_0')) | $spl58),
  inference(unit_resulting_resolution,[],[f217,f8146,f167])).
fof(f8247,plain,(
  $false | ($spl1 | $spl58)),
  inference(subsumption_resolution,[],[f8162,f8146])).
fof(f8162,plain,(
  p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f560,f8146,f167])).
fof(f8246,plain,(
  $false | $spl58),
  inference(subsumption_resolution,[],[f8170,f8146])).
fof(f8170,plain,(
  p(s(bool,'F_0')) | $spl58),
  inference(unit_resulting_resolution,[],[f217,f8146,f167])).
fof(f8245,plain,(
  $false | ($spl1 | $spl58)),
  inference(subsumption_resolution,[],[f8172,f8146])).
fof(f8172,plain,(
  p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f560,f8146,f167])).
fof(f8206,plain,(
  $false | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f560,f560,f8146,f272])).
fof(f8204,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f217,f217,f8146,f272])).
fof(f8189,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f126,f8017,f8146,f204])).
fof(f8187,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f126,f2072,f8146,f204])).
fof(f8184,plain,(
  $false | ($spl24 | $spl58)),
  inference(unit_resulting_resolution,[],[f126,f5435,f8146,f204])).
fof(f5435,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl24),
  inference(unit_resulting_resolution,[],[f217,f4239])).
fof(f8182,plain,(
  $false | ($spl20 | $spl58)),
  inference(unit_resulting_resolution,[],[f126,f1762,f8146,f204])).
fof(f8179,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f126,f175,f8146,f204])).
fof(f8177,plain,(
  $false | ($spl30 | $spl58)),
  inference(unit_resulting_resolution,[],[f126,f7092,f8146,f204])).
fof(f8173,plain,(
  $false | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f8146,f560,f8146,f167])).
fof(f8171,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f128,f217,f8146,f167])).
fof(f8169,plain,(
  $false | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f8146,f560,f8146,f167])).
fof(f8167,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f8146,f217,f8146,f167])).
fof(f8163,plain,(
  $false | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f8146,f560,f8146,f167])).
fof(f8161,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f128,f217,f8146,f167])).
fof(f8159,plain,(
  $false | ($spl1 | $spl58)),
  inference(unit_resulting_resolution,[],[f8146,f560,f8146,f167])).
fof(f8157,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f8146,f217,f8146,f167])).
fof(f8153,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f217,f8146,f136])).
fof(f8151,plain,(
  $false | ($spl30 | $spl58)),
  inference(unit_resulting_resolution,[],[f7092,f8146,f131])).
fof(f8149,plain,(
  $false | $spl58),
  inference(unit_resulting_resolution,[],[f126,f8146])).
fof(f7720,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) | $spl60),
  inference(cnf_transformation,[],[f7720_D])).
fof(f7720_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7667,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X39)) ) | ($spl20 | $spl24 | $spl30)),
  inference(duplicate_literal_removal,[],[f7587])).
fof(f7587,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) ) | ($spl20 | $spl24 | $spl30)),
  inference(superposition,[],[f7116,f6271])).
fof(f6271,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f5744,f217])).
fof(f5744,plain,(
  ( ! [X53] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0)))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl20 | $spl24)),
  inference(superposition,[],[f5454,f1762])).
fof(f5454,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f4239,f4239])).
fof(f7116,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl24 | $spl30)),
  inference(superposition,[],[f4515,f4239])).
fof(f6817,plain,(
  $false | ($spl3 | $spl24 | $spl30)),
  inference(global_subsumption,[],[f6816,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f4239,f2562,f2561,f2465,f2449,f2633,f2617,f2616,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2459,f2627,f3494,f3493,f3491,f3490,f3470,f3469,f3467,f3466,f3326,f3323,f3319,f3317,f3315,f3281,f3246,f3245,f2157,f1474,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3325,f1480,f966,f1650,f1683,f1450,f1430,f2234,f1319,f165,f2997,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f5453,f5454,f5504,f5675,f5676,f5677,f5678,f5712,f5746,f5747,f5748,f5749,f5837,f5839,f5841,f5849,f5853,f5855,f5926,f5949,f5951,f6289,f6307,f5996,f5998,f6000,f6008,f6012,f6014,f6344,f6345,f6383,f6382,f6360,f6506,f6343,f6507,f5561,f6520,f6532,f6534,f6535,f6633,f6645,f5552,f6651,f6652,f6653,f6792,f6802,f6805,f6814])).
fof(f6814,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f5605,f6652])).
fof(f5605,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) ) | $spl24),
  inference(condensation,[],[f5604])).
fof(f5604,plain,(
  ( ! [X21,X22] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl24),
  inference(subsumption_resolution,[],[f5482,f2989])).
fof(f5482,plain,(
  ( ! [X21,X22] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) ) | $spl24),
  inference(superposition,[],[f272,f4239])).
fof(f6805,plain,(
  ( ! [X17,X18,X16] : (s(X16,X17) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6688,f6652])).
fof(f6688,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(X16,X17) = s(X16,sK5(sK4('const.option.SOME_1'(s(X16,X17)),X16),X18,X16))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6653,f2217])).
fof(f6802,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(X14,X16) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6663,f6652])).
fof(f6663,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'F_0')) | s(X14,X16) = s(X14,sK5('const.option.option_CASE_3'(s('type.option.option'(X15),'const.option.NONE_0'),s(X14,X16),s(fun(X15,X14),X17)),X18,X14))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6653,f870])).
fof(f6792,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6660,f6652])).
fof(f6660,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6653,f164])).
fof(f6653,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10)) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f5621,f6652])).
fof(f5621,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,'F_0') = s(bool,X10)) ) | $spl24),
  inference(subsumption_resolution,[],[f5528,f129])).
fof(f5528,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(bool,'F_0') = s(bool,X10)) ) | $spl24),
  inference(superposition,[],[f131,f4239])).
fof(f6652,plain,(
  ( ! [X12] : (~p(s(bool,X12))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f5623,f6520])).
fof(f5623,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl24),
  inference(subsumption_resolution,[],[f5530,f135])).
fof(f5530,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X12)) ) | $spl24),
  inference(superposition,[],[f135,f4239])).
fof(f5552,plain,(
  ( ! [X52,X51] : (s(bool,X51) != s(bool,X52) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'F_0') = s(bool,X51)) ) | $spl24),
  inference(superposition,[],[f163,f4239])).
fof(f6645,plain,(
  ( ! [X26,X24,X25] : (s('type.option.option'(X24),'const.option.SOME_1'(s(X24,sK4(X25,X24)))) = s('type.option.option'(X24),sK5('const.option.SOME_1'(s(X24,sK4(X25,X24))),X26,'type.option.option'(X24))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X25)) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6573,f6506])).
fof(f6573,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s('type.option.option'(X24),'const.option.SOME_1'(s(X24,sK4(X25,X24)))) = s('type.option.option'(X24),sK5('const.option.SOME_1'(s(X24,sK4(X25,X24))),X26,'type.option.option'(X24))) | s('type.option.option'(X24),'const.option.NONE_0') = s('type.option.option'(X24),X25)) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f966])).
fof(f6633,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (s(X135,X140) = s(X135,sK5(X140,'const.option.option_CASE_3'(s('type.option.option'(X136),'const.option.SOME_1'(s(X136,X137))),s(fun(X135,bool),X138),s(fun(X136,fun(X135,bool)),X139)),X135))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6558,f6506])).
fof(f6558,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s(X135,X140) = s(X135,sK5(X140,'const.option.option_CASE_3'(s('type.option.option'(X136),'const.option.SOME_1'(s(X136,X137))),s(fun(X135,bool),X138),s(fun(X136,fun(X135,bool)),X139)),X135))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f1430])).
fof(f6520,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl3 | $spl24)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f4239,f2562,f2561,f2465,f2449,f2633,f2617,f2616,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2459,f2627,f3494,f3493,f3491,f3490,f3470,f3469,f3467,f3466,f3326,f3323,f3319,f3317,f3315,f3281,f3246,f3245,f2157,f1474,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3325,f1480,f966,f1650,f1683,f1450,f1430,f2234,f1319,f165,f2997,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f5453,f5454,f5504,f5675,f5676,f5677,f5678,f5712,f5746,f5747,f5748,f5749,f5837,f5839,f5841,f5849,f5853,f5855,f5926,f5949,f5951,f6289,f6307,f5996,f5998,f6000,f6008,f6012,f6014,f6344,f6345,f6383,f6382,f6360,f6506,f6343,f6507,f5561])).
fof(f5561,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl24),
  inference(equality_factoring,[],[f4239])).
fof(f6507,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl24),
  inference(subsumption_resolution,[],[f6062,f128])).
fof(f6062,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'F_0')) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl24),
  inference(superposition,[],[f272,f5454])).
fof(f6343,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | s(bool,'F_0') = s(bool,sK5(X28,X29,bool))) ) | $spl24),
  inference(superposition,[],[f165,f4239])).
fof(f6506,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | ($spl3 | $spl24)),
  inference(global_subsumption,[],[f6360,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f2562,f2561,f2465,f2449,f2633,f2617,f2616,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2459,f2627,f3494,f3493,f3491,f3490,f3470,f3469,f3467,f3466,f3326,f3323,f3319,f3317,f3315,f3281,f3246,f3245,f2157,f1474,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3325,f1480,f966,f1650,f1683,f1450,f1430,f2234,f1319,f165,f2997,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630])).
fof(f6360,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | p(s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK5(X43,X42,X41))))) ) | $spl24),
  inference(superposition,[],[f165,f4239])).
fof(f6345,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X33)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK5(X33,X34,bool))) ) | $spl24),
  inference(superposition,[],[f165,f5454])).
fof(f6344,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X31),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X30)))) | s(bool,X32) = s(bool,sK5(X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl24),
  inference(superposition,[],[f165,f5454])).
fof(f6014,plain,(
  ( ! [X158,X156,X157,X155] : (s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0'))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(bool,X158) = s(bool,sK6(X155,X156,X157,bool)) | s(bool,'F_0') = s(bool,X158)) ) | $spl24),
  inference(superposition,[],[f176,f5454])).
fof(f6012,plain,(
  ( ! [X152,X151] : (s('type.option.option'(bool),X151) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X151) | s(bool,X152) = s(bool,sK4(X151,bool)) | s(bool,'F_0') = s(bool,X152)) ) | $spl24),
  inference(superposition,[],[f154,f5454])).
fof(f6008,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'F_0') != s(bool,f(s(X133,sK6(X134,X135,bool,X133)))) | s(fun(X133,bool),X134) = s(fun(X133,bool),X135) | s(bool,X136) = s(bool,g(s(X133,sK6(X134,X135,bool,X133)))) | s(bool,'F_0') = s(bool,X136)) ) | $spl24),
  inference(superposition,[],[f176,f5454])).
fof(f6000,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,happ(s(fun(bool,bool),X102),s(bool,'F_0')))) | s(bool,X101) = s(bool,sK5(X101,X102,bool)) | s(bool,X101) = s(bool,X103) | s(bool,'F_0') = s(bool,X103)) ) | $spl24),
  inference(superposition,[],[f164,f5454])).
fof(f5998,plain,(
  ( ! [X97,X95,X96] : (s(bool,'F_0') != s(bool,X96) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X96))) | s(bool,X95) = s(bool,X97) | s(bool,'F_0') = s(bool,X97)) ) | $spl24),
  inference(superposition,[],[f163,f5454])).
fof(f5996,plain,(
  ( ! [X90,X91,X89] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X90))) | s(bool,X89) = s(bool,X90) | s(bool,X89) = s(bool,X91) | s(bool,'F_0') = s(bool,X91)) ) | $spl24),
  inference(superposition,[],[f162,f5454])).
fof(f6307,plain,(
  ( ! [X80,X81,X79,X82] : (s(X79,X81) = s(X79,sK5(X81,X80,X79)) | s(bool,X82) = s(bool,happ(s(fun(X79,bool),X80),s(X79,X81))) | s(bool,'F_0') = s(bool,X82)) ) | $spl24),
  inference(subsumption_resolution,[],[f5993,f128])).
fof(f5993,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,'F_0')) | s(X79,X81) = s(X79,sK5(X81,X80,X79)) | s(bool,X82) = s(bool,happ(s(fun(X79,bool),X80),s(X79,X81))) | s(bool,'F_0') = s(bool,X82)) ) | $spl24),
  inference(superposition,[],[f164,f5454])).
fof(f6289,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl24),
  inference(subsumption_resolution,[],[f5963,f128])).
fof(f5963,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl24),
  inference(superposition,[],[f167,f5454])).
fof(f5951,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,happ(s(fun(bool,bool),X77),s(bool,X76)))) | s(bool,'F_0') = s(bool,sK5('F_0',X77,bool)) | s(bool,X76) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl24),
  inference(superposition,[],[f164,f5454])).
fof(f5949,plain,(
  ( ! [X70,X72,X71] : (s(bool,X70) != s(bool,X71) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X71))) | s(bool,X70) = s(bool,X72) | s(bool,'F_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f163,f5454])).
fof(f5926,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl24),
  inference(superposition,[],[f217,f5454])).
fof(f5855,plain,(
  ( ! [X180,X178,X181,X179] : (s(X180,f(s(bool,X181))) != s(X180,g(s(bool,X181))) | s(fun(bool,X180),X178) = s(fun(bool,X180),X179) | s(bool,'F_0') = s(bool,X181) | s(bool,'F_0') = s(bool,sK6(X178,X179,X180,bool))) ) | $spl24),
  inference(superposition,[],[f176,f5454])).
fof(f5853,plain,(
  ( ! [X175,X174] : (s('type.option.option'(bool),X174) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X175))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X174) | s(bool,'F_0') = s(bool,X175) | s(bool,'F_0') = s(bool,sK4(X174,bool))) ) | $spl24),
  inference(superposition,[],[f154,f5454])).
fof(f5849,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X159) != s(bool,f(s(X156,sK6(X157,X158,bool,X156)))) | s(fun(X156,bool),X157) = s(fun(X156,bool),X158) | s(bool,'F_0') = s(bool,X159) | s(bool,'F_0') = s(bool,g(s(X156,sK6(X157,X158,bool,X156))))) ) | $spl24),
  inference(superposition,[],[f176,f5454])).
fof(f5841,plain,(
  ( ! [X125,X123,X124] : (p(s(bool,happ(s(fun(bool,bool),X125),s(bool,X124)))) | s(bool,X123) = s(bool,sK5(X123,X125,bool)) | s(bool,'F_0') = s(bool,X124) | s(bool,'F_0') = s(bool,X123)) ) | $spl24),
  inference(superposition,[],[f164,f5454])).
fof(f5839,plain,(
  ( ! [X118,X119,X117] : (s(bool,X118) != s(bool,X119) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X117))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X119))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,X117)) ) | $spl24),
  inference(superposition,[],[f163,f5454])).
fof(f5837,plain,(
  ( ! [X111,X112,X113] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X112))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X113))) | s(bool,X111) = s(bool,X113) | s(bool,'F_0') = s(bool,X112) | s(bool,'F_0') = s(bool,X111)) ) | $spl24),
  inference(superposition,[],[f162,f5454])).
fof(f5749,plain,(
  ( ! [X72,X73] : (s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,sK4('const.option.SOME_1'(s(bool,X72)),bool)) | s(bool,'F_0') = s(bool,X73)) ) | $spl24),
  inference(superposition,[],[f5454,f2072])).
fof(f5748,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X67),X68),s(bool,X69),s(fun(X67,bool),X70))) | s(bool,X71) = s(bool,happ(s(fun(fun(X67,bool),bool),happ(s(fun(bool,fun(fun(X67,bool),bool)),happ(s(fun('type.option.option'(X67),fun(bool,fun(fun(X67,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X67),X68))),s(bool,X69))),s(fun(X67,bool),X70))) | s(bool,'F_0') = s(bool,X71)) ) | $spl24),
  inference(superposition,[],[f5454,f178])).
fof(f5747,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,happ(s(fun(X62,bool),X65),s(X62,X63))) | s(bool,X66) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X62),'const.option.SOME_1'(s(X62,X63))),s(bool,X64),s(fun(X62,bool),X65))) | s(bool,'F_0') = s(bool,X66)) ) | $spl24),
  inference(superposition,[],[f5454,f177])).
fof(f5746,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X58),'const.option.NONE_0'),s(bool,X59),s(fun(X58,bool),X60))) | s(bool,'F_0') = s(bool,X61)) ) | $spl24),
  inference(superposition,[],[f5454,f175])).
fof(f5712,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f5454,f5454])).
fof(f5678,plain,(
  ( ! [X88,X87] : (s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X87)),bool))) ) | $spl24),
  inference(superposition,[],[f5454,f2072])).
fof(f5677,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,X86) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X82),X83),s(bool,X84),s(fun(X82,bool),X85))) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X82,bool),bool),happ(s(fun(bool,fun(fun(X82,bool),bool)),happ(s(fun('type.option.option'(X82),fun(bool,fun(fun(X82,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X82),X83))),s(bool,X84))),s(fun(X82,bool),X85)))) ) | $spl24),
  inference(superposition,[],[f5454,f178])).
fof(f5676,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X81) = s(bool,happ(s(fun(X77,bool),X80),s(X77,X78))) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X77),'const.option.SOME_1'(s(X77,X78))),s(bool,X79),s(fun(X77,bool),X80)))) ) | $spl24),
  inference(superposition,[],[f5454,f177])).
fof(f5675,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X74) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,X74),s(fun(X73,bool),X75)))) ) | $spl24),
  inference(superposition,[],[f5454,f175])).
fof(f5504,plain,(
  ( ! [X52,X51] : (s(bool,X52) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X51))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X51)) ) | $spl24),
  inference(superposition,[],[f163,f4239])).
fof(f5453,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X40)),bool))) ) | $spl24),
  inference(superposition,[],[f4239,f2072])).
fof(f6816,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl3 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f4515,f6520])).
fof(f6815,plain,(
  $false | ($spl3 | $spl24)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f213,f511,f510,f586,f577,f696,f717,f868,f869,f871,f875,f892,f951,f956,f962,f964,f972,f979,f982,f1387,f1405,f1613,f2072,f2092,f2341,f2344,f2365,f2378,f2381,f2444,f2446,f2463,f2466,f2536,f2559,f2989,f2990,f2612,f2614,f2631,f2634,f3120,f4239,f2562,f2561,f2465,f2449,f2633,f2617,f2616,f2996,f2379,f2342,f2217,f2192,f870,f849,f166,f164,f2459,f2627,f3494,f3493,f3491,f3490,f3470,f3469,f3467,f3466,f3326,f3323,f3319,f3317,f3315,f3281,f3246,f3245,f2157,f1474,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f3325,f1480,f966,f1650,f1683,f1450,f1430,f2234,f1319,f165,f2997,f1437,f1664,f3492,f3468,f3247,f2380,f2343,f1670,f1669,f1630,f5453,f5454,f5504,f5675,f5676,f5677,f5678,f5712,f5746,f5747,f5748,f5749,f5837,f5839,f5841,f5849,f5853,f5855,f5926,f5949,f5951,f6289,f6307,f5996,f5998,f6000,f6008,f6012,f6014,f6344,f6345,f6383,f6382,f6360,f6506,f6343,f6507,f5561,f6520,f6532,f6534,f6535,f6633,f6645,f5552,f6651,f6652,f6653,f6792,f6802,f6805,f6814])).
fof(f6651,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21))) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f5628,f6520])).
fof(f5628,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,X21) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | p(s(bool,X21))) ) | $spl24),
  inference(subsumption_resolution,[],[f5535,f135])).
fof(f5535,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,X21) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | p(s(bool,X21)) | s(bool,'F_0') = s(bool,X20)) ) | $spl24),
  inference(superposition,[],[f204,f4239])).
fof(f6644,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X24,'type.option.option'(bool)))) ) | $spl56),
  inference(cnf_transformation,[],[f6644_D])).
fof(f6644_D,plain,(
  ( ! [X24] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X24,'type.option.option'(bool)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f6643,plain,(
  ( ! [X24,X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X24,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6572,f6506])).
fof(f6572,plain,(
  ( ! [X24,X23] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X24,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3076])).
fof(f6642,plain,(
  ( ! [X48] : (s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool)),X48,'type.option.option'(bool)))) ) | $spl54),
  inference(cnf_transformation,[],[f6642_D])).
fof(f6642_D,plain,(
  ( ! [X48] : (s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool)),X48,'type.option.option'(bool)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6641,plain,(
  ( ! [X47,X48] : (s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool)),X48,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6570,f6506])).
fof(f6570,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s('type.option.option'(bool),sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool))) = s('type.option.option'(bool),sK5(sK5('const.option.SOME_1'(s(bool,'F_0')),X48,'type.option.option'(bool)),X48,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3084])).
fof(f6640,plain,(
  ( ! [X50,X51] : (s('type.option.option'(bool),X51) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X51) = s('type.option.option'(bool),sK5(X51,X50,'type.option.option'(bool)))) ) | $spl52),
  inference(cnf_transformation,[],[f6640_D])).
fof(f6640_D,plain,(
  ( ! [X50,X51] : (s('type.option.option'(bool),X51) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X51) = s('type.option.option'(bool),sK5(X51,X50,'type.option.option'(bool)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6639,plain,(
  ( ! [X50,X51,X49] : (s('type.option.option'(bool),X51) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X51) = s('type.option.option'(bool),sK5(X51,X50,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl3 | $spl24)),
  inference(subsumption_resolution,[],[f6569,f6506])).
fof(f6569,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | s('type.option.option'(bool),X51) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X51) = s('type.option.option'(bool),sK5(X51,X50,'type.option.option'(bool))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3085])).
fof(f6615,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X69,bool),bool),happ(s(fun(bool,fun(fun(X69,bool),bool)),happ(s(fun('type.option.option'(X69),fun(bool,fun(fun(X69,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X69),X70))),s(bool,X71))),s(fun(X69,bool),X72)))) ) | $spl44),
  inference(cnf_transformation,[],[f6615_D])).
fof(f6615_D,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X69,bool),bool),happ(s(fun(bool,fun(fun(X69,bool),bool)),happ(s(fun('type.option.option'(X69),fun(bool,fun(fun(X69,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X69),X70))),s(bool,X71))),s(fun(X69,bool),X72)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6619,plain,(
  ( ! [X54] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X54)) ) | $spl50),
  inference(cnf_transformation,[],[f6619_D])).
fof(f6619_D,plain,(
  ( ! [X54] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X54)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f6535,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X54))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3492])).
fof(f6612,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | $spl40),
  inference(cnf_transformation,[],[f6612_D])).
fof(f6612_D,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6618,plain,(
  ( ! [X52] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'F_0') = s(bool,X52)) ) | $spl48),
  inference(cnf_transformation,[],[f6618_D])).
fof(f6618_D,plain,(
  ( ! [X52] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'F_0') = s(bool,X52)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6534,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X52))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X48,bool),bool),happ(s(fun(bool,fun(fun(X48,bool),bool)),happ(s(fun('type.option.option'(X48),fun(bool,fun(fun(X48,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X48),X49))),s(bool,X50))),s(fun(X48,bool),X51)))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3468])).
fof(f6616,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X73,bool),X74)))) ) | $spl46),
  inference(cnf_transformation,[],[f6616_D])).
fof(f6616_D,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X73,bool),X74)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6532,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X73),'const.option.NONE_0'),s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))),s(fun(X73,bool),X74))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X69,bool),bool),happ(s(fun(bool,fun(fun(X69,bool),bool)),happ(s(fun('type.option.option'(X69),fun(bool,fun(fun(X69,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X69),X70))),s(bool,X71))),s(fun(X69,bool),X72)))) ) | ($spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3247])).
fof(f6531,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | ($spl1 | $spl3 | $spl24)),
  inference(backward_demodulation,[],[f6520,f3062])).
fof(f6504,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK5(X54,X53,X52))))) ) | $spl36),
  inference(cnf_transformation,[],[f6504_D])).
fof(f6504_D,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK5(X54,X53,X52))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6382,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X51)) | p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK5(X50,X49,X48))))) ) | $spl24),
  inference(subsumption_resolution,[],[f6362,f135])).
fof(f6362,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X51)) | p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) | s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),X49),s(X48,sK5(X50,X49,X48))))) ) | $spl24),
  inference(superposition,[],[f165,f5454])).
fof(f6383,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X55)) | p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK5(X54,X53,X52))))) ) | $spl24),
  inference(subsumption_resolution,[],[f6363,f135])).
fof(f6363,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X55)) | p(s(bool,happ(s(fun(X52,bool),X53),s(X52,X54)))) | s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK5(X54,X53,X52)))) | s(bool,'F_0') = s(bool,X55)) ) | $spl24),
  inference(superposition,[],[f165,f5454])).
fof(f6493,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X24)) ) | $spl34),
  inference(cnf_transformation,[],[f6493_D])).
fof(f6493_D,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f6465,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl2 | $spl24)),
  inference(duplicate_literal_removal,[],[f6439])).
fof(f6439,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f5521,f2823])).
fof(f2823,plain,(
  ( ! [X37,X38] : (s(bool,X37) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl2),
  inference(trivial_inequality_removal,[],[f2585])).
fof(f2585,plain,(
  ( ! [X37,X38] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X37) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl2),
  inference(superposition,[],[f1163,f239])).
fof(f1163,plain,(
  ( ! [X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X5))) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl2),
  inference(superposition,[],[f162,f1130])).
fof(f1130,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | $spl2),
  inference(unit_resulting_resolution,[],[f212,f132])).
fof(f212,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | $spl2),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f5521,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f4083,f4239])).
fof(f4083,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(condensation,[],[f4061])).
fof(f4061,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl2),
  inference(equality_factoring,[],[f2823])).
fof(f6492,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X15)) ) | $spl32),
  inference(cnf_transformation,[],[f6492_D])).
fof(f6492_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'F_0') = s(bool,X15)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6466,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15)) ) | ($spl1 | $spl2 | $spl24)),
  inference(duplicate_literal_removal,[],[f6436])).
fof(f6436,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2 | $spl24)),
  inference(superposition,[],[f5521,f640])).
fof(f4926,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4925,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889,f4903,f4909,f4914])).
fof(f4914,plain,(
  ( ! [X21,X19,X20] : (s('type.option.option'(X19),'const.option.SOME_1'(s(X19,sK4(X20,X19)))) = s('type.option.option'(X19),sK5('const.option.SOME_1'(s(X19,sK4(X20,X19))),X21,'type.option.option'(X19))) | s('type.option.option'(X19),'const.option.NONE_0') = s('type.option.option'(X19),X20)) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4563,f4514])).
fof(f4563,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(X19),'const.option.SOME_1'(s(X19,sK4(X20,X19)))) = s('type.option.option'(X19),sK5('const.option.SOME_1'(s(X19,sK4(X20,X19))),X21,'type.option.option'(X19))) | s('type.option.option'(X19),'const.option.NONE_0') = s('type.option.option'(X19),X20)) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1293])).
fof(f4909,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(X225,X230) = s(X225,sK5(X230,happ(s(fun(fun(X226,fun(X225,bool)),fun(X225,bool)),happ(s(fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool))),happ(s(fun('type.option.option'(X226),fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X226),X227))),s(fun(X225,bool),X228))),s(fun(X226,fun(X225,bool)),X229)),X225))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4553,f4514])).
fof(f4553,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(X225,X230) = s(X225,sK5(X230,happ(s(fun(fun(X226,fun(X225,bool)),fun(X225,bool)),happ(s(fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool))),happ(s(fun('type.option.option'(X226),fun(fun(X225,bool),fun(fun(X226,fun(X225,bool)),fun(X225,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X226),X227))),s(fun(X225,bool),X228))),s(fun(X226,fun(X225,bool)),X229)),X225))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1654])).
fof(f4903,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (s(X129,X134) = s(X129,sK5(X134,'const.option.option_CASE_3'(s('type.option.option'(X130),'const.option.SOME_1'(s(X130,X131))),s(fun(X129,bool),X132),s(fun(X130,fun(X129,bool)),X133)),X129))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4548,f4514])).
fof(f4548,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(X129,X134) = s(X129,sK5(X134,'const.option.option_CASE_3'(s('type.option.option'(X130),'const.option.SOME_1'(s(X130,X131))),s(fun(X129,bool),X132),s(fun(X130,fun(X129,bool)),X133)),X129))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1429])).
fof(f4889,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X48,happ(s(fun(X49,X48),X52),s(X49,X50))) = s(X48,sK5('const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.SOME_1'(s(X49,X50))),s(X48,X51),s(fun(X49,X48),X52)),X53,X48))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4538,f4514])).
fof(f4538,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(X48,happ(s(fun(X49,X48),X52),s(X49,X50))) = s(X48,sK5('const.option.option_CASE_3'(s('type.option.option'(X49),'const.option.SOME_1'(s(X49,X50))),s(X48,X51),s(fun(X49,X48),X52)),X53,X48))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1439])).
fof(f4885,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(X71,'const.option.option_CASE_3'(s('type.option.option'(X72),X73),s(X71,X74),s(fun(X72,X71),X75))) = s(X71,sK5(happ(s(fun(fun(X72,X71),X71),happ(s(fun(X71,fun(fun(X72,X71),X71)),happ(s(fun('type.option.option'(X72),fun(X71,fun(fun(X72,X71),X71))),'const.option.option_CASE_0'),s('type.option.option'(X72),X73))),s(X71,X74))),s(fun(X72,X71),X75)),X76,X71))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4534,f4514])).
fof(f4534,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(X71,'const.option.option_CASE_3'(s('type.option.option'(X72),X73),s(X71,X74),s(fun(X72,X71),X75))) = s(X71,sK5(happ(s(fun(fun(X72,X71),X71),happ(s(fun(X71,fun(fun(X72,X71),X71)),happ(s(fun('type.option.option'(X72),fun(X71,fun(fun(X72,X71),X71))),'const.option.option_CASE_0'),s('type.option.option'(X72),X73))),s(X71,X74))),s(fun(X72,X71),X75)),X76,X71))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1667])).
fof(f4866,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10)) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4705,f4620])).
fof(f4705,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(X9),'const.option.NONE_0'),s(bool,X10),s(fun(X9,bool),X11))) | s(bool,'F_0') = s(bool,X10)) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f3315])).
fof(f4864,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool))) ) | ($spl1 | $spl2 | $spl28)),
  inference(forward_demodulation,[],[f4701,f4620])).
fof(f4701,plain,(
  ( ! [X35,X36] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) | s(bool,X36) = s(bool,sK4('const.option.SOME_1'(s(bool,X35)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X35)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f3189])).
fof(f4862,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4700,f929])).
fof(f4700,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK5(X33,X32,X31))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f3188])).
fof(f4853,plain,(
  ( ! [X59,X58] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58))) | s(bool,X58) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X58)),bool),X59,bool))) ) | ($spl1 | $spl2 | $spl28)),
  inference(forward_demodulation,[],[f4852,f4620])).
fof(f4852,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X58)),bool),X59,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4678,f4620])).
fof(f4678,plain,(
  ( ! [X59,X58] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0'))) | s(bool,X58) = s(bool,sK5(sK4('const.option.SOME_1'(s(bool,X58)),bool),X59,bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X58)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2315])).
fof(f4847,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(condensation,[],[f4846])).
fof(f4846,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4675,f4620])).
fof(f4675,plain,(
  ( ! [X33,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X33))) | s(bool,'T_0') != s(bool,X33)) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2310])).
fof(f4809,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(duplicate_literal_removal,[],[f4674])).
fof(f4674,plain,(
  ( ! [X26] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X26)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2309])).
fof(f4672,plain,(
  ( ! [X12,X11] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X11))) | p(s(bool,X12)) | s(bool,X11) != s(bool,X12)) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2306])).
fof(f4671,plain,(
  ( ! [X10,X9] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X9))) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10)) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2304])).
fof(f4669,plain,(
  ( ! [X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | p(s(bool,X1))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2301])).
fof(f4838,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4668,f163])).
fof(f4668,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2292])).
fof(f4667,plain,(
  ( ! [X28,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X28))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,sK4('const.option.SOME_1'(s(bool,X28)),bool))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2274])).
fof(f4835,plain,(
  ( ! [X47,X45,X46] : (s(X45,X46) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4666,f4620])).
fof(f4666,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,happ(s(fun(X45,bool),X47),s(X45,X46))) | s(X45,X46) = s(X45,sK5(sK4('const.option.SOME_1'(s(X45,X46)),X45),X47,X45))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2220])).
fof(f4833,plain,(
  ( ! [X101,X99,X100] : (s(X99,X101) = s(X99,sK5(X101,sK4('const.option.SOME_1'(s(fun(X99,bool),X100)),fun(X99,bool)),X99))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4665,f4620])).
fof(f4665,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,happ(s(fun(X99,bool),X100),s(X99,X101))) | s(X99,X101) = s(X99,sK5(X101,sK4('const.option.SOME_1'(s(fun(X99,bool),X100)),fun(X99,bool)),X99))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2196])).
fof(f4823,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) = s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4644,f4620])).
fof(f4644,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X4),s(X0,X2))) | s(X0,X2) = s(X0,sK5('const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(X0,X2),s(fun(X1,X0),X3)),X4,X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f1301])).
fof(f4819,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X4) = s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4641,f4620])).
fof(f4641,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))) | s(X0,X4) = s(X0,sK5(X4,'const.option.option_CASE_3'(s('type.option.option'(X1),'const.option.NONE_0'),s(fun(X0,bool),X2),s(fun(X1,fun(X0,bool)),X3)),X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f1284])).
fof(f4817,plain,(
  ( ! [X14,X15,X13] : (s(X13,X15) = s(X13,sK5(X15,X14,X13))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4639,f877])).
fof(f4639,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,happ(s(fun(X13,bool),X14),s(X13,X15))) | s(X13,X15) = s(X13,sK5(X15,X14,X13))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f948])).
fof(f4632,plain,(
  ( ! [X7] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f630])).
fof(f4812,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4622,f4620])).
fof(f4622,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f135])).
fof(f4620,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f4619,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f403,f409,f951,f956,f962,f964,f966,f870,f849,f166,f165,f164,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1282,f1319,f1387,f1405,f1437,f1450,f1430,f1474,f1480,f1496,f1613,f1664,f1669,f1630,f1670,f1650,f1683,f1709,f1710,f1711,f2072,f2092,f2157,f2217,f2190,f2192,f2234,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2459,f2463,f2465,f2466,f2536,f2540,f2541,f2559,f2561,f2562,f2989,f2990,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2627,f2631,f2633,f2634,f3120,f3245,f3246,f3247,f3315,f3317,f3319,f3281,f3323,f3325,f3326,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f3595,f3596,f3597,f4167,f3620,f3673,f3813,f3815,f4113,f4441,f3903,f3917,f3921,f3924,f3932,f3961,f4486,f4098,f4031,f4033,f4047,f4051,f4054,f4055,f4083,f4514])).
fof(f4619,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4530,f4514])).
fof(f4530,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X54) = s(bool,happ(s(fun(fun(X50,bool),bool),happ(s(fun(bool,fun(fun(X50,bool),bool)),happ(s(fun('type.option.option'(X50),fun(bool,fun(fun(X50,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X50),X51))),s(bool,X52))),s(fun(X50,bool),X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f3442])).
fof(f4617,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4616,f398])).
fof(f4616,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | ($spl1 | $spl2 | $spl28)),
  inference(forward_demodulation,[],[f4528,f1130])).
fof(f4528,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X26,bool),bool),happ(s(fun(bool,fun(fun(X26,bool),bool)),happ(s(fun('type.option.option'(X26),fun(bool,fun(fun(X26,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X26),X27))),s(bool,X28))),s(fun(X26,bool),X29)))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f3062])).
fof(f4610,plain,(
  ( ! [X107,X105,X106,X104] : (s(fun(X104,bool),X107) = s(fun(X104,bool),sK5(X107,happ(s(fun(bool,fun(fun(X104,bool),bool)),happ(s(fun('type.option.option'(X104),fun(bool,fun(fun(X104,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X104),X105))),s(bool,X106)),fun(X104,bool)))) ) | ($spl1 | $spl28)),
  inference(subsumption_resolution,[],[f4521,f4514])).
fof(f4521,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(fun(X104,bool),X107) = s(fun(X104,bool),sK5(X107,happ(s(fun(bool,fun(fun(X104,bool),bool)),happ(s(fun('type.option.option'(X104),fun(bool,fun(fun(X104,bool),bool))),'const.option.option_CASE_0'),s('type.option.option'(X104),X105))),s(bool,X106)),fun(X104,bool)))) ) | ($spl1 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1629])).
fof(f4054,plain,(
  ( ! [X243,X246,X244,X245] : (s(X245,f(s(bool,'T_0'))) != s(X245,g(s(bool,'T_0'))) | s(fun(bool,X245),X243) = s(fun(bool,X245),X244) | s(bool,X246) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X246) = s(bool,sK6(X243,X244,X245,bool))) ) | $spl2),
  inference(superposition,[],[f176,f2823])).
fof(f4051,plain,(
  ( ! [X237,X236] : (s('type.option.option'(bool),X236) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X236) | s(bool,X237) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X237) = s(bool,sK4(X236,bool))) ) | $spl2),
  inference(superposition,[],[f154,f2823])).
fof(f4047,plain,(
  ( ! [X218,X220,X219,X221] : (s(bool,'T_0') != s(bool,f(s(X218,sK6(X219,X220,bool,X218)))) | s(fun(X218,bool),X219) = s(fun(X218,bool),X220) | s(bool,X221) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X221) = s(bool,g(s(X218,sK6(X219,X220,bool,X218))))) ) | $spl2),
  inference(superposition,[],[f176,f2823])).
fof(f4043,plain,(
  ( ! [X204,X205,X203] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X204),s(bool,'T_0'))) | s(bool,X203) = s(bool,sK5(X203,X204,bool)) | s(bool,X205) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X203) = s(bool,X205)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f2823])).
fof(f4042,plain,(
  ( ! [X202,X200,X201] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X201),s(bool,'T_0'))) | s(bool,X200) = s(bool,sK5(X200,X201,bool)) | s(bool,X202) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X200) = s(bool,X202)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f2823])).
fof(f4033,plain,(
  ( ! [X167,X166,X168] : (s(bool,'T_0') != s(bool,X167) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X166))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X167))) | s(bool,X168) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X166) = s(bool,X168)) ) | $spl2),
  inference(superposition,[],[f163,f2823])).
fof(f4031,plain,(
  ( ! [X161,X162,X160] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X161))) | s(bool,X160) = s(bool,X161) | s(bool,X162) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X160) = s(bool,X162)) ) | $spl2),
  inference(superposition,[],[f162,f2823])).
fof(f4493,plain,(
  ( ! [X66,X67,X65] : (s(bool,X65) = s(bool,X66) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X66))) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X65) = s(bool,X67)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3998,f126])).
fof(f3998,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | s(bool,X65) = s(bool,X66) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X66))) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X65) = s(bool,X67)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2304,f2823])).
fof(f4098,plain,(
  ( ! [X56,X55] : (s(bool,X55) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X55) = s(bool,X56)) ) | $spl2),
  inference(trivial_inequality_removal,[],[f3993])).
fof(f3993,plain,(
  ( ! [X56,X55] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,X55) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X55) = s(bool,X56)) ) | $spl2),
  inference(superposition,[],[f1163,f2823])).
fof(f4491,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X45) = s(bool,happ(s(fun(X46,bool),X47),s(X46,X48))) | s(X46,X48) = s(X46,sK5(X48,X47,X46)) | s(bool,X49) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X45) = s(bool,X49)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3991,f126])).
fof(f3991,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,X45) = s(bool,happ(s(fun(X46,bool),X47),s(X46,X48))) | s(X46,X48) = s(X46,sK5(X48,X47,X46)) | s(bool,X49) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X45) = s(bool,X49)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f949,f2823])).
fof(f4486,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X6) = s(bool,X7)) ) | $spl2),
  inference(subsumption_resolution,[],[f3978,f126])).
fof(f3978,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X6) = s(bool,X7)) ) | $spl2),
  inference(superposition,[],[f135,f2823])).
fof(f3971,plain,(
  ( ! [X140,X141,X139] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X140),s(bool,X139))) | s(bool,'T_0') = s(bool,sK5('T_0',X140,bool)) | s(bool,X141) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X139) = s(bool,X141)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f2823])).
fof(f3970,plain,(
  ( ! [X138,X136,X137] : (s(bool,X136) = s(bool,happ(s(fun(bool,bool),X137),s(bool,X136))) | s(bool,'T_0') = s(bool,sK5('T_0',X137,bool)) | s(bool,X138) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X136) = s(bool,X138)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f2823])).
fof(f3961,plain,(
  ( ! [X103,X102,X104] : (s(bool,X102) != s(bool,X103) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X103))) | s(bool,X104) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X102) = s(bool,X104)) ) | $spl2),
  inference(superposition,[],[f163,f2823])).
fof(f4478,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X83)),bool)) | s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X83) = s(bool,X84)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3954,f398])).
fof(f3954,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,sK4('const.option.SOME_1'(s(bool,X83)),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X84) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X83) = s(bool,X84)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2302,f2823])).
fof(f3932,plain,(
  ( ! [X23,X21,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X21))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X22))) | s(bool,X22) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X21) = s(bool,X23)) ) | $spl2),
  inference(superposition,[],[f1163,f2823])).
fof(f3924,plain,(
  ( ! [X261,X259,X260,X262] : (s(X261,f(s(bool,X262))) != s(X261,g(s(bool,X262))) | s(fun(bool,X261),X259) = s(fun(bool,X261),X260) | s(bool,X262) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK6(X259,X260,X261,bool))) ) | $spl2),
  inference(superposition,[],[f176,f2823])).
fof(f3921,plain,(
  ( ! [X253,X252] : (s('type.option.option'(bool),X252) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X253))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X252) | s(bool,X253) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK4(X252,bool))) ) | $spl2),
  inference(superposition,[],[f154,f2823])).
fof(f3917,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,X237) != s(bool,f(s(X234,sK6(X235,X236,bool,X234)))) | s(fun(X234,bool),X235) = s(fun(X234,bool),X236) | s(bool,X237) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,g(s(X234,sK6(X235,X236,bool,X234))))) ) | $spl2),
  inference(superposition,[],[f176,f2823])).
fof(f3913,plain,(
  ( ! [X218,X220,X219] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X220),s(bool,X219))) | s(bool,X218) = s(bool,sK5(X218,X220,bool)) | s(bool,X219) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X218)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f2823])).
fof(f3912,plain,(
  ( ! [X216,X215,X217] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X217),s(bool,X216))) | s(bool,X215) = s(bool,sK5(X215,X217,bool)) | s(bool,X216) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X215)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f2823])).
fof(f3903,plain,(
  ( ! [X182,X183,X181] : (s(bool,X182) != s(bool,X183) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X181))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X183))) | s(bool,X182) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X181)) ) | $spl2),
  inference(superposition,[],[f163,f2823])).
fof(f4444,plain,(
  ( ! [X109,X110] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,'T_0') = s(bool,X110) | s(bool,X109) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3878,f217])).
fof(f3878,plain,(
  ( ! [X109,X110] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X109))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) | s(bool,'T_0') = s(bool,X110) | s(bool,X109) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2292,f2823])).
fof(f4441,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,X99) | s(bool,X99) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f3876,f217])).
fof(f3876,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,X99) | s(bool,X99) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl2),
  inference(superposition,[],[f217,f2823])).
fof(f3825,plain,(
  ( ! [X218,X220,X219] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X220),s(bool,X219))) | s(bool,X218) = s(bool,sK5(X218,X220,bool)) | s(bool,X218) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X219)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f2823])).
fof(f3824,plain,(
  ( ! [X216,X215,X217] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X217),s(bool,X216))) | s(bool,X215) = s(bool,sK5(X215,X217,bool)) | s(bool,X215) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X216)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f2823])).
fof(f3815,plain,(
  ( ! [X182,X183,X181] : (s(bool,X182) != s(bool,X183) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X181))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X183))) | s(bool,X181) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X182)) ) | $spl2),
  inference(superposition,[],[f163,f2823])).
fof(f3813,plain,(
  ( ! [X177,X175,X176] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X176))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X177))) | s(bool,X175) = s(bool,X177) | s(bool,X175) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X176)) ) | $spl2),
  inference(superposition,[],[f162,f2823])).
fof(f3746,plain,(
  ( ! [X146,X147,X148] : (p(s(bool,happ(s(fun(bool,bool),X147),s(bool,X146)))) | s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X147,bool)) = s(bool,sK5(sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X147,bool),X147,bool)) | s(bool,X146) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1321,f2823])).
fof(f4350,plain,(
  ( ! [X144,X145,X143] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X144,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X144),s(bool,X143))) | s(bool,X143) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3745,f1130])).
fof(f3745,plain,(
  ( ! [X144,X145,X143] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X144),s(bool,X143))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X144,bool)) | s(bool,X143) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f2823])).
fof(f4349,plain,(
  ( ! [X142,X140,X141] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X141,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X141),s(bool,X140))) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3744,f1130])).
fof(f3744,plain,(
  ( ! [X142,X140,X141] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X141),s(bool,X140))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X141,bool)) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f2823])).
fof(f3673,plain,(
  ( ! [X146,X144,X145] : (s(bool,X145) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X144))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X145))) | s(bool,X144) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) ) | $spl2),
  inference(superposition,[],[f163,f2823])).
fof(f3620,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X56) = s(bool,sK4('const.option.SOME_1'(s(bool,X55)),bool))) ) | $spl2),
  inference(superposition,[],[f2823,f2072])).
fof(f3597,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl2),
  inference(superposition,[],[f2823,f239])).
fof(f3596,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X0) = s(bool,X3)) ) | $spl2),
  inference(superposition,[],[f2823,f2823])).
fof(f3595,plain,(
  ( ! [X70,X71] : (s(bool,X70) = s(bool,X71) | s(bool,X71) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,sK4('const.option.SOME_1'(s(bool,X70)),bool))) ) | $spl2),
  inference(superposition,[],[f2823,f2072])).
fof(f4242,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X38) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3584,f560])).
fof(f3584,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X38) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2823,f558])).
fof(f3528,plain,(
  ( ! [X61,X62] : (s(bool,X61) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X62) = s(bool,sK4('const.option.SOME_1'(s(bool,X61)),bool)) | s(bool,'T_0') = s(bool,X62)) ) | $spl2),
  inference(superposition,[],[f2823,f2072])).
fof(f3185,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3165,f126])).
fof(f3165,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s(bool,X27) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X27)))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2304,f1130])).
fof(f2541,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X17,bool)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl2),
  inference(superposition,[],[f1282,f239])).
fof(f2540,plain,(
  ( ! [X14,X15,X13] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X13))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X14))) | s(bool,X14) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl2),
  inference(superposition,[],[f1163,f239])).
fof(f2190,plain,(
  ( ! [X83] : (p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),sK4('const.option.SOME_1'(s(fun(bool,bool),X83)),fun(bool,bool)),bool))) ) | $spl2),
  inference(superposition,[],[f1282,f2072])).
fof(f1711,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),'const.option.option_CASE_3'(s('type.option.option'(X7),X8),s(fun(bool,bool),X9),s(fun(X7,fun(bool,bool)),X10))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),happ(s(fun(fun(X7,fun(bool,bool)),fun(bool,bool)),happ(s(fun(fun(bool,bool),fun(fun(X7,fun(bool,bool)),fun(bool,bool))),happ(s(fun('type.option.option'(X7),fun(fun(bool,bool),fun(fun(X7,fun(bool,bool)),fun(bool,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X7),X8))),s(fun(bool,bool),X9))),s(fun(X7,fun(bool,bool)),X10)),bool))) ) | $spl2),
  inference(superposition,[],[f1282,f178])).
fof(f1710,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X3,fun(bool,bool)),X6),s(X3,X4))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(bool,bool),X5),s(fun(X3,fun(bool,bool)),X6)),bool))) ) | $spl2),
  inference(superposition,[],[f1282,f177])).
fof(f1709,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(bool,bool),X1),s(fun(X0,fun(bool,bool)),X2)),bool))) ) | $spl2),
  inference(superposition,[],[f1282,f175])).
fof(f1496,plain,(
  ( ! [X0] : (s(sK0,sK3) != s(sK0,X0) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1482,f126])).
fof(f1482,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(sK0,sK3) != s(sK0,X0) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0))))) ) | $spl2),
  inference(superposition,[],[f166,f1130])).
fof(f1360,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) | s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool)) = s(bool,sK5(sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool),X14,bool))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1321,f1130])).
fof(f1305,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool)) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1291,f1130])).
fof(f1291,plain,(
  ( ! [X14] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f948,f1130])).
fof(f1282,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool))) ) | $spl2),
  inference(forward_demodulation,[],[f1269,f1130])).
fof(f1269,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool))) ) | $spl2),
  inference(superposition,[],[f164,f1130])).
fof(f1260,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1232,f1130])).
fof(f1232,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0'))) | s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) = s(bool,sK5(happ(s(fun(sK0,bool),sK2),s(sK0,sK3)),X14,bool))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f877,f1130])).
fof(f4925,plain,(
  ( ! [X6,X4,X5,X3] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(sK0,bool),X5),s(fun(X3,fun(sK0,bool)),X6)),sK0))) ) | ($spl1 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4578,f4514])).
fof(f4578,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(sK0,bool),X5),s(fun(X3,fun(sK0,bool)),X6)),sK0))) ) | ($spl1 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1453])).
fof(f1453,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(X3,fun(sK0,bool)),X6),s(X3,X4))),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X3),'const.option.SOME_1'(s(X3,X4))),s(fun(sK0,bool),X5),s(fun(X3,fun(sK0,bool)),X6)),sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f1306,f177])).
fof(f1306,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f1294,f873])).
fof(f1294,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))) | s(sK0,sK5(sK3,sK2,sK0)) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f948,f873])).
fof(f4924,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4923,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889,f4903,f4909,f4914])).
fof(f4923,plain,(
  ( ! [X185,X186,X184,X183] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),happ(s(fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X183),fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X183),X184))),s(fun(sK0,bool),X185))),s(fun(X183,fun(sK0,bool)),X186)),sK0))) ) | ($spl1 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4577,f4514])).
fof(f4577,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),happ(s(fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X183),fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X183),X184))),s(fun(sK0,bool),X185))),s(fun(X183,fun(sK0,bool)),X186)),sK0))) ) | ($spl1 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4514,f1646])).
fof(f1646,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),'const.option.option_CASE_3'(s('type.option.option'(X183),X184),s(fun(sK0,bool),X185),s(fun(X183,fun(sK0,bool)),X186))),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),happ(s(fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)),happ(s(fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool))),happ(s(fun('type.option.option'(X183),fun(fun(sK0,bool),fun(fun(X183,fun(sK0,bool)),fun(sK0,bool)))),'const.option.option_CASE_0'),s('type.option.option'(X183),X184))),s(fun(sK0,bool),X185))),s(fun(X183,fun(sK0,bool)),X186)),sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f1306,f178])).
fof(f4915,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889,f4903,f4909,f4914])).
fof(f4910,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889,f4903,f4909])).
fof(f4904,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889,f4903])).
fof(f4890,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885,f4889])).
fof(f4886,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866,f4885])).
fof(f4867,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864,f4866])).
fof(f4865,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862,f4864])).
fof(f4863,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853,f4862])).
fof(f4861,plain,(
  $false | ($spl1 | $spl2 | $spl20 | $spl28)),
  inference(global_subsumption,[],[f4860,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853])).
fof(f4860,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) ) | ($spl1 | $spl2 | $spl20 | $spl28)),
  inference(subsumption_resolution,[],[f4699,f1798])).
fof(f4699,plain,(
  ( ! [X30] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X30))) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) ) | ($spl1 | $spl2 | $spl20 | $spl28)),
  inference(backward_demodulation,[],[f4620,f3187])).
fof(f4859,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4858,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853])).
fof(f4858,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X28,sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4698,f1341])).
fof(f4698,plain,(
  ( ! [X28,X29] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X29))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X28,sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4620,f3186])).
fof(f4854,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847,f4853])).
fof(f4848,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809,f4847])).
fof(f4845,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672,f4809])).
fof(f4843,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671,f4672])).
fof(f4842,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669,f4671])).
fof(f4840,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838,f4669])).
fof(f4839,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667,f4838])).
fof(f4837,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835,f4667])).
fof(f4836,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833,f4835])).
fof(f4834,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823,f4833])).
fof(f4832,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4831,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823])).
fof(f4831,plain,(
  ( ! [X80] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK4('const.option.SOME_1'(s(fun(sK0,bool),X80)),fun(sK0,bool)),sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4664,f4620])).
fof(f4664,plain,(
  ( ! [X80] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK4('const.option.SOME_1'(s(fun(sK0,bool),X80)),fun(sK0,bool)),sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4620,f2187])).
fof(f2187,plain,(
  ( ! [X80] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK4('const.option.SOME_1'(s(fun(sK0,bool),X80)),fun(sK0,bool)),sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f1306,f2072])).
fof(f4830,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4829,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823])).
fof(f4829,plain,(
  ( ! [X2,X0,X1] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4654,f4620])).
fof(f4654,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4620,f1452])).
fof(f1452,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(fun(sK0,bool),X1),s(fun(X0,fun(sK0,bool)),X2)),sK0))) ) | ($spl1 | $spl12)),
  inference(superposition,[],[f1306,f175])).
fof(f4826,plain,(
  $false | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(global_subsumption,[],[f4825,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823])).
fof(f4825,plain,(
  ( ! [X22] : (s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(subsumption_resolution,[],[f4649,f4620])).
fof(f4649,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK3))) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),X22,sK0))) ) | ($spl1 | $spl2 | $spl12 | $spl28)),
  inference(backward_demodulation,[],[f4620,f1306])).
fof(f4824,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819,f4823])).
fof(f4820,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817,f4819])).
fof(f4818,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632,f4817])).
fof(f4816,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812,f4632])).
fof(f4815,plain,(
  $false | ($spl1 | $spl2 | $spl10 | $spl28)),
  inference(global_subsumption,[],[f4630,f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812])).
fof(f4630,plain,(
  ( ! [X6] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2 | $spl10 | $spl28)),
  inference(backward_demodulation,[],[f4620,f495])).
fof(f4813,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(global_subsumption,[],[f125,f124,f123,f126,f129,f132,f131,f154,f159,f163,f162,f170,f167,f175,f176,f177,f178,f204,f218,f217,f225,f239,f272,f337,f344,f403,f417,f404,f405,f419,f409,f422,f426,f427,f450,f440,f458,f465,f492,f483,f500,f211,f951,f956,f962,f964,f929,f877,f966,f968,f950,f949,f946,f933,f931,f903,f894,f870,f849,f166,f165,f164,f969,f948,f940,f898,f864,f853,f445,f441,f414,f406,f398,f303,f136,f135,f128,f212,f1130,f1163,f1260,f1282,f1284,f1301,f1302,f1288,f1305,f1292,f1293,f1319,f1320,f1321,f1353,f1354,f1370,f1357,f1360,f1362,f1363,f1372,f1387,f1434,f1405,f1437,f1438,f1439,f1413,f1427,f1450,f1429,f1430,f1431,f1461,f1462,f1463,f1491,f1474,f1479,f1480,f1496,f1498,f1531,f1538,f1539,f1551,f1559,f1560,f1580,f1565,f1570,f1571,f1586,f1609,f1661,f1613,f1664,f1666,f1667,f1624,f1626,f1669,f1628,f1629,f1630,f1670,f1649,f1650,f1683,f1653,f1654,f1655,f1656,f1709,f1710,f1711,f1359,f1358,f1304,f1303,f943,f942,f867,f866,f797,f775,f756,f640,f630,f560,f558,f2072,f2116,f2092,f2156,f2157,f2217,f2219,f2220,f2171,f2190,f2191,f2192,f2234,f2195,f2196,f2197,f2198,f2291,f2292,f2301,f2302,f2304,f2306,f2307,f2309,f2274,f2310,f2312,f2313,f2314,f2315,f2286,f2316,f2341,f2342,f2343,f2344,f2365,f2378,f2379,f2380,f2381,f2444,f2446,f2449,f2454,f2455,f2456,f2457,f2459,f2463,f2465,f2466,f2536,f2537,f2538,f2539,f2540,f2541,f2559,f2561,f2562,f2567,f2568,f2569,f2570,f2989,f2990,f2991,f2823,f2996,f2997,f2612,f2614,f2616,f2617,f2622,f2623,f2624,f2625,f2627,f2631,f2633,f2634,f3041,f3060,f3061,f3062,f3066,f3072,f3074,f3076,f3077,f3082,f3083,f3084,f3085,f3088,f3095,f3097,f3099,f3100,f3105,f3106,f3107,f3108,f3120,f3174,f3185,f3188,f3189,f3245,f3246,f3247,f3315,f3317,f3319,f3320,f3322,f3281,f3323,f3325,f3326,f3327,f3328,f3294,f3329,f3366,f3440,f3441,f3442,f3466,f3467,f3468,f3469,f3470,f3490,f3491,f3492,f3493,f3494,f3528,f4196,f4183,f4176,f4242,f3595,f3596,f3597,f4167,f3620,f3673,f4349,f4350,f3746,f3813,f3815,f3824,f3825,f4113,f4441,f4444,f3903,f3912,f3913,f3917,f3921,f3924,f3932,f4478,f3961,f3970,f3971,f4486,f4491,f4098,f4493,f4031,f4033,f4042,f4043,f4047,f4051,f4054,f4055,f4083,f4514,f4610,f4617,f4619,f4620,f4812])).
fof(f4811,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(subsumption_resolution,[],[f4621,f126])).
fof(f4621,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f128])).
fof(f4810,plain,(
  $false | ($spl1 | $spl2 | $spl28)),
  inference(trivial_inequality_removal,[],[f4625])).
fof(f4625,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl1 | $spl2 | $spl28)),
  inference(backward_demodulation,[],[f4620,f398])).
fof(f4515,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl30),
  inference(cnf_transformation,[],[f4515_D])).
fof(f4515_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4055,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl2),
  inference(equality_factoring,[],[f2823])).
fof(f4429,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X43)) ) | $spl26),
  inference(cnf_transformation,[],[f4429_D])).
fof(f4429_D,plain,(
  ( ! [X43,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X43)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4113,plain,(
  ( ! [X43,X41,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3852])).
fof(f3852,plain,(
  ( ! [X43,X41,X42] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X42))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X43))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X41)) ) | $spl2),
  inference(superposition,[],[f406,f2823])).
fof(f4224,plain,(
  ( ! [X21,X19] : (s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl22),
  inference(cnf_transformation,[],[f4224_D])).
fof(f4224_D,plain,(
  ( ! [X21,X19] : (s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4167,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21) | s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3612])).
fof(f3612,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21) | s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X20)) ) | $spl2),
  inference(superposition,[],[f2823,f2823])).
fof(f4176,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,X24) | s(bool,X24) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3577])).
fof(f3577,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) = s(bool,X24) | s(bool,X24) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X22)) ) | $spl2),
  inference(superposition,[],[f2823,f2823])).
fof(f4239,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f4239_D])).
fof(f4239_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4183,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3569])).
fof(f3569,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f2823,f129])).
fof(f4196,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f3549])).
fof(f3549,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))) | s(bool,X19) = s(bool,X21)) ) | $spl2),
  inference(superposition,[],[f2823,f2823])).
fof(f1752,plain,(
  $false | ($spl1 | $spl18)),
  inference(subsumption_resolution,[],[f1731,f175])).
fof(f1731,plain,(
  s(bool,'T_0') != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl18)),
  inference(backward_demodulation,[],[f1725,f560])).
fof(f1751,plain,(
  $false | ($spl1 | $spl18)),
  inference(subsumption_resolution,[],[f1750,f217])).
fof(f1750,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl18)),
  inference(forward_demodulation,[],[f1730,f175])).
fof(f1730,plain,(
  s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | ($spl1 | $spl18)),
  inference(backward_demodulation,[],[f1725,f558])).
fof(f1727,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) | $spl20),
  inference(cnf_transformation,[],[f1727_D])).
fof(f1727_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1719,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(sK1,sK0))))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | $spl6),
  inference(equality_resolution,[],[f960])).
fof(f960,plain,(
  ( ! [X9] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),X9) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK4(X9,sK0))))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),X9)) ) | $spl6),
  inference(superposition,[],[f216,f154])).
fof(f1098,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f984,f1011])).
fof(f1011,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f874,f874,f170])).
fof(f984,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f874,f135])).
fof(f1097,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f986,f1011])).
fof(f986,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f128,f874,f204])).
fof(f1096,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f998,f126])).
fof(f998,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f217,f874,f170])).
fof(f1095,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1000,f126])).
fof(f1000,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f511,f874,f170])).
fof(f1094,plain,(
  $false | ($spl1 | $spl14)),
  inference(subsumption_resolution,[],[f1002,f126])).
fof(f1002,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f560,f874,f170])).
fof(f1093,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1004,f874])).
fof(f1004,plain,(
  ~p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f217,f874,f170])).
fof(f1092,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1006,f874])).
fof(f1006,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f511,f874,f170])).
fof(f1091,plain,(
  $false | ($spl1 | $spl14)),
  inference(subsumption_resolution,[],[f1008,f874])).
fof(f1008,plain,(
  ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f560,f874,f170])).
fof(f1090,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1012,f126])).
fof(f1012,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f217,f874,f170])).
fof(f1089,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1014,f126])).
fof(f1014,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f511,f874,f170])).
fof(f1088,plain,(
  $false | ($spl1 | $spl14)),
  inference(subsumption_resolution,[],[f1016,f126])).
fof(f1016,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f560,f874,f170])).
fof(f1087,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1018,f874])).
fof(f1018,plain,(
  ~p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f217,f874,f170])).
fof(f1086,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1020,f874])).
fof(f1020,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f511,f874,f170])).
fof(f1085,plain,(
  $false | ($spl1 | $spl14)),
  inference(subsumption_resolution,[],[f1022,f874])).
fof(f1022,plain,(
  ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f560,f874,f170])).
fof(f1082,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1081,f874])).
fof(f1081,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f1037,f1024])).
fof(f1024,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f874,f132])).
fof(f1037,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f1024,f135])).
fof(f1080,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1036,f126])).
fof(f1036,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f1024,f128])).
fof(f1079,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f1040])).
fof(f1040,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | $spl14),
  inference(backward_demodulation,[],[f1024,f398])).
fof(f1025,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f217,f874,f132])).
fof(f1023,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f874,f560,f874,f170])).
fof(f1021,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f874,f511,f874,f170])).
fof(f1019,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f874,f217,f874,f170])).
fof(f1017,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f126,f560,f874,f170])).
fof(f1015,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f126,f511,f874,f170])).
fof(f1013,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f126,f217,f874,f170])).
fof(f1009,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f874,f560,f874,f170])).
fof(f1007,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f874,f511,f874,f170])).
fof(f1005,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f874,f217,f874,f170])).
fof(f1003,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f126,f560,f874,f170])).
fof(f1001,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f126,f511,f874,f170])).
fof(f999,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f126,f217,f874,f170])).
fof(f995,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f128,f510,f874,f204])).
fof(f993,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f128,f175,f874,f204])).
fof(f991,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f128,f558,f874,f204])).
fof(f985,plain,(
  $false | ($spl1 | $spl14)),
  inference(unit_resulting_resolution,[],[f558,f874,f135])).
fof(f983,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f128,f874])).
fof(f981,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0)) | $spl16),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0)) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f979,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f973,f131])).
fof(f973,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK3) = s(sK0,sK5(sK5(sK3,sK2,sK0),sK2,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f892,f586])).
fof(f873,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0)) | $spl12),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0)) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f871,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0))) ) | $spl3),
  inference(subsumption_resolution,[],[f859,f131])).
fof(f859,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK3) = s(sK0,sK5(sK3,sK2,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f164,f586])).
fof(f538,plain,(
  $false | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f537,f126])).
fof(f537,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f211,f494])).
fof(f534,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f125,f124,f123,f126,f128,f129,f132,f131,f136,f135,f154,f159,f163,f162,f166,f165,f164,f170,f167,f175,f176,f177,f178,f204,f217,f239,f272,f303,f403,f406,f409,f441,f445,f414,f398,f494,f533])).
fof(f533,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f518,f494])).
fof(f518,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f494,f135])).
fof(f532,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f517,f126])).
fof(f517,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f494,f128])).
fof(f531,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f521])).
fof(f521,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | $spl8),
  inference(backward_demodulation,[],[f494,f398])).
fof(f516,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f502,f215])).
fof(f502,plain,(
  s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f213,f216])).
fof(f495,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6)))) ) | $spl10),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f492,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(duplicate_literal_removal,[],[f481])).
fof(f481,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f450,f337])).
fof(f210,plain,(
  p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f216,plain,(
  ( ! [X4] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X4))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4))))) ) | $spl6),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  ( ! [X4] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X4))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f123,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4)))) | s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X4))) | p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | ! [X4] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X4))) | p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4)))))) & (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) & ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f91])).
fof(f91,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) | ! [X4] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X4))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) | ? [X3] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f90])).
fof(f90,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) | ! [X3] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) | ? [X3] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) <~> ! [X3] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) <=> ! [X3] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X0),X1),s(bool,'T_0'),s(fun(X0,bool),X2)))) <=> ! [X3] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X3))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X6),X7),s(bool,'T_0'),s(fun(X6,bool),X8)))) <=> ! [X9] : (s('type.option.option'(X6),X7) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X9))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(X6),X7),s(bool,'T_0'),s(fun(X6,bool),X8)))) <=> ! [X9] : (s('type.option.option'(X6),X7) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X9))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))))),
  file('thm.clos_relation.clos_relation_87_prove_ltb','thm.clos_relation.clos_relation_87_prove_ltb')).
fof(f211,plain,(
  ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f215,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) | $spl4),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f124,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK3))) | ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f92])).
fof(f213,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | $spl3),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f125,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) | ~p(s(bool,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK1),s(bool,'T_0'),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f92])).
% SZS output end Proof for thm.clos_relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 178
% Passive clauses: 2060
% Generated clauses: 14890
% Final active clauses: 30
% Final passive clauses: 309
% Input formulas: 30
% Initial clauses: 70
% 
% Duplicate literals: 474
% Trivial inequalities: 213
% Fw subsumption resolutions: 957
% Fw demodulations: 2173
% Bw demodulations: 1267
% Condensations: 706
% Global subsumptions: 115
% 
% Simple tautologies: 193
% Equational tautologies: 131
% Forward subsumptions: 7508
% Fw demodulations to eq. taut.: 150
% Bw demodulations to eq. taut.: 67
% 
% Unit resulting resolution: 679
% Factoring: 1
% Forward superposition: 2989
% Backward superposition: 3077
% Self superposition: 119
% Equality factoring: 84
% Equality resolution: 39
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 2443
% SAT solver unit clauses: 51
% SAT solver binary clauses: 329
% 
% Sat splits: 49
% Sat splitting refutations: 168
% 
% Memory used [KB]: 8187
% Time elapsed: 0.919 s
% ------------------------------
% ------------------------------
% Success in time 2.26 s
