Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.decProps
% (129176)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 926
% Passive clauses: 6980
% Generated clauses: 26505
% Final active clauses: 784
% Final passive clauses: 5326
% Input formulas: 78
% Initial clauses: 201
% 
% Duplicate literals: 331
% Trivial inequalities: 138
% Fw subsumption resolutions: 2178
% Fw demodulations: 3257
% Bw demodulations: 141
% 
% Simple tautologies: 255
% Equational tautologies: 57
% Forward subsumptions: 13415
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 7225
% Factoring: 1369
% Forward superposition: 7172
% Backward superposition: 4318
% Self superposition: 52
% Equality factoring: 26
% Equality resolution: 22
% 
% Split clauses: 310
% Split components: 621
% SAT solver clauses: 626
% SAT solver unit clauses: 7
% SAT solver binary clauses: 101
% 
% Sat splits: 310
% 
% TWLsolver clauses: 626
% TWLsolver calls for satisfiability: 455
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.decProps
% (129197)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 784
% Passive clauses: 2598
% Generated clauses: 23786
% Final active clauses: 621
% Final passive clauses: 1593
% Input formulas: 78
% Initial clauses: 201
% 
% Duplicate literals: 138
% Trivial inequalities: 110
% Fw subsumption resolutions: 1381
% Fw demodulations: 2224
% Bw demodulations: 305
% Condensations: 135
% Global subsumptions: 13
% 
% Simple tautologies: 231
% Equational tautologies: 8
% Forward subsumptions: 13963
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 2853
% Factoring: 350
% Forward superposition: 4685
% Backward superposition: 8333
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 65
% Split components: 130
% SAT solver clauses: 3003
% SAT solver unit clauses: 253
% SAT solver binary clauses: 82
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 15
% 
% Sat splits: 65
% Sat splitting refutations: 65
% 
% TWLsolver clauses: 2948
% TWLsolver calls for satisfiability: 1820
% 
% Memory used [KB]: 16375
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.decProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.decProps
% SZS output start Proof for thm.decProps
fof(f6376,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1874,f1885,f1878,f1889,f1879,f1890,f1880,f1891,f1882,f1894,f1892,f1883,f1896,f1884,f1900,f1898,f3254,f2286,f3264,f2205,f1887,f3266,f6353,f6357,f6359,f6361,f6314,f6369,f6367])).
fof(f6367,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) | $spl24),
  inference(cnf_transformation,[],[f6367_D])).
fof(f6367_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6369,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl27),
  inference(cnf_transformation,[],[f6369_D])).
fof(f6369_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f6314,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) | $spl14),
  inference(superposition,[],[f3292,f5744])).
fof(f5744,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))),
  inference(unit_resulting_resolution,[],[f297,f378])).
fof(f378,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X1),s('type.decSem.state'(X0),X2),s('type.list.list'('type.conLang.exp'),X3))) != s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X5))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),X3)))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X1),s('type.decSem.state'(X0),X2),s('type.list.list'('type.conLang.exp'),X3))) != s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X5))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),X3))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X1),s('type.decSem.state'(X0),X2),s('type.list.list'('type.conLang.exp'),X3))) = s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X5))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),X3))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X1),s('type.decSem.state'(X0),X2),s('type.list.list'('type.conLang.exp'),X3))) = s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X4),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X5))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),X5))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),X3))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X17] : ! [X10,X9,X18,X8,X7] : (s('type.pair.prod'('type.decSem.state'(X17),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X10),s('type.decSem.state'(X17),X9),s('type.list.list'('type.conLang.exp'),X18))) = s('type.pair.prod'('type.decSem.state'(X17),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X17),X8),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X7))))) => s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),X7))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),X18))))),
  file('thm.decProps.evaluate_sing|split|0','thm.decProps.evaluate_length|split|0')).
fof(f297,plain,(
  s('type.pair.prod'('type.decSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.pair.prod'('type.decSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',sK16),s('type.decSem.state'(sK12),sK15),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  s('type.pair.prod'('type.decSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(sK12),sK14),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.pair.prod'('type.decSem.state'(sK12),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',sK16),s('type.decSem.state'(sK12),sK15),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) & ! [X6] : s('type.list.list'('type.conSem.v'),sK13) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17])],[f179])).
fof(f179,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X1))))) = s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X4),s('type.decSem.state'(X0),X3),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',X5),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) & ! [X6] : s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X1))))) = s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X4),s('type.decSem.state'(X0),X3),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',X5),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) => ? [X6] : s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5] : (s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X0),X2),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X1))))) = s('type.pair.prod'('type.decSem.state'(X0),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X4),s('type.decSem.state'(X0),X3),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',X5),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) => ? [X6] : s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11] : (s('type.pair.prod'('type.decSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X6),X8),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X7))))) = s('type.pair.prod'('type.decSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X10),s('type.decSem.state'(X6),X9),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',X11),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) => ? [X12] : s('type.list.list'('type.conSem.v'),X7) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X12),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11] : (s('type.pair.prod'('type.decSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.pair.,_2'(s('type.decSem.state'(X6),X8),s('type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.conSem.v'),X7))))) = s('type.pair.prod'('type.decSem.state'(X6),'type.semanticPrimitives.result'('type.list.list'('type.conSem.v'),'type.conSem.v')),'const.decSem.evaluate_3'(s('type.decSem.environment',X10),s('type.decSem.state'(X6),X9),s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',X11),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0'))))) => ? [X12] : s('type.list.list'('type.conSem.v'),X7) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X12),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'))))),
  file('thm.decProps.evaluate_sing|split|0','thm.decProps.evaluate_sing|split|0')).
fof(f3292,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl14),
  inference(backward_demodulation,[],[f3274,f346])).
fof(f346,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X24] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X24))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|18')).
fof(f3274,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0') | $spl14),
  inference(backward_demodulation,[],[f1896,f302])).
fof(f302,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f43])).
fof(f43,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.decProps.evaluate_sing|split|0','thm.prim_rec.PRE|split|0')).
fof(f6361,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6360,f5679])).
fof(f5679,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f298,f3297])).
fof(f3297,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl14),
  inference(backward_demodulation,[],[f3274,f508])).
fof(f508,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f374,f303])).
fof(f303,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|13')).
fof(f374,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X23] : (s('type.list.list'(X6),X23) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,'const.list.HD_1'(s('type.list.list'(X6),X23))),s('type.list.list'(X6),'const.list.NIL_0'))) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X23))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  file('thm.decProps.evaluate_sing|split|0','thm.misc.SING_HD|split|0')).
fof(f298,plain,(
  ( ! [X6] : (s('type.list.list'('type.conSem.v'),sK13) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f226])).
fof(f6360,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) | $spl14),
  inference(forward_demodulation,[],[f6274,f3276])).
fof(f3276,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl14),
  inference(backward_demodulation,[],[f3274,f310])).
fof(f310,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.decProps.evaluate_sing|split|0','thm.list.LENGTH|split|0')).
fof(f6274,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))),
  inference(superposition,[],[f377,f5744])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X22,X15] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X15))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X22),s('type.list.list'(X6),X15)))))),
  file('thm.decProps.evaluate_sing|split|0','thm.list.LENGTH|split|1')).
fof(f6359,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6358,f5679])).
fof(f6358,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) | $spl14),
  inference(forward_demodulation,[],[f6273,f3276])).
fof(f6273,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))),
  inference(superposition,[],[f5744,f377])).
fof(f6357,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6356,f5686])).
fof(f5686,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(forward_demodulation,[],[f5685,f3423])).
fof(f3423,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl14),
  inference(backward_demodulation,[],[f3274,f2113])).
fof(f2113,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f323,f1916])).
fof(f1916,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f1903,f319])).
fof(f319,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X24] : s('type.num.num',X24) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0')))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|7')).
fof(f1903,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f319,f347])).
fof(f347,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f233])).
fof(f323,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X24] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|14')).
fof(f5685,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(unit_resulting_resolution,[],[f5679,f396])).
fof(f396,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X24,X25] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))) <=> s('type.num.num',X24) = s('type.num.num',X25))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|19')).
fof(f6356,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(forward_demodulation,[],[f6355,f3423])).
fof(f6355,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(forward_demodulation,[],[f6354,f3276])).
fof(f6354,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))))),
  inference(forward_demodulation,[],[f6271,f377])).
fof(f6271,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))))),
  inference(unit_resulting_resolution,[],[f5744,f397])).
fof(f397,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f250])).
fof(f6353,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6352,f5686])).
fof(f6352,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(forward_demodulation,[],[f6351,f3423])).
fof(f6351,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) | $spl14),
  inference(forward_demodulation,[],[f6350,f3276])).
fof(f6350,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))))),
  inference(forward_demodulation,[],[f6272,f377])).
fof(f6272,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conSem.v'),sK13))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.conLang.exp'),'const.list.CONS_2'(s('type.conLang.exp',sK17),s('type.list.list'('type.conLang.exp'),'const.list.NIL_0')))))))),
  inference(unit_resulting_resolution,[],[f5744,f397])).
fof(f3266,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))) ) | $spl22),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1887,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f1887_D])).
fof(f1887_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2205,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))) )),
  inference(superposition,[],[f346,f325])).
fof(f325,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X24] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24)))))))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|10')).
fof(f3264,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))))) ) | $spl20),
  inference(cnf_transformation,[],[f3264_D])).
fof(f3264_D,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2286,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))))) )),
  inference(superposition,[],[f346,f326])).
fof(f326,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X24] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24)))))))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|9')).
fof(f3254,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f3253,f299])).
fof(f299,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.decProps.evaluate_sing|split|0',aHL_TRUTH)).
fof(f3253,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f3252,f313])).
fof(f313,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X24] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|26')).
fof(f3252,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl3),
  inference(backward_demodulation,[],[f3210,f1982])).
fof(f1982,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f1887,f350])).
fof(f350,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X24))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|30')).
fof(f3210,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f380,f1916])).
fof(f380,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X24,X25] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num',X24))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|31')).
fof(f1898,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl16),
  inference(cnf_transformation,[],[f1898_D])).
fof(f1898_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1900,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(cnf_transformation,[],[f1900_D])).
fof(f1900_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1884,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f346,f303])).
fof(f1896,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(cnf_transformation,[],[f1896_D])).
fof(f1896_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1883,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f346,f302])).
fof(f1892,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f1892_D])).
fof(f1892_D,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1894,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl13),
  inference(cnf_transformation,[],[f1894_D])).
fof(f1894_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1882,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f346,f507])).
fof(f507,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f324,f303])).
fof(f324,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X24] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0')))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|11')).
fof(f1891,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) | $spl8),
  inference(cnf_transformation,[],[f1891_D])).
fof(f1891_D,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1880,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) )),
  inference(superposition,[],[f346,f315])).
fof(f315,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X24] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|6')).
fof(f1890,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(cnf_transformation,[],[f1890_D])).
fof(f1890_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1879,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f346,f320])).
fof(f320,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X24] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0')))),
  file('thm.decProps.evaluate_sing|split|0','thm.numeral.numeral_distrib|split|4')).
fof(f1889,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) ) | $spl4),
  inference(cnf_transformation,[],[f1889_D])).
fof(f1889_D,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1878,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) )),
  inference(superposition,[],[f346,f316])).
fof(f316,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X24] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))),
  file('thm.decProps.evaluate_sing|split|0','thm.arithmetic.MULT|split|0')).
fof(f1885,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f1885_D])).
fof(f1885_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1874,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f346,f310])).
% SZS output end Proof for thm.decProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 413
% Passive clauses: 844
% Generated clauses: 6075
% Final active clauses: 283
% Final passive clauses: 299
% Input formulas: 78
% Initial clauses: 201
% 
% Duplicate literals: 5
% Trivial inequalities: 32
% Fw subsumption resolutions: 79
% Fw demodulations: 678
% Bw demodulations: 258
% Condensations: 5
% 
% Simple tautologies: 36
% Equational tautologies: 17
% Forward subsumptions: 4290
% Fw demodulations to eq. taut.: 74
% 
% Unit resulting resolution: 1241
% Factoring: 1
% Forward superposition: 1760
% Backward superposition: 1750
% Self superposition: 34
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 871
% SAT solver unit clauses: 287
% SAT solver binary clauses: 313
% 
% Sat splits: 10
% Sat splitting refutations: 5
% 
% Memory used [KB]: 4989
% Time elapsed: 0.128 s
% ------------------------------
% ------------------------------
% Success in time 1.48 s
