Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (43233)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 318
% Passive clauses: 5343
% Generated clauses: 34259
% Final active clauses: 223
% Final passive clauses: 4652
% Input formulas: 81
% Initial clauses: 131
% 
% Duplicate literals: 1102
% Trivial inequalities: 388
% Fw subsumption resolutions: 2214
% Fw demodulations: 7386
% Bw demodulations: 88
% 
% Simple tautologies: 518
% Equational tautologies: 99
% Forward subsumptions: 16957
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 126
% Factoring: 1
% Forward superposition: 11876
% Backward superposition: 10397
% Self superposition: 204
% Equality factoring: 150
% Equality resolution: 69
% 
% Split clauses: 333
% Split components: 666
% SAT solver clauses: 559
% SAT solver unit clauses: 15
% SAT solver binary clauses: 95
% 
% Sat splits: 333
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 559
% TWLsolver calls for satisfiability: 92
% 
% Memory used [KB]: 17014
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f22504,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f22093,f214,f349,f347,f518,f215,f350,f348,f762,f776,f253,f394,f781,f783,f761,f774,f772,f786,f771,f784,f788,f9953,f9955,f251,f393,f296,f418,f11074,f11124,f330,f450,f21923,f1669,f1671,f221,f360,f406,f407,f487,f512,f218,f354,f271,f411,f9941,f9957,f297,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl88),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : ! [X11,X12] : s(X9,X11) = s(X9,'const.bool.COND_3'(s(bool,'T_0'),s(X9,X11),s(X9,X12)))),
  file('thm.closProps.Boolv_|49||49|','thm.bool.COND_CLAUSES|split|0')).
fof(f9957,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.true_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | $spl1160),
  inference(cnf_transformation,[],[f9957_D])).
fof(f9957_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.true_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) <=> ~$spl1160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1160])])).
fof(f9941,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.true_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | ($spl62 | $spl88)),
  inference(superposition,[],[f393,f419])).
fof(f411,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) ) | $spl80),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f271,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X28] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X28))),
  file('thm.closProps.Boolv_|49||49|','thm.numeral.numeral_distrib|split|18')).
fof(f354,plain,(
  s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.num.|48|_0') | $spl6),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.num.|48|_0') <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f218,plain,(
  s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.num.|48|_0')),
  inference(cnf_transformation,[],[f37])).
fof(f37,axiom,(
  s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.num.|48|_0')),
  file('thm.closProps.Boolv_|49||49|','thm.conLang.true_tag_def')).
fof(f512,plain,(
  ( ! [X0] : (s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | $spl210),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  ( ! [X0] : (s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f487,plain,(
  ( ! [X0] : (s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) ) | ($spl6 | $spl80)),
  inference(backward_demodulation,[],[f354,f411])).
fof(f407,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) ) | $spl76),
  inference(cnf_transformation,[],[f407_D])).
fof(f407_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f406,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(global_subsumption,[],[f267,f224,f216])).
fof(f216,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.closProps.Boolv_|49||49|',aHL_TRUTH)).
fof(f224,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.closProps.Boolv_|49||49|','thm.bool.NOT_CLAUSES|split|2')).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X28] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))) <=> p(s(bool,'F_0')))),
  file('thm.closProps.Boolv_|49||49|','thm.numeral.numeral_eq|split|1')).
fof(f360,plain,(
  s('type.num.num','const.conLang.false_tag_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl12),
  inference(cnf_transformation,[],[f360_D])).
fof(f360_D,plain,(
  s('type.num.num','const.conLang.false_tag_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f221,plain,(
  s('type.num.num','const.conLang.false_tag_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f38])).
fof(f38,axiom,(
  s('type.num.num','const.conLang.false_tag_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.closProps.Boolv_|49||49|','thm.conLang.false_tag_def')).
fof(f1671,plain,(
  s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.conLang.false_tag_0') | $spl323),
  inference(cnf_transformation,[],[f1671_D])).
fof(f1671_D,plain,(
  s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.conLang.false_tag_0') <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f1669,plain,(
  s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.conLang.false_tag_0') | ($spl12 | $spl76 | $spl210)),
  inference(subsumption_resolution,[],[f1655,f407])).
fof(f1655,plain,(
  s('type.num.num','const.conLang.true_tag_0') != s('type.num.num','const.conLang.false_tag_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl12 | $spl210)),
  inference(superposition,[],[f512,f360])).
fof(f21923,plain,(
  $false | ($spl323 | $spl1160 | $spl1266)),
  inference(subsumption_resolution,[],[f21922,f1671])).
fof(f21922,plain,(
  s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.conLang.false_tag_0') | ($spl1160 | $spl1266)),
  inference(trivial_inequality_removal,[],[f21913])).
fof(f21913,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) | s('type.num.num','const.conLang.true_tag_0') = s('type.num.num','const.conLang.false_tag_0') | ($spl1160 | $spl1266)),
  inference(superposition,[],[f11124,f9957])).
fof(f450,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) ) | $spl142),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f330,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X0) = s('type.num.num',X2) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3)))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2,X3] : ((s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.closSem.v'),X1) != s('type.list.list'('type.closSem.v'),X3) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3)))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2,X3] : ((s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.closSem.v'),X1) != s('type.list.list'('type.closSem.v'),X3)) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3)))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),X1))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X2),s('type.list.list'('type.closSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.closSem.v'),X1) = s('type.list.list'('type.closSem.v'),X3)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X17,X18,X19,X20] : (s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X17),s('type.list.list'('type.closSem.v'),X18))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X19),s('type.list.list'('type.closSem.v'),X20))) <=> (s('type.num.num',X17) = s('type.num.num',X19) & s('type.list.list'('type.closSem.v'),X18) = s('type.list.list'('type.closSem.v'),X20)))),
  file('thm.closProps.Boolv_|49||49|','thm.closSem.v_|49||49||split|2')).
fof(f11124,plain,(
  ( ! [X4,X3] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X3),s('type.list.list'('type.closSem.v'),X4))) | s('type.num.num','const.conLang.false_tag_0') = s('type.num.num',X3)) ) | $spl1266),
  inference(cnf_transformation,[],[f11124_D])).
fof(f11124_D,plain,(
  ( ! [X4,X3] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X3),s('type.list.list'('type.closSem.v'),X4))) | s('type.num.num','const.conLang.false_tag_0') = s('type.num.num',X3)) ) <=> ~$spl1266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1266])])).
fof(f11074,plain,(
  ( ! [X4,X3] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X3),s('type.list.list'('type.closSem.v'),X4))) | s('type.num.num','const.conLang.false_tag_0') = s('type.num.num',X3)) ) | ($spl142 | $spl1158)),
  inference(superposition,[],[f450,f9955])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) | $spl86),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ! [X11,X12] : s(X9,X12) = s(X9,'const.bool.COND_3'(s(bool,'F_0'),s(X9,X11),s(X9,X12)))),
  file('thm.closProps.Boolv_|49||49|','thm.bool.COND_CLAUSES|split|1')).
fof(f393,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) ) | $spl62),
  inference(cnf_transformation,[],[f393_D])).
fof(f393_D,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f251,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X27] : s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X27))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X27),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  file('thm.closProps.Boolv_|49||49|','thm.closSem.Boolv_def')).
fof(f9955,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.false_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | $spl1158),
  inference(cnf_transformation,[],[f9955_D])).
fof(f9955_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.false_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) <=> ~$spl1158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1158])])).
fof(f9953,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.false_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | ($spl62 | $spl86 | $spl240)),
  inference(forward_demodulation,[],[f9940,f788])).
fof(f9940,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.conLang.false_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | ($spl62 | $spl86)),
  inference(superposition,[],[f393,f418])).
fof(f788,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) | $spl240),
  inference(cnf_transformation,[],[f788_D])).
fof(f788_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f784,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) | ($spl232 | $spl238)),
  inference(backward_demodulation,[],[f783,f774])).
fof(f771,plain,(
  s(bool,'T_0') != s(bool,sK1) | $spl231),
  inference(cnf_transformation,[],[f771_D])).
fof(f771_D,plain,(
  s(bool,'T_0') != s(bool,sK1) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f786,plain,(
  s(bool,'T_0') != s(bool,sK1) | ($spl3 | $spl238)),
  inference(backward_demodulation,[],[f783,f350])).
fof(f772,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl230),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  s(bool,'T_0') = s(bool,sK1) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f774,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) | $spl232),
  inference(cnf_transformation,[],[f774_D])).
fof(f774_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f761,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) | s(bool,'T_0') = s(bool,sK1) | ($spl0 | $spl64)),
  inference(superposition,[],[f347,f394])).
fof(f783,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl238),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  s(bool,'T_0') = s(bool,sK0) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f781,plain,(
  s(bool,'T_0') = s(bool,sK0) | ($spl64 | $spl235)),
  inference(trivial_inequality_removal,[],[f780])).
fof(f780,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK0) | ($spl64 | $spl235)),
  inference(superposition,[],[f776,f394])).
fof(f394,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl64),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.closProps.Boolv_|49||49|',aHL_BOOL_CASES)).
fof(f776,plain,(
  s(bool,'F_0') != s(bool,sK0) | $spl235),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s(bool,'F_0') != s(bool,sK0) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f762,plain,(
  s(bool,'F_0') != s(bool,sK0) | s(bool,'T_0') = s(bool,sK1) | ($spl3 | $spl64)),
  inference(superposition,[],[f350,f394])).
fof(f348,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) | $spl1),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f350,plain,(
  s(bool,sK0) != s(bool,sK1) | $spl3),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  s(bool,sK0) != s(bool,sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f215,plain,(
  s(bool,sK0) != s(bool,sK1) | s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1)))),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) | s(bool,sK0) = s(bool,sK1)) & (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) | s(bool,sK0) != s(bool,sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f171])).
fof(f171,plain,(
  ? [X0,X1] : ((s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1)) & (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) | s(bool,X0) != s(bool,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ? [X0,X1] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) <~> s(bool,X0) = s(bool,X1))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ~! [X0,X1] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) <=> s(bool,X0) = s(bool,X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X6))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X7))) <=> s(bool,X6) = s(bool,X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X6))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X7))) <=> s(bool,X6) = s(bool,X7))),
  file('thm.closProps.Boolv_|49||49|','thm.closProps.Boolv_|49||49|')).
fof(f518,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f517])).
fof(f517,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f348,f349])).
fof(f347,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) | $spl0),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f349,plain,(
  s(bool,sK0) = s(bool,sK1) | $spl2),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  s(bool,sK0) = s(bool,sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f214,plain,(
  s(bool,sK0) = s(bool,sK1) | s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1)))),
  inference(cnf_transformation,[],[f172])).
fof(f22093,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))) | ($spl0 | $spl230 | $spl234)),
  inference(forward_demodulation,[],[f22092,f775])).
fof(f775,plain,(
  s(bool,'F_0') = s(bool,sK0) | $spl234),
  inference(cnf_transformation,[],[f775_D])).
fof(f775_D,plain,(
  s(bool,'F_0') = s(bool,sK0) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f22092,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))) | ($spl0 | $spl230)),
  inference(forward_demodulation,[],[f347,f772])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 230
% Passive clauses: 1040
% Generated clauses: 26260
% Final active clauses: 147
% Final passive clauses: 550
% Input formulas: 81
% Initial clauses: 131
% 
% Duplicate literals: 712
% Trivial inequalities: 376
% Fw subsumption resolutions: 1328
% Fw demodulations: 4556
% Bw demodulations: 233
% Condensations: 747
% Global subsumptions: 21
% 
% Simple tautologies: 347
% Equational tautologies: 50
% Forward subsumptions: 14730
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 47
% 
% Binary resolution: 69
% Factoring: 1
% Forward superposition: 7400
% Backward superposition: 5260
% Self superposition: 135
% Equality factoring: 48
% Equality resolution: 64
% 
% Split clauses: 352
% Split components: 710
% SAT solver clauses: 2513
% SAT solver unit clauses: 216
% SAT solver binary clauses: 107
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 16
% 
% Sat splits: 352
% Sat splitting refutations: 29
% 
% TWLsolver clauses: 2239
% TWLsolver calls for satisfiability: 694
% 
% Memory used [KB]: 12920
% Time elapsed: 0.666 s
% ------------------------------
% ------------------------------
% Success in time 1.3 s
