Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.ml_translator
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_translator
% SZS output start Proof for thm.ml_translator
fof(f8382,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f8381])).
fof(f8381,plain,(
  s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK4))) != s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK4)))),
  inference(superposition,[],[f43,f7478])).
fof(f7478,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X3)),X4),s('type.pair.prod'(X0,X1),X2))) = s(X3,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X2),s(fun(X0,fun(X1,X3)),X4)))) )),
  inference(forward_demodulation,[],[f3116,f2816])).
fof(f2816,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X3)),X4),s('type.pair.prod'(X0,X1),X2))) = s(X3,happ(s(fun(X1,X3),happ(s(fun(X0,fun(X1,X3)),X4),s(X0,sK5(X2,X1,X0)))),s(X1,sK6(X2,X1,X0))))) )),
  inference(superposition,[],[f56,f50])).
fof(f50,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK5(X2,X1,X0)),s(X1,sK6(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK5(X2,X1,X0)),s(X1,sK6(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X10] : ? [X12,X13] : s('type.pair.prod'(X6,X7),X10) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13)))),
  file('thm.ml_translator.ml_translator_437_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f56,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X6,X7] : ! [X9,X10,X14] : s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X9),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X10),s(X7,X14))))) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X9),s(X6,X10))),s(X7,X14)))),
  file('thm.ml_translator.ml_translator_437_prove_ltb','thm.pair.UNCURRY_DEF')).
fof(f3116,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X2),s(fun(X0,fun(X1,X3)),X4))) = s(X3,happ(s(fun(X1,X3),happ(s(fun(X0,fun(X1,X3)),X4),s(X0,sK5(X2,X1,X0)))),s(X1,sK6(X2,X1,X0))))) )),
  inference(superposition,[],[f57,f50])).
fof(f57,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X8] : ! [X14,X10,X9] : s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X8),'const.pair.,_2'(s(X7,X10),s(X8,X14))),s(fun(X7,fun(X8,X6)),X9))) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X9),s(X7,X10))),s(X8,X14)))),
  file('thm.ml_translator.ml_translator_437_prove_ltb','thm.pair.pair_case_thm')).
fof(f43,plain,(
  s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK4))) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,sK2)),sK3)))),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK4))) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,sK2)),sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f34])).
fof(f34,plain,(
  ? [X0,X1,X2,X3,X4] : s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),X3)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1,X2,X3,X4] : s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),X3)))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2] : ! [X3,X4] : s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X5] : s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X9),s('type.pair.prod'(X6,X7),X5))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X5),s(fun(X6,fun(X7,X8)),X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X5] : s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X9),s('type.pair.prod'(X6,X7),X5))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X5),s(fun(X6,fun(X7,X8)),X9)))),
  file('thm.ml_translator.ml_translator_437_prove_ltb','thm.ml_translator.ml_translator_437_prove_ltb')).
% SZS output end Proof for thm.ml_translator
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 50
% Passive clauses: 1949
% Generated clauses: 7962
% Final active clauses: 47
% Final passive clauses: 1856
% Input formulas: 13
% Initial clauses: 15
% 
% Duplicate literals: 637
% Trivial inequalities: 88
% Fw subsumption resolutions: 834
% Fw demodulations: 127
% Bw demodulations: 1
% 
% Simple tautologies: 193
% Equational tautologies: 94
% Forward subsumptions: 2992
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 16
% Forward superposition: 3255
% Backward superposition: 2493
% Self superposition: 302
% Equality factoring: 250
% Equality resolution: 9
% 
% Split clauses: 343
% Split components: 691
% SAT solver clauses: 1089
% SAT solver unit clauses: 243
% SAT solver binary clauses: 382
% 
% Sat splits: 343
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1085
% TWLsolver calls for satisfiability: 45
% 
% Memory used [KB]: 5117
% Time elapsed: 0.374 s
% ------------------------------
% ------------------------------
% Success in time 0.38 s
