Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ImplicitQueue
% (127256)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 3535
% Generated clauses: 18034
% Final active clauses: 138
% Final passive clauses: 3205
% Input formulas: 68
% Initial clauses: 103
% 
% Duplicate literals: 544
% Trivial inequalities: 303
% Fw subsumption resolutions: 1425
% Fw demodulations: 2597
% Bw demodulations: 18
% 
% Simple tautologies: 260
% Equational tautologies: 46
% Forward subsumptions: 8879
% 
% Binary resolution: 468
% Factoring: 1
% Forward superposition: 6220
% Backward superposition: 6165
% Self superposition: 152
% Equality factoring: 70
% Equality resolution: 31
% 
% Split clauses: 232
% Split components: 466
% SAT solver clauses: 619
% SAT solver unit clauses: 9
% SAT solver binary clauses: 142
% 
% Sat splits: 232
% 
% TWLsolver clauses: 619
% TWLsolver calls for satisfiability: 175
% 
% Memory used [KB]: 17910
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ImplicitQueue
% (127282)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 1613
% Generated clauses: 13227
% Final active clauses: 131
% Final passive clauses: 1412
% Input formulas: 68
% Initial clauses: 103
% 
% Duplicate literals: 236
% Trivial inequalities: 294
% Fw subsumption resolutions: 281
% Fw demodulations: 1545
% Bw demodulations: 66
% Condensations: 698
% Global subsumptions: 8
% 
% Simple tautologies: 98
% Equational tautologies: 38
% Forward subsumptions: 5932
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 169
% Factoring: 1
% Forward superposition: 4055
% Backward superposition: 3905
% Self superposition: 129
% Equality factoring: 54
% Equality resolution: 28
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 2565
% SAT solver unit clauses: 149
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 16
% 
% Sat splits: 49
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2533
% TWLsolver calls for satisfiability: 414
% 
% Memory used [KB]: 24818
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ImplicitQueue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ImplicitQueue
% SZS output start Proof for thm.ImplicitQueue
fof(f45160,plain,(
  $false),
  inference(subsumption_resolution,[],[f45159,f45158])).
fof(f45158,plain,(
  ( ! [X360] : (~p(s(bool,X360))) )),
  inference(forward_demodulation,[],[f45157,f315])).
fof(f315,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X19] : ! [X34,X21,X35] : s(X19,X34) = s(X19,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0'),s(X19,X34),s(fun('type.ImplicitQueue.times'(X6),X19),X21),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X19)),X35)))),
  file('thm.ImplicitQueue.digit_distinct|split|1','thm.ImplicitQueue.digit_case_def|split|0')).
fof(f45157,plain,(
  ( ! [X360,X362,X361] : (~p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0'),s(bool,X360),s(fun('type.ImplicitQueue.times'(sK0),bool),X361),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X362))))) )),
  inference(subsumption_resolution,[],[f44576,f44437])).
fof(f44437,plain,(
  ( ! [X4,X7,X5] : (s(X4,X5) = s(X4,happ(s(fun('type.ImplicitQueue.times'(sK0),X4),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X4)),X7),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) )),
  inference(superposition,[],[f27063,f315])).
fof(f27063,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0'),s(X0,X2),s(fun('type.ImplicitQueue.times'(sK0),X0),X3),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X1),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) )),
  inference(superposition,[],[f317,f229])).
fof(f229,plain,(
  s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),sK2),s('type.ImplicitQueue.times'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),sK2),s('type.ImplicitQueue.times'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f182])).
fof(f182,plain,(
  ? [X0,X1,X2] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7)))),
  file('thm.ImplicitQueue.digit_distinct|split|1','thm.ImplicitQueue.digit_distinct|split|1')).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X19,X6] : ! [X8,X7,X34,X21,X35] : s(X19,happ(s(fun('type.ImplicitQueue.times'(X6),X19),happ(s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X19)),X35),s('type.ImplicitQueue.times'(X6),X8))),s('type.ImplicitQueue.times'(X6),X7))) = s(X19,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7))),s(X19,X34),s(fun('type.ImplicitQueue.times'(X6),X19),X21),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X19)),X35)))),
  file('thm.ImplicitQueue.digit_distinct|split|1','thm.ImplicitQueue.digit_case_def|split|2')).
fof(f44576,plain,(
  ( ! [X360,X362,X361] : (s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X362),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))) | ~p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0'),s(bool,X360),s(fun('type.ImplicitQueue.times'(sK0),bool),X361),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X362))))) )),
  inference(superposition,[],[f244,f27063])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X29] : (s(bool,'F_0') = s(bool,X29) <=> ~p(s(bool,X29)))),
  file('thm.ImplicitQueue.digit_distinct|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f45159,plain,(
  ( ! [X401,X400] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X400),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(bool,X401))))) )),
  inference(subsumption_resolution,[],[f44586,f45158])).
fof(f44586,plain,(
  ( ! [X399,X401,X398,X400] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X400),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(bool,X401)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Zero_0'),s(bool,X398),s(fun('type.ImplicitQueue.times'(sK0),bool),X399),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X400))))) )),
  inference(superposition,[],[f305,f27063])).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X45,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X45))),s(bool,X46)))) <=> (p(s(bool,X45)) => p(s(bool,X46))))),
  file('thm.ImplicitQueue.digit_distinct|split|1','==>_HL_0')).
% SZS output end Proof for thm.ImplicitQueue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 269
% Passive clauses: 3988
% Generated clauses: 45302
% Final active clauses: 188
% Final passive clauses: 3013
% Input formulas: 68
% Initial clauses: 103
% 
% Duplicate literals: 1033
% Trivial inequalities: 158
% Fw subsumption resolutions: 2439
% Fw demodulations: 8674
% Bw demodulations: 728
% Condensations: 1713
% Global subsumptions: 43
% 
% Simple tautologies: 232
% Equational tautologies: 194
% Forward subsumptions: 25220
% Fw demodulations to eq. taut.: 1042
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 3485
% Factoring: 3
% Forward superposition: 16031
% Backward superposition: 9844
% Self superposition: 248
% Equality factoring: 80
% Equality resolution: 27
% 
% Split clauses: 67
% Split components: 145
% SAT solver clauses: 4304
% SAT solver unit clauses: 1043
% SAT solver binary clauses: 1288
% 
% Sat splits: 67
% Sat splitting refutations: 137
% 
% Memory used [KB]: 50916
% Time elapsed: 2.543 s
% ------------------------------
% ------------------------------
% Success in time 3.88 s
