Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.infer_eComplete
% (133478)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 2174
% Generated clauses: 22543
% Final active clauses: 147
% Final passive clauses: 1856
% Input formulas: 31
% Initial clauses: 75
% 
% Duplicate literals: 1044
% Trivial inequalities: 457
% Fw subsumption resolutions: 4075
% Fw demodulations: 2473
% Bw demodulations: 15
% 
% Simple tautologies: 572
% Equational tautologies: 58
% Forward subsumptions: 9265
% 
% Binary resolution: 494
% Forward superposition: 5079
% Backward superposition: 8107
% Self superposition: 213
% Equality factoring: 173
% Equality resolution: 44
% 
% Split clauses: 1525
% Split components: 3232
% SAT solver clauses: 2451
% SAT solver unit clauses: 31
% SAT solver binary clauses: 251
% 
% Sat splits: 1525
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2420
% TWLsolver calls for satisfiability: 185
% 
% Memory used [KB]: 11769
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.infer_eComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.infer_eComplete
% SZS output start Proof for thm.infer_eComplete
fof(f13060,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f13055,f121,f215,f124,f221,f126,f223,f167,f238,f254,f258,f120,f217,f129,f225,f271,f276,f123,f212,f133,f227,f3154,f3155,f172,f239,f190,f245,f252,f256,f119,f219,f13011,f13013,f4787,f4788,f122,f213,f612,f625,f191,f244])).
fof(f244,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X3,X2,X1,X0))))))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl38),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X3,X2,X1,X0))))))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f191,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X3,X2,X1,X0))))))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK5(X3,X2,X1,X0)),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X3,X2,X1,X0)))))))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4))))))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4))))))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3))))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4))))))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X19,X8,X9,X7] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X10),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X10))))))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X19),s('type.infer_t.infer_t',X7))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X19),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t',X7)))))))),
  file('thm.infer_eComplete.sub_completion_completes','thm.infer_eComplete.check_t_less|split|0')).
fof(f625,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X5,X4,X3,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))))))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))),s('type.infer_t.infer_t',X5)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X5)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))) ) | $spl142),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X5,X4,X3,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))))))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))),s('type.infer_t.infer_t',X5)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X5)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f612,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X5)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))),s('type.infer_t.infer_t',X5)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X4),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK5(X5,X4,X3,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))))))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3))))) ) | ($spl26 | $spl38)),
  inference(superposition,[],[f244,f238])).
fof(f213,plain,(
  ( ! [X4] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK3))))) ) | $spl2),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ( ! [X4] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK3))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f122,plain,(
  ( ! [X4] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK3))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',sK3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4))))))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f69])).
fof(f69,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))),s('type.infer_t.infer_t',X1)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4))))))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))),s('type.infer_t.infer_t',X1)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4)))))))))) & ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))),s('type.infer_t.infer_t',X1)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))) & ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num',X3)))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X4)))))))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X9))),s('type.infer_t.infer_t',X7)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X9))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8))) & ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X9)))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X6))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X10)))))))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X6))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X9))),s('type.infer_t.infer_t',X7)))) & s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X9))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8))) & ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X9)))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X6))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X10)))))))))) => p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X6))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t',X7)))))))),
  file('thm.infer_eComplete.sub_completion_completes','thm.infer_eComplete.sub_completion_completes')).
fof(f4788,plain,(
  ( ! [X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) | $spl554),
  inference(cnf_transformation,[],[f4788_D])).
fof(f4788_D,plain,(
  ( ! [X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f4787,plain,(
  ( ! [X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',X1)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) | ($spl2 | $spl4 | $spl10 | $spl50 | $spl142)),
  inference(subsumption_resolution,[],[f4786,f221])).
fof(f4786,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',X1)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) | ($spl2 | $spl4 | $spl50 | $spl142)),
  inference(forward_demodulation,[],[f4785,f256])).
fof(f4785,plain,(
  ( ! [X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',X1)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) | ($spl2 | $spl4 | $spl142)),
  inference(forward_demodulation,[],[f4770,f215])).
fof(f4770,plain,(
  ( ! [X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',X1)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3))))) ) | ($spl2 | $spl142)),
  inference(resolution,[],[f625,f213])).
fof(f13013,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) | $spl1513),
  inference(cnf_transformation,[],[f13013_D])).
fof(f13013_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) <=> ~$spl1513),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1513])])).
fof(f13011,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) | ($spl10 | $spl13 | $spl52 | $spl54 | $spl554)),
  inference(subsumption_resolution,[],[f13010,f221])).
fof(f13010,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) | ($spl13 | $spl52 | $spl54 | $spl554)),
  inference(forward_demodulation,[],[f13009,f258])).
fof(f13009,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) | ($spl13 | $spl54 | $spl554)),
  inference(subsumption_resolution,[],[f12964,f223])).
fof(f12964,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(sK1,'const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0)),sK2,'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),s('type.num.num',sK3)))) | ($spl54 | $spl554)),
  inference(superposition,[],[f4788,f276])).
fof(f219,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))) | $spl8),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f119,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))),
  inference(cnf_transformation,[],[f84])).
fof(f256,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) | $spl50),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f252,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) | ($spl8 | $spl16)),
  inference(resolution,[],[f225,f219])).
fof(f245,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK5(X3,X2,X1,X0)),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl40),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK5(X3,X2,X1,X0)),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f190,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X0),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X0),s('type.infer_t.infer_t',X3)))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',sK5(X3,X2,X1,X0)),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f239,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X1)))))) ) | $spl28),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X1)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f172,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X21,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X9))) = s(bool,'const.bool.IN_2'(s('type.num.num',X21),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',X9)))))),
  file('thm.infer_eComplete.sub_completion_completes','thm.pred_set.IN_COUNT')).
fof(f3155,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0))))) ) | $spl342),
  inference(cnf_transformation,[],[f3155_D])).
fof(f3155_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0))))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f3154,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0))))) ) | ($spl4 | $spl10 | $spl28 | $spl40 | $spl50)),
  inference(subsumption_resolution,[],[f3153,f221])).
fof(f3153,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0))))) ) | ($spl4 | $spl28 | $spl40 | $spl50)),
  inference(forward_demodulation,[],[f3152,f256])).
fof(f3152,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))) ) | ($spl4 | $spl28 | $spl40)),
  inference(forward_demodulation,[],[f2371,f215])).
fof(f2371,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s('type.num.num',sK3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))) ) | ($spl4 | $spl28 | $spl40)),
  inference(forward_demodulation,[],[f2364,f239])).
fof(f2364,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK5(X0,X1,sK2,X2)),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),X2),s('type.infer_t.infer_t',X0)))) | p(s(bool,'const.infer.check_t_3'(s('type.num.num',X1),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),X2),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',X0)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))))) ) | ($spl4 | $spl40)),
  inference(superposition,[],[f245,f215])).
fof(f227,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.infer_eComplete.sub_completion_completes','thm.bool.EQ_CLAUSES|split|3')).
fof(f212,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f123,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1))))))),
  inference(cnf_transformation,[],[f84])).
fof(f276,plain,(
  s(bool,'F_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1))))) | $spl54),
  inference(cnf_transformation,[],[f276_D])).
fof(f276_D,plain,(
  s(bool,'F_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f271,plain,(
  s(bool,'F_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1))))) | ($spl1 | $spl20)),
  inference(resolution,[],[f227,f212])).
fof(f225,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.infer_eComplete.sub_completion_completes','thm.bool.EQ_CLAUSES|split|1')).
fof(f217,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) | $spl6),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f120,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1))))),
  inference(cnf_transformation,[],[f84])).
fof(f258,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1))) | $spl52),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f254,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1))) | ($spl6 | $spl16)),
  inference(resolution,[],[f225,f217])).
fof(f238,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1)))))) ) | $spl26),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f167,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : ! [X12] : s(fun(X11,bool),'const.pred_set.EMPTY_0') = s(fun(X11,bool),'const.pred_set.INTER_2'(s(fun(X11,bool),X12),s(fun(X11,bool),'const.pred_set.COMPL_1'(s(fun(X11,bool),X12)))))),
  file('thm.infer_eComplete.sub_completion_completes','thm.pred_set.COMPL_INTER|split|0')).
fof(f223,plain,(
  ~p(s(bool,'F_0')) | $spl13),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f126,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.infer_eComplete.sub_completion_completes',aHL_FALSITY)).
fof(f221,plain,(
  p(s(bool,'T_0')) | $spl10),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f124,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.infer_eComplete.sub_completion_completes',aHL_TRUTH)).
fof(f215,plain,(
  s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) | $spl4),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f121,plain,(
  s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))) = s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2)))),
  inference(cnf_transformation,[],[f84])).
fof(f13055,plain,(
  $false | ($spl4 | $spl10 | $spl13 | $spl26 | $spl52 | $spl54 | $spl342 | $spl1513)),
  inference(subsumption_resolution,[],[f13054,f221])).
fof(f13054,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl13 | $spl26 | $spl52 | $spl54 | $spl342 | $spl1513)),
  inference(forward_demodulation,[],[f13053,f258])).
fof(f13053,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s('type.infer_t.infer_t',sK1)))) | ($spl4 | $spl13 | $spl26 | $spl54 | $spl342 | $spl1513)),
  inference(forward_demodulation,[],[f13052,f215])).
fof(f13052,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',sK1)))) | ($spl4 | $spl13 | $spl26 | $spl54 | $spl342 | $spl1513)),
  inference(subsumption_resolution,[],[f13051,f223])).
fof(f13051,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',sK1)))) | ($spl4 | $spl26 | $spl54 | $spl342 | $spl1513)),
  inference(forward_demodulation,[],[f13050,f276])).
fof(f13050,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',sK1)))) | ($spl4 | $spl26 | $spl342 | $spl1513)),
  inference(forward_demodulation,[],[f13049,f238])).
fof(f13049,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',sK1)))) | ($spl4 | $spl342 | $spl1513)),
  inference(forward_demodulation,[],[f13034,f215])).
fof(f13034,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK0))),s(fun('type.num.num',bool),'const.pred_set.INTER_2'(s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s(fun('type.num.num',bool),'const.pred_set.COMPL_1'(s(fun('type.num.num',bool),'const.pred_set.count_1'(s('type.num.num',sK3))))))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2),s('type.infer_t.infer_t',sK1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK2))),s('type.infer_t.infer_t',sK1)))) | ($spl342 | $spl1513)),
  inference(resolution,[],[f13013,f3155])).
% SZS output end Proof for thm.infer_eComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 222
% Passive clauses: 787
% Generated clauses: 13048
% Final active clauses: 156
% Final passive clauses: 474
% Input formulas: 31
% Initial clauses: 75
% 
% Duplicate literals: 632
% Trivial inequalities: 264
% Fw subsumption resolutions: 1504
% Fw demodulations: 1814
% Bw demodulations: 115
% Condensations: 723
% Global subsumptions: 8
% 
% Simple tautologies: 264
% Equational tautologies: 25
% Forward subsumptions: 5397
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 424
% Forward superposition: 3841
% Backward superposition: 2535
% Self superposition: 73
% Equality factoring: 49
% Equality resolution: 47
% 
% Split clauses: 452
% Split components: 915
% SAT solver clauses: 1640
% SAT solver unit clauses: 87
% SAT solver binary clauses: 45
% SAT solver learnt clauses: 17
% 
% Sat splits: 452
% Sat splitting refutations: 21
% 
% TWLsolver clauses: 1607
% TWLsolver calls for satisfiability: 450
% 
% Memory used [KB]: 10618
% Time elapsed: 0.595 s
% ------------------------------
% ------------------------------
% Success in time 1.24 s
