Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.weakening
% (110470)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 634
% Passive clauses: 7228
% Generated clauses: 15926
% Final active clauses: 590
% Final passive clauses: 6221
% Input formulas: 37
% Initial clauses: 149
% 
% Duplicate literals: 346
% Trivial inequalities: 22
% Fw subsumption resolutions: 1537
% Fw demodulations: 145
% Bw demodulations: 43
% 
% Simple tautologies: 184
% Equational tautologies: 60
% Forward subsumptions: 5710
% 
% Binary resolution: 8686
% Factoring: 1142
% Forward superposition: 2045
% Backward superposition: 1599
% Self superposition: 26
% Equality factoring: 28
% Equality resolution: 5
% 
% Split clauses: 617
% Split components: 1252
% SAT solver clauses: 936
% SAT solver unit clauses: 12
% SAT solver binary clauses: 91
% 
% Sat splits: 617
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 933
% TWLsolver calls for satisfiability: 370
% 
% Memory used [KB]: 8315
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.weakening
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.weakening
% SZS output start Proof for thm.weakening
fof(f1769,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1768,f210,f370,f473,f477,f208,f366,f224,f378,f475,f479,f207,f368,f1728,f1729,f371,f373,f297,f410,f492,f497,f209,f364,f228,f380])).
fof(f380,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f380_D])).
fof(f380_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f228,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.weakening.type_e_weakening|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f364,plain,(
  ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) | $spl1),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f209,plain,(
  ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))))),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) & ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f94])).
fof(f94,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) & ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0))))) & ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0))))) => p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6))))) => p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8)))) & p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6))))) => p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8)))))),
  file('thm.weakening.type_e_weakening|split|1','thm.weakening.type_e_weakening|split|1')).
fof(f497,plain,(
  s(bool,'F_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) | $spl178),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s(bool,'F_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f492,plain,(
  s(bool,'F_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK13),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) | ($spl1 | $spl16)),
  inference(resolution,[],[f380,f364])).
fof(f410,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2))))) ) | $spl58),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f297,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) | ! [X3] : (~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) | p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2))))))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X0),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2)))) => ! [X3] : (p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X3),s('type.typeSystem.type_environment',X0)))) => p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X3),s('type.list.list'('type.ast.exp'),X1),s('type.list.list'('type.ast.t'),X2))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X6),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8)))) => ! [X9] : (p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',X9),s('type.typeSystem.type_environment',X6)))) => p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X9),s('type.list.list'('type.ast.exp'),X7),s('type.list.list'('type.ast.t'),X8))))))),
  file('thm.weakening.type_e_weakening|split|1','thm.weakening.weakening_29_prove_ltb|split|1')).
fof(f373,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f371,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f213,f210])).
fof(f213,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.weakening.type_e_weakening|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f1729,plain,(
  ( ! [X1] : (~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X1),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',X1))))) ) | $spl658),
  inference(cnf_transformation,[],[f1729_D])).
fof(f1729_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X1),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',X1))))) ) <=> ~$spl658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl658])])).
fof(f1728,plain,(
  ( ! [X1] : (~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',X1)))) | ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X1),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))))) ) | ($spl9 | $spl58 | $spl178)),
  inference(subsumption_resolution,[],[f1717,f373])).
fof(f1717,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',X1)))) | ~p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',X1),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))))) ) | ($spl58 | $spl178)),
  inference(superposition,[],[f410,f497])).
fof(f368,plain,(
  p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) | $spl4),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f207,plain,(
  p(s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))))),
  inference(cnf_transformation,[],[f132])).
fof(f479,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) | $spl176),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f475,plain,(
  s(bool,'T_0') = s(bool,'const.typeSystem.type_es_3'(s('type.typeSystem.type_environment',sK10),s('type.list.list'('type.ast.exp'),sK11),s('type.list.list'('type.ast.t'),sK12))) | ($spl4 | $spl12)),
  inference(resolution,[],[f378,f368])).
fof(f378,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f378_D])).
fof(f378_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f224,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.weakening.type_e_weakening|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f366,plain,(
  p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) | $spl2),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f208,plain,(
  p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10))))),
  inference(cnf_transformation,[],[f132])).
fof(f477,plain,(
  s(bool,'T_0') = s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10))) | $spl174),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  s(bool,'T_0') = s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f473,plain,(
  s(bool,'T_0') = s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10))) | ($spl2 | $spl12)),
  inference(resolution,[],[f378,f366])).
fof(f370,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f210,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.weakening.type_e_weakening|split|1',aHL_TRUTH)).
fof(f1768,plain,(
  $false | ($spl6 | $spl174 | $spl176 | $spl658)),
  inference(subsumption_resolution,[],[f1767,f370])).
fof(f1767,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl174 | $spl176 | $spl658)),
  inference(forward_demodulation,[],[f1766,f477])).
fof(f1766,plain,(
  ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) | ($spl6 | $spl176 | $spl658)),
  inference(subsumption_resolution,[],[f1748,f370])).
fof(f1748,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.weakening.weak_2'(s('type.typeSystem.type_environment',sK13),s('type.typeSystem.type_environment',sK10)))) | ($spl176 | $spl658)),
  inference(superposition,[],[f1729,f479])).
% SZS output end Proof for thm.weakening
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 165
% Passive clauses: 322
% Generated clauses: 1528
% Final active clauses: 162
% Final passive clauses: 157
% Input formulas: 37
% Initial clauses: 149
% 
% Duplicate literals: 57
% Trivial inequalities: 3
% Fw subsumption resolutions: 138
% Fw demodulations: 26
% Bw demodulations: 3
% Condensations: 21
% Global subsumptions: 13
% 
% Simple tautologies: 134
% Equational tautologies: 3
% Forward subsumptions: 421
% 
% Binary resolution: 575
% Factoring: 38
% Forward superposition: 135
% Backward superposition: 40
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 574
% SAT solver unit clauses: 162
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 2
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 521
% TWLsolver calls for satisfiability: 1393
% 
% Memory used [KB]: 2046
% Time elapsed: 0.088 s
% ------------------------------
% ------------------------------
% Success in time 0.696 s
