Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (96719)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 3357
% Generated clauses: 19866
% Final active clauses: 104
% Final passive clauses: 3061
% Input formulas: 44
% Initial clauses: 80
% 
% Duplicate literals: 476
% Trivial inequalities: 148
% Fw subsumption resolutions: 1891
% Fw demodulations: 2744
% Bw demodulations: 4
% 
% Simple tautologies: 199
% Equational tautologies: 73
% Forward subsumptions: 9488
% 
% Binary resolution: 18
% Forward superposition: 7116
% Backward superposition: 6799
% Self superposition: 314
% Equality factoring: 200
% Equality resolution: 25
% 
% Split clauses: 175
% Split components: 355
% SAT solver clauses: 384
% SAT solver unit clauses: 1
% SAT solver binary clauses: 39
% 
% Sat splits: 175
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 384
% TWLsolver calls for satisfiability: 90
% 
% Memory used [KB]: 15607
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (96739)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 86
% Passive clauses: 818
% Generated clauses: 9194
% Final active clauses: 82
% Final passive clauses: 719
% Input formulas: 44
% Initial clauses: 80
% 
% Duplicate literals: 357
% Trivial inequalities: 44
% Fw subsumption resolutions: 269
% Fw demodulations: 796
% Bw demodulations: 13
% Condensations: 873
% Global subsumptions: 7
% 
% Simple tautologies: 73
% Equational tautologies: 23
% Forward subsumptions: 4422
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 8
% Forward superposition: 3373
% Backward superposition: 2416
% Self superposition: 127
% Equality factoring: 60
% Equality resolution: 5
% 
% Split clauses: 43
% Split components: 91
% SAT solver clauses: 1364
% SAT solver unit clauses: 123
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 10
% 
% Sat splits: 43
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1336
% TWLsolver calls for satisfiability: 629
% 
% Memory used [KB]: 15735
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f94864,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f161,f241,f240,f160,f243,f159,f245,f239,f272,f275,f273,f976,f977,f993,f978,f979,f994,f980,f995,f3720,f3721,f3722,f3723,f3724,f3744,f3745,f3746,f3747,f3748,f3752,f3754,f3756,f3757,f3758,f3759,f3760,f3761,f3801,f3850,f3803,f3851,f3805,f3806,f3807,f3808,f3809,f3853,f3849,f3810,f3852,f3848,f3854,f3855,f3856,f3870,f3857,f3873,f8610,f8611,f8612,f8613,f8614,f8615,f8616,f8617,f8618,f8619,f8620,f8621,f8622,f8623,f8624,f8625,f8626,f8627,f8628,f8629,f8630,f8631,f8632,f8633,f8634,f8635,f8636,f8637,f8638,f8639,f8640,f8641,f8642,f8643,f8644,f8645,f8646,f8647,f8648,f8649,f8650,f8651,f8652,f8653,f8654,f8655,f8659,f8664,f8660,f8663,f8658,f8669,f8812,f11574,f11579,f11701,f30084,f30102,f30085,f30086,f30087,f30104,f30146,f3869,f30147,f30148,f30169,f30149,f30170,f3871,f54446,f54431,f54486,f54396,f54391,f54382,f54381,f54508,f54546,f54335,f54330,f54327,f54614,f30106,f31192,f30103,f992,f83396,f83381,f83426,f54487,f83354,f54506,f83349,f83348,f83339,f54454,f83337,f83441,f83445,f54509,f83294,f54571,f83291,f54572,f8661,f89234,f89251,f89252,f89253,f91347,f91348,f91355,f91356,f93774,f93775,f94858,f94863])).
fof(f94863,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f94857,f91357])).
fof(f91357,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | $spl0),
  inference(unit_resulting_resolution,[],[f239,f222])).
fof(f222,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,sK12(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK13(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,sK12(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK13(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,sK12(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK13(X8,X7,X6,X5,X4,X3,X2,X1,X0)))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X11,X12] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X11))),s(X3,X12))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X11))),s(X3,X12))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(rectify,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X19,X20,X18] : (s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X18))),s(X9,X20))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.infer.st_ex_bind_3'(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X9,X19))) <=> ? [X15,X16] : (s('type.pair.prod'('type.infer.exc'(X6,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X6,X7),'const.infer.Success_1'(s(X6,X15))),s(X9,X16))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(X9,X19))) & s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X18))),s(X9,X20))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9)),happ(s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X6,X15))),s(X9,X16)))))),
  file('thm.inferProps.success_eqns|split|1','thm.inferProps.inferProps_25_prove_ltb')).
fof(f94857,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f94856])).
fof(f94856,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f241,f91358])).
fof(f91358,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | $spl0),
  inference(unit_resulting_resolution,[],[f239,f221])).
fof(f221,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,sK12(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK13(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f148])).
fof(f94858,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f94843,f91357])).
fof(f94843,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f91358,f241])).
fof(f93775,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f93235,f91357])).
fof(f93235,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f93144])).
fof(f93144,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2 | $spl6)),
  inference(superposition,[],[f89254,f92367])).
fof(f92367,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f91358,f245])).
fof(f89254,plain,(
  ( ! [X10,X9] : (s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,X9))),s(sK3,X10))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,X9))),s(sK3,X10)))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f241,f245])).
fof(f93774,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f92779,f91357])).
fof(f92779,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl0 | $spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f92367,f89254])).
fof(f91356,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f91345,f243])).
fof(f91345,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))) | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f89254])).
fof(f91355,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f91346,f89255])).
fof(f89255,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f89232,f222])).
fof(f89232,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f245,f243,f223])).
fof(f223,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f91346,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl2 | $spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f91344])).
fof(f91344,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl2 | $spl4 | $spl6)),
  inference(superposition,[],[f89254,f90267])).
fof(f90267,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f89256,f245])).
fof(f89256,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f89232,f221])).
fof(f91348,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f91328,f89255])).
fof(f91328,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK12(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))),s(sK3,sK13(sK8,sK7,sK6,sK5,sK4,sK3,sK2,sK1,sK0)))) | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f90267,f89254])).
fof(f91347,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f91329])).
fof(f91329,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f243,f89254])).
fof(f89253,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f89231,f243])).
fof(f89231,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f245,f240,f223])).
fof(f89252,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f89232,f240])).
fof(f89251,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f89233,f245])).
fof(f89233,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f240,f243,f223])).
fof(f89234,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f245,f240,f243,f223])).
fof(f8661,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(cnf_transformation,[],[f8661_D])).
fof(f8661_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f54572,plain,(
  ( ! [X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) | $spl68),
  inference(cnf_transformation,[],[f54572_D])).
fof(f54572_D,plain,(
  ( ! [X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f83291,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'('type.infer.exc'(bool,X89),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X89),bool),'const.infer.read_1'(s(bool,X88))) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,happ(s(fun(bool,bool),X90),s(bool,X88)))) ) | ($spl28 | $spl36)),
  inference(duplicate_literal_removal,[],[f82959])).
fof(f82959,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'('type.infer.exc'(bool,X89),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X89),bool),'const.infer.read_1'(s(bool,X88))) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,happ(s(fun(bool,bool),X90),s(bool,X88))) | s(bool,'T_0') = s(bool,X87)) ) | ($spl28 | $spl36)),
  inference(superposition,[],[f64870,f58491])).
fof(f58491,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl28),
  inference(superposition,[],[f57587,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f166,f166])).
fof(f166,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.inferProps.success_eqns|split|1',aHL_BOOL_CASES)).
fof(f57587,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X28),s(bool,'T_0')))) ) | $spl28),
  inference(superposition,[],[f56366,f55730])).
fof(f55730,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X0),s('type.one.one',X1)))) ) | $spl28),
  inference(unit_resulting_resolution,[],[f3852,f269])).
fof(f269,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f163,f166])).
fof(f163,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.inferProps.success_eqns|split|1',aHL_FALSITY)).
fof(f56366,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,happ(s(fun('type.one.one',X24),X26),s('type.one.one',X27)))))) ) | $spl28),
  inference(superposition,[],[f55730,f205])).
fof(f205,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X40,X41,X42] : ! [X44,X45,X47] : s(X40,happ(s(fun(X41,X40),X44),s(X41,happ(s(fun(X42,X41),X45),s(X42,X47))))) = s(X40,happ(s(fun(X42,X40),'_dst_x0x1_2'(s(fun(X41,X40),X44),s(fun(X42,X41),X45))),s(X42,X47)))),
  file('thm.inferProps.success_eqns|split|1','_dst_x0x1_ax')).
fof(f64870,plain,(
  ( ! [X2,X3] : (s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl36),
  inference(superposition,[],[f1668,f31918])).
fof(f31918,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))) ) | $spl36),
  inference(forward_demodulation,[],[f31917,f176])).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X9] : ! [X5] : s('type.pair.prod'('type.infer.exc'(X9,X6),X9),'const.infer.read_1'(s(X9,X5))) = s('type.pair.prod'('type.infer.exc'(X9,X6),X9),'const.pair.,_2'(s('type.infer.exc'(X9,X6),'const.infer.Success_1'(s(X9,X5))),s(X9,X5)))),
  file('thm.inferProps.success_eqns|split|1','thm.infer.read_def')).
fof(f31917,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl36),
  inference(forward_demodulation,[],[f31916,f3871])).
fof(f31916,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,g(s('type.one.one',X1))))),s(bool,'T_0')))) ) | $spl36),
  inference(forward_demodulation,[],[f30767,f3871])).
fof(f30767,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,g(s('type.one.one',X1))))),s(bool,g(s('type.one.one',X2)))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f3871,f3871,f211])).
fof(f211,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & ((s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6)) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X9,X6] : ! [X18,X19,X15,X20] : (s('type.pair.prod'('type.infer.exc'(X9,X7),X6),'const.infer.st_ex_return_2'(s(X9,X18),s(X6,X19))) = s('type.pair.prod'('type.infer.exc'(X9,X7),X6),'const.pair.,_2'(s('type.infer.exc'(X9,X7),'const.infer.Success_1'(s(X9,X15))),s(X6,X20))) <=> (s(X9,X15) = s(X9,X18) & s(X6,X19) = s(X6,X20)))),
  file('thm.inferProps.success_eqns|split|1','thm.inferProps.inferProps_24_prove_ltb')).
fof(f1668,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.read_1'(s(X0,X2))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.st_ex_return_2'(s(X0,X3),s(X0,X4))) | s(X0,X2) = s(X0,X4)) )),
  inference(superposition,[],[f210,f176])).
fof(f210,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X2,X4) = s(X2,X6)) )),
  inference(cnf_transformation,[],[f141])).
fof(f54571,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54)))) ) | $spl66),
  inference(cnf_transformation,[],[f54571_D])).
fof(f54571_D,plain,(
  ( ! [X54,X55] : (~p(s(bool,X54)) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f83294,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82951])).
fof(f82951,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,'T_0') = s(bool,X53)) ) | $spl28),
  inference(superposition,[],[f269,f58491])).
fof(f54509,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10)))) ) | $spl64),
  inference(cnf_transformation,[],[f54509_D])).
fof(f54509_D,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f83445,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | $spl28),
  inference(subsumption_resolution,[],[f82824,f269])).
fof(f82824,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (p(s(bool,X11)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | $spl28),
  inference(superposition,[],[f194,f58491])).
fof(f194,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & ((~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X6] : ! [X35,X31,X19,X18,X20] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X9),X6),'const.infer.guard_3'(s(bool,X35),s(X9,X31),s(X6,X19))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X9),X6),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X9),'const.infer.Success_1'(s('type.one.one',X18))),s(X6,X20))) <=> (p(s(bool,X35)) & s('type.one.one','const.one.one_0') = s('type.one.one',X18) & s(X6,X19) = s(X6,X20)))),
  file('thm.inferProps.success_eqns|split|1','thm.inferProps.inferProps_38_prove_ltb')).
fof(f83441,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,f(s('type.one.one',X20))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X21)))) ) | $spl70),
  inference(cnf_transformation,[],[f83441_D])).
fof(f83441_D,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,f(s('type.one.one',X20))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X21)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f83337,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X36))) | s(bool,X36) = s(bool,f(s('type.one.one',X34)))) ) | ($spl28 | $spl48)),
  inference(duplicate_literal_removal,[],[f82738])).
fof(f82738,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X36))) | s(bool,X36) = s(bool,f(s('type.one.one',X34))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl28 | $spl48)),
  inference(superposition,[],[f58491,f54761])).
fof(f54761,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,f(s('type.one.one',X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(trivial_inequality_removal,[],[f54756])).
fof(f54756,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,f(s('type.one.one',X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(superposition,[],[f30103,f267])).
fof(f54454,plain,(
  ( ! [X37,X35,X38,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37)))) ) | $spl58),
  inference(cnf_transformation,[],[f54454_D])).
fof(f54454_D,plain,(
  ( ! [X37,X35,X38,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f83339,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X20))) | s(bool,X18) = s(bool,X20) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X18)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82735])).
fof(f82735,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X20))) | s(bool,X18) = s(bool,X20) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) ) | $spl28),
  inference(superposition,[],[f58491,f58491])).
fof(f83348,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X34) = s(bool,X35) | s(bool,X35) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X35))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X34)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82683])).
fof(f82683,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X34) = s(bool,X35) | s(bool,X35) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X35))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X34))) | s(bool,'T_0') = s(bool,X33)) ) | $spl28),
  inference(superposition,[],[f58491,f58491])).
fof(f83349,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,f(s('type.one.one',X20))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X21))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl28 | $spl48)),
  inference(duplicate_literal_removal,[],[f82680])).
fof(f82680,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,f(s('type.one.one',X20))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X22),s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | ($spl28 | $spl48)),
  inference(superposition,[],[f58491,f54761])).
fof(f54506,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) | $spl62),
  inference(cnf_transformation,[],[f54506_D])).
fof(f54506_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f83354,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82675])).
fof(f82675,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(superposition,[],[f58491,f166])).
fof(f54487,plain,(
  ( ! [X151,X153] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))) = s(bool,happ(s(fun(bool,bool),X153),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151)))) ) | $spl60),
  inference(cnf_transformation,[],[f54487_D])).
fof(f54487_D,plain,(
  ( ! [X151,X153] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))) = s(bool,happ(s(fun(bool,bool),X153),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f83426,plain,(
  ( ! [X132,X133,X131] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X131))) = s(bool,happ(s(fun(bool,bool),X133),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X131))))) | s(bool,'T_0') = s(bool,X132) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X131)))) ) | $spl28),
  inference(forward_demodulation,[],[f83361,f167])).
fof(f167,plain,(
  ( ! [X0,X1] : (s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X38] : s(X9,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X9),X38))) = s(X9,happ(s(fun('type.infer.infer_st'(X9),X9),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X9),X38)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_1_const.infer.infer_st_subst')).
fof(f83361,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') = s(bool,X132) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X131))) = s(bool,happ(s(fun(bool,bool),X133),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X131))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X131)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82657])).
fof(f82657,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') = s(bool,X132) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X131))) = s(bool,happ(s(fun(bool,bool),X133),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X131))))) | s(bool,'T_0') = s(bool,X132) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X131)))) ) | $spl28),
  inference(superposition,[],[f58491,f286])).
fof(f286,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X2))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X2)))) )),
  inference(superposition,[],[f167,f166])).
fof(f83381,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X23))) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X26))) | s(bool,X23) = s(bool,X26)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82635])).
fof(f82635,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X26))) | s(bool,X23) = s(bool,X26)) ) | $spl28),
  inference(superposition,[],[f58491,f58491])).
fof(f83396,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X38) = s(bool,happ(s(fun(bool,bool),X39),s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,happ(s(fun(bool,bool),X42),s(bool,X41))) | s(bool,X38) = s(bool,X41)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f82580])).
fof(f82580,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X38) = s(bool,happ(s(fun(bool,bool),X39),s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,happ(s(fun(bool,bool),X42),s(bool,X41))) | s(bool,X38) = s(bool,X41)) ) | $spl28),
  inference(superposition,[],[f58491,f58491])).
fof(f992,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) | $spl12),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f30103,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl48),
  inference(cnf_transformation,[],[f30103_D])).
fof(f30103_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f31192,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X49))) | s(fun('type.one.one',bool),X50) = s(fun('type.one.one',bool),X51)) ) | $spl36),
  inference(superposition,[],[f562,f3871])).
fof(f562,plain,(
  ( ! [X21,X19,X22,X20] : (s(X21,f(s('type.one.one',X22))) != s(X21,g(s('type.one.one',X22))) | s(fun('type.one.one',X21),X19) = s(fun('type.one.one',X21),X20)) )),
  inference(superposition,[],[f181,f253])).
fof(f253,plain,(
  ( ! [X0,X1] : (s('type.one.one',X0) = s('type.one.one',X1)) )),
  inference(superposition,[],[f164,f164])).
fof(f164,plain,(
  ( ! [X0] : (s('type.one.one','const.one.one_0') = s('type.one.one',X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.one.one','const.one.one_0') = s('type.one.one',X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X18] : s('type.one.one','const.one.one_0') = s('type.one.one',X18)),
  file('thm.inferProps.success_eqns|split|1','thm.one.one')).
fof(f181,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.inferProps.success_eqns|split|1',aHL_EXT)).
fof(f30106,plain,(
  s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl53),
  inference(cnf_transformation,[],[f30106_D])).
fof(f30106_D,plain,(
  s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f54614,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(fun('type.one.one',bool),X52) = s(fun('type.one.one',bool),X53)) ) | $spl36),
  inference(forward_demodulation,[],[f31193,f164])).
fof(f31193,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,f(s('type.one.one',sK11(X52,X53,bool,'type.one.one')))) | s(fun('type.one.one',bool),X52) = s(fun('type.one.one',bool),X53)) ) | $spl36),
  inference(superposition,[],[f181,f3871])).
fof(f54327,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X80) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53986])).
fof(f53986,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X80) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f24503,f12703])).
fof(f12703,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f11988,f267])).
fof(f11988,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f8155,f3849])).
fof(f8155,plain,(
  ( ! [X14,X17,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,happ(s(fun('type.one.one',X13),X15),s('type.one.one',X17)))))) ) | ($spl12 | $spl22)),
  inference(forward_demodulation,[],[f7974,f3849])).
fof(f7974,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,happ(s(fun('type.one.one',bool),X16),s('type.one.one',X17))) = s(bool,happ(s(fun(X13,bool),X14),s(X13,happ(s(fun('type.one.one',X13),X15),s('type.one.one',X17)))))) ) | $spl12),
  inference(superposition,[],[f205,f992])).
fof(f24503,plain,(
  ( ! [X2,X3] : (s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(superposition,[],[f1668,f9554])).
fof(f9554,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))) ) | $spl22),
  inference(forward_demodulation,[],[f9553,f176])).
fof(f9553,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl22),
  inference(forward_demodulation,[],[f9552,f3849])).
fof(f9552,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,'T_0')))) ) | $spl22),
  inference(forward_demodulation,[],[f9429,f3849])).
fof(f9429,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,happ(s(fun('type.one.one',bool),X3),s('type.one.one',X4)))))) ) | $spl22),
  inference(unit_resulting_resolution,[],[f3849,f3849,f211])).
fof(f54330,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53979])).
fof(f53979,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f269,f12703])).
fof(f54335,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53970])).
fof(f53970,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f166,f12703])).
fof(f54546,plain,(
  ( ! [X282,X283,X284] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X282))) = s(bool,happ(s(fun(bool,bool),X284),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X282))))) | s(bool,'T_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X282)))) ) | ($spl12 | $spl22)),
  inference(forward_demodulation,[],[f54344,f167])).
fof(f54344,plain,(
  ( ! [X282,X283,X284] : (s(bool,'T_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X282))) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X282))) = s(bool,happ(s(fun(bool,bool),X284),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X282)))))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53902])).
fof(f53902,plain,(
  ( ! [X282,X283,X284] : (s(bool,'T_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X282))) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X282))) = s(bool,happ(s(fun(bool,bool),X284),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X282))))) | s(bool,'T_0') = s(bool,X283)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f286,f12703])).
fof(f54508,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl12 | $spl22)),
  inference(subsumption_resolution,[],[f53846,f269])).
fof(f53846,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (p(s(bool,X11)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f194,f12703])).
fof(f54381,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53760])).
fof(f53760,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X28))) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X27)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f166])).
fof(f54382,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X14)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53758])).
fof(f53758,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f12703])).
fof(f54391,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X32),s(bool,X31))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X30)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53700])).
fof(f53700,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X32),s(bool,X31))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f12703])).
fof(f54396,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53693])).
fof(f53693,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f166])).
fof(f54486,plain,(
  ( ! [X152,X151,X153] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))) = s(bool,happ(s(fun(bool,bool),X153),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151))))) | s(bool,'T_0') = s(bool,X152) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151)))) ) | ($spl12 | $spl22)),
  inference(forward_demodulation,[],[f54405,f167])).
fof(f54405,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') = s(bool,X152) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X151))) = s(bool,happ(s(fun(bool,bool),X153),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X151))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151)))) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53672])).
fof(f53672,plain,(
  ( ! [X152,X151,X153] : (s(bool,'T_0') = s(bool,X152) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X151))) = s(bool,happ(s(fun(bool,bool),X153),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X151))))) | s(bool,'T_0') = s(bool,X152) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X151)))) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f286])).
fof(f54431,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X19))) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X22))) | s(bool,X19) = s(bool,X22)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53646])).
fof(f53646,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X22))) | s(bool,X19) = s(bool,X22)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f12703])).
fof(f54446,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))) | s(bool,X34) = s(bool,X37)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f53584])).
fof(f53584,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))) | s(bool,X34) = s(bool,X37)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f12703,f12703])).
fof(f3871,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl36),
  inference(cnf_transformation,[],[f3871_D])).
fof(f3871_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f30170,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl56),
  inference(cnf_transformation,[],[f30170_D])).
fof(f30170_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f30149,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl16),
  inference(superposition,[],[f994,f164])).
fof(f30169,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl54),
  inference(cnf_transformation,[],[f30169_D])).
fof(f30169_D,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f30148,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) != s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl16),
  inference(superposition,[],[f994,f253])).
fof(f30147,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) != s(bool,f(s('type.one.one',X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,g(s('type.one.one',X2)))) ) | $spl16),
  inference(superposition,[],[f994,f253])).
fof(f3869,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) | $spl32),
  inference(cnf_transformation,[],[f3869_D])).
fof(f3869_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f30146,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl16),
  inference(superposition,[],[f994,f164])).
fof(f30104,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) | $spl50),
  inference(cnf_transformation,[],[f30104_D])).
fof(f30104_D,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f30087,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,X9) = s(bool,g(s('type.one.one',X8)))) ) | $spl14),
  inference(superposition,[],[f993,f164])).
fof(f30086,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,X7) = s(bool,g(s('type.one.one',X5)))) ) | $spl14),
  inference(superposition,[],[f993,f253])).
fof(f30085,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) | $spl14),
  inference(superposition,[],[f993,f253])).
fof(f30102,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl46),
  inference(cnf_transformation,[],[f30102_D])).
fof(f30102_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f30084,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl14),
  inference(superposition,[],[f993,f164])).
fof(f11701,plain,(
  $false | ($spl20 | $spl22)),
  inference(subsumption_resolution,[],[f11584,f162])).
fof(f162,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.inferProps.success_eqns|split|1',aHL_TRUTH)).
fof(f11584,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl22)),
  inference(backward_demodulation,[],[f11565,f163])).
fof(f11565,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | ($spl20 | $spl22)),
  inference(duplicate_literal_removal,[],[f11437])).
fof(f11437,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | ($spl20 | $spl22)),
  inference(superposition,[],[f3848,f3849])).
fof(f11579,plain,(
  $false | ($spl11 | $spl20 | $spl22)),
  inference(subsumption_resolution,[],[f11578,f275])).
fof(f11578,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl22)),
  inference(condensation,[],[f11577])).
fof(f11577,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X21)) ) | ($spl20 | $spl22)),
  inference(forward_demodulation,[],[f11567,f3849])).
fof(f11567,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X22),s('type.one.one',X23))) | s(bool,'T_0') = s(bool,X21)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f11415])).
fof(f11415,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X22),s('type.one.one',X23))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl20),
  inference(superposition,[],[f3848,f166])).
fof(f11574,plain,(
  $false | ($spl11 | $spl20 | $spl22)),
  inference(subsumption_resolution,[],[f11573,f275])).
fof(f11573,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl20 | $spl22)),
  inference(forward_demodulation,[],[f11409,f3849])).
fof(f11409,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X0),s('type.one.one',X1)))) ) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f275,f3848])).
fof(f8812,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f8695,f162])).
fof(f8695,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(backward_demodulation,[],[f8661,f163])).
fof(f8669,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) | $spl44),
  inference(cnf_transformation,[],[f8669_D])).
fof(f8669_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8658,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8657])).
fof(f8657,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f8602,f359])).
fof(f359,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f162,f267])).
fof(f8602,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f248,f267])).
fof(f248,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(subsumption_resolution,[],[f197,f164])).
fof(f197,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | s(X1,X4) != s(X1,X6) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f136])).
fof(f8663,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8662])).
fof(f8662,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8603,f359])).
fof(f8603,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,X47)) | s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f248,f267])).
fof(f8660,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) | $spl40),
  inference(cnf_transformation,[],[f8660_D])).
fof(f8660_D,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8664,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8663,f3886])).
fof(f3886,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,'T_0'),s(X29,X31),s(X30,X32))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34)))) ) | ($spl22 | $spl26)),
  inference(backward_demodulation,[],[f3849,f3851])).
fof(f8659,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8658,f3886])).
fof(f8655,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8551,f3886])).
fof(f8551,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f167,f248])).
fof(f8654,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8552,f3886])).
fof(f8552,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f162,f205,f248])).
fof(f8653,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8553,f3886])).
fof(f8553,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun(sK3,X1),'_dst_x0x1_2'(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X1),X3),s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4))),s(sK3,sK6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X1),X3),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10)))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f162,f1831,f248])).
fof(f1831,plain,(
  ( ! [X243,X242] : (s(X242,happ(s(fun(sK3,X242),'_dst_x0x1_2'(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X242),X243),s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4))),s(sK3,sK6))) = s(X242,happ(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X242),X243),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10)))))) ) | $spl6),
  inference(superposition,[],[f205,f245])).
fof(f8652,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8554,f3886])).
fof(f8554,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f162,f556,f248])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,'type.one.one'),X1) = s(fun(X0,'type.one.one'),X2)) )),
  inference(unit_resulting_resolution,[],[f253,f181])).
fof(f8651,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8555,f3886])).
fof(f8555,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f162,f7957,f248])).
fof(f7957,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun('type.one.one',bool)),X1) = s(fun(X0,fun('type.one.one',bool)),X2)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f992,f181])).
fof(f8650,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8556,f3886])).
fof(f8556,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f162,f671,f248])).
fof(f671,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,'type.one.one')),X2) = s(fun(X0,fun(X1,'type.one.one')),X3)) )),
  inference(unit_resulting_resolution,[],[f556,f181])).
fof(f8649,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8557,f3886])).
fof(f8557,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f162,f699,f248])).
fof(f699,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X3) = s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X4)) )),
  inference(unit_resulting_resolution,[],[f671,f181])).
fof(f8648,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8558,f3886])).
fof(f8558,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f162,f734,f248])).
fof(f734,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X4) = s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X5)) )),
  inference(unit_resulting_resolution,[],[f699,f181])).
fof(f8647,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8559,f3886])).
fof(f8559,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f162,f762,f248])).
fof(f762,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X5) = s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X6)) )),
  inference(unit_resulting_resolution,[],[f734,f181])).
fof(f8646,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8560,f3886])).
fof(f8560,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f162,f992,f248])).
fof(f8645,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8561,f3886])).
fof(f8561,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f162,f245,f248])).
fof(f8644,plain,(
  $false | ($spl4 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8562,f3886])).
fof(f8562,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f243,f248])).
fof(f8643,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8563,f3886])).
fof(f8563,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f176,f248])).
fof(f8642,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8564,f3886])).
fof(f8564,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f177,f248])).
fof(f177,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X9] : ! [X38] : s('type.pair.prod'('type.infer.exc'(X9,X6),X9),'const.infer.read_1'(s(X9,X38))) = s('type.pair.prod'('type.infer.exc'(X9,X6),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X9,X6),X9)),'const.infer.read_0'),s(X9,X38)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_1_const.infer.read')).
fof(f8641,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8565,f3886])).
fof(f8565,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f178,f248])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9,X6] : ! [X38] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X9),'type.infer.infer_st'(X6)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X6),X38))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X9),'type.infer.infer_st'(X6)),happ(s(fun('type.infer.infer_st'(X6),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X9),'type.infer.infer_st'(X6))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X6),X38)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_1_const.infer.fresh_uvar')).
fof(f8640,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8566,f3886])).
fof(f8566,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f203,f248])).
fof(f203,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X7,X9,X6] : ! [X38,X39] : s('type.pair.prod'('type.infer.exc'(X7,X9),X6),'const.infer.failwith_2'(s(X9,X38),s(X6,X39))) = s('type.pair.prod'('type.infer.exc'(X7,X9),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X7,X9),X6)),'const.infer.failwith_1'(s(X9,X38))),s(X6,X39)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_1_2_const.infer.failwith')).
fof(f8639,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8567,f3886])).
fof(f8567,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f182,f248])).
fof(f182,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X9] : ! [X38,X39,X48] : s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X9),'const.infer.guard_3'(s(bool,X38),s(X6,X39),s(X9,X48))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'('type.one.one',X6),X9)),'const.infer.guard_2'(s(bool,X38),s(X6,X39))),s(X9,X48)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_2_3_const.infer.guard')).
fof(f8638,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8568,f3886])).
fof(f8568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f174,f248])).
fof(f174,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X38,X39] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X39))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X38))),s('type.num.num',X39)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_1_2_const.arithmetic.+')).
fof(f8637,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8569,f3886])).
fof(f8569,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f162,f165,f248])).
fof(f165,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X38] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X38))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X38)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_1_const.infer_t.Infer_Tuvar')).
fof(f8636,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8570,f3886])).
fof(f8570,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f175,f248])).
fof(f175,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X38,X39] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38),s('type.infer_t.infer_t',X39))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X38))),s('type.infer_t.infer_t',X39)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_2_const.unify.t_walkstar')).
fof(f8635,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8571,f3886])).
fof(f8571,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f162,f179,f248])).
fof(f179,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X9] : ! [X38,X39] : s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X9,X6),X38),s('type.list.list'(X9),X39))) = s('type.list.list'(X6),happ(s(fun('type.list.list'(X9),'type.list.list'(X6)),happ(s(fun(fun(X9,X6),fun('type.list.list'(X9),'type.list.list'(X6))),'const.list.MAP_0'),s(fun(X9,X6),X38))),s('type.list.list'(X9),X39)))),
  file('thm.inferProps.success_eqns|split|1','ARITY_0_2_const.list.MAP')).
fof(f8634,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8572,f3886])).
fof(f8572,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one','const.one.one_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.one.one',X3)))) )),
  inference(unit_resulting_resolution,[],[f162,f164,f248])).
fof(f8633,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8573,f3886])).
fof(f8573,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f162,f253,f248])).
fof(f8632,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8574,f3886])).
fof(f8574,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f167,f248])).
fof(f8631,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8575,f3886])).
fof(f8575,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X8),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f162,f205,f248])).
fof(f8630,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8576,f3886])).
fof(f8576,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X1),X3),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun(sK3,X1),'_dst_x0x1_2'(s(fun('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),X1),X3),s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4))),s(sK3,sK6)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f162,f1831,f248])).
fof(f8629,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8577,f3886])).
fof(f8577,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f162,f556,f248])).
fof(f8628,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8578,f3886])).
fof(f8578,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f162,f7957,f248])).
fof(f8627,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8579,f3886])).
fof(f8579,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f162,f671,f248])).
fof(f8626,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8580,f3886])).
fof(f8580,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f162,f699,f248])).
fof(f8625,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8581,f3886])).
fof(f8581,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f162,f734,f248])).
fof(f8624,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8582,f3886])).
fof(f8582,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f162,f762,f248])).
fof(f8623,plain,(
  $false | ($spl12 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8583,f3886])).
fof(f8583,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f162,f992,f248])).
fof(f8622,plain,(
  $false | ($spl6 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8584,f3886])).
fof(f8584,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f162,f245,f248])).
fof(f8621,plain,(
  $false | ($spl4 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8585,f3886])).
fof(f8585,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f162,f243,f248])).
fof(f8620,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8586,f3886])).
fof(f8586,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f176,f248])).
fof(f8619,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8587,f3886])).
fof(f8587,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f177,f248])).
fof(f8618,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8588,f3886])).
fof(f8588,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f162,f178,f248])).
fof(f8617,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8589,f3886])).
fof(f8589,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f203,f248])).
fof(f8616,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8590,f3886])).
fof(f8590,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f162,f182,f248])).
fof(f8615,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8591,f3886])).
fof(f8591,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f174,f248])).
fof(f8614,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8592,f3886])).
fof(f8592,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f162,f165,f248])).
fof(f8613,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8593,f3886])).
fof(f8593,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f175,f248])).
fof(f8612,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8594,f3886])).
fof(f8594,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f162,f179,f248])).
fof(f8611,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8595,f3886])).
fof(f8595,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one','const.one.one_0')))) )),
  inference(unit_resulting_resolution,[],[f162,f164,f248])).
fof(f8610,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f8596,f3886])).
fof(f8596,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f162,f253,f248])).
fof(f3873,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl39),
  inference(cnf_transformation,[],[f3873_D])).
fof(f3873_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f3857,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl18),
  inference(superposition,[],[f995,f164])).
fof(f3870,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl34),
  inference(cnf_transformation,[],[f3870_D])).
fof(f3870_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3856,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X3)))) ) | $spl18),
  inference(superposition,[],[f995,f253])).
fof(f3855,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X2))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl18),
  inference(superposition,[],[f995,f253])).
fof(f3854,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl18),
  inference(superposition,[],[f995,f164])).
fof(f3848,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl20),
  inference(cnf_transformation,[],[f3848_D])).
fof(f3848_D,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3852,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) | $spl28),
  inference(cnf_transformation,[],[f3852_D])).
fof(f3852_D,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3810,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl12),
  inference(superposition,[],[f359,f3521])).
fof(f3521,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,happ(s(fun('type.one.one',bool),X270),s('type.one.one',X268))) = s(bool,happ(s(fun('type.one.one',bool),X271),s('type.one.one',X269)))) ) | $spl12),
  inference(superposition,[],[f1939,f253])).
fof(f1939,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,happ(s(fun('type.one.one',bool),X125),s('type.one.one',X126))) = s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun('type.one.one',X122),X124),s('type.one.one',X126)))))) ) | $spl12),
  inference(superposition,[],[f205,f992])).
fof(f3849,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) | $spl22),
  inference(cnf_transformation,[],[f3849_D])).
fof(f3849_D,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3853,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) | $spl30),
  inference(cnf_transformation,[],[f3853_D])).
fof(f3853_D,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3809,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl12),
  inference(superposition,[],[f269,f3521])).
fof(f3808,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3807,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3806,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3805,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3851,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl26),
  inference(cnf_transformation,[],[f3851_D])).
fof(f3851_D,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3803,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl12),
  inference(superposition,[],[f194,f3521])).
fof(f3850,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) | $spl24),
  inference(cnf_transformation,[],[f3850_D])).
fof(f3850_D,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3801,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl12),
  inference(superposition,[],[f166,f3521])).
fof(f3761,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl12),
  inference(superposition,[],[f359,f3521])).
fof(f3760,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl12),
  inference(superposition,[],[f269,f3521])).
fof(f3759,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3758,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3757,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3756,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl12),
  inference(superposition,[],[f267,f3521])).
fof(f3754,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl12),
  inference(superposition,[],[f194,f3521])).
fof(f3752,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl12),
  inference(superposition,[],[f166,f3521])).
fof(f3748,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl12),
  inference(superposition,[],[f3521,f166])).
fof(f3747,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3746,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3745,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3744,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3724,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl12),
  inference(superposition,[],[f3521,f166])).
fof(f3723,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3722,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3721,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f3720,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl12),
  inference(superposition,[],[f3521,f267])).
fof(f995,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) | $spl18),
  inference(cnf_transformation,[],[f995_D])).
fof(f995_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f980,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(fun('type.one.one',bool),X17) = s(fun('type.one.one',bool),X18) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) )),
  inference(superposition,[],[f562,f166])).
fof(f994,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl16),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f979,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,f(s('type.one.one',X12))) | s(fun('type.one.one',bool),X14) = s(fun('type.one.one',bool),X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,g(s('type.one.one',X12)))) )),
  inference(superposition,[],[f562,f267])).
fof(f978,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(fun('type.one.one',bool),X10) = s(fun('type.one.one',bool),X11) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f562,f267])).
fof(f993,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) | $spl14),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f977,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X4))) | s(fun('type.one.one',bool),X5) = s(fun('type.one.one',bool),X6) | s(bool,X7) = s(bool,g(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f562,f267])).
fof(f976,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2) | s(bool,X3) = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f562,f267])).
fof(f273,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f275,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f272,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f166])).
fof(f239,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) | $spl0),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f245,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | $spl6),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f159,plain,(
  s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  (s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) | (s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,sK9))),s(sK3,sK10))) = s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) & s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))))) & (s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) | ! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,X9))),s(sK3,X10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,X9))),s(sK3,X10)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f124])).
fof(f124,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X11,X12] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X11))),s(X3,X12))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X11))),s(X3,X12))))) & (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))))))),
  inference(rectify,[],[f123])).
fof(f123,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))))) & (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <~> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14] : (s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X14))),s(X9,X13))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.infer.st_ex_bind_3'(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X9,X12))) <=> ? [X15,X16] : (s('type.pair.prod'('type.infer.exc'(X6,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X6,X7),'const.infer.Success_1'(s(X6,X15))),s(X9,X16))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(X9,X12))) & s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X14))),s(X9,X13))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9)),happ(s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X6,X15))),s(X9,X16)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14] : (s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X14))),s(X9,X13))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.infer.st_ex_bind_3'(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X9,X12))) <=> ? [X15,X16] : (s('type.pair.prod'('type.infer.exc'(X6,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X6,X7),'const.infer.Success_1'(s(X6,X15))),s(X9,X16))) = s('type.pair.prod'('type.infer.exc'(X6,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X6,X7),X9)),X10),s(X9,X12))) & s('type.pair.prod'('type.infer.exc'(X8,X7),X9),'const.pair.,_2'(s('type.infer.exc'(X8,X7),'const.infer.Success_1'(s(X8,X14))),s(X9,X13))) = s('type.pair.prod'('type.infer.exc'(X8,X7),X9),happ(s(fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9)),happ(s(fun(X6,fun(X9,'type.pair.prod'('type.infer.exc'(X8,X7),X9))),X11),s(X6,X15))),s(X9,X16)))))),
  file('thm.inferProps.success_eqns|split|1','thm.inferProps.success_eqns|split|1')).
fof(f243,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))) | $spl4),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f160,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,sK9))),s(sK3,sK10))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) = s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6)))),
  inference(cnf_transformation,[],[f125])).
fof(f240,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) | $spl1),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f241,plain,(
  ( ! [X10,X9] : (s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,X9))),s(sK3,X10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,X9))),s(sK3,X10)))) ) | $spl2),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ( ! [X10,X9] : (s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,X9))),s(sK3,X10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,X9))),s(sK3,X10)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f161,plain,(
  ( ! [X10,X9] : (s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3)),happ(s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK0,X9))),s(sK3,X10))) | s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK0,sK1),'const.infer.Success_1'(s(sK0,X9))),s(sK3,X10))) != s('type.pair.prod'('type.infer.exc'(sK0,sK1),sK3),happ(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(sK3,sK6))) | s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.pair.,_2'(s('type.infer.exc'(sK2,sK1),'const.infer.Success_1'(s(sK2,sK8))),s(sK3,sK7))) != s('type.pair.prod'('type.infer.exc'(sK2,sK1),sK3),'const.infer.st_ex_bind_3'(s(fun(sK3,'type.pair.prod'('type.infer.exc'(sK0,sK1),sK3)),sK4),s(fun(sK0,fun(sK3,'type.pair.prod'('type.infer.exc'(sK2,sK1),sK3))),sK5),s(sK3,sK6)))) )),
  inference(cnf_transformation,[],[f125])).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 212
% Passive clauses: 5055
% Generated clauses: 94327
% Final active clauses: 114
% Final passive clauses: 895
% Input formulas: 44
% Initial clauses: 80
% Discarded non-redundant clauses: 17140
% 
% Duplicate literals: 548
% Trivial inequalities: 100
% Fw subsumption resolutions: 239
% Fw demodulations: 40677
% Bw demodulations: 255
% Condensations: 841
% 
% Simple tautologies: 41
% Equational tautologies: 54
% Forward subsumptions: 34073
% Fw demodulations to eq. taut.: 1638
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 33883
% Forward superposition: 8157
% Backward superposition: 8750
% Self superposition: 252
% Equality factoring: 58
% Equality resolution: 14
% 
% Split clauses: 76
% Split components: 152
% SAT solver clauses: 5271
% SAT solver unit clauses: 715
% SAT solver binary clauses: 2024
% 
% Sat splits: 76
% Sat splitting refutations: 62
% 
% Memory used [KB]: 79572
% Time elapsed: 5.231 s
% ------------------------------
% ------------------------------
% Success in time 6.59 s
