Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f3138,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3137,f581,f292,f444,f293,f446,f442,f3084,f294,f447,f445,f443])).
fof(f443,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl1),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f445,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | $spl3),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f447,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',sK0) | $spl5),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',sK0) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f294,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',sK0) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f263])).
fof(f263,plain,(
  ? [X0,X1] : ((s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X1) | s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',X0)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ? [X0,X1] : ((s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X1) | s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor',X0))))),
  inference(nnf_transformation,[],[f257])).
fof(f257,plain,(
  ? [X0,X1] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) <~> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X0)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ~! [X0,X1] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X7),s('type.ast.tctor',X6))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X7) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X7),s('type.ast.tctor',X6))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X7) & s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',X6)))),
  file('thm.evalProps.Tword_simp|split|19','thm.evalProps.Tword_simp|split|19')).
fof(f3084,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f3083,f445])).
fof(f3083,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f3077])).
fof(f3077,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword64_0') | s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | ($spl0 | $spl4)),
  inference(superposition,[],[f2129,f580])).
fof(f580,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0')))),
  inference(forward_demodulation,[],[f574,f299])).
fof(f299,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0')))),
  inference(cnf_transformation,[],[f101])).
fof(f101,axiom,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0')))),
  file('thm.evalProps.Tword_simp|split|19','thm.ast.Tword64_def')).
fof(f574,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0')))),
  inference(superposition,[],[f375,f298])).
fof(f298,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W64_0')))),
  inference(cnf_transformation,[],[f98])).
fof(f98,axiom,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W64_0')))),
  file('thm.evalProps.Tword_simp|split|19','thm.ast.TC_word_def|split|1')).
fof(f375,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X16] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X16))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X16)))))),
  file('thm.evalProps.Tword_simp|split|19','thm.ast.Tword_def')).
fof(f2129,plain,(
  ( ! [X12,X11] : (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X11),s('type.ast.tctor',X12))) | s('type.list.list'('type.ast.t'),sK1) = s('type.list.list'('type.ast.t'),X11)) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f437,f583])).
fof(f583,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word64_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f582,f299])).
fof(f582,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word64_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f442,f446])).
fof(f437,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & (s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & ((s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3)) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) <=> (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X12,X13,X14,X15] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X12),s('type.ast.tctor',X13))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X14),s('type.ast.tctor',X15))) <=> (s('type.list.list'('type.ast.t'),X12) = s('type.list.list'('type.ast.t'),X14) & s('type.ast.tctor',X13) = s('type.ast.tctor',X15)))),
  file('thm.evalProps.Tword_simp|split|19','thm.ast.t_|49||49||split|2')).
fof(f442,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl0),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f446,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0) | $spl4),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f293,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f444,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | $spl2),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f292,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f581,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f580,f524])).
fof(f524,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0'))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f523,f299])).
fof(f523,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0'))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f522,f444])).
fof(f522,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word64_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f443,f446])).
fof(f3137,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f3135,f447])).
fof(f3135,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0) | $spl0),
  inference(trivial_inequality_removal,[],[f3129])).
fof(f3129,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword64_0') | s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor',sK0) | $spl0),
  inference(superposition,[],[f3111,f580])).
fof(f3111,plain,(
  ( ! [X6,X7] : (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X6),s('type.ast.tctor',X7))) | s('type.ast.tctor',sK0) = s('type.ast.tctor',X7)) ) | $spl0),
  inference(superposition,[],[f438,f3106])).
fof(f3106,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl0),
  inference(forward_demodulation,[],[f442,f299])).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | s('type.ast.tctor',X1) = s('type.ast.tctor',X3)) )),
  inference(cnf_transformation,[],[f291])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 216
% Passive clauses: 575
% Generated clauses: 2744
% Final active clauses: 168
% Final passive clauses: 212
% Input formulas: 120
% Initial clauses: 148
% 
% Duplicate literals: 214
% Trivial inequalities: 46
% Fw subsumption resolutions: 279
% Fw demodulations: 107
% Bw demodulations: 20
% 
% Simple tautologies: 120
% Equational tautologies: 51
% Forward subsumptions: 1194
% 
% Binary resolution: 35
% Factoring: 1
% Forward superposition: 839
% Backward superposition: 758
% Self superposition: 82
% Equality factoring: 86
% Equality resolution: 13
% 
% Split clauses: 104
% Split components: 209
% SAT solver clauses: 201
% SAT solver unit clauses: 37
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 104
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 198
% TWLsolver calls for satisfiability: 49
% 
% Memory used [KB]: 1791
% Time elapsed: 0.115 s
% ------------------------------
% ------------------------------
% Success in time 0.14 s
