Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f13435,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f13434,f510,f1003,f1001,f12974,f509,f1002,f1000])).
fof(f1000,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) | $spl0),
  inference(cnf_transformation,[],[f1000_D])).
fof(f1000_D,plain,(
  p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1002,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),sK51) | $spl2),
  inference(cnf_transformation,[],[f1002_D])).
fof(f1002_D,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),sK51) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f509,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),sK51) | p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51))))),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),sK51)) & (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),sK51))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK50,sK51])],[f275])).
fof(f275,plain,(
  ? [X0,X1] : ((p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X1)))) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1)) & (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X1)))) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ? [X0,X1] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X1)))) <~> s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X1)))) <=> s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X7)))) <=> s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),X7)))) <=> s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X7))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|4','thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|4')).
fof(f12974,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12973,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X40] : ! [X14,X13] : ~s('type.list.list'(X40),'const.list.NIL_0') = s('type.list.list'(X40),'const.list.CONS_2'(s(X40,X13),s('type.list.list'(X40),X14)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|4','thm.list.list_distinct')).
fof(f12973,plain,(
  s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',sK183('const.list.NIL_0',sK51,sK50)),s('type.list.list'('type.conSem.v'),sK185('const.list.NIL_0',sK51,sK50)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12972,f1003])).
fof(f12972,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),sK51) | s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',sK183('const.list.NIL_0',sK51,sK50)),s('type.list.list'('type.conSem.v'),sK185('const.list.NIL_0',sK51,sK50)))) | $spl0),
  inference(resolution,[],[f756,f5353])).
fof(f5353,plain,(
  sP25(sK50,sK51,'const.list.NIL_0') | $spl0),
  inference(resolution,[],[f765,f1000])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) | sP25(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) | sP25(X0,X2,X1)) & (~sP25(X0,X2,X1) | p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2))))))),
  inference(nnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) <=> sP25(X0,X2,X1))),
  inference(definition_folding,[],[f144,f243])).
fof(f243,plain,(
  ! [X0,X2,X1] : (sP25(X0,X2,X1) <=> ((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X2)) | ? [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) & s('type.list.list'('type.exhSem.v'),X2) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP25])])).
fof(f144,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X1),s('type.list.list'('type.exhSem.v'),X2)))) <=> ((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X2)) | ? [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) & s('type.list.list'('type.exhSem.v'),X2) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X33,X34,X35] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.conSem.v'),X34),s('type.list.list'('type.exhSem.v'),X35)))) <=> ((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X34) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X35)) | ? [X28,X29,X7,X18] : (s('type.list.list'('type.conSem.v'),X34) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X28),s('type.list.list'('type.conSem.v'),X7))) & s('type.list.list'('type.exhSem.v'),X35) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X29),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.conSem.v',X28),s('type.exhSem.v',X29)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.list.list'('type.conSem.v'),X7),s('type.list.list'('type.exhSem.v'),X18)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|4','thm.dec_to_exhProof.v_rel_cases|split|1')).
fof(f756,plain,(
  ( ! [X2,X0,X1] : (~sP25(X0,X1,X2) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1) | s('type.list.list'('type.conSem.v'),X2) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',sK183(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK185(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X2) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1)) | (s('type.list.list'('type.conSem.v'),X2) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',sK183(X2,X1,X0)),s('type.list.list'('type.conSem.v'),sK185(X2,X1,X0)))) & s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',sK184(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK186(X2,X1,X0)))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',sK183(X2,X1,X0)),s('type.exhSem.v',sK184(X2,X1,X0))))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),sK185(X2,X1,X0)),s('type.list.list'('type.exhSem.v'),sK186(X2,X1,X0))))))) & (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X2) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),X1)) & ! [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X2) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) | s('type.list.list'('type.exhSem.v'),X1) != s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) | sP25(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK183,sK184,sK185,sK186])],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X2) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X1)) | ? [X7,X8,X9,X10] : (s('type.list.list'('type.conSem.v'),X2) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X7),s('type.list.list'('type.conSem.v'),X9))) & s('type.list.list'('type.exhSem.v'),X1) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X8),s('type.list.list'('type.exhSem.v'),X10))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X7),s('type.exhSem.v',X8)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X9),s('type.list.list'('type.exhSem.v'),X10)))))) & (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X2) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),X1)) & ! [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X2) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) | s('type.list.list'('type.exhSem.v'),X1) != s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) | sP25(X0,X1,X2)))),
  inference(rectify,[],[f407])).
fof(f407,plain,(
  ! [X0,X2,X1] : ((~sP25(X0,X2,X1) | (s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X2)) | ? [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) & s('type.list.list'('type.exhSem.v'),X2) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) & (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X1) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),X2)) & ! [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) | s('type.list.list'('type.exhSem.v'),X2) != s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) | sP25(X0,X2,X1)))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0,X2,X1] : ((~sP25(X0,X2,X1) | ((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X1) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X2)) | ? [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) & s('type.list.list'('type.exhSem.v'),X2) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) & p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) & p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6))))))) & (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') != s('type.list.list'('type.conSem.v'),X1) | s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),X2)) & ! [X3,X4,X5,X6] : (s('type.list.list'('type.conSem.v'),X1) != s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X3),s('type.list.list'('type.conSem.v'),X5))) | s('type.list.list'('type.exhSem.v'),X2) != s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X4),s('type.list.list'('type.exhSem.v'),X6))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X3),s('type.exhSem.v',X4)))) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),X5),s('type.list.list'('type.exhSem.v'),X6)))))) | sP25(X0,X2,X1)))),
  inference(nnf_transformation,[],[f243])).
fof(f1001,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) | $spl1),
  inference(cnf_transformation,[],[f1001_D])).
fof(f1001_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1003,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),sK51) | $spl3),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),sK51) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f510,plain,(
  s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') != s('type.list.list'('type.exhSem.v'),sK51) | ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),sK51))))),
  inference(cnf_transformation,[],[f276])).
fof(f13434,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f13433,f518])).
fof(f518,plain,(
  ( ! [X0] : (p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),'const.list.NIL_0'))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),'const.list.NIL_0'))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|4','thm.dec_to_exhProof.v_rel_rules|split|6')).
fof(f13433,plain,(
  ~p(s(bool,'const.dec_to_exhProof.vs_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK50),s('type.list.list'('type.conSem.v'),'const.list.NIL_0'),s('type.list.list'('type.exhSem.v'),'const.list.NIL_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1001,f1002])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 652
% Passive clauses: 7012
% Generated clauses: 12920
% Final active clauses: 649
% Final passive clauses: 5929
% Input formulas: 69
% Initial clauses: 485
% 
% Duplicate literals: 282
% Trivial inequalities: 32
% Fw subsumption resolutions: 572
% Fw demodulations: 12
% 
% Simple tautologies: 204
% Equational tautologies: 60
% Forward subsumptions: 4685
% 
% Binary resolution: 9368
% Factoring: 1588
% Forward superposition: 376
% Backward superposition: 82
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 33
% 
% Split clauses: 198
% Split components: 397
% SAT solver clauses: 406
% SAT solver unit clauses: 18
% SAT solver binary clauses: 67
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 198
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 404
% TWLsolver calls for satisfiability: 348
% 
% Memory used [KB]: 6012
% Time elapsed: 0.452 s
% ------------------------------
% ------------------------------
% Success in time 0.46 s
