Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% (47532)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1112
% Passive clauses: 6775
% Generated clauses: 28664
% Final active clauses: 1054
% Final passive clauses: 5223
% Input formulas: 261
% Initial clauses: 532
% 
% Duplicate literals: 206
% Trivial inequalities: 53
% Fw subsumption resolutions: 1957
% Fw demodulations: 2677
% Bw demodulations: 18
% 
% Simple tautologies: 321
% Equational tautologies: 32
% Forward subsumptions: 16428
% 
% Binary resolution: 6267
% Factoring: 546
% Forward superposition: 8470
% Backward superposition: 7836
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 381
% Split components: 764
% SAT solver clauses: 518
% SAT solver unit clauses: 23
% SAT solver binary clauses: 70
% 
% Sat splits: 381
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 518
% TWLsolver calls for satisfiability: 380
% 
% Memory used [KB]: 15607
% Time elapsed: 0.607 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_knownProof
% (47554)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 696
% Passive clauses: 3078
% Generated clauses: 26646
% Final active clauses: 632
% Final passive clauses: 1999
% Input formulas: 261
% Initial clauses: 532
% 
% Duplicate literals: 111
% Trivial inequalities: 81
% Fw subsumption resolutions: 983
% Fw demodulations: 3276
% Bw demodulations: 407
% Condensations: 96
% Global subsumptions: 27
% 
% Simple tautologies: 303
% Equational tautologies: 16
% Forward subsumptions: 14818
% Fw demodulations to eq. taut.: 45
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2235
% Factoring: 208
% Forward superposition: 8501
% Backward superposition: 6917
% Self superposition: 89
% Equality factoring: 4
% Equality resolution: 20
% 
% Split clauses: 43
% Split components: 86
% SAT solver clauses: 4308
% SAT solver unit clauses: 697
% SAT solver binary clauses: 162
% SAT solver learnt clauses: 55
% SAT solver learnt literals: 19
% 
% Sat splits: 43
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 4199
% TWLsolver calls for satisfiability: 3875
% 
% Memory used [KB]: 22259
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f25096,plain,(
  $false),
  inference(subsumption_resolution,[],[f25095,f1455])).
fof(f1455,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f907,f896])).
fof(f896,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.val_approx_better_approx',aHL_TRUTH)).
fof(f907,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f702])).
fof(f702,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.bool.NOT_CLAUSES|split|2')).
fof(f25095,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f24896,f1689])).
fof(f1689,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK18),s('type.closSem.v',sK17)))),
  inference(unit_resulting_resolution,[],[f895,f984])).
fof(f984,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f711])).
fof(f711,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f332])).
fof(f332,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.bool.EQ_CLAUSES|split|3')).
fof(f895,plain,(
  ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK18),s('type.closSem.v',sK17))))),
  inference(cnf_transformation,[],[f701])).
fof(f701,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',sK18)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK16),s('type.closSem.v',sK17)))) & ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK18),s('type.closSem.v',sK17))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18])],[f613])).
fof(f613,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1)))) & ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X2),s('type.closSem.v',X1)))))),
  inference(flattening,[],[f612])).
fof(f612,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1))))) & ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X2),s('type.closSem.v',X1)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1))))) => p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X2),s('type.closSem.v',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx',X8)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X6),s('type.closSem.v',X7))))) => p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X8),s('type.closSem.v',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx',X8)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X6),s('type.closSem.v',X7))))) => p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X8),s('type.closSem.v',X7)))))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.clos_knownProof.val_approx_better_approx')).
fof(f24896,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK18),s('type.closSem.v',sK17))))),
  inference(superposition,[],[f20045,f15544])).
fof(f15544,plain,(
  s('type.clos_known.val_approx',sK18) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',sK18)))),
  inference(subsumption_resolution,[],[f15476,f896])).
fof(f15476,plain,(
  ~p(s(bool,'T_0')) | s('type.clos_known.val_approx',sK18) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',sK18)))),
  inference(superposition,[],[f1151,f1640])).
fof(f1640,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',sK18)))),
  inference(unit_resulting_resolution,[],[f893,f980])).
fof(f980,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f709])).
fof(f709,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.bool.EQ_CLAUSES|split|1')).
fof(f893,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',sK18))))),
  inference(cnf_transformation,[],[f701])).
fof(f1151,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f781])).
fof(f781,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) & (s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))))),
  inference(nnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) <=> s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))),
  inference(rectify,[],[f202])).
fof(f202,axiom,(
  ! [X6,X8] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx',X8)))) <=> s('type.clos_known.val_approx',X8) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx',X8))))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.clos_knownProof.subapprox_def')).
fof(f20045,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',X0))),s('type.closSem.v',sK17))))) )),
  inference(subsumption_resolution,[],[f19981,f896])).
fof(f19981,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK16),s('type.clos_known.val_approx',X0))),s('type.closSem.v',sK17))))) )),
  inference(superposition,[],[f1243,f1642])).
fof(f1642,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK16),s('type.closSem.v',sK17)))),
  inference(unit_resulting_resolution,[],[f894,f980])).
fof(f894,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK16),s('type.closSem.v',sK17))))),
  inference(cnf_transformation,[],[f701])).
fof(f1243,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2))),s('type.closSem.v',X1))))) )),
  inference(cnf_transformation,[],[f666])).
fof(f666,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1)))) & ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X2),s('type.closSem.v',X1))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2))),s('type.closSem.v',X1)))))),
  inference(ennf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X2),s('type.closSem.v',X1))))) => p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X2))),s('type.closSem.v',X1)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7,X8] : ((p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X6),s('type.closSem.v',X7)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X8),s('type.closSem.v',X7))))) => p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X6),s('type.clos_known.val_approx',X8))),s('type.closSem.v',X7)))))),
  file('thm.clos_knownProof.val_approx_better_approx','thm.clos_knownProof.val_approx_val_merge_I')).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 715
% Passive clauses: 9683
% Generated clauses: 24194
% Final active clauses: 703
% Final passive clauses: 8936
% Input formulas: 261
% Initial clauses: 532
% 
% Duplicate literals: 12
% Trivial inequalities: 8
% Fw subsumption resolutions: 121
% Fw demodulations: 2615
% Bw demodulations: 37
% Condensations: 19
% 
% Simple tautologies: 90
% Equational tautologies: 27
% Forward subsumptions: 11184
% Fw demodulations to eq. taut.: 429
% Bw demodulations to eq. taut.: 7
% 
% Unit resulting resolution: 11592
% Factoring: 1
% Forward superposition: 5681
% Backward superposition: 3496
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 14
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 9693
% SAT solver unit clauses: 8877
% SAT solver binary clauses: 650
% 
% Sat splits: 4
% 
% Memory used [KB]: 33517
% Time elapsed: 0.922 s
% ------------------------------
% ------------------------------
% Success in time 2.27 s
