Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (124111)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 175
% Passive clauses: 1658
% Generated clauses: 19568
% Final active clauses: 148
% Final passive clauses: 1125
% Input formulas: 63
% Initial clauses: 106
% 
% Duplicate literals: 462
% Trivial inequalities: 500
% Fw subsumption resolutions: 2188
% Fw demodulations: 1811
% Bw demodulations: 21
% 
% Simple tautologies: 384
% Equational tautologies: 24
% Forward subsumptions: 12057
% 
% Binary resolution: 95
% Forward superposition: 7312
% Backward superposition: 6660
% Self superposition: 211
% Equality factoring: 52
% Equality resolution: 52
% 
% Split clauses: 282
% Split components: 576
% SAT solver clauses: 637
% SAT solver unit clauses: 19
% SAT solver binary clauses: 133
% 
% Sat splits: 282
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 635
% TWLsolver calls for satisfiability: 118
% 
% Memory used [KB]: 13688
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (124141)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 288
% Generated clauses: 16623
% Final active clauses: 106
% Final passive clauses: 151
% Input formulas: 63
% Initial clauses: 106
% 
% Duplicate literals: 472
% Trivial inequalities: 181
% Fw subsumption resolutions: 1670
% Fw demodulations: 1554
% Bw demodulations: 25
% Condensations: 2474
% Global subsumptions: 8
% 
% Simple tautologies: 188
% Equational tautologies: 23
% Forward subsumptions: 9246
% 
% Binary resolution: 20
% Forward superposition: 5922
% Backward superposition: 3638
% Self superposition: 198
% Equality factoring: 51
% Equality resolution: 35
% 
% Split clauses: 46
% Split components: 92
% SAT solver clauses: 634
% SAT solver unit clauses: 147
% SAT solver binary clauses: 37
% SAT solver learnt clauses: 16
% 
% Sat splits: 46
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 602
% TWLsolver calls for satisfiability: 349
% 
% Memory used [KB]: 12537
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% (124159)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 664
% Passive clauses: 6087
% Generated clauses: 200143
% Final active clauses: 150
% Final passive clauses: 532
% Input formulas: 63
% Initial clauses: 106
% 
% Duplicate literals: 4146
% Trivial inequalities: 1311
% Fw subsumption resolutions: 15538
% Fw demodulations: 28711
% Bw demodulations: 4707
% Condensations: 20672
% Global subsumptions: 429
% 
% Simple tautologies: 1312
% Equational tautologies: 1282
% Forward subsumptions: 119432
% Fw demodulations to eq. taut.: 61
% Bw demodulations to eq. taut.: 112
% 
% Unit resulting resolution: 10975
% Forward superposition: 60141
% Backward superposition: 46160
% Self superposition: 942
% Equality factoring: 204
% Equality resolution: 218
% 
% Split clauses: 263
% Split components: 537
% SAT solver clauses: 7647
% SAT solver unit clauses: 219
% SAT solver binary clauses: 689
% 
% Sat splits: 263
% Sat splitting refutations: 603
% 
% Memory used [KB]: 84817
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f9937,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f314,f315,f316,f317,f318,f320,f340,f321,f331,f338,f384,f392,f388,f349,f350,f351,f352,f353,f341,f394,f354,f356,f367,f427,f418,f447,f468,f469,f497,f471,f501,f494,f493,f535,f537,f534,f533,f578,f575,f574,f573,f597,f598,f599,f600,f708,f710,f711,f712,f715,f725,f811,f813,f814,f815,f818,f828,f930,f934,f935,f1143,f1144,f1153,f1155,f1156,f1157,f1160,f1170,f1210,f1212,f1214,f1215,f1216,f1219,f1229,f1275,f1276,f1274,f1280,f1281,f1278,f1348,f1359,f1360,f1314,f1313,f1366,f1266,f1265,f1367,f932,f1072,f933,f1368,f1390,f1392,f1393,f1394,f1397,f1407,f1618,f1639,f1646,f255,f253,f1696,f234,f1695,f1697,f1699,f1700,f1702,f1704,f1710,f1717,f1718,f1720,f1721,f1723,f1725,f1731,f1356,f1779,f1781,f1782,f1783,f1786,f1794,f2015,f1999,f2050,f2052,f2053,f2055,f2057,f2064,f2132,f2133,f2134,f2135,f2155,f2178,f2179,f2186,f2188,f2189,f2190,f2193,f2201,f2204,f2206,f2207,f2208,f2211,f2219,f2249,f2254,f2256,f2257,f2259,f2261,f2268,f2315,f2317,f2318,f2319,f2322,f2330,f2354,f2356,f2357,f2358,f2361,f2369,f755,f2417,f2419,f2420,f2422,f2424,f2431,f393,f2618,f2620,f2621,f2622,f2625,f2633,f2679,f2634,f2636,f2637,f2639,f2641,f2648,f2680,f3626,f3662,f3663,f3667,f3669,f3672,f3675,f3678,f3691,f3696,f3701,f3704,f3709,f3713,f3718,f3730,f3734,f3736,f3739,f3742,f3745,f3758,f3763,f3768,f3771,f3776,f3780,f3785,f3796,f3798,f3799,f3800,f3803,f3811,f3865,f3812,f3814,f3815,f3817,f3819,f3826,f3866,f3871,f3873,f3927,f3928,f2350,f3936,f3938,f3937,f3946,f3948,f3949,f3951,f3953,f3960,f3999,f4167,f390,f386,f4017,f4019,f4020,f4021,f4024,f4032,f4168,f4035,f4037,f4038,f4040,f4042,f4049,f4169,f4056,f4058,f4059,f4061,f4063,f4070,f4170,f4182,f4184,f4307,f4309,f4310,f4311,f4314,f4322,f4354,f4355,f4361,f4363,f4364,f4365,f4368,f4376,f1201,f4413,f4415,f4416,f4417,f4420,f4428,f4458,f4460,f4461,f3719,f4585,f4588,f4591,f4593,f4596,f4600,f2101,f4610,f4611,f4639,f4641,f4642,f4644,f4646,f4653,f4829,f4832,f4835,f4837,f4840,f4844,f2309,f4854,f4855,f4860,f4862,f4863,f4864,f4867,f4875,f4908,f4910,f4912,f4907,f4916,f4915,f4988,f4990,f4991,f4993,f4995,f5002,f5048,f5116,f5118,f6522,f6529,f6519,f6543,f6518,f6515,f6514,f6513,f6512,f6511,f6510,f6509,f6508,f6505,f6550,f6504,f6501,f6557,f6500,f6564,f6497,f6569,f6571,f6494,f6491,f6545,f6578,f6579,f6490,f6584,f6489,f6488,f6589,f6484,f6651,f6654,f6656,f6658,f6660,f6706,f6632,f6712,f4352,f6713,f6473,f6471,f6468,f6723,f6724,f6725,f6726,f6727,f6728,f6729,f6730,f6462,f6731,f5119,f6461,f6732,f5117,f6733,f6734,f6735,f6736,f6737,f6740,f6745,f6746,f6747,f6750,f6752,f6753,f6757,f6758,f6759,f6761,f6762,f6460,f6764,f6459,f6458,f6765,f6766,f6768,f6769,f6770,f6771,f6772,f6774,f6775,f6776,f6777,f6778,f6780,f6782,f6783,f6784,f6788,f6790,f6793,f6794,f6795,f6796,f6445,f6444,f6556,f6814,f6815,f6816,f6817,f6818,f6819,f6820,f6826,f6830,f6832,f6834,f6839,f6841,f6844,f6847,f6850,f6852,f6853,f6855,f6443,f6857,f6442,f6441,f6861,f6864,f6866,f6440,f6439,f6869,f6870,f6871,f6876,f6879,f6880,f6885,f1365,f6886,f6436,f6789,f6888,f4911,f5861,f6903,f5863,f6905,f5865,f6907,f6908,f6909,f6910,f6911,f5872,f6913,f5874,f6915,f5879,f5880,f5881,f5882,f5883,f5884,f5885,f6916,f6919,f6920,f6921,f6922,f6924,f6925,f6926,f6927,f6928,f6930,f6931,f6933,f6937,f6938,f6431,f6939,f6940,f6941,f6942,f6429,f6943,f6945,f6948,f6951,f6952,f6953,f6956,f6957,f6958,f6961,f6962,f6965,f6966,f6917,f6967,f6970,f6791,f6972,f6984,f6985,f6419,f6987,f6988,f6417,f6989,f6990,f6991,f7001,f7004,f7006,f7009,f7011,f7017,f7022,f7023,f7024,f7027,f7039,f7044,f7049,f7052,f7057,f7064,f4697,f7075,f7076,f7082,f6527,f6378,f7098,f7100,f7101,f7103,f7105,f7112,f7166,f7171,f7173,f7174,f7149,f6546,f7175,f7919,f7918,f7917,f7914,f7913,f7926,f7911,f7908,f7907,f7929,f7930,f7904,f7901,f7900,f7934,f7899,f7937,f7939,f7942,f7943,f7944,f7946,f7948,f7949,f7952,f7953,f7954,f7955,f7956,f7898,f7957,f7897,f7896,f7958,f7959,f7961,f7962,f7963,f7964,f7965,f7966,f7968,f7969,f7970,f7971,f7972,f7974,f7975,f7976,f7893,f7980,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7990,f8000,f8002,f8003,f8004,f8005,f8006,f8008,f8007,f8009,f8012,f8015,f8016,f8017,f8020,f8021,f8022,f8025,f8026,f8027,f8028,f8029,f7891,f8030,f7890,f7889,f8031,f8032,f8033,f7888,f7887,f8034,f8036,f8038,f8039,f8040,f8041,f8042,f8044,f7884,f8046,f7578,f6902,f8055,f7580,f6904,f8056,f7582,f6906,f7583,f4183,f8057,f8058,f6572,f8059,f8060,f6570,f7589,f6912,f7590,f3627,f7591,f6914,f7592,f3872,f7596,f7597,f7598,f7599,f7600,f7601,f7602,f8061,f8063,f8066,f8067,f8068,f8069,f8071,f8072,f8073,f8074,f8075,f8077,f8078,f8080,f8082,f8084,f8086,f8087,f8088,f7878,f8089,f8090,f8091,f7876,f8092,f8093,f8094,f8095,f8098,f8101,f8102,f8103,f8106,f8107,f8108,f8111,f8112,f8115,f8116,f8064,f8117,f8121,f8122,f7922,f8124,f8062,f8126,f7865,f7920,f8130,f8131,f7864,f8133,f931,f8134,f7862,f8135,f8136,f8137,f6944,f7857,f536,f8148,f8150,f8151,f8153,f8155,f8162,f8222,f6530,f6531,f9373,f9385,f9387,f9389,f9392,f9393,f9427,f7172,f9439,f6653,f9441,f6655,f9443,f6657,f9446,f6659,f9447,f6661,f9448,f9453,f9456,f9461,f9465,f9472,f9480,f9481,f9483,f9498,f9503,f9508,f9511,f9516,f9523,f9533,f9545,f9550,f9553,f9558,f9562,f9569,f9577,f9578,f496,f9580,f6763,f9595,f9600,f9605,f9608,f9613,f9620,f6741,f9630,f6779,f9716,f6773,f9701,f7025,f1277,f9729,f9731,f9732,f9734,f9736,f9743,f9803,f9808,f1355,f9810,f9811,f9812,f9813,f9809,f9823,f9825,f9826,f9827,f9830,f9838,f1239,f9868,f9871,f9873,f9874,f9875,f9870,f9879,f9881,f9882,f339,f9883,f343,f9886,f342,f9894,f1418,f9925,f9928,f9930,f344,f9931,f9932,f6889,f9927,f9935,f9936])).
fof(f9936,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f390,f311])).
fof(f311,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(resolution,[],[f205,f275])).
fof(f275,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3)))) | (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) & s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3)))))))) & (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) | s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) != s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3)))))) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3)))) | (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) & s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3)))))))) & ((s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) | s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) != s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3))))))) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3)))) <=> (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) & s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3))))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3)))) <=> (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) & s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3))))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9),s('type.decSem.state'(X6),X8),s('type.exhSem.state'(X6),X7)))) <=> (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X6),X8))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X6),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X6),X8))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X6),X7)))))) & s('type.ffi.ffi_state'(X6),'const.decSem.state_ffi_1'(s('type.decSem.state'(X6),X8))) = s('type.ffi.ffi_state'(X6),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X6),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X6),X8))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X6),X7))))))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.dec_to_exhProof.state_rel_def')).
fof(f205,plain,(
  p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))))),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) & ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f162])).
fof(f162,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X1)))) & ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.decSem.state'(X0),X2))),s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.exhSem.state'(X0),X1)))))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X1)))) => p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.decSem.state'(X0),X2))),s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.exhSem.state'(X0),X1)))))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X1)))) => p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.decSem.state'(X0),X2))),s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X4))),s('type.exhSem.state'(X0),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9),s('type.decSem.state'(X6),X8),s('type.exhSem.state'(X6),X7)))) => p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9),s('type.decSem.state'(X6),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.decSem.state'(X6),X8))),s('type.exhSem.state'(X6),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.exhSem.state'(X6),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9),s('type.decSem.state'(X6),X8),s('type.exhSem.state'(X6),X7)))) => p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9),s('type.decSem.state'(X6),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.decSem.state'(X6),X8))),s('type.exhSem.state'(X6),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X10))),s('type.exhSem.state'(X6),X7)))))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.dec_to_exhProof.state_rel_with_clock')).
fof(f9935,plain,(
  $false | $spl15),
  inference(subsumption_resolution,[],[f386,f313])).
fof(f313,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(resolution,[],[f205,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3)))))) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f9927,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | $spl242),
  inference(cnf_transformation,[],[f9927_D])).
fof(f9927_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f6889,plain,(
  ( ! [X436,X439] : (p(s(bool,X436)) | s(bool,X436) = s(bool,X439)) ) | $spl184),
  inference(cnf_transformation,[],[f6889_D])).
fof(f6889_D,plain,(
  ( ! [X436,X439] : (p(s(bool,X436)) | s(bool,X436) = s(bool,X439)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f9932,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9915,f2198])).
fof(f2198,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,X48) = s(bool,X49) | p(s(bool,X48))) )),
  inference(resolution,[],[f1796,f263])).
fof(f263,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.EQ_CLAUSES|split|2')).
fof(f1796,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(superposition,[],[f264,f264])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f9915,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f313,f2782])).
fof(f2782,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f2328,f306])).
fof(f306,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.dec_to_exhProof.state_rel_with_clock',aHL_BOOL_CASES)).
fof(f2328,plain,(
  ( ! [X50,X51] : (s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X51)) )),
  inference(resolution,[],[f1797,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.EQ_CLAUSES|split|1')).
fof(f1797,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f264,f306])).
fof(f9931,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl244),
  inference(cnf_transformation,[],[f9931_D])).
fof(f9931_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f344,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl12),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f9930,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f9914,f267])).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f9914,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f313,f2782])).
fof(f9928,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9907,f267])).
fof(f9907,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f313,f2328])).
fof(f9925,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9906,f267])).
fof(f9906,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f313,f2328])).
fof(f1418,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) ) | $spl56),
  inference(cnf_transformation,[],[f1418_D])).
fof(f1418_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f9894,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f221])).
fof(f221,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13,X18,X14,X19] : ((s(bool,X13) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X14) = s(bool,X19))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CONG')).
fof(f342,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl8),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9886,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f313,f224])).
fof(f224,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f343,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl10),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f9883,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f223])).
fof(f223,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f339,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f9882,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f9881,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f313,f220])).
fof(f220,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f9879,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f218])).
fof(f218,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f9870,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | $spl238),
  inference(cnf_transformation,[],[f9870_D])).
fof(f9870_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f9875,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9859,f2198])).
fof(f9859,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X15)) )),
  inference(superposition,[],[f311,f2782])).
fof(f9874,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl240),
  inference(cnf_transformation,[],[f9874_D])).
fof(f9874_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f9873,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f9858,f267])).
fof(f9858,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f311,f2782])).
fof(f9871,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9851,f267])).
fof(f9851,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f311,f2328])).
fof(f9868,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f9850,f267])).
fof(f9850,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f311,f2328])).
fof(f1239,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) ) | $spl38),
  inference(cnf_transformation,[],[f1239_D])).
fof(f1239_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9838,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f307])).
fof(f9830,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f311,f224])).
fof(f9827,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f223])).
fof(f9826,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f222])).
fof(f9825,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f311,f220])).
fof(f9823,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f218])).
fof(f9809,plain,(
  ( ! [X68,X66,X67,X65,X63] : (s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s(bool,'F_0') = s(bool,X63) | s('type.ffi.ffi_state'(X66),'const.decSem.state_ffi_1'(s('type.decSem.state'(X66),X67))) = s('type.ffi.ffi_state'(X66),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X66),X68)))) ) | $spl234),
  inference(cnf_transformation,[],[f9809_D])).
fof(f9809_D,plain,(
  ( ! [X68,X66,X67,X65,X63] : (s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s(bool,'F_0') = s(bool,X63) | s('type.ffi.ffi_state'(X66),'const.decSem.state_ffi_1'(s('type.decSem.state'(X66),X67))) = s('type.ffi.ffi_state'(X66),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X66),X68)))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f9813,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,X104)) | s(bool,X103) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X105),s('type.decSem.state'(X106),X107),s('type.exhSem.state'(X106),X108))) | s('type.ffi.ffi_state'(X106),'const.decSem.state_ffi_1'(s('type.decSem.state'(X106),X107))) = s('type.ffi.ffi_state'(X106),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X106),X108))) | s(bool,'F_0') = s(bool,X103)) ) | $spl68),
  inference(subsumption_resolution,[],[f9783,f263])).
fof(f9783,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,X104)) | s(bool,X103) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X105),s('type.decSem.state'(X106),X107),s('type.exhSem.state'(X106),X108))) | s('type.ffi.ffi_state'(X106),'const.decSem.state_ffi_1'(s('type.decSem.state'(X106),X107))) = s('type.ffi.ffi_state'(X106),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X106),X108))) | s(bool,'F_0') = s(bool,X103) | s(bool,'F_0') = s(bool,X104)) ) | $spl68),
  inference(superposition,[],[f2309,f6418])).
fof(f6418,plain,(
  ( ! [X514,X513] : (s(bool,X513) = s(bool,X514) | s(bool,'F_0') = s(bool,X514) | s(bool,'F_0') = s(bool,X513)) )),
  inference(trivial_inequality_removal,[],[f6095])).
fof(f6095,plain,(
  ( ! [X514,X513] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X513) = s(bool,X514) | s(bool,'F_0') = s(bool,X514) | s(bool,'F_0') = s(bool,X513)) )),
  inference(superposition,[],[f992,f2782])).
fof(f992,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f329,f327])).
fof(f327,plain,(
  s(bool,'F_0') != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))),
  inference(resolution,[],[f205,f263])).
fof(f329,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))),
  inference(resolution,[],[f205,f268])).
fof(f9812,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X98)) | s(bool,X97) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X99),s('type.decSem.state'(X100),X101),s('type.exhSem.state'(X100),X102))) | s('type.ffi.ffi_state'(X100),'const.decSem.state_ffi_1'(s('type.decSem.state'(X100),X101))) = s('type.ffi.ffi_state'(X100),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X100),X102))) | s(bool,'F_0') = s(bool,X97)) ) | $spl68),
  inference(subsumption_resolution,[],[f9782,f263])).
fof(f9782,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X98)) | s(bool,X97) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X99),s('type.decSem.state'(X100),X101),s('type.exhSem.state'(X100),X102))) | s('type.ffi.ffi_state'(X100),'const.decSem.state_ffi_1'(s('type.decSem.state'(X100),X101))) = s('type.ffi.ffi_state'(X100),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X100),X102))) | s(bool,'F_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X97)) ) | $spl68),
  inference(superposition,[],[f2309,f6418])).
fof(f9811,plain,(
  ( ! [X70,X76,X74,X72,X75,X73] : (s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,X70) = s(bool,X76) | s('type.ffi.ffi_state'(X73),'const.decSem.state_ffi_1'(s('type.decSem.state'(X73),X74))) = s('type.ffi.ffi_state'(X73),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X73),X75)))) ) | $spl236),
  inference(cnf_transformation,[],[f9811_D])).
fof(f9811_D,plain,(
  ( ! [X70,X76,X74,X72,X75,X73] : (s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,X70) = s(bool,X76) | s('type.ffi.ffi_state'(X73),'const.decSem.state_ffi_1'(s('type.decSem.state'(X73),X74))) = s('type.ffi.ffi_state'(X73),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X73),X75)))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f9810,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s('type.ffi.ffi_state'(X73),'const.decSem.state_ffi_1'(s('type.decSem.state'(X73),X74))) = s('type.ffi.ffi_state'(X73),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X73),X75))) | s(bool,X70) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) ) | $spl68),
  inference(subsumption_resolution,[],[f9778,f263])).
fof(f9778,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s('type.ffi.ffi_state'(X73),'const.decSem.state_ffi_1'(s('type.decSem.state'(X73),X74))) = s('type.ffi.ffi_state'(X73),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X73),X75))) | s(bool,X70) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X71)) ) | $spl68),
  inference(superposition,[],[f2309,f2782])).
fof(f1355,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | $spl46),
  inference(cnf_transformation,[],[f1355_D])).
fof(f1355_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9808,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s('type.ffi.ffi_state'(X66),'const.decSem.state_ffi_1'(s('type.decSem.state'(X66),X67))) = s('type.ffi.ffi_state'(X66),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X66),X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X63)) ) | $spl68),
  inference(subsumption_resolution,[],[f9777,f806])).
fof(f806,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f785,f234])).
fof(f785,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,'F_0')) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f368,f306])).
fof(f368,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(resolution,[],[f206,f307])).
fof(f206,plain,(
  ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))))),
  inference(cnf_transformation,[],[f172])).
fof(f9777,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s('type.ffi.ffi_state'(X66),'const.decSem.state_ffi_1'(s('type.decSem.state'(X66),X67))) = s('type.ffi.ffi_state'(X66),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X66),X68))) | s(bool,X64) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X63)) ) | $spl68),
  inference(superposition,[],[f2309,f2782])).
fof(f9803,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s('type.ffi.ffi_state'(X2),'const.decSem.state_ffi_1'(s('type.decSem.state'(X2),X3))) = s('type.ffi.ffi_state'(X2),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X2),X4)))) ) | $spl232),
  inference(cnf_transformation,[],[f9803_D])).
fof(f9803_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s('type.ffi.ffi_state'(X2),'const.decSem.state_ffi_1'(s('type.decSem.state'(X2),X3))) = s('type.ffi.ffi_state'(X2),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X2),X4)))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f9743,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,X94) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X95),s('type.decSem.state'(X96),X97),s('type.exhSem.state'(X96),X98))) | s('type.ffi.ffi_state'(X96),'const.decSem.state_ffi_1'(s('type.decSem.state'(X96),X97))) = s('type.ffi.ffi_state'(X96),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X96),X98))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X94) != s(bool,X100)) ) | $spl68),
  inference(resolution,[],[f2309,f307])).
fof(f9736,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X57),s('type.decSem.state'(X58),X59),s('type.exhSem.state'(X58),X60))) | s('type.ffi.ffi_state'(X58),'const.decSem.state_ffi_1'(s('type.decSem.state'(X58),X59))) = s('type.ffi.ffi_state'(X58),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X58),X60))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X56) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl68),
  inference(resolution,[],[f2309,f224])).
fof(f9734,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X41),s('type.decSem.state'(X42),X43),s('type.exhSem.state'(X42),X44))) | s('type.ffi.ffi_state'(X42),'const.decSem.state_ffi_1'(s('type.decSem.state'(X42),X43))) = s('type.ffi.ffi_state'(X42),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X42),X44))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X40) != s(bool,X46)) ) | $spl68),
  inference(resolution,[],[f2309,f223])).
fof(f9732,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28))) | s('type.ffi.ffi_state'(X26),'const.decSem.state_ffi_1'(s('type.decSem.state'(X26),X27))) = s('type.ffi.ffi_state'(X26),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X26),X28))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X24) != s(bool,X30)) ) | $spl68),
  inference(resolution,[],[f2309,f222])).
fof(f9731,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X17),s('type.decSem.state'(X18),X19),s('type.exhSem.state'(X18),X20))) | s('type.ffi.ffi_state'(X18),'const.decSem.state_ffi_1'(s('type.decSem.state'(X18),X19))) = s('type.ffi.ffi_state'(X18),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X18),X20))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X16) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl68),
  inference(resolution,[],[f2309,f220])).
fof(f9729,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s('type.ffi.ffi_state'(X2),'const.decSem.state_ffi_1'(s('type.decSem.state'(X2),X3))) = s('type.ffi.ffi_state'(X2),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X2),X4))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X0) != s(bool,X6)) ) | $spl68),
  inference(resolution,[],[f2309,f218])).
fof(f1277,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) | $spl42),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7025,plain,(
  ( ! [X372,X369,X371] : (s(bool,'T_0') != s(bool,X369) | s(bool,'T_0') = s(bool,X372) | s(bool,X371) = s(bool,X372) | s(bool,X369) = s(bool,X371)) ) | $spl202),
  inference(cnf_transformation,[],[f7025_D])).
fof(f7025_D,plain,(
  ( ! [X372,X369,X371] : (s(bool,'T_0') != s(bool,X369) | s(bool,'T_0') = s(bool,X372) | s(bool,X371) = s(bool,X372) | s(bool,X369) = s(bool,X371)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f9701,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(duplicate_literal_removal,[],[f9689])).
fof(f9689,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f6894,f2782])).
fof(f6894,plain,(
  ( ! [X528,X529] : (s(bool,'T_0') != s(bool,X528) | s(bool,X528) = s(bool,X529) | s(bool,'F_0') = s(bool,X529)) )),
  inference(subsumption_resolution,[],[f5854,f992])).
fof(f5854,plain,(
  ( ! [X528,X529] : (s(bool,'T_0') != s(bool,X528) | s(bool,X528) = s(bool,X529) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X529)) )),
  inference(superposition,[],[f992,f2782])).
fof(f6773,plain,(
  ( ! [X313,X315,X316] : (s(bool,'F_0') = s(bool,X315) | s(bool,'F_0') = s(bool,X316) | s(bool,X313) = s(bool,X316) | s(bool,X313) != s(bool,X315)) ) | $spl164),
  inference(cnf_transformation,[],[f6773_D])).
fof(f6773_D,plain,(
  ( ! [X313,X315,X316] : (s(bool,'F_0') = s(bool,X315) | s(bool,'F_0') = s(bool,X316) | s(bool,X313) = s(bool,X316) | s(bool,X313) != s(bool,X315)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f9716,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X48)) )),
  inference(duplicate_literal_removal,[],[f9657])).
fof(f9657,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f6894,f2782])).
fof(f6779,plain,(
  ( ! [X354,X351,X353] : (p(s(bool,X353)) | s(bool,'F_0') = s(bool,X354) | s(bool,X351) = s(bool,X354) | s(bool,X351) = s(bool,X353)) ) | $spl166),
  inference(cnf_transformation,[],[f6779_D])).
fof(f6779_D,plain,(
  ( ! [X354,X351,X353] : (p(s(bool,X353)) | s(bool,'F_0') = s(bool,X354) | s(bool,X351) = s(bool,X354) | s(bool,X351) = s(bool,X353)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f9630,plain,(
  ( ! [X300,X302,X301,X299] : (p(s(bool,X301)) | p(s(bool,X300)) | s(bool,X299) = s(bool,X301) | s(bool,X299) = s(bool,X302) | s(bool,'F_0') = s(bool,X302)) ) | $spl66),
  inference(subsumption_resolution,[],[f9126,f264])).
fof(f9126,plain,(
  ( ! [X300,X302,X301,X299] : (s(bool,'F_0') != s(bool,X300) | p(s(bool,X301)) | p(s(bool,X300)) | s(bool,X299) = s(bool,X301) | s(bool,X299) = s(bool,X302) | s(bool,'F_0') = s(bool,X302)) ) | $spl66),
  inference(superposition,[],[f2249,f6418])).
fof(f6741,plain,(
  ( ! [X24,X29,X27] : (~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X29) | s(bool,X24) = s(bool,X29) | s(bool,X24) != s(bool,X27)) ) | $spl158),
  inference(cnf_transformation,[],[f6741_D])).
fof(f6741_D,plain,(
  ( ! [X24,X29,X27] : (~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X29) | s(bool,X24) = s(bool,X29) | s(bool,X24) != s(bool,X27)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f9620,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(subsumption_resolution,[],[f9102,f271])).
fof(f271,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.dec_to_exhProof.state_rel_with_clock',aHL_FALSITY)).
fof(f9102,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,'F_0')) | p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(superposition,[],[f307,f6418])).
fof(f9613,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f9078,f271])).
fof(f9078,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,'F_0')) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f224,f6418])).
fof(f9608,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f9070,f271])).
fof(f9070,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f223,f6418])).
fof(f9605,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f9065,f271])).
fof(f9065,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'F_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f222,f6418])).
fof(f9600,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f9058,f271])).
fof(f9058,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f220,f6418])).
fof(f9595,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f9051,f271])).
fof(f9051,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f218,f6418])).
fof(f6763,plain,(
  ( ! [X242,X240] : (s(bool,'T_0') != s(bool,X240) | s(bool,'F_0') = s(bool,X242) | s(bool,X240) = s(bool,X242)) ) | $spl160),
  inference(cnf_transformation,[],[f6763_D])).
fof(f6763_D,plain,(
  ( ! [X242,X240] : (s(bool,'T_0') != s(bool,X240) | s(bool,'F_0') = s(bool,X242) | s(bool,X240) = s(bool,X242)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f9580,plain,(
  ( ! [X325,X323,X324] : (s(bool,'T_0') != s(bool,X323) | p(s(bool,X324)) | s(bool,X323) = s(bool,X325) | s(bool,'F_0') = s(bool,X325)) )),
  inference(subsumption_resolution,[],[f9021,f264])).
fof(f9021,plain,(
  ( ! [X325,X323,X324] : (s(bool,'T_0') != s(bool,X323) | s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | s(bool,X323) = s(bool,X325) | s(bool,'F_0') = s(bool,X325)) )),
  inference(superposition,[],[f2028,f6418])).
fof(f2028,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | p(s(bool,X0))) )),
  inference(superposition,[],[f1729,f264])).
fof(f1729,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | s(bool,'T_0') != s(bool,X35)) )),
  inference(resolution,[],[f263,f267])).
fof(f496,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) | $spl26),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9578,plain,(
  ( ! [X300,X301,X299] : (s(bool,'T_0') != s(bool,X299) | s(bool,'F_0') != s(bool,X300) | s(bool,X299) = s(bool,X301) | s(bool,'F_0') = s(bool,X301)) ) | $spl44),
  inference(subsumption_resolution,[],[f9013,f1729])).
fof(f9013,plain,(
  ( ! [X300,X301,X299] : (s(bool,'T_0') != s(bool,X299) | s(bool,'T_0') = s(bool,X300) | s(bool,'F_0') != s(bool,X300) | s(bool,X299) = s(bool,X301) | s(bool,'F_0') = s(bool,X301)) ) | $spl44),
  inference(superposition,[],[f1278,f6418])).
fof(f9577,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,X293) != s(bool,X294) | s(bool,'F_0') = s(bool,X294) | s(bool,'T_0') = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296)) )),
  inference(subsumption_resolution,[],[f9011,f306])).
fof(f9011,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,X293) != s(bool,X294) | s(bool,'F_0') = s(bool,X294) | s(bool,'F_0') != s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296)) )),
  inference(superposition,[],[f906,f6418])).
fof(f906,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f824,f306])).
fof(f824,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X50) | s(bool,X50) != s(bool,X51)) )),
  inference(resolution,[],[f796,f263])).
fof(f796,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f768,f308])).
fof(f308,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f231])).
fof(f231,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X13] : (s(X6,X13) = s(X6,X13) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.REFL_CLAUSE')).
fof(f768,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f368,f306])).
fof(f9569,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f8999,f271])).
fof(f8999,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | p(s(bool,'F_0')) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(superposition,[],[f307,f6418])).
fof(f9562,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(subsumption_resolution,[],[f8974,f271])).
fof(f8974,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,'F_0')) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f224,f6418])).
fof(f9558,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f8967,f271])).
fof(f8967,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | p(s(bool,'F_0')) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f223,f6418])).
fof(f9553,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f8960,f271])).
fof(f8960,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | p(s(bool,'F_0')) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(superposition,[],[f222,f6418])).
fof(f9550,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8955,f271])).
fof(f8955,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,'F_0')) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f220,f6418])).
fof(f9545,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8945,f271])).
fof(f8945,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,'F_0')) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f218,f6418])).
fof(f9533,plain,(
  ( ! [X300,X302,X301,X299] : (p(s(bool,X301)) | p(s(bool,X300)) | s(bool,X299) = s(bool,X301) | s(bool,X299) = s(bool,X302) | s(bool,'F_0') = s(bool,X302)) ) | $spl66),
  inference(subsumption_resolution,[],[f8883,f264])).
fof(f8883,plain,(
  ( ! [X300,X302,X301,X299] : (s(bool,'F_0') != s(bool,X300) | p(s(bool,X301)) | p(s(bool,X300)) | s(bool,X299) = s(bool,X301) | s(bool,X299) = s(bool,X302) | s(bool,'F_0') = s(bool,X302)) ) | $spl66),
  inference(superposition,[],[f2249,f6418])).
fof(f9523,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(subsumption_resolution,[],[f8859,f271])).
fof(f8859,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,'F_0')) | p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X218) != s(bool,X220) | s(bool,X218) = s(bool,X221) | s(bool,'F_0') = s(bool,X221)) )),
  inference(superposition,[],[f307,f6418])).
fof(f9516,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f8835,f271])).
fof(f8835,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,'F_0')) | ~p(s(bool,X156)) | ~p(s(bool,X157)) | s(bool,X155) != s(bool,X157) | s(bool,X156) != s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f224,f6418])).
fof(f9511,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f8827,f271])).
fof(f8827,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,'F_0')) | p(s(bool,X116)) | ~p(s(bool,X117)) | ~p(s(bool,X118)) | s(bool,X115) != s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f223,f6418])).
fof(f9508,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f8822,f271])).
fof(f8822,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,'F_0')) | p(s(bool,X91)) | ~p(s(bool,X92)) | s(bool,X91) != s(bool,X93) | s(bool,X90) != s(bool,X92) | s(bool,X90) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f222,f6418])).
fof(f9503,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f8815,f271])).
fof(f8815,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,'F_0')) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X55) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f220,f6418])).
fof(f9498,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f8808,f271])).
fof(f8808,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22) | s(bool,X20) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f218,f6418])).
fof(f9483,plain,(
  ( ! [X325,X323,X324] : (s(bool,'T_0') != s(bool,X323) | p(s(bool,X324)) | s(bool,X323) = s(bool,X325) | s(bool,'F_0') = s(bool,X325)) )),
  inference(subsumption_resolution,[],[f8778,f264])).
fof(f8778,plain,(
  ( ! [X325,X323,X324] : (s(bool,'T_0') != s(bool,X323) | s(bool,'F_0') != s(bool,X324) | p(s(bool,X324)) | s(bool,X323) = s(bool,X325) | s(bool,'F_0') = s(bool,X325)) )),
  inference(superposition,[],[f2028,f6418])).
fof(f9481,plain,(
  ( ! [X300,X301,X299] : (s(bool,'T_0') != s(bool,X299) | s(bool,'F_0') != s(bool,X300) | s(bool,X299) = s(bool,X301) | s(bool,'F_0') = s(bool,X301)) ) | $spl44),
  inference(subsumption_resolution,[],[f8770,f1729])).
fof(f8770,plain,(
  ( ! [X300,X301,X299] : (s(bool,'T_0') != s(bool,X299) | s(bool,'T_0') = s(bool,X300) | s(bool,'F_0') != s(bool,X300) | s(bool,X299) = s(bool,X301) | s(bool,'F_0') = s(bool,X301)) ) | $spl44),
  inference(superposition,[],[f1278,f6418])).
fof(f9480,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,X293) != s(bool,X294) | s(bool,'F_0') = s(bool,X294) | s(bool,'T_0') = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296)) )),
  inference(subsumption_resolution,[],[f8768,f306])).
fof(f8768,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,X293) != s(bool,X294) | s(bool,'F_0') = s(bool,X294) | s(bool,'F_0') != s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296)) )),
  inference(superposition,[],[f906,f6418])).
fof(f9472,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f8756,f271])).
fof(f8756,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,X253) != s(bool,X254) | p(s(bool,X255)) | ~p(s(bool,X254)) | p(s(bool,'F_0')) | s(bool,X253) = s(bool,X256) | s(bool,'F_0') = s(bool,X256)) )),
  inference(superposition,[],[f307,f6418])).
fof(f9465,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(subsumption_resolution,[],[f8731,f271])).
fof(f8731,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X186) != s(bool,X187) | p(s(bool,'F_0')) | ~p(s(bool,X188)) | ~p(s(bool,X187)) | s(bool,X188) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f224,f6418])).
fof(f9461,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f8724,f271])).
fof(f8724,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,X151) != s(bool,X152) | p(s(bool,X153)) | ~p(s(bool,X152)) | ~p(s(bool,X154)) | p(s(bool,'F_0')) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f223,f6418])).
fof(f9456,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f8717,f271])).
fof(f8717,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,X116) != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X117)) | p(s(bool,'F_0')) | s(bool,X118) != s(bool,X119) | s(bool,X116) = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) )),
  inference(superposition,[],[f222,f6418])).
fof(f9453,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8712,f271])).
fof(f8712,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,'F_0')) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f220,f6418])).
fof(f9448,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8702,f271])).
fof(f8702,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) != s(bool,X42) | p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,'F_0')) | s(bool,X43) != s(bool,X44) | s(bool,X41) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f218,f6418])).
fof(f6661,plain,(
  ( ! [X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) | $spl146),
  inference(cnf_transformation,[],[f6661_D])).
fof(f6661_D,plain,(
  ( ! [X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f9447,plain,(
  ( ! [X528,X530,X526,X529,X527] : (~p(s(bool,X530)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526),s('type.decSem.state'(X527),X528),s('type.exhSem.state'(X527),X529)))) | s('type.ffi.ffi_state'(X527),'const.decSem.state_ffi_1'(s('type.decSem.state'(X527),X528))) != s('type.ffi.ffi_state'(X527),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X527),X529))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X527),X529)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X527),X528))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X527),X529))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X527),X529)))))) )),
  inference(subsumption_resolution,[],[f8685,f263])).
fof(f8685,plain,(
  ( ! [X528,X530,X526,X529,X527] : (~p(s(bool,X530)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526),s('type.decSem.state'(X527),X528),s('type.exhSem.state'(X527),X529)))) | s('type.ffi.ffi_state'(X527),'const.decSem.state_ffi_1'(s('type.decSem.state'(X527),X528))) != s('type.ffi.ffi_state'(X527),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X527),X529))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X527),X529)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X527),X528))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X527),X529))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X527),X529))))) | s(bool,'F_0') = s(bool,X530)) )),
  inference(superposition,[],[f278,f6418])).
fof(f278,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X3)))))) | s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) != s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X3)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3)))) )),
  inference(cnf_transformation,[],[f203])).
fof(f6659,plain,(
  ( ! [X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) | $spl144),
  inference(cnf_transformation,[],[f6659_D])).
fof(f6659_D,plain,(
  ( ! [X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f9446,plain,(
  ( ! [X520,X516,X518,X517,X519] : (~p(s(bool,X520)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516),s('type.decSem.state'(X517),X518),s('type.exhSem.state'(X517),X519)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X517),X519)))))) | s('type.ffi.ffi_state'(X517),'const.decSem.state_ffi_1'(s('type.decSem.state'(X517),X518))) != s('type.ffi.ffi_state'(X517),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X517),X519))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X517),X518))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X517),X519))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X517),X519)))))) )),
  inference(subsumption_resolution,[],[f8683,f263])).
fof(f8683,plain,(
  ( ! [X520,X516,X518,X517,X519] : (~p(s(bool,X520)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516),s('type.decSem.state'(X517),X518),s('type.exhSem.state'(X517),X519)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X517),X519)))))) | s('type.ffi.ffi_state'(X517),'const.decSem.state_ffi_1'(s('type.decSem.state'(X517),X518))) != s('type.ffi.ffi_state'(X517),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X517),X519))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X517),X518))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X517),X519))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X517),X519))))) | s(bool,'F_0') = s(bool,X520)) )),
  inference(superposition,[],[f278,f6418])).
fof(f6657,plain,(
  ( ! [X573,X571,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) | $spl142),
  inference(cnf_transformation,[],[f6657_D])).
fof(f6657_D,plain,(
  ( ! [X573,X571,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f9443,plain,(
  ( ! [X484,X486,X482,X485,X483] : (~p(s(bool,X486)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X483),X484))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X483),X485)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482),s('type.decSem.state'(X483),X484),s('type.exhSem.state'(X483),X485)))) )),
  inference(subsumption_resolution,[],[f8676,f263])).
fof(f8676,plain,(
  ( ! [X484,X486,X482,X485,X483] : (~p(s(bool,X486)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X483),X484))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X483),X485)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482),s('type.decSem.state'(X483),X484),s('type.exhSem.state'(X483),X485))) | s(bool,'F_0') = s(bool,X486)) )),
  inference(superposition,[],[f277,f6418])).
fof(f6655,plain,(
  ( ! [X561,X559,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) | $spl140),
  inference(cnf_transformation,[],[f6655_D])).
fof(f6655_D,plain,(
  ( ! [X561,X559,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f9441,plain,(
  ( ! [X476,X472,X474,X475,X473] : (~p(s(bool,X476)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X473),X474))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X473),X475)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472),s('type.decSem.state'(X473),X474),s('type.exhSem.state'(X473),X475)))) )),
  inference(subsumption_resolution,[],[f8674,f263])).
fof(f8674,plain,(
  ( ! [X476,X472,X474,X475,X473] : (~p(s(bool,X476)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X473),X474))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X473),X475)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472),s('type.decSem.state'(X473),X474),s('type.exhSem.state'(X473),X475))) | s(bool,'F_0') = s(bool,X476)) )),
  inference(superposition,[],[f275,f6418])).
fof(f6653,plain,(
  s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | $spl138),
  inference(cnf_transformation,[],[f6653_D])).
fof(f6653_D,plain,(
  s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9439,plain,(
  ( ! [X466] : (~p(s(bool,X466)) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f8672,f263])).
fof(f8672,plain,(
  ( ! [X466] : (~p(s(bool,X466)) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X466)) )),
  inference(superposition,[],[f206,f6418])).
fof(f7172,plain,(
  ( ! [X68,X66,X67,X65,X63] : (s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s(bool,'F_0') = s(bool,X63) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X66),X67))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X66),X68)))) ) | $spl208),
  inference(cnf_transformation,[],[f7172_D])).
fof(f7172_D,plain,(
  ( ! [X68,X66,X67,X65,X63] : (s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s(bool,'F_0') = s(bool,X63) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X66),X67))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X66),X68)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f9427,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,X293)) | s(bool,X292) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X294),s('type.decSem.state'(X295),X296),s('type.exhSem.state'(X295),X297))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'F_0') = s(bool,X292)) ) | $spl64),
  inference(subsumption_resolution,[],[f8615,f263])).
fof(f8615,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,X293)) | s(bool,X292) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X294),s('type.decSem.state'(X295),X296),s('type.exhSem.state'(X295),X297))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'F_0') = s(bool,X292) | s(bool,'F_0') = s(bool,X293)) ) | $spl64),
  inference(superposition,[],[f2101,f6418])).
fof(f9393,plain,(
  ( ! [X528,X530,X526,X529,X527] : (~p(s(bool,X530)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526),s('type.decSem.state'(X527),X528),s('type.exhSem.state'(X527),X529)))) | s('type.ffi.ffi_state'(X527),'const.decSem.state_ffi_1'(s('type.decSem.state'(X527),X528))) != s('type.ffi.ffi_state'(X527),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X527),X529))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X527),X529)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X527),X528))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X527),X529))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X527),X529)))))) )),
  inference(subsumption_resolution,[],[f8531,f263])).
fof(f8531,plain,(
  ( ! [X528,X530,X526,X529,X527] : (~p(s(bool,X530)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526),s('type.decSem.state'(X527),X528),s('type.exhSem.state'(X527),X529)))) | s('type.ffi.ffi_state'(X527),'const.decSem.state_ffi_1'(s('type.decSem.state'(X527),X528))) != s('type.ffi.ffi_state'(X527),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X527),X529))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X527),X529)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X527),X528))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X527),X529))) | s(bool,'F_0') = s(bool,X530) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X526))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X527),X528))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X527),X529)))))) )),
  inference(superposition,[],[f278,f6418])).
fof(f9392,plain,(
  ( ! [X520,X516,X518,X517,X519] : (~p(s(bool,X520)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516),s('type.decSem.state'(X517),X518),s('type.exhSem.state'(X517),X519)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X517),X519)))))) | s('type.ffi.ffi_state'(X517),'const.decSem.state_ffi_1'(s('type.decSem.state'(X517),X518))) != s('type.ffi.ffi_state'(X517),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X517),X519))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X517),X518))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X517),X519))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X517),X519)))))) )),
  inference(subsumption_resolution,[],[f8529,f263])).
fof(f8529,plain,(
  ( ! [X520,X516,X518,X517,X519] : (~p(s(bool,X520)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516),s('type.decSem.state'(X517),X518),s('type.exhSem.state'(X517),X519)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X517),X519)))))) | s('type.ffi.ffi_state'(X517),'const.decSem.state_ffi_1'(s('type.decSem.state'(X517),X518))) != s('type.ffi.ffi_state'(X517),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X517),X519))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X517),X518))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X517),X519))) | s(bool,'F_0') = s(bool,X520) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X516))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X517),X518))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X517),X519)))))) )),
  inference(superposition,[],[f278,f6418])).
fof(f9389,plain,(
  ( ! [X484,X486,X482,X485,X483] : (~p(s(bool,X486)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X483),X484))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X483),X485)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482),s('type.decSem.state'(X483),X484),s('type.exhSem.state'(X483),X485)))) )),
  inference(subsumption_resolution,[],[f8522,f263])).
fof(f8522,plain,(
  ( ! [X484,X486,X482,X485,X483] : (~p(s(bool,X486)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X483),X484))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X483),X485)))))) | s(bool,'F_0') = s(bool,X486) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X482),s('type.decSem.state'(X483),X484),s('type.exhSem.state'(X483),X485)))) )),
  inference(superposition,[],[f277,f6418])).
fof(f9387,plain,(
  ( ! [X476,X472,X474,X475,X473] : (~p(s(bool,X476)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X473),X474))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X473),X475)))))) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472),s('type.decSem.state'(X473),X474),s('type.exhSem.state'(X473),X475)))) )),
  inference(subsumption_resolution,[],[f8520,f263])).
fof(f8520,plain,(
  ( ! [X476,X472,X474,X475,X473] : (~p(s(bool,X476)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X473),X474))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X473),X475)))))) | s(bool,'F_0') = s(bool,X476) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X472),s('type.decSem.state'(X473),X474),s('type.exhSem.state'(X473),X475)))) )),
  inference(superposition,[],[f275,f6418])).
fof(f9385,plain,(
  ( ! [X466] : (~p(s(bool,X466)) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f8518,f263])).
fof(f8518,plain,(
  ( ! [X466] : (~p(s(bool,X466)) | s(bool,'F_0') = s(bool,X466) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f206,f6418])).
fof(f9373,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,X293)) | s(bool,X292) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X294),s('type.decSem.state'(X295),X296),s('type.exhSem.state'(X295),X297))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'F_0') = s(bool,X292)) ) | $spl64),
  inference(subsumption_resolution,[],[f8461,f263])).
fof(f8461,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,X293)) | s(bool,X292) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X294),s('type.decSem.state'(X295),X296),s('type.exhSem.state'(X295),X297))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'F_0') = s(bool,X293) | s(bool,'F_0') = s(bool,X292)) ) | $spl64),
  inference(superposition,[],[f2101,f6418])).
fof(f6531,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl124),
  inference(cnf_transformation,[],[f6531_D])).
fof(f6531_D,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6530,plain,(
  ( ! [X55,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) | $spl122),
  inference(cnf_transformation,[],[f6530_D])).
fof(f6530_D,plain,(
  ( ! [X55,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f8222,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X30)) ) | $spl128),
  inference(subsumption_resolution,[],[f8196,f1797])).
fof(f8196,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X30)) ) | $spl128),
  inference(superposition,[],[f6546,f2782])).
fof(f8162,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X52) != s(bool,X55)) ) | $spl128),
  inference(resolution,[],[f6546,f307])).
fof(f8155,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X35) != s(bool,X38) | s(bool,X37) != s(bool,X39)) ) | $spl128),
  inference(resolution,[],[f6546,f224])).
fof(f8153,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X25) != s(bool,X28)) ) | $spl128),
  inference(resolution,[],[f6546,f223])).
fof(f8151,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X15) != s(bool,X18)) ) | $spl128),
  inference(resolution,[],[f6546,f222])).
fof(f8150,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | ~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X10) != s(bool,X13) | s(bool,X12) != s(bool,X14)) ) | $spl128),
  inference(resolution,[],[f6546,f220])).
fof(f8148,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X3)) ) | $spl128),
  inference(resolution,[],[f6546,f218])).
fof(f536,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) | $spl28),
  inference(cnf_transformation,[],[f536_D])).
fof(f536_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7857,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl124),
  inference(equality_factoring,[],[f6531])).
fof(f6944,plain,(
  ( ! [X478,X479] : (s(bool,X478) != s(bool,X479) | s(bool,'T_0') != s(bool,X479)) ) | $spl198),
  inference(cnf_transformation,[],[f6944_D])).
fof(f6944_D,plain,(
  ( ! [X478,X479] : (s(bool,X478) != s(bool,X479) | s(bool,'T_0') != s(bool,X479)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f8137,plain,(
  ( ! [X447,X446,X448] : (s(bool,X446) != s(bool,X447) | s(bool,'T_0') != s(bool,X447) | s(bool,'F_0') = s(bool,X448)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7831,f1278])).
fof(f7831,plain,(
  ( ! [X447,X446,X448] : (s(bool,'T_0') != s(bool,X446) | s(bool,X446) != s(bool,X447) | s(bool,'T_0') != s(bool,X447) | s(bool,'F_0') = s(bool,X448)) ) | $spl124),
  inference(superposition,[],[f4878,f6531])).
fof(f4878,plain,(
  ( ! [X85,X84] : (s(bool,'F_0') != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'T_0') != s(bool,X85)) )),
  inference(resolution,[],[f3958,f271])).
fof(f3958,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,X61)) | s(bool,X60) != s(bool,X61) | s(bool,X60) != s(bool,X62) | s(bool,'T_0') != s(bool,X62)) )),
  inference(resolution,[],[f2388,f267])).
fof(f2388,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) )),
  inference(superposition,[],[f2028,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.EQ_CLAUSES|split|0')).
fof(f8136,plain,(
  ( ! [X441,X440,X439] : (s(bool,'T_0') != s(bool,X439) | s(bool,'T_0') = s(bool,X440) | s(bool,'F_0') = s(bool,X441)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7829,f1278])).
fof(f7829,plain,(
  ( ! [X441,X440,X439] : (s(bool,'T_0') != s(bool,X439) | s(bool,'T_0') = s(bool,X440) | s(bool,X439) = s(bool,X440) | s(bool,'F_0') = s(bool,X441)) ) | $spl124),
  inference(superposition,[],[f2327,f6531])).
fof(f2327,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,X48) = s(bool,X49)) )),
  inference(resolution,[],[f1797,f263])).
fof(f8135,plain,(
  ( ! [X436,X438,X437] : (s(bool,'T_0') != s(bool,X436) | p(s(bool,X437)) | s(bool,'F_0') = s(bool,X438)) ) | $spl124),
  inference(subsumption_resolution,[],[f7828,f2028])).
fof(f7828,plain,(
  ( ! [X436,X438,X437] : (s(bool,'T_0') != s(bool,X436) | s(bool,X436) = s(bool,X437) | p(s(bool,X437)) | s(bool,'F_0') = s(bool,X438)) ) | $spl124),
  inference(superposition,[],[f2198,f6531])).
fof(f7862,plain,(
  ( ! [X434,X435] : (s(bool,'T_0') != s(bool,X434) | s(bool,'F_0') = s(bool,X435)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7827])).
fof(f7827,plain,(
  ( ! [X434,X435] : (s(bool,'T_0') != s(bool,X434) | s(bool,'T_0') != s(bool,X434) | s(bool,'F_0') = s(bool,X435)) ) | $spl124),
  inference(superposition,[],[f1729,f6531])).
fof(f8134,plain,(
  ( ! [X432,X431,X433] : (~p(s(bool,X432)) | s(bool,X431) != s(bool,X432) | s(bool,'F_0') = s(bool,X433)) ) | ($spl48 | $spl124)),
  inference(subsumption_resolution,[],[f7826,f806])).
fof(f7826,plain,(
  ( ! [X432,X431,X433] : (s(bool,'T_0') != s(bool,X431) | ~p(s(bool,X432)) | s(bool,X431) != s(bool,X432) | s(bool,'F_0') = s(bool,X433)) ) | ($spl48 | $spl124)),
  inference(superposition,[],[f1356,f6531])).
fof(f931,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(cnf_transformation,[],[f931_D])).
fof(f931_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8133,plain,(
  ( ! [X428,X429,X427] : (s(bool,'F_0') = s(bool,X428) | s(bool,X427) != s(bool,X428) | s(bool,'F_0') = s(bool,X429)) ) | $spl124),
  inference(subsumption_resolution,[],[f7824,f826])).
fof(f826,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X55)) )),
  inference(resolution,[],[f796,f268])).
fof(f7824,plain,(
  ( ! [X428,X429,X427] : (s(bool,'T_0') != s(bool,X427) | s(bool,'F_0') = s(bool,X428) | s(bool,X427) != s(bool,X428) | s(bool,'F_0') = s(bool,X429)) ) | $spl124),
  inference(superposition,[],[f824,f6531])).
fof(f7864,plain,(
  ( ! [X426,X425] : (s(bool,'T_0') = s(bool,X425) | s(bool,'F_0') = s(bool,X426)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7823])).
fof(f7823,plain,(
  ( ! [X426,X425] : (s(bool,'T_0') = s(bool,X425) | s(bool,'T_0') = s(bool,X425) | s(bool,'F_0') = s(bool,X426)) ) | $spl124),
  inference(superposition,[],[f306,f6531])).
fof(f8131,plain,(
  ( ! [X422,X423] : (p(s(bool,X422)) | s(bool,'F_0') = s(bool,X423)) ) | $spl124),
  inference(subsumption_resolution,[],[f7821,f267])).
fof(f7821,plain,(
  ( ! [X422,X423] : (s(bool,'T_0') = s(bool,X422) | p(s(bool,X422)) | s(bool,'F_0') = s(bool,X423)) ) | $spl124),
  inference(superposition,[],[f264,f6531])).
fof(f8130,plain,(
  ( ! [X420,X421] : (~p(s(bool,X420)) | s(bool,'F_0') = s(bool,X421)) ) | $spl124),
  inference(subsumption_resolution,[],[f7820,f268])).
fof(f7820,plain,(
  ( ! [X420,X421] : (s(bool,'T_0') != s(bool,X420) | ~p(s(bool,X420)) | s(bool,'F_0') = s(bool,X421)) ) | $spl124),
  inference(superposition,[],[f263,f6531])).
fof(f7920,plain,(
  ( ! [X24,X25,X22] : (s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,X22) = s(bool,X25)) ) | $spl214),
  inference(cnf_transformation,[],[f7920_D])).
fof(f7920_D,plain,(
  ( ! [X24,X25,X22] : (s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,X22) = s(bool,X25)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f7865,plain,(
  ( ! [X341,X339,X340,X338] : (s(bool,'T_0') = s(bool,X339) | s(bool,X338) = s(bool,X340) | s(bool,'F_0') = s(bool,X340) | s(bool,X338) = s(bool,X341)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7792])).
fof(f7792,plain,(
  ( ! [X341,X339,X340,X338] : (s(bool,'T_0') = s(bool,X339) | s(bool,X338) = s(bool,X340) | s(bool,'T_0') = s(bool,X339) | s(bool,'F_0') = s(bool,X340) | s(bool,X338) = s(bool,X341)) ) | $spl124),
  inference(superposition,[],[f2782,f6531])).
fof(f8126,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | p(s(bool,X304)) | p(s(bool,X302)) | s(bool,X302) = s(bool,X305)) ) | ($spl66 | $spl124)),
  inference(subsumption_resolution,[],[f7783,f2028])).
fof(f7783,plain,(
  ( ! [X302,X304,X303,X305] : (s(bool,'T_0') != s(bool,X303) | p(s(bool,X304)) | p(s(bool,X302)) | s(bool,X303) = s(bool,X304) | s(bool,X302) = s(bool,X305)) ) | ($spl66 | $spl124)),
  inference(superposition,[],[f2249,f6531])).
fof(f8062,plain,(
  ( ! [X43,X46,X44] : (s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X46)) ) | $spl228),
  inference(cnf_transformation,[],[f8062_D])).
fof(f8062_D,plain,(
  ( ! [X43,X46,X44] : (s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X46)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8124,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,X268) != s(bool,X269) | s(bool,'T_0') = s(bool,X270) | s(bool,X268) = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(subsumption_resolution,[],[f8123,f308])).
fof(f8123,plain,(
  ( ! [X269,X271,X268,X270] : (~p(s(bool,'T_0')) | s(bool,X268) != s(bool,X269) | s(bool,'T_0') = s(bool,X270) | s(bool,X268) = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(subsumption_resolution,[],[f7773,f6531])).
fof(f7773,plain,(
  ( ! [X269,X271,X268,X270] : (~p(s(bool,'T_0')) | s(bool,X268) != s(bool,X269) | s(bool,'T_0') = s(bool,X270) | s(bool,X269) != s(bool,X270) | s(bool,X268) = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(superposition,[],[f1367,f6531])).
fof(f7922,plain,(
  ( ! [X48,X46] : (s(bool,X46) = s(bool,X48)) ) | $spl216),
  inference(cnf_transformation,[],[f7922_D])).
fof(f7922_D,plain,(
  ( ! [X48,X46] : (s(bool,X46) = s(bool,X48)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f8122,plain,(
  ( ! [X263,X262,X264] : (s(bool,'T_0') = s(bool,X263) | s(bool,X262) = s(bool,X264)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7869,f6531])).
fof(f7869,plain,(
  ( ! [X263,X262,X264] : (s(bool,'T_0') = s(bool,X263) | s(bool,X262) != s(bool,X263) | s(bool,X262) = s(bool,X264)) ) | ($spl44 | $spl124)),
  inference(trivial_inequality_removal,[],[f7771])).
fof(f7771,plain,(
  ( ! [X263,X262,X264] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X263) | s(bool,X262) != s(bool,X263) | s(bool,X262) = s(bool,X264)) ) | ($spl44 | $spl124)),
  inference(superposition,[],[f1278,f6531])).
fof(f8121,plain,(
  ( ! [X255,X257,X254,X256] : (s(bool,'T_0') != s(bool,X255) | s(bool,'T_0') = s(bool,X256) | s(bool,X254) = s(bool,X257)) ) | $spl124),
  inference(subsumption_resolution,[],[f8120,f6531])).
fof(f8120,plain,(
  ( ! [X255,X257,X254,X256] : (s(bool,'T_0') != s(bool,X255) | s(bool,X254) != s(bool,X256) | s(bool,'T_0') = s(bool,X256) | s(bool,X254) = s(bool,X257)) ) | $spl124),
  inference(subsumption_resolution,[],[f7769,f1729])).
fof(f7769,plain,(
  ( ! [X255,X257,X254,X256] : (s(bool,'T_0') != s(bool,X255) | s(bool,'F_0') = s(bool,X255) | s(bool,X254) != s(bool,X256) | s(bool,'T_0') = s(bool,X256) | s(bool,X254) = s(bool,X257)) ) | $spl124),
  inference(superposition,[],[f906,f6531])).
fof(f8117,plain,(
  ( ! [X233,X234,X232] : (s(bool,'T_0') != s(bool,X233) | p(s(bool,X232)) | s(bool,X232) = s(bool,X234)) ) | $spl124),
  inference(subsumption_resolution,[],[f7762,f267])).
fof(f7762,plain,(
  ( ! [X233,X234,X232] : (s(bool,'T_0') != s(bool,X233) | p(s(bool,X232)) | ~p(s(bool,X233)) | s(bool,X232) = s(bool,X234)) ) | $spl124),
  inference(superposition,[],[f368,f6531])).
fof(f8064,plain,(
  ( ! [X80,X83,X79] : (s(bool,X79) != s(bool,X80) | s(bool,X79) = s(bool,X83) | p(s(bool,X80))) ) | $spl230),
  inference(cnf_transformation,[],[f8064_D])).
fof(f8064_D,plain,(
  ( ! [X80,X83,X79] : (s(bool,X79) != s(bool,X80) | s(bool,X79) = s(bool,X83) | p(s(bool,X80))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f8116,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228)) ) | $spl124),
  inference(subsumption_resolution,[],[f7760,f308])).
fof(f7760,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f8115,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'T_0') != s(bool,X214) | p(s(bool,X215)) | p(s(bool,X213)) | s(bool,X213) = s(bool,X216)) ) | $spl124),
  inference(subsumption_resolution,[],[f7757,f267])).
fof(f7757,plain,(
  ( ! [X216,X215,X213,X214] : (s(bool,'T_0') != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X214)) | p(s(bool,X213)) | s(bool,X213) = s(bool,X216)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f8112,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169)) ) | $spl124),
  inference(subsumption_resolution,[],[f7737,f308])).
fof(f7737,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8111,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X145)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,X148) | s(bool,X145) = s(bool,X149)) ) | $spl124),
  inference(subsumption_resolution,[],[f7733,f267])).
fof(f7733,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X146) | p(s(bool,X145)) | ~p(s(bool,X147)) | ~p(s(bool,X146)) | s(bool,X147) != s(bool,X148) | s(bool,X145) = s(bool,X149)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8108,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129)) ) | $spl124),
  inference(subsumption_resolution,[],[f7729,f308])).
fof(f7729,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8107,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X113)) | p(s(bool,X110)) | s(bool,X110) = s(bool,X114)) ) | $spl124),
  inference(subsumption_resolution,[],[f7726,f267])).
fof(f7726,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') != s(bool,X111) | p(s(bool,X112)) | ~p(s(bool,X111)) | ~p(s(bool,X113)) | p(s(bool,X110)) | s(bool,X110) = s(bool,X114)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8106,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104)) ) | $spl124),
  inference(subsumption_resolution,[],[f7724,f308])).
fof(f7724,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8103,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X77) != s(bool,X78) | s(bool,X75) = s(bool,X79)) ) | $spl124),
  inference(subsumption_resolution,[],[f7719,f267])).
fof(f7719,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') != s(bool,X76) | p(s(bool,X77)) | ~p(s(bool,X76)) | p(s(bool,X75)) | s(bool,X77) != s(bool,X78) | s(bool,X75) = s(bool,X79)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8102,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69)) ) | $spl124),
  inference(subsumption_resolution,[],[f7717,f308])).
fof(f7717,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8101,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,X53) | s(bool,X50) = s(bool,X54)) ) | $spl124),
  inference(subsumption_resolution,[],[f7714,f267])).
fof(f7714,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X52)) | ~p(s(bool,X51)) | s(bool,X52) != s(bool,X53) | s(bool,X50) = s(bool,X54)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8098,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34)) ) | $spl124),
  inference(subsumption_resolution,[],[f7710,f308])).
fof(f7710,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8095,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) = s(bool,X4)) ) | $spl124),
  inference(subsumption_resolution,[],[f7704,f267])).
fof(f7704,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) = s(bool,X4)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8094,plain,(
  ( ! [X414,X416,X413,X415] : (p(s(bool,X413)) | s(bool,X414) != s(bool,X415) | s(bool,'T_0') != s(bool,X415) | s(bool,X413) = s(bool,X416)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7701,f1278])).
fof(f7701,plain,(
  ( ! [X414,X416,X413,X415] : (p(s(bool,X413)) | s(bool,'T_0') != s(bool,X414) | s(bool,X414) != s(bool,X415) | s(bool,'T_0') != s(bool,X415) | s(bool,X413) = s(bool,X416)) ) | $spl124),
  inference(superposition,[],[f3958,f6531])).
fof(f8093,plain,(
  ( ! [X408,X407,X409,X406] : (s(bool,'T_0') = s(bool,X408) | s(bool,'T_0') = s(bool,X407) | s(bool,X406) = s(bool,X409)) ) | $spl124),
  inference(subsumption_resolution,[],[f7875,f6531])).
fof(f7875,plain,(
  ( ! [X408,X407,X409,X406] : (s(bool,X406) != s(bool,X407) | s(bool,'T_0') = s(bool,X408) | s(bool,'T_0') = s(bool,X407) | s(bool,X406) = s(bool,X409)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7699])).
fof(f7699,plain,(
  ( ! [X408,X407,X409,X406] : (s(bool,X406) != s(bool,X407) | s(bool,'T_0') = s(bool,X408) | s(bool,'T_0') = s(bool,X408) | s(bool,'T_0') = s(bool,X407) | s(bool,X406) = s(bool,X409)) ) | $spl124),
  inference(superposition,[],[f3625,f6531])).
fof(f3625,plain,(
  ( ! [X344,X343,X345] : (s(bool,X343) != s(bool,X344) | s(bool,'T_0') = s(bool,X345) | s(bool,X344) = s(bool,X345) | s(bool,'T_0') = s(bool,X343)) )),
  inference(subsumption_resolution,[],[f2925,f992])).
fof(f2925,plain,(
  ( ! [X344,X343,X345] : (s(bool,X343) != s(bool,X344) | s(bool,'T_0') = s(bool,X345) | s(bool,X344) = s(bool,X345) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X343)) )),
  inference(superposition,[],[f2327,f2328])).
fof(f8092,plain,(
  ( ! [X405,X403,X404,X402] : (p(s(bool,X402)) | s(bool,'T_0') = s(bool,X403) | s(bool,'T_0') = s(bool,X404) | s(bool,X402) = s(bool,X405)) ) | $spl124),
  inference(subsumption_resolution,[],[f7698,f6531])).
fof(f7698,plain,(
  ( ! [X405,X403,X404,X402] : (p(s(bool,X402)) | s(bool,'T_0') = s(bool,X403) | s(bool,X403) != s(bool,X404) | s(bool,'T_0') = s(bool,X404) | s(bool,X402) = s(bool,X405)) ) | $spl124),
  inference(superposition,[],[f3624,f6531])).
fof(f3624,plain,(
  ( ! [X341,X340,X342] : (p(s(bool,X342)) | s(bool,X341) = s(bool,X342) | s(bool,X340) != s(bool,X341) | s(bool,'T_0') = s(bool,X340)) )),
  inference(subsumption_resolution,[],[f2924,f992])).
fof(f2924,plain,(
  ( ! [X341,X340,X342] : (s(bool,X340) != s(bool,X341) | s(bool,X341) = s(bool,X342) | p(s(bool,X342)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f2198,f2328])).
fof(f7876,plain,(
  ( ! [X389,X387,X388,X390] : (s(bool,X387) = s(bool,X388) | s(bool,'T_0') = s(bool,X389) | s(bool,'F_0') = s(bool,X388) | s(bool,X387) = s(bool,X390)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7694])).
fof(f7694,plain,(
  ( ! [X389,X387,X388,X390] : (s(bool,X387) = s(bool,X388) | s(bool,'T_0') = s(bool,X389) | s(bool,'T_0') = s(bool,X389) | s(bool,'F_0') = s(bool,X388) | s(bool,X387) = s(bool,X390)) ) | $spl124),
  inference(superposition,[],[f2782,f6531])).
fof(f8091,plain,(
  ( ! [X374,X376,X373,X375] : (p(s(bool,X373)) | s(bool,'T_0') != s(bool,X374) | s(bool,'T_0') = s(bool,X375) | s(bool,X373) = s(bool,X376)) ) | ($spl44 | $spl70 | $spl124)),
  inference(subsumption_resolution,[],[f7691,f1278])).
fof(f7691,plain,(
  ( ! [X374,X376,X373,X375] : (p(s(bool,X373)) | s(bool,'T_0') != s(bool,X374) | s(bool,X374) = s(bool,X375) | s(bool,'T_0') = s(bool,X375) | s(bool,X373) = s(bool,X376)) ) | ($spl70 | $spl124)),
  inference(superposition,[],[f2350,f6531])).
fof(f8090,plain,(
  ( ! [X358,X360,X359,X361] : (s(bool,X358) != s(bool,X359) | p(s(bool,X360)) | p(s(bool,X359)) | s(bool,X358) = s(bool,X361)) ) | ($spl66 | $spl124)),
  inference(subsumption_resolution,[],[f7687,f267])).
fof(f7687,plain,(
  ( ! [X358,X360,X359,X361] : (s(bool,X358) != s(bool,X359) | p(s(bool,X360)) | p(s(bool,X359)) | s(bool,'T_0') = s(bool,X360) | s(bool,X358) = s(bool,X361)) ) | ($spl66 | $spl124)),
  inference(superposition,[],[f2249,f6531])).
fof(f8089,plain,(
  ( ! [X346,X345,X347] : (p(s(bool,X345)) | s(bool,'T_0') != s(bool,X346) | s(bool,X345) = s(bool,X347)) ) | ($spl62 | $spl124)),
  inference(subsumption_resolution,[],[f7684,f1729])).
fof(f7684,plain,(
  ( ! [X346,X345,X347] : (p(s(bool,X345)) | s(bool,'T_0') != s(bool,X346) | s(bool,'F_0') = s(bool,X346) | s(bool,X345) = s(bool,X347)) ) | ($spl62 | $spl124)),
  inference(superposition,[],[f1999,f6531])).
fof(f7878,plain,(
  ( ! [X344,X343,X342] : (p(s(bool,X342)) | s(bool,'T_0') = s(bool,X343) | s(bool,X342) = s(bool,X344)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7683])).
fof(f7683,plain,(
  ( ! [X344,X343,X342] : (p(s(bool,X342)) | s(bool,'T_0') = s(bool,X343) | s(bool,'T_0') = s(bool,X343) | s(bool,X342) = s(bool,X344)) ) | $spl124),
  inference(superposition,[],[f1797,f6531])).
fof(f8088,plain,(
  ( ! [X341,X339,X340] : (p(s(bool,X339)) | p(s(bool,X340)) | s(bool,X339) = s(bool,X341)) ) | $spl124),
  inference(subsumption_resolution,[],[f7682,f267])).
fof(f7682,plain,(
  ( ! [X341,X339,X340] : (p(s(bool,X339)) | p(s(bool,X340)) | s(bool,'T_0') = s(bool,X340) | s(bool,X339) = s(bool,X341)) ) | $spl124),
  inference(superposition,[],[f1796,f6531])).
fof(f8087,plain,(
  ( ! [X337,X336,X338] : (p(s(bool,X336)) | p(s(bool,X337)) | s(bool,X336) = s(bool,X338)) ) | $spl124),
  inference(subsumption_resolution,[],[f7681,f267])).
fof(f7681,plain,(
  ( ! [X337,X336,X338] : (p(s(bool,X336)) | p(s(bool,X337)) | s(bool,'T_0') = s(bool,X337) | s(bool,X336) = s(bool,X338)) ) | $spl124),
  inference(superposition,[],[f1796,f6531])).
fof(f8086,plain,(
  ( ! [X325,X323,X324,X326] : (s(bool,X323) != s(bool,X324) | s(bool,'T_0') = s(bool,X325) | s(bool,X323) = s(bool,X326)) ) | $spl124),
  inference(subsumption_resolution,[],[f8085,f992])).
fof(f8085,plain,(
  ( ! [X325,X323,X324,X326] : (s(bool,X323) != s(bool,X324) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X325) | s(bool,X323) = s(bool,X326)) ) | $spl124),
  inference(subsumption_resolution,[],[f7678,f6531])).
fof(f7678,plain,(
  ( ! [X325,X323,X324,X326] : (s(bool,X323) != s(bool,X324) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X324) != s(bool,X325) | s(bool,'T_0') = s(bool,X325) | s(bool,X323) = s(bool,X326)) ) | $spl124),
  inference(superposition,[],[f906,f6531])).
fof(f8084,plain,(
  ( ! [X315,X316,X314] : (s(bool,'T_0') = s(bool,X315) | s(bool,X314) = s(bool,X316)) ) | $spl124),
  inference(subsumption_resolution,[],[f8083,f992])).
fof(f8083,plain,(
  ( ! [X315,X316,X314] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X315) | s(bool,X314) = s(bool,X316)) ) | $spl124),
  inference(subsumption_resolution,[],[f7675,f6531])).
fof(f7675,plain,(
  ( ! [X315,X316,X314] : (s(bool,X314) != s(bool,X315) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X315) | s(bool,X314) = s(bool,X316)) ) | $spl124),
  inference(superposition,[],[f826,f6531])).
fof(f8082,plain,(
  ( ! [X310,X312,X311] : (s(bool,'T_0') = s(bool,X311) | s(bool,X310) = s(bool,X312)) ) | $spl124),
  inference(subsumption_resolution,[],[f8081,f308])).
fof(f8081,plain,(
  ( ! [X310,X312,X311] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X311) | s(bool,X310) = s(bool,X312)) ) | $spl124),
  inference(subsumption_resolution,[],[f7674,f6531])).
fof(f7674,plain,(
  ( ! [X310,X312,X311] : (s(bool,X310) != s(bool,X311) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X311) | s(bool,X310) = s(bool,X312)) ) | $spl124),
  inference(superposition,[],[f806,f6531])).
fof(f8080,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,X304)) | s(bool,X305) != s(bool,X306) | s(bool,'F_0') = s(bool,X306) | s(bool,X304) = s(bool,X307)) ) | ($spl30 | $spl124)),
  inference(subsumption_resolution,[],[f7672,f826])).
fof(f7672,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,X304)) | s(bool,'T_0') != s(bool,X305) | s(bool,X305) != s(bool,X306) | s(bool,'F_0') = s(bool,X306) | s(bool,X304) = s(bool,X307)) ) | ($spl30 | $spl124)),
  inference(superposition,[],[f755,f6531])).
fof(f8078,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,X286)) | p(s(bool,X287)) | s(bool,'T_0') != s(bool,X288) | s(bool,X286) = s(bool,X289)) ) | $spl124),
  inference(subsumption_resolution,[],[f7667,f267])).
fof(f7667,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,X286)) | p(s(bool,X287)) | ~p(s(bool,X288)) | s(bool,'T_0') != s(bool,X288) | s(bool,X286) = s(bool,X289)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f8077,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X278) != s(bool,X279) | p(s(bool,X280)) | p(s(bool,X279)) | s(bool,X278) = s(bool,X281)) ) | $spl124),
  inference(subsumption_resolution,[],[f7665,f308])).
fof(f7665,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X278) != s(bool,X279) | p(s(bool,X280)) | ~p(s(bool,'T_0')) | p(s(bool,X279)) | s(bool,X278) = s(bool,X281)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f8075,plain,(
  ( ! [X231,X229,X233,X230,X232] : (p(s(bool,X229)) | ~p(s(bool,X230)) | s(bool,'T_0') != s(bool,X231) | s(bool,X230) != s(bool,X232) | s(bool,X229) = s(bool,X233)) ) | $spl124),
  inference(subsumption_resolution,[],[f7647,f267])).
fof(f7647,plain,(
  ( ! [X231,X229,X233,X230,X232] : (p(s(bool,X229)) | ~p(s(bool,X230)) | ~p(s(bool,X231)) | s(bool,'T_0') != s(bool,X231) | s(bool,X230) != s(bool,X232) | s(bool,X229) = s(bool,X233)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8074,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(bool,X224) != s(bool,X225) | p(s(bool,X225)) | ~p(s(bool,X226)) | s(bool,X226) != s(bool,X227) | s(bool,X224) = s(bool,X228)) ) | $spl124),
  inference(subsumption_resolution,[],[f7646,f308])).
fof(f7646,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s(bool,X224) != s(bool,X225) | p(s(bool,X225)) | ~p(s(bool,X226)) | ~p(s(bool,'T_0')) | s(bool,X226) != s(bool,X227) | s(bool,X224) = s(bool,X228)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8073,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X189)) | p(s(bool,X190)) | ~p(s(bool,X192)) | s(bool,'T_0') != s(bool,X191) | s(bool,X189) = s(bool,X193)) ) | $spl124),
  inference(subsumption_resolution,[],[f7639,f267])).
fof(f7639,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X189)) | p(s(bool,X190)) | ~p(s(bool,X191)) | ~p(s(bool,X192)) | s(bool,'T_0') != s(bool,X191) | s(bool,X189) = s(bool,X193)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8072,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,X182)) | p(s(bool,X180)) | s(bool,X179) = s(bool,X183)) ) | $spl124),
  inference(subsumption_resolution,[],[f7637,f308])).
fof(f7637,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,'T_0')) | ~p(s(bool,X182)) | p(s(bool,X180)) | s(bool,X179) = s(bool,X183)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8071,plain,(
  ( ! [X167,X165,X166,X164,X168] : (p(s(bool,X164)) | p(s(bool,X165)) | s(bool,X165) != s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,X164) = s(bool,X168)) ) | $spl124),
  inference(subsumption_resolution,[],[f7634,f267])).
fof(f7634,plain,(
  ( ! [X167,X165,X166,X164,X168] : (p(s(bool,X164)) | p(s(bool,X165)) | ~p(s(bool,X166)) | s(bool,X165) != s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,X164) = s(bool,X168)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8069,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,X144) != s(bool,X145) | p(s(bool,X146)) | p(s(bool,X145)) | s(bool,X146) != s(bool,X147) | s(bool,X144) = s(bool,X148)) ) | $spl124),
  inference(subsumption_resolution,[],[f7630,f308])).
fof(f7630,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,X144) != s(bool,X145) | p(s(bool,X146)) | ~p(s(bool,'T_0')) | p(s(bool,X145)) | s(bool,X146) != s(bool,X147) | s(bool,X144) = s(bool,X148)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8068,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | s(bool,'T_0') != s(bool,X131) | s(bool,X130) != s(bool,X132) | s(bool,X129) = s(bool,X133)) ) | $spl124),
  inference(subsumption_resolution,[],[f7627,f267])).
fof(f7627,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,X129)) | ~p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,'T_0') != s(bool,X131) | s(bool,X130) != s(bool,X132) | s(bool,X129) = s(bool,X133)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8067,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,X119) != s(bool,X120) | p(s(bool,X120)) | ~p(s(bool,X121)) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,X123)) ) | $spl124),
  inference(subsumption_resolution,[],[f7625,f308])).
fof(f7625,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,X119) != s(bool,X120) | p(s(bool,X120)) | ~p(s(bool,X121)) | ~p(s(bool,'T_0')) | s(bool,X121) != s(bool,X122) | s(bool,X119) = s(bool,X123)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8066,plain,(
  ( ! [X94,X97,X95,X98,X96] : (p(s(bool,X94)) | p(s(bool,X95)) | s(bool,X95) != s(bool,X97) | s(bool,'T_0') != s(bool,X96) | s(bool,X94) = s(bool,X98)) ) | $spl124),
  inference(subsumption_resolution,[],[f7620,f267])).
fof(f7620,plain,(
  ( ! [X94,X97,X95,X98,X96] : (p(s(bool,X94)) | p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,X95) != s(bool,X97) | s(bool,'T_0') != s(bool,X96) | s(bool,X94) = s(bool,X98)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8063,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X81)) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83)) ) | $spl124),
  inference(subsumption_resolution,[],[f7617,f308])).
fof(f7617,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,'T_0')) | p(s(bool,X80)) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8061,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X43) != s(bool,X44) | s(bool,'T_0') = s(bool,X45) | s(bool,X43) = s(bool,X46)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7611,f6531])).
fof(f7611,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X43) != s(bool,X44) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) != s(bool,X45) | s(bool,X43) = s(bool,X46)) ) | ($spl44 | $spl124)),
  inference(superposition,[],[f1278,f6531])).
fof(f7602,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) = s(bool,X20)) ) | $spl124),
  inference(superposition,[],[f261,f6531])).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) & p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.AND_CLAUSES|split|0')).
fof(f7601,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) = s(bool,X17)) ) | $spl124),
  inference(superposition,[],[f259,f6531])).
fof(f259,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : ((p(s(bool,X12)) & p(s(bool,'T_0'))) <=> p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.AND_CLAUSES|split|1')).
fof(f7600,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) = s(bool,X14)) ) | $spl124),
  inference(superposition,[],[f246,f6531])).
fof(f246,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CLAUSES|split|2')).
fof(f7599,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X11)) ) | $spl124),
  inference(superposition,[],[f243,f6531])).
fof(f243,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CLAUSES|split|3')).
fof(f7598,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) = s(bool,X8)) ) | $spl124),
  inference(superposition,[],[f242,f6531])).
fof(f242,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f7597,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5)) ) | $spl124),
  inference(superposition,[],[f239,f6531])).
fof(f239,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CLAUSES|split|1')).
fof(f7596,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X2)) ) | $spl124),
  inference(superposition,[],[f236,f6531])).
fof(f236,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) => p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CLAUSES|split|0')).
fof(f3872,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) ) | $spl84),
  inference(cnf_transformation,[],[f3872_D])).
fof(f3872_D,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7592,plain,(
  ( ! [X532,X528,X530,X529,X531] : (p(s(bool,X532)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X528),s('type.decSem.state'(X529),X530),s('type.exhSem.state'(X529),X531)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X528))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X529),X530))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X529),X531)))))) ) | $spl124),
  inference(superposition,[],[f277,f6531])).
fof(f6914,plain,(
  ( ! [X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) | $spl194),
  inference(cnf_transformation,[],[f6914_D])).
fof(f6914_D,plain,(
  ( ! [X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7591,plain,(
  ( ! [X524,X526,X525,X527,X523] : (~p(s(bool,X527)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X523),s('type.decSem.state'(X524),X525),s('type.exhSem.state'(X524),X526)))) | s('type.ffi.ffi_state'(X524),'const.decSem.state_ffi_1'(s('type.decSem.state'(X524),X525))) != s('type.ffi.ffi_state'(X524),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X524),X526))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X523))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X524),X525))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X524),X526)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X524),X525))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X524),X526))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X523))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X524),X525))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X524),X526)))))) ) | $spl124),
  inference(superposition,[],[f278,f6531])).
fof(f3627,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370)))))) ) | $spl76),
  inference(cnf_transformation,[],[f3627_D])).
fof(f3627_D,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7590,plain,(
  ( ! [X520,X522,X518,X521,X519] : (p(s(bool,X522)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518),s('type.decSem.state'(X519),X520),s('type.exhSem.state'(X519),X521)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X519),X520))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X519),X521)))))) ) | $spl124),
  inference(superposition,[],[f275,f6531])).
fof(f6912,plain,(
  ( ! [X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) | $spl192),
  inference(cnf_transformation,[],[f6912_D])).
fof(f6912_D,plain,(
  ( ! [X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7589,plain,(
  ( ! [X516,X514,X517,X513,X515] : (~p(s(bool,X517)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X513),s('type.decSem.state'(X514),X515),s('type.exhSem.state'(X514),X516)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X513))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X514),X515))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X514),X516)))))) | s('type.ffi.ffi_state'(X514),'const.decSem.state_ffi_1'(s('type.decSem.state'(X514),X515))) != s('type.ffi.ffi_state'(X514),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X514),X516))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X514),X515))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X514),X516))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X513))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X514),X515))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X514),X516)))))) ) | $spl124),
  inference(superposition,[],[f278,f6531])).
fof(f6570,plain,(
  ( ! [X94,X95,X93] : (s('type.ffi.ffi_state'(X93),'const.decSem.state_ffi_1'(s('type.decSem.state'(X93),X94))) = s('type.ffi.ffi_state'(X93),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X93),X95)))) ) | $spl132),
  inference(cnf_transformation,[],[f6570_D])).
fof(f6570_D,plain,(
  ( ! [X94,X95,X93] : (s('type.ffi.ffi_state'(X93),'const.decSem.state_ffi_1'(s('type.decSem.state'(X93),X94))) = s('type.ffi.ffi_state'(X93),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X93),X95)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8060,plain,(
  ( ! [X505,X507,X508,X506] : (s(bool,'T_0') != s(bool,X508) | s('type.ffi.ffi_state'(X505),'const.decSem.state_ffi_1'(s('type.decSem.state'(X505),X506))) = s('type.ffi.ffi_state'(X505),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X505),X507)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7587,f2266])).
fof(f2266,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X71),s('type.decSem.state'(X68),X69),s('type.exhSem.state'(X68),X70))) | s('type.ffi.ffi_state'(X68),'const.decSem.state_ffi_1'(s('type.decSem.state'(X68),X69))) = s('type.ffi.ffi_state'(X68),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X68),X70)))) )),
  inference(resolution,[],[f276,f267])).
fof(f276,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X2))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f7587,plain,(
  ( ! [X505,X507,X508,X504,X506] : (s(bool,'T_0') != s(bool,X508) | s('type.ffi.ffi_state'(X505),'const.decSem.state_ffi_1'(s('type.decSem.state'(X505),X506))) = s('type.ffi.ffi_state'(X505),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X505),X507))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X504),s('type.decSem.state'(X505),X506),s('type.exhSem.state'(X505),X507)))) ) | $spl124),
  inference(superposition,[],[f2266,f6531])).
fof(f8059,plain,(
  ( ! [X500,X502,X501,X503] : (s(bool,'F_0') = s(bool,X503) | s('type.ffi.ffi_state'(X500),'const.decSem.state_ffi_1'(s('type.decSem.state'(X500),X501))) = s('type.ffi.ffi_state'(X500),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X500),X502)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7586,f2266])).
fof(f7586,plain,(
  ( ! [X499,X500,X502,X501,X503] : (s(bool,'F_0') = s(bool,X503) | s('type.ffi.ffi_state'(X500),'const.decSem.state_ffi_1'(s('type.decSem.state'(X500),X501))) = s('type.ffi.ffi_state'(X500),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X500),X502))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X499),s('type.decSem.state'(X500),X501),s('type.exhSem.state'(X500),X502)))) ) | $spl124),
  inference(superposition,[],[f2265,f6531])).
fof(f2265,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X67),s('type.decSem.state'(X64),X65),s('type.exhSem.state'(X64),X66))) | s('type.ffi.ffi_state'(X64),'const.decSem.state_ffi_1'(s('type.decSem.state'(X64),X65))) = s('type.ffi.ffi_state'(X64),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X64),X66)))) )),
  inference(resolution,[],[f276,f264])).
fof(f6572,plain,(
  ( ! [X101,X99,X100] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X99),X100))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X99),X101)))) ) | $spl134),
  inference(cnf_transformation,[],[f6572_D])).
fof(f6572_D,plain,(
  ( ! [X101,X99,X100] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X99),X100))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X99),X101)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f8058,plain,(
  ( ! [X496,X498,X495,X497] : (s(bool,'T_0') != s(bool,X498) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X495),X496))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X495),X497)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7585,f2062])).
fof(f2062,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X71),s('type.decSem.state'(X68),X69),s('type.exhSem.state'(X68),X70))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X68),X69))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X68),X70)))) )),
  inference(resolution,[],[f274,f267])).
fof(f274,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X3))) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X0),X2),s('type.exhSem.state'(X0),X3))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f7585,plain,(
  ( ! [X494,X496,X498,X495,X497] : (s(bool,'T_0') != s(bool,X498) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X495),X496))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X495),X497))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X494),s('type.decSem.state'(X495),X496),s('type.exhSem.state'(X495),X497)))) ) | $spl124),
  inference(superposition,[],[f2062,f6531])).
fof(f8057,plain,(
  ( ! [X492,X490,X493,X491] : (s(bool,'F_0') = s(bool,X493) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X490),X491))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X490),X492)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7584,f2062])).
fof(f7584,plain,(
  ( ! [X492,X490,X493,X489,X491] : (s(bool,'F_0') = s(bool,X493) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X490),X491))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X490),X492))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X489),s('type.decSem.state'(X490),X491),s('type.exhSem.state'(X490),X492)))) ) | $spl124),
  inference(superposition,[],[f2061,f6531])).
fof(f2061,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X67),s('type.decSem.state'(X64),X65),s('type.exhSem.state'(X64),X66))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X64),X65))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X64),X66)))) )),
  inference(resolution,[],[f274,f264])).
fof(f4183,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X21),X22))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X21),X23))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.decSem.state_ffi_1'(s('type.decSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X21),X23)))) ) | $spl96),
  inference(cnf_transformation,[],[f4183_D])).
fof(f4183_D,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X21),X22))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X21),X23))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.decSem.state_ffi_1'(s('type.decSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X21),X23)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7583,plain,(
  ( ! [X484,X486,X488,X485,X487] : (p(s(bool,X488)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.ffi.ffi_state'(X485),'const.decSem.state_ffi_1'(s('type.decSem.state'(X485),X486))) != s('type.ffi.ffi_state'(X485),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X485),X487))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X485),X486))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X485),X487))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484),s('type.decSem.state'(X485),X486),s('type.exhSem.state'(X485),X487)))) ) | $spl124),
  inference(superposition,[],[f278,f6531])).
fof(f6906,plain,(
  ( ! [X573,X571,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) | $spl190),
  inference(cnf_transformation,[],[f6906_D])).
fof(f6906_D,plain,(
  ( ! [X573,X571,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7582,plain,(
  ( ! [X480,X482,X479,X481,X483] : (~p(s(bool,X483)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X479))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X480),X481))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X480),X482)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X479),s('type.decSem.state'(X480),X481),s('type.exhSem.state'(X480),X482)))) ) | $spl124),
  inference(superposition,[],[f277,f6531])).
fof(f8056,plain,(
  ( ! [X476,X478,X475,X477] : (~p(s(bool,X478)) | s('type.ffi.ffi_state'(X475),'const.decSem.state_ffi_1'(s('type.decSem.state'(X475),X476))) = s('type.ffi.ffi_state'(X475),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X475),X477)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7581,f2266])).
fof(f7581,plain,(
  ( ! [X476,X478,X474,X475,X477] : (~p(s(bool,X478)) | s('type.ffi.ffi_state'(X475),'const.decSem.state_ffi_1'(s('type.decSem.state'(X475),X476))) = s('type.ffi.ffi_state'(X475),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X475),X477))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X474),s('type.decSem.state'(X475),X476),s('type.exhSem.state'(X475),X477)))) ) | $spl124),
  inference(superposition,[],[f276,f6531])).
fof(f6904,plain,(
  ( ! [X561,X559,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) | $spl188),
  inference(cnf_transformation,[],[f6904_D])).
fof(f6904_D,plain,(
  ( ! [X561,X559,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7580,plain,(
  ( ! [X472,X469,X471,X473,X470] : (~p(s(bool,X473)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X470),X471))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X470),X472)))))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469),s('type.decSem.state'(X470),X471),s('type.exhSem.state'(X470),X472)))) ) | $spl124),
  inference(superposition,[],[f275,f6531])).
fof(f8055,plain,(
  ( ! [X466,X465,X467,X468] : (~p(s(bool,X468)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X465),X466))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X465),X467)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7579,f2062])).
fof(f7579,plain,(
  ( ! [X466,X465,X467,X468,X464] : (~p(s(bool,X468)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X465),X466))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X465),X467))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X464),s('type.decSem.state'(X465),X466),s('type.exhSem.state'(X465),X467)))) ) | $spl124),
  inference(superposition,[],[f274,f6531])).
fof(f6902,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | $spl186),
  inference(cnf_transformation,[],[f6902_D])).
fof(f6902_D,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7578,plain,(
  ( ! [X463] : (~p(s(bool,X463)) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) ) | $spl124),
  inference(superposition,[],[f206,f6531])).
fof(f8046,plain,(
  ( ! [X372,X374,X373,X375] : (p(s(bool,X373)) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') != s(bool,X375) | s(bool,'T_0') = s(bool,X372)) ) | $spl124),
  inference(subsumption_resolution,[],[f7543,f6531])).
fof(f7543,plain,(
  ( ! [X372,X374,X373,X375] : (p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') != s(bool,X375) | s(bool,'T_0') = s(bool,X372)) ) | $spl124),
  inference(superposition,[],[f3958,f6531])).
fof(f7884,plain,(
  ( ! [X352,X354,X353,X355] : (s(bool,X353) = s(bool,X354) | s(bool,X354) = s(bool,X355) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,X355)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7538])).
fof(f7538,plain,(
  ( ! [X352,X354,X353,X355] : (s(bool,X353) = s(bool,X354) | s(bool,X354) = s(bool,X355) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,X355) | s(bool,'T_0') = s(bool,X352)) ) | $spl124),
  inference(superposition,[],[f2782,f6531])).
fof(f8044,plain,(
  ( ! [X329,X331,X328] : (~p(s(bool,X329)) | p(s(bool,X331)) | s(bool,'T_0') = s(bool,X328)) ) | $spl124),
  inference(subsumption_resolution,[],[f8043,f6531])).
fof(f8043,plain,(
  ( ! [X329,X331,X328,X330] : (~p(s(bool,X329)) | p(s(bool,X331)) | s(bool,X328) != s(bool,X330) | s(bool,'T_0') = s(bool,X328)) ) | $spl124),
  inference(subsumption_resolution,[],[f7532,f3624])).
fof(f7532,plain,(
  ( ! [X329,X331,X328,X330] : (~p(s(bool,X329)) | s(bool,X330) != s(bool,X331) | p(s(bool,X331)) | s(bool,X328) != s(bool,X330) | s(bool,'T_0') = s(bool,X328)) ) | $spl124),
  inference(superposition,[],[f2388,f6531])).
fof(f8042,plain,(
  ( ! [X308,X306,X307] : (s(bool,X306) != s(bool,X307) | s(bool,X307) = s(bool,X308) | p(s(bool,X308))) ) | $spl226),
  inference(cnf_transformation,[],[f8042_D])).
fof(f8042_D,plain,(
  ( ! [X308,X306,X307] : (s(bool,X306) != s(bool,X307) | s(bool,X307) = s(bool,X308) | p(s(bool,X308))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8041,plain,(
  ( ! [X308,X306,X305,X307] : (s(bool,X306) != s(bool,X307) | p(s(bool,X308)) | p(s(bool,X305)) | s(bool,X307) = s(bool,X308)) ) | ($spl66 | $spl124)),
  inference(subsumption_resolution,[],[f7525,f267])).
fof(f7525,plain,(
  ( ! [X308,X306,X305,X307] : (s(bool,X306) != s(bool,X307) | p(s(bool,X308)) | p(s(bool,X305)) | s(bool,X307) = s(bool,X308) | s(bool,'T_0') = s(bool,X305)) ) | ($spl66 | $spl124)),
  inference(superposition,[],[f2249,f6531])).
fof(f8040,plain,(
  ( ! [X292,X296,X293,X295,X297] : (~p(s(bool,X293)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'T_0') = s(bool,X292)) ) | ($spl64 | $spl124)),
  inference(subsumption_resolution,[],[f7522,f6531])).
fof(f7522,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (~p(s(bool,X293)) | s(bool,X292) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X294),s('type.decSem.state'(X295),X296),s('type.exhSem.state'(X295),X297))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X295),X296))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X295),X297))) | s(bool,'T_0') = s(bool,X292)) ) | ($spl64 | $spl124)),
  inference(superposition,[],[f2101,f6531])).
fof(f8039,plain,(
  ( ! [X290,X289,X291] : (s(bool,'T_0') != s(bool,X290) | s(bool,X289) != s(bool,X291) | p(s(bool,X291))) ) | $spl124),
  inference(subsumption_resolution,[],[f7521,f2028])).
fof(f7521,plain,(
  ( ! [X290,X289,X291] : (s(bool,'T_0') != s(bool,X290) | s(bool,X289) != s(bool,X291) | p(s(bool,X291)) | s(bool,'T_0') = s(bool,X289)) ) | $spl124),
  inference(superposition,[],[f2028,f6531])).
fof(f8038,plain,(
  ( ! [X271,X272,X274] : (~p(s(bool,X272)) | s(bool,'T_0') = s(bool,X274) | s(bool,'T_0') = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(subsumption_resolution,[],[f8037,f6531])).
fof(f8037,plain,(
  ( ! [X271,X273,X272,X274] : (~p(s(bool,X272)) | s(bool,X271) != s(bool,X273) | s(bool,'T_0') = s(bool,X274) | s(bool,'T_0') = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(subsumption_resolution,[],[f7515,f3625])).
fof(f7515,plain,(
  ( ! [X271,X273,X272,X274] : (~p(s(bool,X272)) | s(bool,X271) != s(bool,X273) | s(bool,'T_0') = s(bool,X274) | s(bool,X273) != s(bool,X274) | s(bool,'T_0') = s(bool,X271)) ) | ($spl52 | $spl124)),
  inference(superposition,[],[f1367,f6531])).
fof(f8036,plain,(
  ( ! [X267,X266] : (s(bool,'T_0') != s(bool,X266) | s(bool,'T_0') = s(bool,X267)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f8035,f6531])).
fof(f8035,plain,(
  ( ! [X265,X267,X266] : (s(bool,'T_0') != s(bool,X266) | s(bool,'T_0') = s(bool,X267) | s(bool,X265) != s(bool,X267)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7513,f1278])).
fof(f7513,plain,(
  ( ! [X265,X267,X266] : (s(bool,'T_0') != s(bool,X266) | s(bool,'T_0') = s(bool,X267) | s(bool,X265) != s(bool,X267) | s(bool,'T_0') = s(bool,X265)) ) | ($spl44 | $spl124)),
  inference(superposition,[],[f1278,f6531])).
fof(f8034,plain,(
  ( ! [X257,X259,X260,X258] : (s(bool,X258) != s(bool,X259) | s(bool,'F_0') = s(bool,X259) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,X257)) ) | $spl124),
  inference(subsumption_resolution,[],[f7511,f6531])).
fof(f7511,plain,(
  ( ! [X257,X259,X260,X258] : (s(bool,X258) != s(bool,X259) | s(bool,'F_0') = s(bool,X259) | s(bool,X257) != s(bool,X260) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,X257)) ) | $spl124),
  inference(superposition,[],[f906,f6531])).
fof(f7887,plain,(
  ( ! [X255,X254,X256] : (s(bool,X255) != s(bool,X256) | s(bool,'F_0') = s(bool,X256) | s(bool,'T_0') = s(bool,X254)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7510])).
fof(f7510,plain,(
  ( ! [X255,X254,X256] : (s(bool,X255) != s(bool,X256) | s(bool,'F_0') = s(bool,X256) | s(bool,'T_0') = s(bool,X254) | s(bool,'T_0') = s(bool,X254)) ) | $spl124),
  inference(superposition,[],[f826,f6531])).
fof(f7888,plain,(
  ( ! [X243,X244,X242] : (s(bool,X243) != s(bool,X244) | ~p(s(bool,X244)) | s(bool,'T_0') = s(bool,X242)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7506])).
fof(f7506,plain,(
  ( ! [X243,X244,X242] : (s(bool,X243) != s(bool,X244) | ~p(s(bool,X244)) | s(bool,'T_0') = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) ) | $spl124),
  inference(superposition,[],[f806,f6531])).
fof(f8033,plain,(
  ( ! [X237,X235,X236] : (s(bool,X236) != s(bool,X237) | p(s(bool,X235)) | ~p(s(bool,X237))) ) | $spl124),
  inference(subsumption_resolution,[],[f7504,f267])).
fof(f7504,plain,(
  ( ! [X237,X235,X236] : (s(bool,X236) != s(bool,X237) | p(s(bool,X235)) | ~p(s(bool,X237)) | s(bool,'T_0') = s(bool,X235)) ) | $spl124),
  inference(superposition,[],[f368,f6531])).
fof(f8032,plain,(
  ( ! [X231,X229,X230,X228] : (~p(s(bool,X229)) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X228) != s(bool,X231)) ) | $spl124),
  inference(subsumption_resolution,[],[f7502,f2028])).
fof(f7502,plain,(
  ( ! [X231,X229,X230,X228] : (~p(s(bool,X229)) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X228) != s(bool,X231) | s(bool,'T_0') = s(bool,X228)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f8031,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,X217) != s(bool,X218) | p(s(bool,X219)) | ~p(s(bool,X218)) | p(s(bool,X216))) ) | $spl124),
  inference(subsumption_resolution,[],[f7499,f267])).
fof(f7499,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,X217) != s(bool,X218) | p(s(bool,X219)) | ~p(s(bool,X218)) | p(s(bool,X216)) | s(bool,'T_0') = s(bool,X216)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f7889,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,X211) | s(bool,'T_0') = s(bool,X210)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7497])).
fof(f7497,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X210)) ) | $spl124),
  inference(superposition,[],[f306,f6531])).
fof(f7890,plain,(
  ( ! [X206,X207] : (~p(s(bool,X207)) | s(bool,'T_0') = s(bool,X206)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7495])).
fof(f7495,plain,(
  ( ! [X206,X207] : (~p(s(bool,X207)) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) ) | $spl124),
  inference(superposition,[],[f270,f6531])).
fof(f8030,plain,(
  ( ! [X200,X201] : (s(bool,'T_0') != s(bool,X201) | p(s(bool,X200))) ) | $spl124),
  inference(subsumption_resolution,[],[f7492,f267])).
fof(f7492,plain,(
  ( ! [X200,X201] : (s(bool,'T_0') != s(bool,X201) | p(s(bool,X200)) | s(bool,'T_0') = s(bool,X200)) ) | $spl124),
  inference(superposition,[],[f269,f6531])).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f7891,plain,(
  ( ! [X198,X199] : (~p(s(bool,X199)) | s(bool,'T_0') = s(bool,X198)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7491])).
fof(f7491,plain,(
  ( ! [X198,X199] : (~p(s(bool,X199)) | s(bool,'T_0') = s(bool,X198) | s(bool,'T_0') = s(bool,X198)) ) | $spl124),
  inference(superposition,[],[f268,f6531])).
fof(f8029,plain,(
  ( ! [X192,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X192))) ) | $spl124),
  inference(subsumption_resolution,[],[f7488,f267])).
fof(f7488,plain,(
  ( ! [X192,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X192)) | s(bool,'T_0') = s(bool,X192)) ) | $spl124),
  inference(superposition,[],[f267,f6531])).
fof(f8028,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') = s(bool,X189) | p(s(bool,X188))) ) | $spl124),
  inference(subsumption_resolution,[],[f7486,f267])).
fof(f7486,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') = s(bool,X189) | p(s(bool,X188)) | s(bool,'T_0') = s(bool,X188)) ) | $spl124),
  inference(superposition,[],[f264,f6531])).
fof(f8027,plain,(
  ( ! [X187,X186] : (~p(s(bool,X187)) | s(bool,'T_0') = s(bool,X186)) ) | $spl124),
  inference(subsumption_resolution,[],[f7485,f6531])).
fof(f7485,plain,(
  ( ! [X187,X186] : (~p(s(bool,X187)) | s(bool,'F_0') != s(bool,X186) | s(bool,'T_0') = s(bool,X186)) ) | $spl124),
  inference(superposition,[],[f263,f6531])).
fof(f8026,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | p(s(bool,X167)) | ~p(s(bool,X168)) | s(bool,X165) != s(bool,X167) | s(bool,X168) != s(bool,X169)) ) | $spl124),
  inference(subsumption_resolution,[],[f7476,f2028])).
fof(f7476,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X166)) | p(s(bool,X167)) | ~p(s(bool,X168)) | s(bool,X165) != s(bool,X167) | s(bool,X168) != s(bool,X169) | s(bool,'T_0') = s(bool,X165)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8025,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,X147) | p(s(bool,X145)) | ~p(s(bool,X148)) | ~p(s(bool,X147)) | s(bool,X148) != s(bool,X149)) ) | $spl124),
  inference(subsumption_resolution,[],[f7472,f267])).
fof(f7472,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,X147) | p(s(bool,X145)) | ~p(s(bool,X148)) | ~p(s(bool,X147)) | s(bool,X148) != s(bool,X149) | s(bool,'T_0') = s(bool,X145)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f8022,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,X126)) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X125) != s(bool,X129)) ) | $spl124),
  inference(subsumption_resolution,[],[f7468,f2028])).
fof(f7468,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,X126)) | p(s(bool,X127)) | ~p(s(bool,X128)) | p(s(bool,X129)) | s(bool,X125) != s(bool,X129) | s(bool,'T_0') = s(bool,X125)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8021,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,X111) != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | ~p(s(bool,X114)) | p(s(bool,X110))) ) | $spl124),
  inference(subsumption_resolution,[],[f7465,f267])).
fof(f7465,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,X111) != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | ~p(s(bool,X114)) | p(s(bool,X110)) | s(bool,'T_0') = s(bool,X110)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f8020,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X100) != s(bool,X103)) ) | $spl124),
  inference(subsumption_resolution,[],[f7463,f2028])).
fof(f7463,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | s(bool,X102) != s(bool,X104) | s(bool,X100) != s(bool,X103) | s(bool,'T_0') = s(bool,X100)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8017,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X78) != s(bool,X79)) ) | $spl124),
  inference(subsumption_resolution,[],[f7458,f267])).
fof(f7458,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | p(s(bool,X75)) | s(bool,X78) != s(bool,X79) | s(bool,'T_0') = s(bool,X75)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f8016,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X65) != s(bool,X67) | s(bool,X68) != s(bool,X69)) ) | $spl124),
  inference(subsumption_resolution,[],[f7456,f2028])).
fof(f7456,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X65) != s(bool,X67) | s(bool,X68) != s(bool,X69) | s(bool,'T_0') = s(bool,X65)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8015,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X51) != s(bool,X52) | p(s(bool,X50)) | ~p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X53) != s(bool,X54)) ) | $spl124),
  inference(subsumption_resolution,[],[f7453,f267])).
fof(f7453,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X51) != s(bool,X52) | p(s(bool,X50)) | ~p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X53) != s(bool,X54) | s(bool,'T_0') = s(bool,X50)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f8012,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X30) != s(bool,X33)) ) | $spl124),
  inference(subsumption_resolution,[],[f7449,f2028])).
fof(f7449,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8009,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4)) ) | $spl124),
  inference(subsumption_resolution,[],[f7443,f267])).
fof(f7443,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X0)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f8007,plain,(
  ( ! [X520,X518,X521,X519] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518),s('type.decSem.state'(X519),X520),s('type.exhSem.state'(X519),X521))))) ) | $spl224),
  inference(cnf_transformation,[],[f8007_D])).
fof(f8007_D,plain,(
  ( ! [X520,X518,X521,X519] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518),s('type.decSem.state'(X519),X520),s('type.exhSem.state'(X519),X521))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8008,plain,(
  ( ! [X532,X528,X530,X529,X531] : (p(s(bool,X532)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X528),s('type.decSem.state'(X529),X530),s('type.exhSem.state'(X529),X531))))) ) | $spl124),
  inference(subsumption_resolution,[],[f7439,f267])).
fof(f7439,plain,(
  ( ! [X532,X528,X530,X529,X531] : (p(s(bool,X532)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X528),s('type.decSem.state'(X529),X530),s('type.exhSem.state'(X529),X531)))) | s(bool,'T_0') = s(bool,X532)) ) | $spl124),
  inference(superposition,[],[f277,f6531])).
fof(f8006,plain,(
  ( ! [X520,X522,X518,X521,X519] : (p(s(bool,X522)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518),s('type.decSem.state'(X519),X520),s('type.exhSem.state'(X519),X521))))) ) | $spl124),
  inference(subsumption_resolution,[],[f7437,f267])).
fof(f7437,plain,(
  ( ! [X520,X522,X518,X521,X519] : (p(s(bool,X522)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X518),s('type.decSem.state'(X519),X520),s('type.exhSem.state'(X519),X521)))) | s(bool,'T_0') = s(bool,X522)) ) | $spl124),
  inference(superposition,[],[f275,f6531])).
fof(f8005,plain,(
  ( ! [X484,X486,X485,X487] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X485),X486))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X485),X487))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.ffi.ffi_state'(X485),'const.decSem.state_ffi_1'(s('type.decSem.state'(X485),X486))) != s('type.ffi.ffi_state'(X485),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X485),X487)))) ) | $spl222),
  inference(cnf_transformation,[],[f8005_D])).
fof(f8005_D,plain,(
  ( ! [X484,X486,X485,X487] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X485),X486))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X485),X487))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.ffi.ffi_state'(X485),'const.decSem.state_ffi_1'(s('type.decSem.state'(X485),X486))) != s('type.ffi.ffi_state'(X485),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X485),X487)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f8004,plain,(
  ( ! [X484,X486,X488,X485,X487] : (p(s(bool,X488)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.ffi.ffi_state'(X485),'const.decSem.state_ffi_1'(s('type.decSem.state'(X485),X486))) != s('type.ffi.ffi_state'(X485),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X485),X487))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X485),X486))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X485),X487)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7430,f267])).
fof(f7430,plain,(
  ( ! [X484,X486,X488,X485,X487] : (p(s(bool,X488)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.ffi.ffi_state'(X485),'const.decSem.state_ffi_1'(s('type.decSem.state'(X485),X486))) != s('type.ffi.ffi_state'(X485),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X485),X487))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X484))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X485),X486))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X485),X487)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X485),X486))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X485),X487))) | s(bool,'T_0') = s(bool,X488)) ) | $spl124),
  inference(superposition,[],[f278,f6531])).
fof(f8003,plain,(
  ( ! [X461,X459,X460] : (s(bool,X460) != s(bool,X461) | s(bool,'T_0') != s(bool,X461) | s(bool,'T_0') = s(bool,X459)) ) | $spl124),
  inference(subsumption_resolution,[],[f7423,f6531])).
fof(f7423,plain,(
  ( ! [X461,X459,X460] : (s(bool,X459) != s(bool,X460) | s(bool,X460) != s(bool,X461) | s(bool,'T_0') != s(bool,X461) | s(bool,'T_0') = s(bool,X459)) ) | $spl124),
  inference(superposition,[],[f4878,f6531])).
fof(f8002,plain,(
  ( ! [X412,X410] : (p(s(bool,X412)) | s(bool,'T_0') = s(bool,X410)) ) | $spl124),
  inference(subsumption_resolution,[],[f8001,f6531])).
fof(f8001,plain,(
  ( ! [X412,X410,X411] : (s(bool,X410) != s(bool,X411) | p(s(bool,X412)) | s(bool,'T_0') = s(bool,X410)) ) | $spl124),
  inference(subsumption_resolution,[],[f7408,f3624])).
fof(f7408,plain,(
  ( ! [X412,X410,X411] : (s(bool,X410) != s(bool,X411) | s(bool,X411) != s(bool,X412) | p(s(bool,X412)) | s(bool,'T_0') = s(bool,X410)) ) | $spl124),
  inference(superposition,[],[f2028,f6531])).
fof(f8000,plain,(
  ( ! [X407,X409] : (s(bool,'T_0') = s(bool,X409) | s(bool,'T_0') = s(bool,X407)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7999,f6531])).
fof(f7999,plain,(
  ( ! [X408,X407,X409] : (s(bool,X407) != s(bool,X408) | s(bool,'T_0') = s(bool,X409) | s(bool,'T_0') = s(bool,X407)) ) | ($spl44 | $spl124)),
  inference(subsumption_resolution,[],[f7407,f3625])).
fof(f7407,plain,(
  ( ! [X408,X407,X409] : (s(bool,X407) != s(bool,X408) | s(bool,'T_0') = s(bool,X409) | s(bool,X408) != s(bool,X409) | s(bool,'T_0') = s(bool,X407)) ) | ($spl44 | $spl124)),
  inference(superposition,[],[f1278,f6531])).
fof(f7990,plain,(
  ( ! [X391,X392] : (p(s(bool,X391)) | ~p(s(bool,X392))) ) | $spl124),
  inference(subsumption_resolution,[],[f7398,f267])).
fof(f7398,plain,(
  ( ! [X391,X392] : (p(s(bool,X391)) | ~p(s(bool,X392)) | s(bool,'T_0') = s(bool,X391)) ) | $spl124),
  inference(superposition,[],[f261,f6531])).
fof(f7989,plain,(
  ( ! [X389,X390] : (p(s(bool,X389)) | ~p(s(bool,X390))) ) | $spl124),
  inference(subsumption_resolution,[],[f7397,f267])).
fof(f7397,plain,(
  ( ! [X389,X390] : (p(s(bool,X389)) | ~p(s(bool,X390)) | s(bool,'T_0') = s(bool,X389)) ) | $spl124),
  inference(superposition,[],[f259,f6531])).
fof(f7988,plain,(
  ( ! [X387,X388] : (p(s(bool,X387)) | ~p(s(bool,X388))) ) | $spl124),
  inference(subsumption_resolution,[],[f7396,f267])).
fof(f7396,plain,(
  ( ! [X387,X388] : (p(s(bool,X387)) | ~p(s(bool,X388)) | s(bool,'T_0') = s(bool,X387)) ) | $spl124),
  inference(superposition,[],[f246,f6531])).
fof(f7987,plain,(
  ( ! [X385,X386] : (p(s(bool,X385)) | ~p(s(bool,X386))) ) | $spl124),
  inference(subsumption_resolution,[],[f7395,f267])).
fof(f7395,plain,(
  ( ! [X385,X386] : (p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,'T_0') = s(bool,X385)) ) | $spl124),
  inference(superposition,[],[f243,f6531])).
fof(f7986,plain,(
  ( ! [X383,X384] : (p(s(bool,X383)) | p(s(bool,X384))) ) | $spl124),
  inference(subsumption_resolution,[],[f7394,f267])).
fof(f7394,plain,(
  ( ! [X383,X384] : (p(s(bool,X383)) | p(s(bool,X384)) | s(bool,'T_0') = s(bool,X383)) ) | $spl124),
  inference(superposition,[],[f242,f6531])).
fof(f7985,plain,(
  ( ! [X382,X381] : (p(s(bool,X381)) | p(s(bool,X382))) ) | $spl124),
  inference(subsumption_resolution,[],[f7393,f267])).
fof(f7393,plain,(
  ( ! [X382,X381] : (p(s(bool,X381)) | p(s(bool,X382)) | s(bool,'T_0') = s(bool,X381)) ) | $spl124),
  inference(superposition,[],[f239,f6531])).
fof(f7984,plain,(
  ( ! [X379,X380] : (p(s(bool,X379)) | p(s(bool,X380))) ) | $spl124),
  inference(subsumption_resolution,[],[f7392,f267])).
fof(f7392,plain,(
  ( ! [X379,X380] : (p(s(bool,X379)) | p(s(bool,X380)) | s(bool,'T_0') = s(bool,X379)) ) | $spl124),
  inference(superposition,[],[f236,f6531])).
fof(f7983,plain,(
  ( ! [X377,X376,X378] : (s(bool,'F_0') != s(bool,X377) | s(bool,X376) != s(bool,X378) | s(bool,'T_0') != s(bool,X378)) ) | $spl124),
  inference(subsumption_resolution,[],[f7391,f1729])).
fof(f7391,plain,(
  ( ! [X377,X376,X378] : (s(bool,'F_0') != s(bool,X377) | s(bool,X376) != s(bool,X378) | s(bool,'T_0') != s(bool,X378) | s(bool,'T_0') = s(bool,X377)) ) | $spl124),
  inference(superposition,[],[f4878,f6531])).
fof(f7982,plain,(
  ( ! [X372,X374,X375] : (s(bool,X372) != s(bool,X374) | s(bool,'T_0') != s(bool,X375) | s(bool,X374) != s(bool,X375)) ) | $spl220),
  inference(cnf_transformation,[],[f7982_D])).
fof(f7982_D,plain,(
  ( ! [X372,X374,X375] : (s(bool,X372) != s(bool,X374) | s(bool,'T_0') != s(bool,X375) | s(bool,X374) != s(bool,X375)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f7981,plain,(
  ( ! [X372,X374,X373,X375] : (p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') != s(bool,X375)) ) | $spl124),
  inference(subsumption_resolution,[],[f7390,f267])).
fof(f7390,plain,(
  ( ! [X372,X374,X373,X375] : (p(s(bool,X373)) | s(bool,X372) != s(bool,X374) | s(bool,X374) != s(bool,X375) | s(bool,'T_0') != s(bool,X375) | s(bool,'T_0') = s(bool,X373)) ) | $spl124),
  inference(superposition,[],[f3958,f6531])).
fof(f7980,plain,(
  ( ! [X360,X362,X361,X363] : (p(s(bool,X361)) | s(bool,X360) = s(bool,X362) | s(bool,'T_0') = s(bool,X363)) ) | $spl124),
  inference(subsumption_resolution,[],[f7979,f6531])).
fof(f7979,plain,(
  ( ! [X360,X362,X361,X363] : (p(s(bool,X361)) | s(bool,X360) = s(bool,X362) | s(bool,X362) != s(bool,X363) | s(bool,'T_0') = s(bool,X363)) ) | $spl124),
  inference(subsumption_resolution,[],[f7387,f267])).
fof(f7387,plain,(
  ( ! [X360,X362,X361,X363] : (p(s(bool,X361)) | s(bool,X360) = s(bool,X362) | s(bool,X362) != s(bool,X363) | s(bool,'T_0') = s(bool,X363) | s(bool,'T_0') = s(bool,X361)) ) | $spl124),
  inference(superposition,[],[f3624,f6531])).
fof(f7893,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,'T_0') = s(bool,X333) | s(bool,X334) = s(bool,X335) | s(bool,X332) = s(bool,X335) | s(bool,'F_0') = s(bool,X334)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7380])).
fof(f7380,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,'T_0') = s(bool,X333) | s(bool,X334) = s(bool,X335) | s(bool,X332) = s(bool,X335) | s(bool,'F_0') = s(bool,X334) | s(bool,'T_0') = s(bool,X333)) ) | $spl124),
  inference(superposition,[],[f2782,f6531])).
fof(f7976,plain,(
  ( ! [X302,X304,X301,X303] : (s(bool,X302) != s(bool,X303) | p(s(bool,X304)) | p(s(bool,X303)) | s(bool,X301) = s(bool,X304)) ) | ($spl66 | $spl124)),
  inference(subsumption_resolution,[],[f7371,f2028])).
fof(f7371,plain,(
  ( ! [X302,X304,X301,X303] : (s(bool,X302) != s(bool,X303) | p(s(bool,X304)) | p(s(bool,X303)) | s(bool,X301) = s(bool,X304) | s(bool,'T_0') = s(bool,X302)) ) | ($spl66 | $spl124)),
  inference(superposition,[],[f2249,f6531])).
fof(f7975,plain,(
  ( ! [X300,X298,X299] : (s(bool,X298) = s(bool,X300) | p(s(bool,X300)) | s(bool,'T_0') = s(bool,X299)) ) | $spl124),
  inference(subsumption_resolution,[],[f7370,f6531])).
fof(f7370,plain,(
  ( ! [X300,X298,X299] : (s(bool,'F_0') != s(bool,X299) | s(bool,X298) = s(bool,X300) | p(s(bool,X300)) | s(bool,'T_0') = s(bool,X299)) ) | $spl124),
  inference(superposition,[],[f2198,f6531])).
fof(f7974,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X287)) | s(bool,X286) != s(bool,X288) | s(bool,'F_0') = s(bool,X288)) ) | ($spl62 | $spl124)),
  inference(subsumption_resolution,[],[f7367,f267])).
fof(f7367,plain,(
  ( ! [X288,X287,X286] : (p(s(bool,X287)) | s(bool,X286) != s(bool,X288) | s(bool,'F_0') = s(bool,X288) | s(bool,'T_0') = s(bool,X287)) ) | ($spl62 | $spl124)),
  inference(superposition,[],[f1999,f6531])).
fof(f7972,plain,(
  ( ! [X280,X282,X281] : (p(s(bool,X281)) | p(s(bool,X282)) | s(bool,X280) = s(bool,X282)) ) | $spl124),
  inference(subsumption_resolution,[],[f7365,f267])).
fof(f7365,plain,(
  ( ! [X280,X282,X281] : (p(s(bool,X281)) | p(s(bool,X282)) | s(bool,X280) = s(bool,X282) | s(bool,'T_0') = s(bool,X281)) ) | $spl124),
  inference(superposition,[],[f1796,f6531])).
fof(f7971,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | s(bool,X277) = s(bool,X279)) ) | $spl124),
  inference(subsumption_resolution,[],[f7364,f267])).
fof(f7364,plain,(
  ( ! [X277,X279,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | s(bool,X277) = s(bool,X279) | s(bool,'T_0') = s(bool,X278)) ) | $spl124),
  inference(superposition,[],[f1796,f6531])).
fof(f7970,plain,(
  ( ! [X275,X276] : (s(bool,'F_0') != s(bool,X276) | s(bool,'T_0') != s(bool,X275)) ) | $spl124),
  inference(subsumption_resolution,[],[f7363,f1729])).
fof(f7363,plain,(
  ( ! [X275,X276] : (s(bool,'F_0') != s(bool,X276) | s(bool,'T_0') != s(bool,X275) | s(bool,'T_0') = s(bool,X276)) ) | $spl124),
  inference(superposition,[],[f1729,f6531])).
fof(f7969,plain,(
  ( ! [X269,X268,X270] : (~p(s(bool,X270)) | s(bool,X268) != s(bool,X270) | s(bool,'T_0') = s(bool,X269)) ) | ($spl48 | $spl124)),
  inference(subsumption_resolution,[],[f7361,f6531])).
fof(f7361,plain,(
  ( ! [X269,X268,X270] : (s(bool,'F_0') != s(bool,X269) | ~p(s(bool,X270)) | s(bool,X268) != s(bool,X270) | s(bool,'T_0') = s(bool,X269)) ) | ($spl48 | $spl124)),
  inference(superposition,[],[f1356,f6531])).
fof(f7968,plain,(
  ( ! [X261,X262,X264] : (s(bool,'F_0') = s(bool,X261) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X262)) ) | $spl124),
  inference(subsumption_resolution,[],[f7967,f6531])).
fof(f7967,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,X262) != s(bool,X263) | s(bool,'F_0') = s(bool,X261) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X262)) ) | $spl124),
  inference(subsumption_resolution,[],[f7359,f3625])).
fof(f7359,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,X262) != s(bool,X263) | s(bool,'F_0') = s(bool,X261) | s(bool,X263) != s(bool,X264) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X262)) ) | $spl124),
  inference(superposition,[],[f906,f6531])).
fof(f7966,plain,(
  ( ! [X253,X251,X252] : (s(bool,'F_0') = s(bool,X251) | s(bool,'T_0') = s(bool,X253) | s(bool,'T_0') = s(bool,X252)) ) | $spl124),
  inference(subsumption_resolution,[],[f7356,f6531])).
fof(f7356,plain,(
  ( ! [X253,X251,X252] : (s(bool,X252) != s(bool,X253) | s(bool,'F_0') = s(bool,X251) | s(bool,'T_0') = s(bool,X253) | s(bool,'T_0') = s(bool,X252)) ) | $spl124),
  inference(superposition,[],[f826,f6531])).
fof(f7965,plain,(
  ( ! [X249,X250,X248] : (s(bool,'F_0') = s(bool,X250) | s(bool,X248) != s(bool,X250) | s(bool,'T_0') = s(bool,X249)) ) | $spl124),
  inference(subsumption_resolution,[],[f7355,f306])).
fof(f7355,plain,(
  ( ! [X249,X250,X248] : (s(bool,'F_0') != s(bool,X249) | s(bool,'F_0') = s(bool,X250) | s(bool,X248) != s(bool,X250) | s(bool,'T_0') = s(bool,X249)) ) | $spl124),
  inference(superposition,[],[f824,f6531])).
fof(f7964,plain,(
  ( ! [X246,X247,X245] : (~p(s(bool,X245)) | s(bool,'T_0') = s(bool,X247) | s(bool,'T_0') = s(bool,X246)) ) | $spl124),
  inference(subsumption_resolution,[],[f7354,f6531])).
fof(f7354,plain,(
  ( ! [X246,X247,X245] : (s(bool,X246) != s(bool,X247) | ~p(s(bool,X245)) | s(bool,'T_0') = s(bool,X247) | s(bool,'T_0') = s(bool,X246)) ) | $spl124),
  inference(superposition,[],[f806,f6531])).
fof(f7963,plain,(
  ( ! [X241,X240,X238] : (s(bool,X238) != s(bool,X240) | s(bool,'F_0') = s(bool,X241) | s(bool,X240) != s(bool,X241)) ) | $spl218),
  inference(cnf_transformation,[],[f7963_D])).
fof(f7963_D,plain,(
  ( ! [X241,X240,X238] : (s(bool,X238) != s(bool,X240) | s(bool,'F_0') = s(bool,X241) | s(bool,X240) != s(bool,X241)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7962,plain,(
  ( ! [X241,X239,X240,X238] : (p(s(bool,X239)) | s(bool,X238) != s(bool,X240) | s(bool,X240) != s(bool,X241) | s(bool,'F_0') = s(bool,X241)) ) | ($spl30 | $spl124)),
  inference(subsumption_resolution,[],[f7352,f267])).
fof(f7352,plain,(
  ( ! [X241,X239,X240,X238] : (p(s(bool,X239)) | s(bool,X238) != s(bool,X240) | s(bool,X240) != s(bool,X241) | s(bool,'F_0') = s(bool,X241) | s(bool,'T_0') = s(bool,X239)) ) | ($spl30 | $spl124)),
  inference(superposition,[],[f755,f6531])).
fof(f7961,plain,(
  ( ! [X233,X234,X232] : (s(bool,X233) != s(bool,X234) | p(s(bool,X234)) | ~p(s(bool,X232))) ) | $spl124),
  inference(subsumption_resolution,[],[f7350,f2028])).
fof(f7350,plain,(
  ( ! [X233,X234,X232] : (s(bool,X233) != s(bool,X234) | p(s(bool,X234)) | ~p(s(bool,X232)) | s(bool,'T_0') = s(bool,X233)) ) | $spl124),
  inference(superposition,[],[f368,f6531])).
fof(f7959,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X220) != s(bool,X223)) ) | $spl124),
  inference(subsumption_resolution,[],[f7347,f267])).
fof(f7347,plain,(
  ( ! [X222,X220,X223,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | ~p(s(bool,X223)) | s(bool,X220) != s(bool,X223) | s(bool,'T_0') = s(bool,X221)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f7958,plain,(
  ( ! [X212,X215,X213,X214] : (s(bool,X213) != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X212)) | p(s(bool,X214))) ) | $spl124),
  inference(subsumption_resolution,[],[f7345,f2028])).
fof(f7345,plain,(
  ( ! [X212,X215,X213,X214] : (s(bool,X213) != s(bool,X214) | p(s(bool,X215)) | ~p(s(bool,X212)) | p(s(bool,X214)) | s(bool,'T_0') = s(bool,X213)) ) | $spl124),
  inference(superposition,[],[f307,f6531])).
fof(f7896,plain,(
  ( ! [X208,X209] : (s(bool,'T_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X208)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7343])).
fof(f7343,plain,(
  ( ! [X208,X209] : (s(bool,'T_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X208) | s(bool,'T_0') = s(bool,X209)) ) | $spl124),
  inference(superposition,[],[f306,f6531])).
fof(f7897,plain,(
  ( ! [X204,X205] : (s(bool,'T_0') = s(bool,X205) | ~p(s(bool,X204))) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7341])).
fof(f7341,plain,(
  ( ! [X204,X205] : (s(bool,'T_0') = s(bool,X205) | ~p(s(bool,X204)) | s(bool,'T_0') = s(bool,X205)) ) | $spl124),
  inference(superposition,[],[f270,f6531])).
fof(f7957,plain,(
  ( ! [X202,X203] : (p(s(bool,X203)) | s(bool,'T_0') != s(bool,X202)) ) | $spl124),
  inference(subsumption_resolution,[],[f7340,f267])).
fof(f7340,plain,(
  ( ! [X202,X203] : (p(s(bool,X203)) | s(bool,'T_0') != s(bool,X202) | s(bool,'T_0') = s(bool,X203)) ) | $spl124),
  inference(superposition,[],[f269,f6531])).
fof(f7898,plain,(
  ( ! [X196,X197] : (s(bool,'T_0') = s(bool,X197) | ~p(s(bool,X196))) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7337])).
fof(f7337,plain,(
  ( ! [X196,X197] : (s(bool,'T_0') = s(bool,X197) | ~p(s(bool,X196)) | s(bool,'T_0') = s(bool,X197)) ) | $spl124),
  inference(superposition,[],[f268,f6531])).
fof(f7956,plain,(
  ( ! [X194,X195] : (p(s(bool,X195)) | s(bool,'T_0') != s(bool,X194)) ) | $spl124),
  inference(subsumption_resolution,[],[f7336,f267])).
fof(f7336,plain,(
  ( ! [X194,X195] : (p(s(bool,X195)) | s(bool,'T_0') != s(bool,X194) | s(bool,'T_0') = s(bool,X195)) ) | $spl124),
  inference(superposition,[],[f267,f6531])).
fof(f7955,plain,(
  ( ! [X191,X190] : (p(s(bool,X191)) | s(bool,'F_0') = s(bool,X190)) ) | $spl124),
  inference(subsumption_resolution,[],[f7334,f267])).
fof(f7334,plain,(
  ( ! [X191,X190] : (p(s(bool,X191)) | s(bool,'F_0') = s(bool,X190) | s(bool,'T_0') = s(bool,X191)) ) | $spl124),
  inference(superposition,[],[f264,f6531])).
fof(f7954,plain,(
  ( ! [X185,X184] : (~p(s(bool,X184)) | s(bool,'T_0') = s(bool,X185)) ) | $spl124),
  inference(subsumption_resolution,[],[f7331,f6531])).
fof(f7331,plain,(
  ( ! [X185,X184] : (s(bool,'F_0') != s(bool,X185) | ~p(s(bool,X184)) | s(bool,'T_0') = s(bool,X185)) ) | $spl124),
  inference(superposition,[],[f263,f6531])).
fof(f7953,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X156)) | ~p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X155) != s(bool,X158) | s(bool,X157) != s(bool,X159)) ) | $spl124),
  inference(subsumption_resolution,[],[f7321,f267])).
fof(f7321,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,X156)) | ~p(s(bool,X157)) | ~p(s(bool,X158)) | s(bool,X155) != s(bool,X158) | s(bool,X157) != s(bool,X159) | s(bool,'T_0') = s(bool,X156)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f7952,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X151) != s(bool,X152) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,X150)) | s(bool,X153) != s(bool,X154)) ) | $spl124),
  inference(subsumption_resolution,[],[f7320,f2028])).
fof(f7320,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,X151) != s(bool,X152) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,X150)) | s(bool,X153) != s(bool,X154) | s(bool,'T_0') = s(bool,X151)) ) | $spl124),
  inference(superposition,[],[f224,f6531])).
fof(f7949,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X115) != s(bool,X118)) ) | $spl124),
  inference(subsumption_resolution,[],[f7313,f267])).
fof(f7313,plain,(
  ( ! [X118,X116,X119,X117,X115] : (p(s(bool,X116)) | p(s(bool,X117)) | ~p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X115) != s(bool,X118) | s(bool,'T_0') = s(bool,X116)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f7948,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X106) != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X105)) | ~p(s(bool,X109)) | p(s(bool,X107))) ) | $spl124),
  inference(subsumption_resolution,[],[f7311,f2028])).
fof(f7311,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X106) != s(bool,X107) | p(s(bool,X108)) | ~p(s(bool,X105)) | ~p(s(bool,X109)) | p(s(bool,X107)) | s(bool,'T_0') = s(bool,X106)) ) | $spl124),
  inference(superposition,[],[f223,f6531])).
fof(f7946,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,X90) != s(bool,X93)) ) | $spl124),
  inference(subsumption_resolution,[],[f7308,f267])).
fof(f7308,plain,(
  ( ! [X94,X92,X90,X93,X91] : (p(s(bool,X91)) | p(s(bool,X92)) | ~p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,X90) != s(bool,X93) | s(bool,'T_0') = s(bool,X91)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f7944,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) != s(bool,X72) | p(s(bool,X73)) | ~p(s(bool,X70)) | p(s(bool,X72)) | s(bool,X73) != s(bool,X74)) ) | $spl124),
  inference(subsumption_resolution,[],[f7304,f2028])).
fof(f7304,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X71) != s(bool,X72) | p(s(bool,X73)) | ~p(s(bool,X70)) | p(s(bool,X72)) | s(bool,X73) != s(bool,X74) | s(bool,'T_0') = s(bool,X71)) ) | $spl124),
  inference(superposition,[],[f222,f6531])).
fof(f7943,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X55) != s(bool,X58) | s(bool,X57) != s(bool,X59)) ) | $spl124),
  inference(subsumption_resolution,[],[f7301,f267])).
fof(f7301,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X55) != s(bool,X58) | s(bool,X57) != s(bool,X59) | s(bool,'T_0') = s(bool,X56)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f7942,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X45)) | s(bool,X48) != s(bool,X49)) ) | $spl124),
  inference(subsumption_resolution,[],[f7299,f2028])).
fof(f7299,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,X46) != s(bool,X47) | p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X45)) | s(bool,X48) != s(bool,X49) | s(bool,'T_0') = s(bool,X46)) ) | $spl124),
  inference(superposition,[],[f220,f6531])).
fof(f7939,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23)) ) | $spl124),
  inference(subsumption_resolution,[],[f7294,f267])).
fof(f7294,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X20) != s(bool,X23) | s(bool,'T_0') = s(bool,X21)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f7937,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | ~p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X8) != s(bool,X9)) ) | $spl124),
  inference(subsumption_resolution,[],[f7291,f2028])).
fof(f7291,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | ~p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X6)) ) | $spl124),
  inference(superposition,[],[f218,f6531])).
fof(f7899,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7285])).
fof(f7285,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) ) | $spl124),
  inference(superposition,[],[f6531,f306])).
fof(f7934,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,X55) | p(s(bool,X54))) ) | $spl124),
  inference(subsumption_resolution,[],[f7284,f267])).
fof(f7284,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55) | p(s(bool,X54))) ) | $spl124),
  inference(superposition,[],[f6531,f264])).
fof(f7900,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7281])).
fof(f7281,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl124),
  inference(superposition,[],[f6531,f6531])).
fof(f7901,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,X28) = s(bool,X31) | s(bool,'F_0') = s(bool,X31)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7279])).
fof(f7279,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X31)) ) | $spl124),
  inference(superposition,[],[f6531,f2782])).
fof(f7904,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X52)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7266])).
fof(f7266,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X52)) ) | $spl124),
  inference(superposition,[],[f6531,f2782])).
fof(f7930,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') = s(bool,X87) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X84),X85))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X84),X86)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7247,f2062])).
fof(f7247,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,X87) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X83),s('type.decSem.state'(X84),X85),s('type.exhSem.state'(X84),X86))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X84),X85))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X84),X86)))) ) | $spl124),
  inference(superposition,[],[f6531,f2061])).
fof(f7929,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'F_0') = s(bool,X82) | s('type.ffi.ffi_state'(X79),'const.decSem.state_ffi_1'(s('type.decSem.state'(X79),X80))) = s('type.ffi.ffi_state'(X79),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X79),X81)))) ) | $spl124),
  inference(subsumption_resolution,[],[f7246,f2266])).
fof(f7246,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'F_0') = s(bool,X82) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X78),s('type.decSem.state'(X79),X80),s('type.exhSem.state'(X79),X81))) | s('type.ffi.ffi_state'(X79),'const.decSem.state_ffi_1'(s('type.decSem.state'(X79),X80))) = s('type.ffi.ffi_state'(X79),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X79),X81)))) ) | $spl124),
  inference(superposition,[],[f6531,f2265])).
fof(f7907,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X52)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7234])).
fof(f7234,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl124),
  inference(superposition,[],[f6531,f6531])).
fof(f7908,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7231])).
fof(f7231,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X45)) ) | $spl124),
  inference(superposition,[],[f6531,f2782])).
fof(f7911,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7221])).
fof(f7221,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl124),
  inference(superposition,[],[f6531,f306])).
fof(f7926,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X0))) ) | $spl124),
  inference(subsumption_resolution,[],[f7217,f267])).
fof(f7217,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | $spl124),
  inference(superposition,[],[f6531,f264])).
fof(f7913,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X48)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7199])).
fof(f7199,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X48)) ) | $spl124),
  inference(superposition,[],[f6531,f6531])).
fof(f7914,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,X22) = s(bool,X25) | s(bool,'F_0') = s(bool,X24)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7193])).
fof(f7193,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,X22) = s(bool,X25) | s(bool,'F_0') = s(bool,X24)) ) | $spl124),
  inference(superposition,[],[f6531,f2782])).
fof(f7917,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7187])).
fof(f7187,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl124),
  inference(superposition,[],[f6531,f306])).
fof(f7918,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7186])).
fof(f7186,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X4))) ) | $spl124),
  inference(superposition,[],[f6531,f270])).
fof(f7919,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl124),
  inference(duplicate_literal_removal,[],[f7185])).
fof(f7185,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl124),
  inference(superposition,[],[f6531,f268])).
fof(f7175,plain,(
  ( ! [X88,X87,X89,X86,X84] : (s(bool,X84) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X86),s('type.decSem.state'(X87),X88),s('type.exhSem.state'(X87),X89))) | s(bool,'T_0') = s(bool,X84) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X87),X88))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X87),X89)))) ) | $spl212),
  inference(cnf_transformation,[],[f7175_D])).
fof(f7175_D,plain,(
  ( ! [X88,X87,X89,X86,X84] : (s(bool,X84) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X86),s('type.decSem.state'(X87),X88),s('type.exhSem.state'(X87),X89))) | s(bool,'T_0') = s(bool,X84) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X87),X88))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X87),X89)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f6546,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | $spl128),
  inference(cnf_transformation,[],[f6546_D])).
fof(f6546_D,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7149,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (~p(s(bool,X85)) | s(bool,X84) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X86),s('type.decSem.state'(X87),X88),s('type.exhSem.state'(X87),X89))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X87),X88))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X87),X89))) | s(bool,X85) = s(bool,X90) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X90)) ) | $spl64),
  inference(superposition,[],[f2101,f2782])).
fof(f7174,plain,(
  ( ! [X70,X76,X74,X72,X75,X73] : (s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,X70) = s(bool,X76) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X73),X74))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X73),X75)))) ) | $spl210),
  inference(cnf_transformation,[],[f7174_D])).
fof(f7174_D,plain,(
  ( ! [X70,X76,X74,X72,X75,X73] : (s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,X70) = s(bool,X76) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X73),X74))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X73),X75)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f7173,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X73),X74))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X73),X75))) | s(bool,X70) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) ) | $spl64),
  inference(subsumption_resolution,[],[f7147,f263])).
fof(f7147,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X72),s('type.decSem.state'(X73),X74),s('type.exhSem.state'(X73),X75))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X73),X74))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X73),X75))) | s(bool,X70) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X71)) ) | $spl64),
  inference(superposition,[],[f2101,f2782])).
fof(f7171,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X66),X67))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X66),X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X63)) ) | $spl64),
  inference(subsumption_resolution,[],[f7146,f806])).
fof(f7146,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X66),X67))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X66),X68))) | s(bool,X64) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X63)) ) | $spl64),
  inference(superposition,[],[f2101,f2782])).
fof(f7166,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X2),X3))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X2),X4)))) ) | $spl206),
  inference(cnf_transformation,[],[f7166_D])).
fof(f7166_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s(bool,X0) != s(bool,X6) | ~p(s(bool,X6)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X2),X3))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X2),X4)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7112,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,X94) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X95),s('type.decSem.state'(X96),X97),s('type.exhSem.state'(X96),X98))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X96),X97))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X96),X98))) | p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X94) != s(bool,X100)) ) | $spl64),
  inference(resolution,[],[f2101,f307])).
fof(f7105,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X57),s('type.decSem.state'(X58),X59),s('type.exhSem.state'(X58),X60))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X58),X59))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X58),X60))) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X56) != s(bool,X62) | s(bool,X61) != s(bool,X63)) ) | $spl64),
  inference(resolution,[],[f2101,f224])).
fof(f7103,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X41),s('type.decSem.state'(X42),X43),s('type.exhSem.state'(X42),X44))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X42),X43))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X42),X44))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X40) != s(bool,X46)) ) | $spl64),
  inference(resolution,[],[f2101,f223])).
fof(f7101,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X26),X27))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X26),X28))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X24) != s(bool,X30)) ) | $spl64),
  inference(resolution,[],[f2101,f222])).
fof(f7100,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X17),s('type.decSem.state'(X18),X19),s('type.exhSem.state'(X18),X20))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X18),X19))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X18),X20))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X16) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl64),
  inference(resolution,[],[f2101,f220])).
fof(f7098,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.decSem.state'(X2),X3),s('type.exhSem.state'(X2),X4))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X2),X3))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X2),X4))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X0) != s(bool,X6)) ) | $spl64),
  inference(resolution,[],[f2101,f218])).
fof(f6378,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f2782])).
fof(f6527,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44)) ) | $spl120),
  inference(cnf_transformation,[],[f6527_D])).
fof(f6527_D,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X44)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7082,plain,(
  ( ! [X422,X424,X421,X423,X425] : (s(bool,'T_0') = s(bool,X423) | s(bool,X421) = s(bool,X423) | s(bool,'T_0') = s(bool,X422) | s(bool,X421) = s(bool,X424) | s(bool,X424) = s(bool,X425) | s(bool,'T_0') = s(bool,X425)) )),
  inference(subsumption_resolution,[],[f6314,f306])).
fof(f6314,plain,(
  ( ! [X422,X424,X421,X423,X425] : (s(bool,'F_0') != s(bool,X422) | s(bool,'T_0') = s(bool,X423) | s(bool,X421) = s(bool,X423) | s(bool,'T_0') = s(bool,X422) | s(bool,X421) = s(bool,X424) | s(bool,X424) = s(bool,X425) | s(bool,'T_0') = s(bool,X425)) )),
  inference(superposition,[],[f3625,f2782])).
fof(f7076,plain,(
  ( ! [X374,X373,X375,X371] : (p(s(bool,X373)) | s(bool,'T_0') = s(bool,X375) | s(bool,X374) = s(bool,X375) | s(bool,X371) = s(bool,X374) | s(bool,X371) = s(bool,X373)) ) | $spl204),
  inference(cnf_transformation,[],[f7076_D])).
fof(f7076_D,plain,(
  ( ! [X374,X373,X375,X371] : (p(s(bool,X373)) | s(bool,'T_0') = s(bool,X375) | s(bool,X374) = s(bool,X375) | s(bool,X371) = s(bool,X374) | s(bool,X371) = s(bool,X373)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7075,plain,(
  ( ! [X372,X374,X373,X375,X371] : (p(s(bool,X373)) | p(s(bool,X372)) | s(bool,X371) = s(bool,X373) | s(bool,X371) = s(bool,X374) | s(bool,X374) = s(bool,X375) | s(bool,'T_0') = s(bool,X375)) ) | $spl66),
  inference(subsumption_resolution,[],[f6303,f264])).
fof(f6303,plain,(
  ( ! [X372,X374,X373,X375,X371] : (s(bool,'F_0') != s(bool,X372) | p(s(bool,X373)) | p(s(bool,X372)) | s(bool,X371) = s(bool,X373) | s(bool,X371) = s(bool,X374) | s(bool,X374) = s(bool,X375) | s(bool,'T_0') = s(bool,X375)) ) | $spl66),
  inference(superposition,[],[f2249,f2782])).
fof(f4697,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | $spl104),
  inference(cnf_transformation,[],[f4697_D])).
fof(f4697_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f7064,plain,(
  ( ! [X277,X273,X275,X276,X274] : (p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X273) != s(bool,X275) | s(bool,X273) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'T_0') = s(bool,X277)) )),
  inference(subsumption_resolution,[],[f6280,f271])).
fof(f6280,plain,(
  ( ! [X277,X273,X275,X276,X274] : (p(s(bool,'F_0')) | p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,X273) != s(bool,X275) | s(bool,X273) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'T_0') = s(bool,X277)) )),
  inference(superposition,[],[f307,f2782])).
fof(f7057,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (~p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X186) != s(bool,X188) | s(bool,X187) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(subsumption_resolution,[],[f6256,f271])).
fof(f6256,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,'F_0')) | ~p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X186) != s(bool,X188) | s(bool,X187) != s(bool,X189) | s(bool,X186) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f224,f2782])).
fof(f7052,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X142) | s(bool,X142) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f6248,f271])).
fof(f6248,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X139)) | ~p(s(bool,X140)) | ~p(s(bool,X141)) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X142) | s(bool,X142) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f223,f2782])).
fof(f7049,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X109) != s(bool,X111) | s(bool,X108) != s(bool,X110) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f6243,f271])).
fof(f6243,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,'F_0')) | p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X109) != s(bool,X111) | s(bool,X108) != s(bool,X110) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f222,f2782])).
fof(f7044,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X66) != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f6236,f271])).
fof(f6236,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,'F_0')) | ~p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X66) != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f220,f2782])).
fof(f7039,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,X24) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f6229,f271])).
fof(f6229,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'F_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26) | s(bool,X24) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f218,f2782])).
fof(f7027,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,'T_0') != s(bool,X401) | p(s(bool,X402)) | s(bool,X401) = s(bool,X403) | s(bool,X403) = s(bool,X404) | s(bool,'T_0') = s(bool,X404)) )),
  inference(subsumption_resolution,[],[f6206,f264])).
fof(f6206,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,'T_0') != s(bool,X401) | s(bool,'F_0') != s(bool,X402) | p(s(bool,X402)) | s(bool,X401) = s(bool,X403) | s(bool,X403) = s(bool,X404) | s(bool,'T_0') = s(bool,X404)) )),
  inference(superposition,[],[f2028,f2782])).
fof(f7024,plain,(
  ( ! [X372,X370,X369,X371] : (s(bool,'T_0') != s(bool,X369) | s(bool,'F_0') != s(bool,X370) | s(bool,X369) = s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372)) ) | $spl44),
  inference(subsumption_resolution,[],[f6198,f1729])).
fof(f6198,plain,(
  ( ! [X372,X370,X369,X371] : (s(bool,'T_0') != s(bool,X369) | s(bool,'T_0') = s(bool,X370) | s(bool,'F_0') != s(bool,X370) | s(bool,X369) = s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372)) ) | $spl44),
  inference(superposition,[],[f1278,f2782])).
fof(f7023,plain,(
  ( ! [X366,X362,X365,X363] : (s(bool,X362) != s(bool,X363) | s(bool,'T_0') = s(bool,X366) | s(bool,X365) = s(bool,X366) | s(bool,X362) = s(bool,X365) | s(bool,'F_0') = s(bool,X363)) ) | $spl200),
  inference(cnf_transformation,[],[f7023_D])).
fof(f7023_D,plain,(
  ( ! [X366,X362,X365,X363] : (s(bool,X362) != s(bool,X363) | s(bool,'T_0') = s(bool,X366) | s(bool,X365) = s(bool,X366) | s(bool,X362) = s(bool,X365) | s(bool,'F_0') = s(bool,X363)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7022,plain,(
  ( ! [X364,X366,X362,X365,X363] : (s(bool,X362) != s(bool,X363) | s(bool,'F_0') = s(bool,X363) | s(bool,'T_0') = s(bool,X364) | s(bool,X362) = s(bool,X365) | s(bool,X365) = s(bool,X366) | s(bool,'T_0') = s(bool,X366)) )),
  inference(subsumption_resolution,[],[f6196,f306])).
fof(f6196,plain,(
  ( ! [X364,X366,X362,X365,X363] : (s(bool,X362) != s(bool,X363) | s(bool,'F_0') = s(bool,X363) | s(bool,'F_0') != s(bool,X364) | s(bool,'T_0') = s(bool,X364) | s(bool,X362) = s(bool,X365) | s(bool,X365) = s(bool,X366) | s(bool,'T_0') = s(bool,X366)) )),
  inference(superposition,[],[f906,f2782])).
fof(f7017,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | p(s(bool,X313)) | ~p(s(bool,X312)) | s(bool,X311) = s(bool,X314) | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(subsumption_resolution,[],[f6184,f271])).
fof(f6184,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,X311) != s(bool,X312) | p(s(bool,X313)) | ~p(s(bool,X312)) | p(s(bool,'F_0')) | s(bool,X311) = s(bool,X314) | s(bool,X314) = s(bool,X315) | s(bool,'T_0') = s(bool,X315)) )),
  inference(superposition,[],[f307,f2782])).
fof(f7011,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | ~p(s(bool,X223)) | ~p(s(bool,X222)) | s(bool,X223) != s(bool,X224) | s(bool,X221) = s(bool,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(subsumption_resolution,[],[f6159,f271])).
fof(f6159,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (s(bool,X221) != s(bool,X222) | p(s(bool,'F_0')) | ~p(s(bool,X223)) | ~p(s(bool,X222)) | s(bool,X223) != s(bool,X224) | s(bool,X221) = s(bool,X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f224,f2782])).
fof(f7009,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,X180)) | ~p(s(bool,X182)) | s(bool,X179) = s(bool,X183) | s(bool,X183) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(subsumption_resolution,[],[f6152,f271])).
fof(f6152,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,X179) != s(bool,X180) | p(s(bool,X181)) | ~p(s(bool,X180)) | ~p(s(bool,X182)) | p(s(bool,'F_0')) | s(bool,X179) = s(bool,X183) | s(bool,X183) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f223,f2782])).
fof(f7006,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X137) != s(bool,X138) | p(s(bool,X139)) | ~p(s(bool,X138)) | s(bool,X139) != s(bool,X140) | s(bool,X137) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(subsumption_resolution,[],[f6145,f271])).
fof(f6145,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X137) != s(bool,X138) | p(s(bool,X139)) | ~p(s(bool,X138)) | p(s(bool,'F_0')) | s(bool,X139) != s(bool,X140) | s(bool,X137) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f222,f2782])).
fof(f7004,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,X107) != s(bool,X108) | ~p(s(bool,X109)) | ~p(s(bool,X108)) | s(bool,X109) != s(bool,X110) | s(bool,X107) = s(bool,X111) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f6140,f271])).
fof(f6140,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,X107) != s(bool,X108) | p(s(bool,'F_0')) | ~p(s(bool,X109)) | ~p(s(bool,X108)) | s(bool,X109) != s(bool,X110) | s(bool,X107) = s(bool,X111) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f220,f2782])).
fof(f7001,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X49) != s(bool,X50) | s(bool,X47) = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f6130,f271])).
fof(f6130,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X47) != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X48)) | p(s(bool,'F_0')) | s(bool,X49) != s(bool,X50) | s(bool,X47) = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6991,plain,(
  ( ! [X532,X530,X533,X531] : (s(bool,X530) != s(bool,X531) | s(bool,'T_0') != s(bool,X531) | s(bool,X532) = s(bool,X533) | s(bool,'F_0') = s(bool,X533) | s(bool,'F_0') = s(bool,X532)) ) | $spl44),
  inference(subsumption_resolution,[],[f6100,f1278])).
fof(f6100,plain,(
  ( ! [X532,X530,X533,X531] : (s(bool,'T_0') != s(bool,X530) | s(bool,X530) != s(bool,X531) | s(bool,'T_0') != s(bool,X531) | s(bool,X532) = s(bool,X533) | s(bool,'F_0') = s(bool,X533) | s(bool,'F_0') = s(bool,X532)) )),
  inference(superposition,[],[f4878,f2782])).
fof(f6990,plain,(
  ( ! [X528,X526,X529,X527] : (s(bool,'T_0') != s(bool,X526) | s(bool,'T_0') = s(bool,X527) | s(bool,X528) = s(bool,X529) | s(bool,'F_0') = s(bool,X529) | s(bool,'F_0') = s(bool,X528)) ) | $spl44),
  inference(subsumption_resolution,[],[f6099,f1278])).
fof(f6099,plain,(
  ( ! [X528,X526,X529,X527] : (s(bool,'T_0') != s(bool,X526) | s(bool,'T_0') = s(bool,X527) | s(bool,X526) = s(bool,X527) | s(bool,X528) = s(bool,X529) | s(bool,'F_0') = s(bool,X529) | s(bool,'F_0') = s(bool,X528)) )),
  inference(superposition,[],[f2327,f2782])).
fof(f6989,plain,(
  ( ! [X524,X522,X525,X523] : (s(bool,'T_0') != s(bool,X522) | p(s(bool,X523)) | s(bool,X524) = s(bool,X525) | s(bool,'F_0') = s(bool,X525) | s(bool,'F_0') = s(bool,X524)) )),
  inference(subsumption_resolution,[],[f6098,f2028])).
fof(f6098,plain,(
  ( ! [X524,X522,X525,X523] : (s(bool,'T_0') != s(bool,X522) | s(bool,X522) = s(bool,X523) | p(s(bool,X523)) | s(bool,X524) = s(bool,X525) | s(bool,'F_0') = s(bool,X525) | s(bool,'F_0') = s(bool,X524)) )),
  inference(superposition,[],[f2198,f2782])).
fof(f6417,plain,(
  ( ! [X520,X521,X519] : (s(bool,'T_0') != s(bool,X519) | s(bool,X520) = s(bool,X521) | s(bool,'F_0') = s(bool,X521) | s(bool,'F_0') = s(bool,X520)) )),
  inference(duplicate_literal_removal,[],[f6097])).
fof(f6097,plain,(
  ( ! [X520,X521,X519] : (s(bool,'T_0') != s(bool,X519) | s(bool,'T_0') != s(bool,X519) | s(bool,X520) = s(bool,X521) | s(bool,'F_0') = s(bool,X521) | s(bool,'F_0') = s(bool,X520)) )),
  inference(superposition,[],[f1729,f2782])).
fof(f6988,plain,(
  ( ! [X516,X518,X517,X515] : (~p(s(bool,X516)) | s(bool,X515) != s(bool,X516) | s(bool,X517) = s(bool,X518) | s(bool,'F_0') = s(bool,X518) | s(bool,'F_0') = s(bool,X517)) ) | $spl48),
  inference(subsumption_resolution,[],[f6096,f806])).
fof(f6096,plain,(
  ( ! [X516,X518,X517,X515] : (s(bool,'T_0') != s(bool,X515) | ~p(s(bool,X516)) | s(bool,X515) != s(bool,X516) | s(bool,X517) = s(bool,X518) | s(bool,'F_0') = s(bool,X518) | s(bool,'F_0') = s(bool,X517)) ) | $spl48),
  inference(superposition,[],[f1356,f2782])).
fof(f6987,plain,(
  ( ! [X509,X511,X512,X510] : (s(bool,'F_0') = s(bool,X510) | s(bool,X509) != s(bool,X510) | s(bool,X511) = s(bool,X512) | s(bool,'F_0') = s(bool,X512) | s(bool,'F_0') = s(bool,X511)) )),
  inference(subsumption_resolution,[],[f6094,f826])).
fof(f6094,plain,(
  ( ! [X509,X511,X512,X510] : (s(bool,'T_0') != s(bool,X509) | s(bool,'F_0') = s(bool,X510) | s(bool,X509) != s(bool,X510) | s(bool,X511) = s(bool,X512) | s(bool,'F_0') = s(bool,X512) | s(bool,'F_0') = s(bool,X511)) )),
  inference(superposition,[],[f824,f2782])).
fof(f6419,plain,(
  ( ! [X507,X508,X506] : (s(bool,'T_0') = s(bool,X506) | s(bool,X507) = s(bool,X508) | s(bool,'F_0') = s(bool,X508) | s(bool,'F_0') = s(bool,X507)) )),
  inference(duplicate_literal_removal,[],[f6093])).
fof(f6093,plain,(
  ( ! [X507,X508,X506] : (s(bool,'T_0') = s(bool,X506) | s(bool,'T_0') = s(bool,X506) | s(bool,X507) = s(bool,X508) | s(bool,'F_0') = s(bool,X508) | s(bool,'F_0') = s(bool,X507)) )),
  inference(superposition,[],[f306,f2782])).
fof(f6985,plain,(
  ( ! [X502,X501,X503] : (p(s(bool,X501)) | s(bool,X502) = s(bool,X503) | s(bool,'F_0') = s(bool,X503) | s(bool,'F_0') = s(bool,X502)) )),
  inference(subsumption_resolution,[],[f6091,f267])).
fof(f6091,plain,(
  ( ! [X502,X501,X503] : (s(bool,'T_0') = s(bool,X501) | p(s(bool,X501)) | s(bool,X502) = s(bool,X503) | s(bool,'F_0') = s(bool,X503) | s(bool,'F_0') = s(bool,X502)) )),
  inference(superposition,[],[f264,f2782])).
fof(f6984,plain,(
  ( ! [X499,X500,X498] : (~p(s(bool,X498)) | s(bool,X499) = s(bool,X500) | s(bool,'F_0') = s(bool,X500) | s(bool,'F_0') = s(bool,X499)) )),
  inference(subsumption_resolution,[],[f6090,f268])).
fof(f6090,plain,(
  ( ! [X499,X500,X498] : (s(bool,'T_0') != s(bool,X498) | ~p(s(bool,X498)) | s(bool,X499) = s(bool,X500) | s(bool,'F_0') = s(bool,X500) | s(bool,'F_0') = s(bool,X499)) )),
  inference(superposition,[],[f263,f2782])).
fof(f6972,plain,(
  ( ! [X377,X374,X376,X378,X375] : (s(bool,'T_0') != s(bool,X375) | p(s(bool,X376)) | p(s(bool,X374)) | s(bool,X377) = s(bool,X378) | s(bool,X374) = s(bool,X378) | s(bool,'F_0') = s(bool,X377)) ) | $spl66),
  inference(subsumption_resolution,[],[f6059,f2028])).
fof(f6059,plain,(
  ( ! [X377,X374,X376,X378,X375] : (s(bool,'T_0') != s(bool,X375) | p(s(bool,X376)) | p(s(bool,X374)) | s(bool,X375) = s(bool,X376) | s(bool,X377) = s(bool,X378) | s(bool,X374) = s(bool,X378) | s(bool,'F_0') = s(bool,X377)) ) | $spl66),
  inference(superposition,[],[f2249,f2782])).
fof(f6791,plain,(
  ( ! [X420,X422,X424,X423] : (s(bool,X420) = s(bool,X422) | s(bool,'F_0') = s(bool,X424) | s(bool,X420) = s(bool,X424) | s(bool,'T_0') = s(bool,X423) | s(bool,X422) != s(bool,X423)) ) | $spl170),
  inference(cnf_transformation,[],[f6791_D])).
fof(f6791_D,plain,(
  ( ! [X420,X422,X424,X423] : (s(bool,X420) = s(bool,X422) | s(bool,'F_0') = s(bool,X424) | s(bool,X420) = s(bool,X424) | s(bool,'T_0') = s(bool,X423) | s(bool,X422) != s(bool,X423)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f6970,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,'T_0') != s(bool,X320) | s(bool,X319) != s(bool,X321) | s(bool,'T_0') = s(bool,X321) | s(bool,X322) = s(bool,X323) | s(bool,X319) = s(bool,X323) | s(bool,'F_0') = s(bool,X322)) )),
  inference(subsumption_resolution,[],[f6046,f1729])).
fof(f6046,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,'T_0') != s(bool,X320) | s(bool,'F_0') = s(bool,X320) | s(bool,X319) != s(bool,X321) | s(bool,'T_0') = s(bool,X321) | s(bool,X322) = s(bool,X323) | s(bool,X319) = s(bool,X323) | s(bool,'F_0') = s(bool,X322)) )),
  inference(superposition,[],[f906,f2782])).
fof(f6967,plain,(
  ( ! [X292,X290,X293,X291] : (s(bool,'T_0') != s(bool,X291) | p(s(bool,X290)) | s(bool,X292) = s(bool,X293) | s(bool,X290) = s(bool,X293) | s(bool,'F_0') = s(bool,X292)) )),
  inference(subsumption_resolution,[],[f6039,f267])).
fof(f6039,plain,(
  ( ! [X292,X290,X293,X291] : (s(bool,'T_0') != s(bool,X291) | p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X292) = s(bool,X293) | s(bool,X290) = s(bool,X293) | s(bool,'F_0') = s(bool,X292)) )),
  inference(superposition,[],[f368,f2782])).
fof(f6917,plain,(
  ( ! [X94,X99,X95,X98] : (s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X98) | s(bool,X94) = s(bool,X99) | s(bool,X98) = s(bool,X99) | p(s(bool,X95))) ) | $spl196),
  inference(cnf_transformation,[],[f6917_D])).
fof(f6917_D,plain,(
  ( ! [X94,X99,X95,X98] : (s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X98) | s(bool,X94) = s(bool,X99) | s(bool,X98) = s(bool,X99) | p(s(bool,X95))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f6966,plain,(
  ( ! [X282,X285,X281,X283,X284] : (p(s(bool,X282)) | p(s(bool,X283)) | s(bool,X281) != s(bool,X283) | s(bool,X284) = s(bool,X285) | s(bool,X281) = s(bool,X285) | s(bool,'F_0') = s(bool,X284)) )),
  inference(subsumption_resolution,[],[f6037,f308])).
fof(f6037,plain,(
  ( ! [X282,X285,X281,X283,X284] : (~p(s(bool,'T_0')) | p(s(bool,X282)) | p(s(bool,X283)) | s(bool,X281) != s(bool,X283) | s(bool,X284) = s(bool,X285) | s(bool,X281) = s(bool,X285) | s(bool,'F_0') = s(bool,X284)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6965,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s(bool,'T_0') != s(bool,X267) | p(s(bool,X268)) | p(s(bool,X266)) | s(bool,X269) = s(bool,X270) | s(bool,X266) = s(bool,X270) | s(bool,'F_0') = s(bool,X269)) )),
  inference(subsumption_resolution,[],[f6034,f267])).
fof(f6034,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s(bool,'T_0') != s(bool,X267) | p(s(bool,X268)) | ~p(s(bool,X267)) | p(s(bool,X266)) | s(bool,X269) = s(bool,X270) | s(bool,X266) = s(bool,X270) | s(bool,'F_0') = s(bool,X269)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6962,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (p(s(bool,X199)) | ~p(s(bool,X200)) | s(bool,X198) != s(bool,X199) | s(bool,X200) != s(bool,X201) | s(bool,X202) = s(bool,X203) | s(bool,X198) = s(bool,X203) | s(bool,'F_0') = s(bool,X202)) )),
  inference(subsumption_resolution,[],[f6014,f308])).
fof(f6014,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X199)) | ~p(s(bool,X200)) | s(bool,X198) != s(bool,X199) | s(bool,X200) != s(bool,X201) | s(bool,X202) = s(bool,X203) | s(bool,X198) = s(bool,X203) | s(bool,'F_0') = s(bool,X202)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6961,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,'T_0') != s(bool,X175) | p(s(bool,X174)) | ~p(s(bool,X176)) | s(bool,X176) != s(bool,X177) | s(bool,X178) = s(bool,X179) | s(bool,X174) = s(bool,X179) | s(bool,'F_0') = s(bool,X178)) )),
  inference(subsumption_resolution,[],[f6010,f267])).
fof(f6010,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,'T_0') != s(bool,X175) | p(s(bool,X174)) | ~p(s(bool,X176)) | ~p(s(bool,X175)) | s(bool,X176) != s(bool,X177) | s(bool,X178) = s(bool,X179) | s(bool,X174) = s(bool,X179) | s(bool,'F_0') = s(bool,X178)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6958,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X150) != s(bool,X153) | s(bool,X154) = s(bool,X155) | s(bool,X150) = s(bool,X155) | s(bool,'F_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f6006,f308])).
fof(f6006,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | ~p(s(bool,X152)) | p(s(bool,X153)) | s(bool,X150) != s(bool,X153) | s(bool,X154) = s(bool,X155) | s(bool,X150) = s(bool,X155) | s(bool,'F_0') = s(bool,X154)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6957,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,X133) | p(s(bool,X134)) | ~p(s(bool,X135)) | p(s(bool,X132)) | s(bool,X136) = s(bool,X137) | s(bool,X132) = s(bool,X137) | s(bool,'F_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f6003,f267])).
fof(f6003,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,X133) | p(s(bool,X134)) | ~p(s(bool,X133)) | ~p(s(bool,X135)) | p(s(bool,X132)) | s(bool,X136) = s(bool,X137) | s(bool,X132) = s(bool,X137) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6956,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X120) != s(bool,X122) | s(bool,X124) = s(bool,X125) | s(bool,X120) = s(bool,X125) | s(bool,'F_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f6001,f308])).
fof(f6001,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,'T_0')) | p(s(bool,X121)) | p(s(bool,X122)) | s(bool,X121) != s(bool,X123) | s(bool,X120) != s(bool,X122) | s(bool,X124) = s(bool,X125) | s(bool,X120) = s(bool,X125) | s(bool,'F_0') = s(bool,X124)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6953,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') != s(bool,X91) | p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X92) != s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,X90) = s(bool,X95) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f5996,f267])).
fof(f5996,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') != s(bool,X91) | p(s(bool,X92)) | ~p(s(bool,X91)) | p(s(bool,X90)) | s(bool,X92) != s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,X90) = s(bool,X95) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6952,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X78) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X82) = s(bool,X83) | s(bool,X78) = s(bool,X83) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f5994,f308])).
fof(f5994,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X78) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X82) = s(bool,X83) | s(bool,X78) = s(bool,X83) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6951,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X60)) | ~p(s(bool,X62)) | s(bool,X62) != s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,X60) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5991,f267])).
fof(f5991,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X60)) | ~p(s(bool,X62)) | ~p(s(bool,X61)) | s(bool,X62) != s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,X60) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6948,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38) | s(bool,X40) = s(bool,X41) | s(bool,X36) = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5987,f308])).
fof(f5987,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X36) != s(bool,X38) | s(bool,X40) = s(bool,X41) | s(bool,X36) = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6945,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5981,f267])).
fof(f5981,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X4) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6943,plain,(
  ( ! [X478,X480,X477,X479,X481] : (p(s(bool,X477)) | s(bool,X478) != s(bool,X479) | s(bool,'T_0') != s(bool,X479) | s(bool,X480) = s(bool,X481) | s(bool,X477) = s(bool,X481) | s(bool,'F_0') = s(bool,X480)) ) | $spl44),
  inference(subsumption_resolution,[],[f5978,f1278])).
fof(f5978,plain,(
  ( ! [X478,X480,X477,X479,X481] : (p(s(bool,X477)) | s(bool,'T_0') != s(bool,X478) | s(bool,X478) != s(bool,X479) | s(bool,'T_0') != s(bool,X479) | s(bool,X480) = s(bool,X481) | s(bool,X477) = s(bool,X481) | s(bool,'F_0') = s(bool,X480)) )),
  inference(superposition,[],[f3958,f2782])).
fof(f6429,plain,(
  ( ! [X472,X469,X471,X473,X470] : (s(bool,X469) != s(bool,X470) | s(bool,'T_0') = s(bool,X471) | s(bool,'T_0') = s(bool,X470) | s(bool,X472) = s(bool,X473) | s(bool,X469) = s(bool,X473) | s(bool,'F_0') = s(bool,X472)) )),
  inference(duplicate_literal_removal,[],[f5976])).
fof(f5976,plain,(
  ( ! [X472,X469,X471,X473,X470] : (s(bool,X469) != s(bool,X470) | s(bool,'T_0') = s(bool,X471) | s(bool,'T_0') = s(bool,X471) | s(bool,'T_0') = s(bool,X470) | s(bool,X472) = s(bool,X473) | s(bool,X469) = s(bool,X473) | s(bool,'F_0') = s(bool,X472)) )),
  inference(superposition,[],[f3625,f2782])).
fof(f6942,plain,(
  ( ! [X466,X465,X467,X468,X464] : (p(s(bool,X464)) | s(bool,'T_0') = s(bool,X465) | s(bool,'T_0') = s(bool,X466) | s(bool,X467) = s(bool,X468) | s(bool,X464) = s(bool,X468) | s(bool,'F_0') = s(bool,X467)) )),
  inference(subsumption_resolution,[],[f5975,f2328])).
fof(f5975,plain,(
  ( ! [X466,X465,X467,X468,X464] : (p(s(bool,X464)) | s(bool,'T_0') = s(bool,X465) | s(bool,X465) != s(bool,X466) | s(bool,'T_0') = s(bool,X466) | s(bool,X467) = s(bool,X468) | s(bool,X464) = s(bool,X468) | s(bool,'F_0') = s(bool,X467)) )),
  inference(superposition,[],[f3624,f2782])).
fof(f6941,plain,(
  ( ! [X453,X449,X451,X452,X450] : (p(s(bool,X449)) | s(bool,'T_0') != s(bool,X450) | s(bool,'T_0') = s(bool,X451) | s(bool,X452) = s(bool,X453) | s(bool,X449) = s(bool,X453) | s(bool,'F_0') = s(bool,X452)) ) | ($spl44 | $spl70)),
  inference(subsumption_resolution,[],[f5972,f1278])).
fof(f5972,plain,(
  ( ! [X453,X449,X451,X452,X450] : (p(s(bool,X449)) | s(bool,'T_0') != s(bool,X450) | s(bool,X450) = s(bool,X451) | s(bool,'T_0') = s(bool,X451) | s(bool,X452) = s(bool,X453) | s(bool,X449) = s(bool,X453) | s(bool,'F_0') = s(bool,X452)) ) | $spl70),
  inference(superposition,[],[f2350,f2782])).
fof(f6940,plain,(
  ( ! [X436,X432,X434,X433,X435] : (s(bool,X432) != s(bool,X433) | p(s(bool,X434)) | p(s(bool,X433)) | s(bool,X435) = s(bool,X436) | s(bool,X432) = s(bool,X436) | s(bool,'F_0') = s(bool,X435)) ) | $spl66),
  inference(subsumption_resolution,[],[f5968,f267])).
fof(f5968,plain,(
  ( ! [X436,X432,X434,X433,X435] : (s(bool,X432) != s(bool,X433) | p(s(bool,X434)) | p(s(bool,X433)) | s(bool,'T_0') = s(bool,X434) | s(bool,X435) = s(bool,X436) | s(bool,X432) = s(bool,X436) | s(bool,'F_0') = s(bool,X435)) ) | $spl66),
  inference(superposition,[],[f2249,f2782])).
fof(f6939,plain,(
  ( ! [X424,X426,X423,X425] : (p(s(bool,X423)) | s(bool,'T_0') != s(bool,X424) | s(bool,X425) = s(bool,X426) | s(bool,X423) = s(bool,X426) | s(bool,'F_0') = s(bool,X425)) ) | $spl62),
  inference(subsumption_resolution,[],[f5966,f1729])).
fof(f5966,plain,(
  ( ! [X424,X426,X423,X425] : (p(s(bool,X423)) | s(bool,'T_0') != s(bool,X424) | s(bool,'F_0') = s(bool,X424) | s(bool,X425) = s(bool,X426) | s(bool,X423) = s(bool,X426) | s(bool,'F_0') = s(bool,X425)) ) | $spl62),
  inference(superposition,[],[f1999,f2782])).
fof(f6431,plain,(
  ( ! [X420,X422,X421,X419] : (p(s(bool,X419)) | s(bool,'T_0') = s(bool,X420) | s(bool,X421) = s(bool,X422) | s(bool,X419) = s(bool,X422) | s(bool,'F_0') = s(bool,X421)) )),
  inference(duplicate_literal_removal,[],[f5965])).
fof(f5965,plain,(
  ( ! [X420,X422,X421,X419] : (p(s(bool,X419)) | s(bool,'T_0') = s(bool,X420) | s(bool,'T_0') = s(bool,X420) | s(bool,X421) = s(bool,X422) | s(bool,X419) = s(bool,X422) | s(bool,'F_0') = s(bool,X421)) )),
  inference(superposition,[],[f1797,f2782])).
fof(f6938,plain,(
  ( ! [X416,X418,X415,X417] : (p(s(bool,X415)) | p(s(bool,X416)) | s(bool,X417) = s(bool,X418) | s(bool,X415) = s(bool,X418) | s(bool,'F_0') = s(bool,X417)) )),
  inference(subsumption_resolution,[],[f5964,f267])).
fof(f5964,plain,(
  ( ! [X416,X418,X415,X417] : (p(s(bool,X415)) | p(s(bool,X416)) | s(bool,'T_0') = s(bool,X416) | s(bool,X417) = s(bool,X418) | s(bool,X415) = s(bool,X418) | s(bool,'F_0') = s(bool,X417)) )),
  inference(superposition,[],[f1796,f2782])).
fof(f6937,plain,(
  ( ! [X412,X414,X413,X411] : (p(s(bool,X411)) | p(s(bool,X412)) | s(bool,X413) = s(bool,X414) | s(bool,X411) = s(bool,X414) | s(bool,'F_0') = s(bool,X413)) )),
  inference(subsumption_resolution,[],[f5963,f267])).
fof(f5963,plain,(
  ( ! [X412,X414,X413,X411] : (p(s(bool,X411)) | p(s(bool,X412)) | s(bool,'T_0') = s(bool,X412) | s(bool,X413) = s(bool,X414) | s(bool,X411) = s(bool,X414) | s(bool,'F_0') = s(bool,X413)) )),
  inference(superposition,[],[f1796,f2782])).
fof(f6933,plain,(
  ( ! [X372,X374,X376,X373,X375] : (p(s(bool,X372)) | s(bool,X373) != s(bool,X374) | s(bool,'F_0') = s(bool,X374) | s(bool,X375) = s(bool,X376) | s(bool,X372) = s(bool,X376) | s(bool,'F_0') = s(bool,X375)) ) | $spl30),
  inference(subsumption_resolution,[],[f5954,f826])).
fof(f5954,plain,(
  ( ! [X372,X374,X376,X373,X375] : (p(s(bool,X372)) | s(bool,'T_0') != s(bool,X373) | s(bool,X373) != s(bool,X374) | s(bool,'F_0') = s(bool,X374) | s(bool,X375) = s(bool,X376) | s(bool,X372) = s(bool,X376) | s(bool,'F_0') = s(bool,X375)) ) | $spl30),
  inference(superposition,[],[f755,f2782])).
fof(f6931,plain,(
  ( ! [X350,X352,X349,X351,X353] : (p(s(bool,X349)) | p(s(bool,X350)) | s(bool,'T_0') != s(bool,X351) | s(bool,X352) = s(bool,X353) | s(bool,X349) = s(bool,X353) | s(bool,'F_0') = s(bool,X352)) )),
  inference(subsumption_resolution,[],[f5949,f267])).
fof(f5949,plain,(
  ( ! [X350,X352,X349,X351,X353] : (p(s(bool,X349)) | p(s(bool,X350)) | ~p(s(bool,X351)) | s(bool,'T_0') != s(bool,X351) | s(bool,X352) = s(bool,X353) | s(bool,X349) = s(bool,X353) | s(bool,'F_0') = s(bool,X352)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6930,plain,(
  ( ! [X341,X343,X339,X340,X342] : (s(bool,X339) != s(bool,X340) | p(s(bool,X341)) | p(s(bool,X340)) | s(bool,X342) = s(bool,X343) | s(bool,X339) = s(bool,X343) | s(bool,'F_0') = s(bool,X342)) )),
  inference(subsumption_resolution,[],[f5947,f308])).
fof(f5947,plain,(
  ( ! [X341,X343,X339,X340,X342] : (s(bool,X339) != s(bool,X340) | p(s(bool,X341)) | ~p(s(bool,'T_0')) | p(s(bool,X340)) | s(bool,X342) = s(bool,X343) | s(bool,X339) = s(bool,X343) | s(bool,'F_0') = s(bool,X342)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6928,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (p(s(bool,X274)) | ~p(s(bool,X275)) | s(bool,'T_0') != s(bool,X276) | s(bool,X275) != s(bool,X277) | s(bool,X278) = s(bool,X279) | s(bool,X274) = s(bool,X279) | s(bool,'F_0') = s(bool,X278)) )),
  inference(subsumption_resolution,[],[f5929,f267])).
fof(f5929,plain,(
  ( ! [X277,X279,X275,X276,X278,X274] : (p(s(bool,X274)) | ~p(s(bool,X275)) | ~p(s(bool,X276)) | s(bool,'T_0') != s(bool,X276) | s(bool,X275) != s(bool,X277) | s(bool,X278) = s(bool,X279) | s(bool,X274) = s(bool,X279) | s(bool,'F_0') = s(bool,X278)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6927,plain,(
  ( ! [X269,X271,X273,X268,X270,X272] : (s(bool,X268) != s(bool,X269) | p(s(bool,X269)) | ~p(s(bool,X270)) | s(bool,X270) != s(bool,X271) | s(bool,X272) = s(bool,X273) | s(bool,X268) = s(bool,X273) | s(bool,'F_0') = s(bool,X272)) )),
  inference(subsumption_resolution,[],[f5928,f308])).
fof(f5928,plain,(
  ( ! [X269,X271,X273,X268,X270,X272] : (s(bool,X268) != s(bool,X269) | p(s(bool,X269)) | ~p(s(bool,X270)) | ~p(s(bool,'T_0')) | s(bool,X270) != s(bool,X271) | s(bool,X272) = s(bool,X273) | s(bool,X268) = s(bool,X273) | s(bool,'F_0') = s(bool,X272)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6926,plain,(
  ( ! [X231,X229,X227,X230,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | ~p(s(bool,X229)) | s(bool,'T_0') != s(bool,X228) | s(bool,X230) = s(bool,X231) | s(bool,X226) = s(bool,X231) | s(bool,'F_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f5921,f267])).
fof(f5921,plain,(
  ( ! [X231,X229,X227,X230,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | ~p(s(bool,X228)) | ~p(s(bool,X229)) | s(bool,'T_0') != s(bool,X228) | s(bool,X230) = s(bool,X231) | s(bool,X226) = s(bool,X231) | s(bool,'F_0') = s(bool,X230)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6925,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,X214) != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,X217)) | p(s(bool,X215)) | s(bool,X218) = s(bool,X219) | s(bool,X214) = s(bool,X219) | s(bool,'F_0') = s(bool,X218)) )),
  inference(subsumption_resolution,[],[f5919,f308])).
fof(f5919,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,X214) != s(bool,X215) | p(s(bool,X216)) | ~p(s(bool,'T_0')) | ~p(s(bool,X217)) | p(s(bool,X215)) | s(bool,X218) = s(bool,X219) | s(bool,X214) = s(bool,X219) | s(bool,'F_0') = s(bool,X218)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6924,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (p(s(bool,X196)) | p(s(bool,X197)) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') != s(bool,X198) | s(bool,X200) = s(bool,X201) | s(bool,X196) = s(bool,X201) | s(bool,'F_0') = s(bool,X200)) )),
  inference(subsumption_resolution,[],[f5916,f267])).
fof(f5916,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (p(s(bool,X196)) | p(s(bool,X197)) | ~p(s(bool,X198)) | s(bool,X197) != s(bool,X199) | s(bool,'T_0') != s(bool,X198) | s(bool,X200) = s(bool,X201) | s(bool,X196) = s(bool,X201) | s(bool,'F_0') = s(bool,X200)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6922,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (s(bool,X172) != s(bool,X173) | p(s(bool,X174)) | p(s(bool,X173)) | s(bool,X174) != s(bool,X175) | s(bool,X176) = s(bool,X177) | s(bool,X172) = s(bool,X177) | s(bool,'F_0') = s(bool,X176)) )),
  inference(subsumption_resolution,[],[f5912,f308])).
fof(f5912,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (s(bool,X172) != s(bool,X173) | p(s(bool,X174)) | ~p(s(bool,'T_0')) | p(s(bool,X173)) | s(bool,X174) != s(bool,X175) | s(bool,X176) = s(bool,X177) | s(bool,X172) = s(bool,X177) | s(bool,'F_0') = s(bool,X176)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6921,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (p(s(bool,X154)) | ~p(s(bool,X155)) | s(bool,'T_0') != s(bool,X156) | s(bool,X155) != s(bool,X157) | s(bool,X158) = s(bool,X159) | s(bool,X154) = s(bool,X159) | s(bool,'F_0') = s(bool,X158)) )),
  inference(subsumption_resolution,[],[f5909,f267])).
fof(f5909,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (p(s(bool,X154)) | ~p(s(bool,X155)) | ~p(s(bool,X156)) | s(bool,'T_0') != s(bool,X156) | s(bool,X155) != s(bool,X157) | s(bool,X158) = s(bool,X159) | s(bool,X154) = s(bool,X159) | s(bool,'F_0') = s(bool,X158)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6920,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,X142) != s(bool,X143) | p(s(bool,X143)) | ~p(s(bool,X144)) | s(bool,X144) != s(bool,X145) | s(bool,X146) = s(bool,X147) | s(bool,X142) = s(bool,X147) | s(bool,'F_0') = s(bool,X146)) )),
  inference(subsumption_resolution,[],[f5907,f308])).
fof(f5907,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,X142) != s(bool,X143) | p(s(bool,X143)) | ~p(s(bool,X144)) | ~p(s(bool,'T_0')) | s(bool,X144) != s(bool,X145) | s(bool,X146) = s(bool,X147) | s(bool,X142) = s(bool,X147) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6919,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | s(bool,X113) != s(bool,X115) | s(bool,'T_0') != s(bool,X114) | s(bool,X116) = s(bool,X117) | s(bool,X112) = s(bool,X117) | s(bool,'F_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f5902,f267])).
fof(f5902,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X113) != s(bool,X115) | s(bool,'T_0') != s(bool,X114) | s(bool,X116) = s(bool,X117) | s(bool,X112) = s(bool,X117) | s(bool,'F_0') = s(bool,X116)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6916,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (s(bool,X94) != s(bool,X95) | p(s(bool,X96)) | p(s(bool,X95)) | s(bool,X96) != s(bool,X97) | s(bool,X98) = s(bool,X99) | s(bool,X94) = s(bool,X99) | s(bool,'F_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f5899,f308])).
fof(f5899,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (s(bool,X94) != s(bool,X95) | p(s(bool,X96)) | ~p(s(bool,'T_0')) | p(s(bool,X95)) | s(bool,X96) != s(bool,X97) | s(bool,X98) = s(bool,X99) | s(bool,X94) = s(bool,X99) | s(bool,'F_0') = s(bool,X98)) )),
  inference(superposition,[],[f218,f2782])).
fof(f5885,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X26) = s(bool,X27) | s(bool,X24) = s(bool,X27) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f261,f2782])).
fof(f5884,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X22) = s(bool,X23) | s(bool,X20) = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f259,f2782])).
fof(f5883,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f246,f2782])).
fof(f5882,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X14) = s(bool,X15) | s(bool,X12) = s(bool,X15) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f243,f2782])).
fof(f5881,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X10) = s(bool,X11) | s(bool,X8) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f242,f2782])).
fof(f5880,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f239,f2782])).
fof(f5879,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f236,f2782])).
fof(f6915,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (p(s(bool,X634)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,X634) = s(bool,X635) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633)))))) )),
  inference(subsumption_resolution,[],[f5875,f2198])).
fof(f5875,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (p(s(bool,X634)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,X634) = s(bool,X635) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633))))) | s(bool,'F_0') = s(bool,X635)) )),
  inference(superposition,[],[f277,f2782])).
fof(f5874,plain,(
  ( ! [X629,X625,X627,X628,X624,X626] : (~p(s(bool,X628)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | s(bool,X628) = s(bool,X629) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s(bool,'F_0') = s(bool,X629)) )),
  inference(superposition,[],[f278,f2782])).
fof(f6913,plain,(
  ( ! [X621,X623,X619,X620,X622,X618] : (p(s(bool,X622)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,X622) = s(bool,X623) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621)))))) )),
  inference(subsumption_resolution,[],[f5873,f2198])).
fof(f5873,plain,(
  ( ! [X621,X623,X619,X620,X622,X618] : (p(s(bool,X622)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,X622) = s(bool,X623) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621))))) | s(bool,'F_0') = s(bool,X623)) )),
  inference(superposition,[],[f275,f2782])).
fof(f5872,plain,(
  ( ! [X617,X613,X615,X616,X612,X614] : (~p(s(bool,X616)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615)))))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s(bool,X616) = s(bool,X617) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s(bool,'F_0') = s(bool,X617)) )),
  inference(superposition,[],[f278,f2782])).
fof(f6911,plain,(
  ( ! [X606,X602,X605,X603,X604] : (s(bool,'T_0') != s(bool,X605) | s('type.ffi.ffi_state'(X602),'const.decSem.state_ffi_1'(s('type.decSem.state'(X602),X603))) = s('type.ffi.ffi_state'(X602),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X602),X604))) | s(bool,X605) = s(bool,X606) | s(bool,'F_0') = s(bool,X606)) )),
  inference(subsumption_resolution,[],[f5870,f2266])).
fof(f5870,plain,(
  ( ! [X606,X602,X605,X601,X603,X604] : (s(bool,'T_0') != s(bool,X605) | s('type.ffi.ffi_state'(X602),'const.decSem.state_ffi_1'(s('type.decSem.state'(X602),X603))) = s('type.ffi.ffi_state'(X602),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X602),X604))) | s(bool,X605) = s(bool,X606) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X601),s('type.decSem.state'(X602),X603),s('type.exhSem.state'(X602),X604))) | s(bool,'F_0') = s(bool,X606)) )),
  inference(superposition,[],[f2266,f2782])).
fof(f6910,plain,(
  ( ! [X600,X596,X598,X597,X599] : (s(bool,'F_0') = s(bool,X599) | s('type.ffi.ffi_state'(X596),'const.decSem.state_ffi_1'(s('type.decSem.state'(X596),X597))) = s('type.ffi.ffi_state'(X596),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X596),X598))) | s(bool,X599) = s(bool,X600) | s(bool,'F_0') = s(bool,X600)) )),
  inference(subsumption_resolution,[],[f5869,f2266])).
fof(f5869,plain,(
  ( ! [X600,X596,X598,X597,X599,X595] : (s(bool,'F_0') = s(bool,X599) | s('type.ffi.ffi_state'(X596),'const.decSem.state_ffi_1'(s('type.decSem.state'(X596),X597))) = s('type.ffi.ffi_state'(X596),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X596),X598))) | s(bool,X599) = s(bool,X600) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X595),s('type.decSem.state'(X596),X597),s('type.exhSem.state'(X596),X598))) | s(bool,'F_0') = s(bool,X600)) )),
  inference(superposition,[],[f2265,f2782])).
fof(f6909,plain,(
  ( ! [X592,X594,X590,X593,X591] : (s(bool,'T_0') != s(bool,X593) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X590),X591))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X590),X592))) | s(bool,X593) = s(bool,X594) | s(bool,'F_0') = s(bool,X594)) )),
  inference(subsumption_resolution,[],[f5868,f2062])).
fof(f5868,plain,(
  ( ! [X592,X594,X590,X593,X589,X591] : (s(bool,'T_0') != s(bool,X593) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X590),X591))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X590),X592))) | s(bool,X593) = s(bool,X594) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X589),s('type.decSem.state'(X590),X591),s('type.exhSem.state'(X590),X592))) | s(bool,'F_0') = s(bool,X594)) )),
  inference(superposition,[],[f2062,f2782])).
fof(f6908,plain,(
  ( ! [X588,X584,X586,X585,X587] : (s(bool,'F_0') = s(bool,X587) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X584),X585))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X584),X586))) | s(bool,X587) = s(bool,X588) | s(bool,'F_0') = s(bool,X588)) )),
  inference(subsumption_resolution,[],[f5867,f2062])).
fof(f5867,plain,(
  ( ! [X588,X584,X586,X585,X587,X583] : (s(bool,'F_0') = s(bool,X587) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X584),X585))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X584),X586))) | s(bool,X587) = s(bool,X588) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X583),s('type.decSem.state'(X584),X585),s('type.exhSem.state'(X584),X586))) | s(bool,'F_0') = s(bool,X588)) )),
  inference(superposition,[],[f2061,f2782])).
fof(f6907,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (p(s(bool,X581)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | s(bool,X581) = s(bool,X582) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580)))) )),
  inference(subsumption_resolution,[],[f5866,f2198])).
fof(f5866,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (p(s(bool,X581)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | s(bool,X581) = s(bool,X582) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580))) | s(bool,'F_0') = s(bool,X582)) )),
  inference(superposition,[],[f278,f2782])).
fof(f5865,plain,(
  ( ! [X573,X575,X571,X576,X572,X574] : (~p(s(bool,X575)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,X575) = s(bool,X576) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574))) | s(bool,'F_0') = s(bool,X576)) )),
  inference(superposition,[],[f277,f2782])).
fof(f6905,plain,(
  ( ! [X569,X567,X568,X570,X566] : (~p(s(bool,X569)) | s('type.ffi.ffi_state'(X566),'const.decSem.state_ffi_1'(s('type.decSem.state'(X566),X567))) = s('type.ffi.ffi_state'(X566),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X566),X568))) | s(bool,X569) = s(bool,X570) | s(bool,'F_0') = s(bool,X570)) )),
  inference(subsumption_resolution,[],[f5864,f2266])).
fof(f5864,plain,(
  ( ! [X569,X565,X567,X568,X570,X566] : (~p(s(bool,X569)) | s('type.ffi.ffi_state'(X566),'const.decSem.state_ffi_1'(s('type.decSem.state'(X566),X567))) = s('type.ffi.ffi_state'(X566),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X566),X568))) | s(bool,X569) = s(bool,X570) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X565),s('type.decSem.state'(X566),X567),s('type.exhSem.state'(X566),X568))) | s(bool,'F_0') = s(bool,X570)) )),
  inference(superposition,[],[f276,f2782])).
fof(f5863,plain,(
  ( ! [X561,X563,X559,X564,X560,X562] : (~p(s(bool,X563)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,X563) = s(bool,X564) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562))) | s(bool,'F_0') = s(bool,X564)) )),
  inference(superposition,[],[f275,f2782])).
fof(f6903,plain,(
  ( ! [X557,X555,X556,X558,X554] : (~p(s(bool,X557)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X554),X555))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X554),X556))) | s(bool,X557) = s(bool,X558) | s(bool,'F_0') = s(bool,X558)) )),
  inference(subsumption_resolution,[],[f5862,f2062])).
fof(f5862,plain,(
  ( ! [X557,X553,X555,X556,X558,X554] : (~p(s(bool,X557)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X554),X555))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X554),X556))) | s(bool,X557) = s(bool,X558) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X553),s('type.decSem.state'(X554),X555),s('type.exhSem.state'(X554),X556))) | s(bool,'F_0') = s(bool,X558)) )),
  inference(superposition,[],[f274,f2782])).
fof(f5861,plain,(
  ( ! [X551,X552] : (~p(s(bool,X551)) | s(bool,X551) = s(bool,X552) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X552)) )),
  inference(superposition,[],[f206,f2782])).
fof(f4911,plain,(
  ( ! [X17,X15,X18] : (s(bool,X15) != s(bool,X17) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') != s(bool,X18) | s(bool,X17) != s(bool,X18)) ) | $spl110),
  inference(cnf_transformation,[],[f4911_D])).
fof(f4911_D,plain,(
  ( ! [X17,X15,X18] : (s(bool,X15) != s(bool,X17) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') != s(bool,X18) | s(bool,X17) != s(bool,X18)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f6888,plain,(
  ( ! [X436,X438,X437,X439,X435] : (p(s(bool,X436)) | s(bool,X435) != s(bool,X437) | s(bool,X437) != s(bool,X438) | s(bool,'T_0') != s(bool,X438) | s(bool,X436) = s(bool,X439) | s(bool,'T_0') = s(bool,X435)) )),
  inference(subsumption_resolution,[],[f5828,f2198])).
fof(f5828,plain,(
  ( ! [X436,X438,X437,X439,X435] : (p(s(bool,X436)) | s(bool,X435) != s(bool,X437) | s(bool,X437) != s(bool,X438) | s(bool,'T_0') != s(bool,X438) | s(bool,X436) = s(bool,X439) | s(bool,'T_0') = s(bool,X435) | s(bool,'F_0') = s(bool,X439)) )),
  inference(superposition,[],[f3958,f2782])).
fof(f6789,plain,(
  ( ! [X408,X405,X407,X409] : (s(bool,X405) != s(bool,X407) | s(bool,'F_0') = s(bool,X409) | s(bool,X405) = s(bool,X409) | s(bool,'T_0') = s(bool,X408) | s(bool,X407) = s(bool,X408)) ) | $spl168),
  inference(cnf_transformation,[],[f6789_D])).
fof(f6789_D,plain,(
  ( ! [X408,X405,X407,X409] : (s(bool,X405) != s(bool,X407) | s(bool,'F_0') = s(bool,X409) | s(bool,X405) = s(bool,X409) | s(bool,'T_0') = s(bool,X408) | s(bool,X407) = s(bool,X408)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6436,plain,(
  ( ! [X430,X432,X434,X431,X433] : (s(bool,X431) != s(bool,X432) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) = s(bool,X433) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X434) | s(bool,'F_0') = s(bool,X434)) )),
  inference(duplicate_literal_removal,[],[f5827])).
fof(f5827,plain,(
  ( ! [X430,X432,X434,X431,X433] : (s(bool,X431) != s(bool,X432) | s(bool,'T_0') = s(bool,X433) | s(bool,X432) = s(bool,X433) | s(bool,'T_0') = s(bool,X430) | s(bool,X431) = s(bool,X434) | s(bool,'T_0') = s(bool,X430) | s(bool,'F_0') = s(bool,X434)) )),
  inference(superposition,[],[f3625,f2782])).
fof(f6886,plain,(
  ( ! [X416,X418,X415,X417,X419] : (~p(s(bool,X416)) | s(bool,'T_0') = s(bool,X418) | s(bool,X415) = s(bool,X417) | s(bool,X416) = s(bool,X419) | s(bool,'T_0') = s(bool,X415) | s(bool,'F_0') = s(bool,X419)) )),
  inference(subsumption_resolution,[],[f5824,f3625])).
fof(f5824,plain,(
  ( ! [X416,X418,X415,X417,X419] : (~p(s(bool,X416)) | s(bool,X417) = s(bool,X418) | s(bool,'T_0') = s(bool,X418) | s(bool,X415) = s(bool,X417) | s(bool,X416) = s(bool,X419) | s(bool,'T_0') = s(bool,X415) | s(bool,'F_0') = s(bool,X419)) )),
  inference(superposition,[],[f2783,f2782])).
fof(f2783,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X11) = s(bool,X12)) )),
  inference(superposition,[],[f2328,f270])).
fof(f1365,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl50),
  inference(cnf_transformation,[],[f1365_D])).
fof(f1365_D,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f6885,plain,(
  ( ! [X412,X414,X410,X413,X411] : (~p(s(bool,X411)) | p(s(bool,X413)) | s(bool,X410) != s(bool,X412) | s(bool,X411) = s(bool,X414) | s(bool,'T_0') = s(bool,X410) | s(bool,'F_0') = s(bool,X414)) )),
  inference(subsumption_resolution,[],[f5823,f3624])).
fof(f5823,plain,(
  ( ! [X412,X414,X410,X413,X411] : (~p(s(bool,X411)) | s(bool,X412) != s(bool,X413) | p(s(bool,X413)) | s(bool,X410) != s(bool,X412) | s(bool,X411) = s(bool,X414) | s(bool,'T_0') = s(bool,X410) | s(bool,'F_0') = s(bool,X414)) )),
  inference(superposition,[],[f2388,f2782])).
fof(f6880,plain,(
  ( ! [X383,X382,X384,X381] : (s(bool,X381) != s(bool,X382) | s(bool,'F_0') = s(bool,X384) | s(bool,X381) = s(bool,X384) | s(bool,X382) = s(bool,X383) | p(s(bool,X383))) ) | $spl182),
  inference(cnf_transformation,[],[f6880_D])).
fof(f6880_D,plain,(
  ( ! [X383,X382,X384,X381] : (s(bool,X381) != s(bool,X382) | s(bool,'F_0') = s(bool,X384) | s(bool,X381) = s(bool,X384) | s(bool,X382) = s(bool,X383) | p(s(bool,X383))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6879,plain,(
  ( ! [X383,X380,X382,X384,X381] : (s(bool,X381) != s(bool,X382) | p(s(bool,X383)) | p(s(bool,X380)) | s(bool,X382) = s(bool,X383) | s(bool,X381) = s(bool,X384) | s(bool,'F_0') = s(bool,X384)) ) | $spl66),
  inference(subsumption_resolution,[],[f5816,f267])).
fof(f5816,plain,(
  ( ! [X383,X380,X382,X384,X381] : (s(bool,X381) != s(bool,X382) | p(s(bool,X383)) | p(s(bool,X380)) | s(bool,X382) = s(bool,X383) | s(bool,X381) = s(bool,X384) | s(bool,'T_0') = s(bool,X380) | s(bool,'F_0') = s(bool,X384)) ) | $spl66),
  inference(superposition,[],[f2249,f2782])).
fof(f6876,plain,(
  ( ! [X368,X370,X367,X369] : (s(bool,'T_0') != s(bool,X368) | s(bool,X367) != s(bool,X369) | p(s(bool,X369)) | s(bool,X368) = s(bool,X370) | s(bool,'F_0') = s(bool,X370)) )),
  inference(subsumption_resolution,[],[f5813,f2028])).
fof(f5813,plain,(
  ( ! [X368,X370,X367,X369] : (s(bool,'T_0') != s(bool,X368) | s(bool,X367) != s(bool,X369) | p(s(bool,X369)) | s(bool,X368) = s(bool,X370) | s(bool,'T_0') = s(bool,X367) | s(bool,'F_0') = s(bool,X370)) )),
  inference(superposition,[],[f2028,f2782])).
fof(f6871,plain,(
  ( ! [X344,X346,X343,X345,X347] : (~p(s(bool,X344)) | s(bool,X343) != s(bool,X345) | s(bool,'T_0') = s(bool,X346) | s(bool,X344) = s(bool,X347) | s(bool,'T_0') = s(bool,X343) | s(bool,'F_0') = s(bool,X347)) ) | $spl52),
  inference(subsumption_resolution,[],[f5807,f3625])).
fof(f5807,plain,(
  ( ! [X344,X346,X343,X345,X347] : (~p(s(bool,X344)) | s(bool,X343) != s(bool,X345) | s(bool,'T_0') = s(bool,X346) | s(bool,X345) != s(bool,X346) | s(bool,X344) = s(bool,X347) | s(bool,'T_0') = s(bool,X343) | s(bool,'F_0') = s(bool,X347)) ) | $spl52),
  inference(superposition,[],[f1367,f2782])).
fof(f6870,plain,(
  ( ! [X335,X337,X336,X338] : (s(bool,'T_0') != s(bool,X336) | s(bool,'T_0') = s(bool,X337) | s(bool,X335) != s(bool,X337) | s(bool,X336) = s(bool,X338) | s(bool,'F_0') = s(bool,X338)) ) | $spl44),
  inference(subsumption_resolution,[],[f5805,f1278])).
fof(f5805,plain,(
  ( ! [X335,X337,X336,X338] : (s(bool,'T_0') != s(bool,X336) | s(bool,'T_0') = s(bool,X337) | s(bool,X335) != s(bool,X337) | s(bool,X336) = s(bool,X338) | s(bool,'T_0') = s(bool,X335) | s(bool,'F_0') = s(bool,X338)) ) | $spl44),
  inference(superposition,[],[f1278,f2782])).
fof(f6869,plain,(
  ( ! [X325,X327,X329,X326,X328] : (s(bool,X326) != s(bool,X327) | s(bool,'F_0') = s(bool,X327) | s(bool,'T_0') = s(bool,X328) | s(bool,X326) = s(bool,X329) | s(bool,'T_0') = s(bool,X325) | s(bool,'F_0') = s(bool,X329)) )),
  inference(subsumption_resolution,[],[f5803,f2328])).
fof(f5803,plain,(
  ( ! [X325,X327,X329,X326,X328] : (s(bool,X326) != s(bool,X327) | s(bool,'F_0') = s(bool,X327) | s(bool,X325) != s(bool,X328) | s(bool,'T_0') = s(bool,X328) | s(bool,X326) = s(bool,X329) | s(bool,'T_0') = s(bool,X325) | s(bool,'F_0') = s(bool,X329)) )),
  inference(superposition,[],[f906,f2782])).
fof(f6439,plain,(
  ( ! [X321,X323,X324,X322] : (s(bool,X322) != s(bool,X323) | s(bool,'F_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X321) | s(bool,X322) = s(bool,X324) | s(bool,'F_0') = s(bool,X324)) )),
  inference(duplicate_literal_removal,[],[f5802])).
fof(f5802,plain,(
  ( ! [X321,X323,X324,X322] : (s(bool,X322) != s(bool,X323) | s(bool,'F_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X321) | s(bool,X322) = s(bool,X324) | s(bool,'T_0') = s(bool,X321) | s(bool,'F_0') = s(bool,X324)) )),
  inference(superposition,[],[f826,f2782])).
fof(f6440,plain,(
  ( ! [X308,X306,X305,X307] : (s(bool,X306) != s(bool,X307) | ~p(s(bool,X307)) | s(bool,'T_0') = s(bool,X305) | s(bool,X306) = s(bool,X308) | s(bool,'F_0') = s(bool,X308)) )),
  inference(duplicate_literal_removal,[],[f5798])).
fof(f5798,plain,(
  ( ! [X308,X306,X305,X307] : (s(bool,X306) != s(bool,X307) | ~p(s(bool,X307)) | s(bool,'T_0') = s(bool,X305) | s(bool,X306) = s(bool,X308) | s(bool,'T_0') = s(bool,X305) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f806,f2782])).
fof(f6866,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,X297) != s(bool,X298) | p(s(bool,X296)) | ~p(s(bool,X298)) | s(bool,X297) = s(bool,X299) | s(bool,'F_0') = s(bool,X299)) )),
  inference(subsumption_resolution,[],[f5796,f267])).
fof(f5796,plain,(
  ( ! [X296,X298,X297,X299] : (s(bool,X297) != s(bool,X298) | p(s(bool,X296)) | ~p(s(bool,X298)) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,X296) | s(bool,'F_0') = s(bool,X299)) )),
  inference(superposition,[],[f368,f2782])).
fof(f6864,plain,(
  ( ! [X288,X290,X287,X289,X291] : (~p(s(bool,X288)) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X287) != s(bool,X290) | s(bool,X288) = s(bool,X291) | s(bool,'F_0') = s(bool,X291)) )),
  inference(subsumption_resolution,[],[f5794,f2028])).
fof(f5794,plain,(
  ( ! [X288,X290,X287,X289,X291] : (~p(s(bool,X288)) | p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X287) != s(bool,X290) | s(bool,X288) = s(bool,X291) | s(bool,'T_0') = s(bool,X287) | s(bool,'F_0') = s(bool,X291)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6861,plain,(
  ( ! [X273,X275,X276,X272,X274] : (s(bool,X273) != s(bool,X274) | p(s(bool,X275)) | ~p(s(bool,X274)) | p(s(bool,X272)) | s(bool,X273) = s(bool,X276) | s(bool,'F_0') = s(bool,X276)) )),
  inference(subsumption_resolution,[],[f5791,f267])).
fof(f5791,plain,(
  ( ! [X273,X275,X276,X272,X274] : (s(bool,X273) != s(bool,X274) | p(s(bool,X275)) | ~p(s(bool,X274)) | p(s(bool,X272)) | s(bool,X273) = s(bool,X276) | s(bool,'T_0') = s(bool,X272) | s(bool,'F_0') = s(bool,X276)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6441,plain,(
  ( ! [X265,X264,X266] : (s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264) | s(bool,X265) = s(bool,X266) | s(bool,'F_0') = s(bool,X266)) )),
  inference(duplicate_literal_removal,[],[f5789])).
fof(f5789,plain,(
  ( ! [X265,X264,X266] : (s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264) | s(bool,X265) = s(bool,X266) | s(bool,'T_0') = s(bool,X264) | s(bool,'F_0') = s(bool,X266)) )),
  inference(superposition,[],[f306,f2782])).
fof(f6442,plain,(
  ( ! [X259,X260,X258] : (~p(s(bool,X259)) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'F_0') = s(bool,X260)) )),
  inference(duplicate_literal_removal,[],[f5787])).
fof(f5787,plain,(
  ( ! [X259,X260,X258] : (~p(s(bool,X259)) | s(bool,'T_0') = s(bool,X258) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X258) | s(bool,'F_0') = s(bool,X260)) )),
  inference(superposition,[],[f270,f2782])).
fof(f6857,plain,(
  ( ! [X251,X249,X250] : (s(bool,'T_0') != s(bool,X250) | p(s(bool,X249)) | s(bool,X250) = s(bool,X251) | s(bool,'F_0') = s(bool,X251)) )),
  inference(subsumption_resolution,[],[f5784,f267])).
fof(f5784,plain,(
  ( ! [X251,X249,X250] : (s(bool,'T_0') != s(bool,X250) | p(s(bool,X249)) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X249) | s(bool,'F_0') = s(bool,X251)) )),
  inference(superposition,[],[f269,f2782])).
fof(f6443,plain,(
  ( ! [X246,X248,X247] : (~p(s(bool,X247)) | s(bool,'T_0') = s(bool,X246) | s(bool,X247) = s(bool,X248) | s(bool,'F_0') = s(bool,X248)) )),
  inference(duplicate_literal_removal,[],[f5783])).
fof(f5783,plain,(
  ( ! [X246,X248,X247] : (~p(s(bool,X247)) | s(bool,'T_0') = s(bool,X246) | s(bool,X247) = s(bool,X248) | s(bool,'T_0') = s(bool,X246) | s(bool,'F_0') = s(bool,X248)) )),
  inference(superposition,[],[f268,f2782])).
fof(f6855,plain,(
  ( ! [X239,X237,X238] : (s(bool,'T_0') != s(bool,X238) | p(s(bool,X237)) | s(bool,X238) = s(bool,X239) | s(bool,'F_0') = s(bool,X239)) )),
  inference(subsumption_resolution,[],[f5780,f267])).
fof(f5780,plain,(
  ( ! [X239,X237,X238] : (s(bool,'T_0') != s(bool,X238) | p(s(bool,X237)) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X237) | s(bool,'F_0') = s(bool,X239)) )),
  inference(superposition,[],[f267,f2782])).
fof(f6853,plain,(
  ( ! [X231,X233,X232] : (s(bool,'F_0') = s(bool,X232) | p(s(bool,X231)) | s(bool,X232) = s(bool,X233) | s(bool,'F_0') = s(bool,X233)) )),
  inference(subsumption_resolution,[],[f5778,f267])).
fof(f5778,plain,(
  ( ! [X231,X233,X232] : (s(bool,'F_0') = s(bool,X232) | p(s(bool,X231)) | s(bool,X232) = s(bool,X233) | s(bool,'T_0') = s(bool,X231) | s(bool,'F_0') = s(bool,X233)) )),
  inference(superposition,[],[f264,f2782])).
fof(f6852,plain,(
  ( ! [X229,X230,X228] : (~p(s(bool,X229)) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f5777,f306])).
fof(f5777,plain,(
  ( ! [X229,X230,X228] : (~p(s(bool,X229)) | s(bool,'F_0') != s(bool,X228) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230)) )),
  inference(superposition,[],[f263,f2782])).
fof(f6850,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X198) != s(bool,X200) | s(bool,X201) != s(bool,X202) | s(bool,X199) = s(bool,X203) | s(bool,'F_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f5768,f2028])).
fof(f5768,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X198) != s(bool,X200) | s(bool,X201) != s(bool,X202) | s(bool,X199) = s(bool,X203) | s(bool,'T_0') = s(bool,X198) | s(bool,'F_0') = s(bool,X203)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6847,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,X175) != s(bool,X176) | p(s(bool,X174)) | ~p(s(bool,X177)) | ~p(s(bool,X176)) | s(bool,X177) != s(bool,X178) | s(bool,X175) = s(bool,X179) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f5764,f267])).
fof(f5764,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(bool,X175) != s(bool,X176) | p(s(bool,X174)) | ~p(s(bool,X177)) | ~p(s(bool,X176)) | s(bool,X177) != s(bool,X178) | s(bool,X175) = s(bool,X179) | s(bool,'T_0') = s(bool,X174) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6844,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,X151)) | p(s(bool,X152)) | ~p(s(bool,X153)) | p(s(bool,X154)) | s(bool,X150) != s(bool,X154) | s(bool,X151) = s(bool,X155) | s(bool,'F_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f5760,f2028])).
fof(f5760,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (~p(s(bool,X151)) | p(s(bool,X152)) | ~p(s(bool,X153)) | p(s(bool,X154)) | s(bool,X150) != s(bool,X154) | s(bool,X151) = s(bool,X155) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X155)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6841,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,X133) != s(bool,X134) | p(s(bool,X135)) | ~p(s(bool,X134)) | ~p(s(bool,X136)) | p(s(bool,X132)) | s(bool,X133) = s(bool,X137) | s(bool,'F_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f5757,f267])).
fof(f5757,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,X133) != s(bool,X134) | p(s(bool,X135)) | ~p(s(bool,X134)) | ~p(s(bool,X136)) | p(s(bool,X132)) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X132) | s(bool,'F_0') = s(bool,X137)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6839,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X122) != s(bool,X124) | s(bool,X120) != s(bool,X123) | s(bool,X121) = s(bool,X125) | s(bool,'F_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f5755,f2028])).
fof(f5755,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (~p(s(bool,X121)) | p(s(bool,X122)) | p(s(bool,X123)) | s(bool,X122) != s(bool,X124) | s(bool,X120) != s(bool,X123) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X125)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6834,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f5750,f267])).
fof(f5750,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X92)) | p(s(bool,X90)) | s(bool,X93) != s(bool,X94) | s(bool,X91) = s(bool,X95) | s(bool,'T_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6832,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X78) != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83) | s(bool,'F_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f5748,f2028])).
fof(f5748,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X78) != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X83) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X83)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6830,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X60)) | ~p(s(bool,X63)) | ~p(s(bool,X62)) | s(bool,X63) != s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f5745,f267])).
fof(f5745,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X60)) | ~p(s(bool,X63)) | ~p(s(bool,X62)) | s(bool,X63) != s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6826,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,X37)) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X36) != s(bool,X39) | s(bool,X37) = s(bool,X41) | s(bool,'F_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5741,f2028])).
fof(f5741,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,X37)) | p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X36) != s(bool,X39) | s(bool,X37) = s(bool,X41) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6820,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,X1) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f5735,f267])).
fof(f5735,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X3) != s(bool,X4) | s(bool,X1) = s(bool,X5) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6819,plain,(
  ( ! [X633,X635,X631,X632,X630] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,'F_0') = s(bool,X635) | s(bool,X635) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633)))))) ) | $spl180),
  inference(cnf_transformation,[],[f6819_D])).
fof(f6819_D,plain,(
  ( ! [X633,X635,X631,X632,X630] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,'F_0') = s(bool,X635) | s(bool,X635) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6818,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (p(s(bool,X634)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,X635) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633))))) | s(bool,'F_0') = s(bool,X635)) )),
  inference(subsumption_resolution,[],[f5731,f267])).
fof(f5731,plain,(
  ( ! [X633,X635,X631,X632,X634,X630] : (p(s(bool,X634)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630),s('type.decSem.state'(X631),X632),s('type.exhSem.state'(X631),X633)))) | s(bool,X635) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X630))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X631),X632))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X631),X633))))) | s(bool,'T_0') = s(bool,X634) | s(bool,'F_0') = s(bool,X635)) )),
  inference(superposition,[],[f277,f2782])).
fof(f6817,plain,(
  ( ! [X621,X623,X619,X620,X618] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,'F_0') = s(bool,X623) | s(bool,X623) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621)))))) ) | $spl178),
  inference(cnf_transformation,[],[f6817_D])).
fof(f6817_D,plain,(
  ( ! [X621,X623,X619,X620,X618] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,'F_0') = s(bool,X623) | s(bool,X623) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621)))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6816,plain,(
  ( ! [X621,X623,X619,X620,X622,X618] : (p(s(bool,X622)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,X623) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621))))) | s(bool,'F_0') = s(bool,X623)) )),
  inference(subsumption_resolution,[],[f5729,f267])).
fof(f5729,plain,(
  ( ! [X621,X623,X619,X620,X622,X618] : (p(s(bool,X622)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618),s('type.decSem.state'(X619),X620),s('type.exhSem.state'(X619),X621)))) | s(bool,X623) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X618))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X619),X620))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X619),X621))))) | s(bool,'T_0') = s(bool,X622) | s(bool,'F_0') = s(bool,X623)) )),
  inference(superposition,[],[f275,f2782])).
fof(f6815,plain,(
  ( ! [X580,X582,X578,X577,X579] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s(bool,'F_0') = s(bool,X582) | s(bool,X582) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580)))) ) | $spl176),
  inference(cnf_transformation,[],[f6815_D])).
fof(f6815_D,plain,(
  ( ! [X580,X582,X578,X577,X579] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s(bool,'F_0') = s(bool,X582) | s(bool,X582) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6814,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (p(s(bool,X581)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | s(bool,X582) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580))) | s(bool,'F_0') = s(bool,X582)) )),
  inference(subsumption_resolution,[],[f5722,f267])).
fof(f5722,plain,(
  ( ! [X580,X582,X578,X581,X577,X579] : (p(s(bool,X581)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.ffi.ffi_state'(X578),'const.decSem.state_ffi_1'(s('type.decSem.state'(X578),X579))) != s('type.ffi.ffi_state'(X578),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X578),X580))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X578),X579))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X578),X580)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X578),X579))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X578),X580))) | s(bool,X582) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X577),s('type.decSem.state'(X578),X579),s('type.exhSem.state'(X578),X580))) | s(bool,'T_0') = s(bool,X581) | s(bool,'F_0') = s(bool,X582)) )),
  inference(superposition,[],[f278,f2782])).
fof(f6556,plain,(
  ( ! [X95] : (s(bool,X95) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'F_0') = s(bool,X95)) ) | $spl130),
  inference(cnf_transformation,[],[f6556_D])).
fof(f6556_D,plain,(
  ( ! [X95] : (s(bool,X95) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'F_0') = s(bool,X95)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6444,plain,(
  ( ! [X549,X550] : (s(bool,'T_0') = s(bool,X549) | s(bool,X550) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'F_0') = s(bool,X550)) )),
  inference(duplicate_literal_removal,[],[f5716])).
fof(f5716,plain,(
  ( ! [X549,X550] : (s(bool,'T_0') = s(bool,X549) | s(bool,X550) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X549) | s(bool,'F_0') = s(bool,X550)) )),
  inference(superposition,[],[f329,f2782])).
fof(f6445,plain,(
  ( ! [X546,X545,X547,X548] : (s(bool,X545) != s(bool,X546) | s(bool,X546) != s(bool,X547) | s(bool,'T_0') != s(bool,X547) | s(bool,'F_0') = s(bool,X548) | s(bool,'T_0') = s(bool,X545)) )),
  inference(duplicate_literal_removal,[],[f5715])).
fof(f5715,plain,(
  ( ! [X546,X545,X547,X548] : (s(bool,X545) != s(bool,X546) | s(bool,X546) != s(bool,X547) | s(bool,'T_0') != s(bool,X547) | s(bool,'F_0') = s(bool,X548) | s(bool,'T_0') = s(bool,X545) | s(bool,'F_0') = s(bool,X548)) )),
  inference(superposition,[],[f4878,f2782])).
fof(f6796,plain,(
  ( ! [X443,X440,X442] : (s(bool,X440) != s(bool,X442) | s(bool,'F_0') = s(bool,X443) | s(bool,X440) = s(bool,X443) | s(bool,'T_0') != s(bool,X442)) ) | $spl174),
  inference(cnf_transformation,[],[f6796_D])).
fof(f6796_D,plain,(
  ( ! [X443,X440,X442] : (s(bool,X440) != s(bool,X442) | s(bool,'F_0') = s(bool,X443) | s(bool,X440) = s(bool,X443) | s(bool,'T_0') != s(bool,X442)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6795,plain,(
  ( ! [X441,X443,X440,X442] : (s(bool,'F_0') != s(bool,X441) | s(bool,X440) != s(bool,X442) | s(bool,'T_0') != s(bool,X442) | s(bool,X440) = s(bool,X443) | s(bool,'F_0') = s(bool,X443)) )),
  inference(subsumption_resolution,[],[f5685,f1729])).
fof(f5685,plain,(
  ( ! [X441,X443,X440,X442] : (s(bool,'F_0') != s(bool,X441) | s(bool,X440) != s(bool,X442) | s(bool,'T_0') != s(bool,X442) | s(bool,X440) = s(bool,X443) | s(bool,'T_0') = s(bool,X441) | s(bool,'F_0') = s(bool,X443)) )),
  inference(superposition,[],[f4878,f2782])).
fof(f6794,plain,(
  ( ! [X438,X437,X439,X435] : (s(bool,X435) != s(bool,X437) | s(bool,'F_0') = s(bool,X439) | s(bool,X435) = s(bool,X439) | s(bool,'T_0') != s(bool,X438) | s(bool,X437) != s(bool,X438)) ) | $spl172),
  inference(cnf_transformation,[],[f6794_D])).
fof(f6794_D,plain,(
  ( ! [X438,X437,X439,X435] : (s(bool,X435) != s(bool,X437) | s(bool,'F_0') = s(bool,X439) | s(bool,X435) = s(bool,X439) | s(bool,'T_0') != s(bool,X438) | s(bool,X437) != s(bool,X438)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f6793,plain,(
  ( ! [X436,X438,X437,X439,X435] : (p(s(bool,X436)) | s(bool,X435) != s(bool,X437) | s(bool,X437) != s(bool,X438) | s(bool,'T_0') != s(bool,X438) | s(bool,X435) = s(bool,X439) | s(bool,'F_0') = s(bool,X439)) )),
  inference(subsumption_resolution,[],[f5684,f267])).
fof(f5684,plain,(
  ( ! [X436,X438,X437,X439,X435] : (p(s(bool,X436)) | s(bool,X435) != s(bool,X437) | s(bool,X437) != s(bool,X438) | s(bool,'T_0') != s(bool,X438) | s(bool,X435) = s(bool,X439) | s(bool,'T_0') = s(bool,X436) | s(bool,'F_0') = s(bool,X439)) )),
  inference(superposition,[],[f3958,f2782])).
fof(f6790,plain,(
  ( ! [X420,X422,X424,X421,X423] : (p(s(bool,X421)) | s(bool,X420) = s(bool,X422) | s(bool,X422) != s(bool,X423) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X424) | s(bool,'F_0') = s(bool,X424)) )),
  inference(subsumption_resolution,[],[f5681,f267])).
fof(f5681,plain,(
  ( ! [X420,X422,X424,X421,X423] : (p(s(bool,X421)) | s(bool,X420) = s(bool,X422) | s(bool,X422) != s(bool,X423) | s(bool,'T_0') = s(bool,X423) | s(bool,X420) = s(bool,X424) | s(bool,'T_0') = s(bool,X421) | s(bool,'F_0') = s(bool,X424)) )),
  inference(superposition,[],[f3624,f2782])).
fof(f6788,plain,(
  ( ! [X408,X405,X407,X409,X406] : (p(s(bool,X406)) | s(bool,X405) != s(bool,X407) | s(bool,X407) = s(bool,X408) | s(bool,'T_0') = s(bool,X408) | s(bool,X405) = s(bool,X409) | s(bool,'F_0') = s(bool,X409)) ) | $spl70),
  inference(subsumption_resolution,[],[f5678,f267])).
fof(f5678,plain,(
  ( ! [X408,X405,X407,X409,X406] : (p(s(bool,X406)) | s(bool,X405) != s(bool,X407) | s(bool,X407) = s(bool,X408) | s(bool,'T_0') = s(bool,X408) | s(bool,X405) = s(bool,X409) | s(bool,'T_0') = s(bool,X406) | s(bool,'F_0') = s(bool,X409)) ) | $spl70),
  inference(superposition,[],[f2350,f2782])).
fof(f6784,plain,(
  ( ! [X377,X379,X376,X378,X375] : (s(bool,X376) != s(bool,X377) | p(s(bool,X378)) | p(s(bool,X377)) | s(bool,X375) = s(bool,X378) | s(bool,X375) = s(bool,X379) | s(bool,'F_0') = s(bool,X379)) ) | $spl66),
  inference(subsumption_resolution,[],[f5671,f2028])).
fof(f5671,plain,(
  ( ! [X377,X379,X376,X378,X375] : (s(bool,X376) != s(bool,X377) | p(s(bool,X378)) | p(s(bool,X377)) | s(bool,X375) = s(bool,X378) | s(bool,X375) = s(bool,X379) | s(bool,'T_0') = s(bool,X376) | s(bool,'F_0') = s(bool,X379)) ) | $spl66),
  inference(superposition,[],[f2249,f2782])).
fof(f6783,plain,(
  ( ! [X372,X374,X373,X371] : (s(bool,X371) = s(bool,X373) | p(s(bool,X373)) | s(bool,X371) = s(bool,X374) | s(bool,'T_0') = s(bool,X372) | s(bool,'F_0') = s(bool,X374)) )),
  inference(subsumption_resolution,[],[f5670,f306])).
fof(f5670,plain,(
  ( ! [X372,X374,X373,X371] : (s(bool,'F_0') != s(bool,X372) | s(bool,X371) = s(bool,X373) | p(s(bool,X373)) | s(bool,X371) = s(bool,X374) | s(bool,'T_0') = s(bool,X372) | s(bool,'F_0') = s(bool,X374)) )),
  inference(superposition,[],[f2198,f2782])).
fof(f6782,plain,(
  ( ! [X364,X366,X365,X363] : (p(s(bool,X364)) | s(bool,X363) != s(bool,X365) | s(bool,'F_0') = s(bool,X365) | s(bool,X363) = s(bool,X366) | s(bool,'F_0') = s(bool,X366)) ) | $spl62),
  inference(subsumption_resolution,[],[f5668,f267])).
fof(f5668,plain,(
  ( ! [X364,X366,X365,X363] : (p(s(bool,X364)) | s(bool,X363) != s(bool,X365) | s(bool,'F_0') = s(bool,X365) | s(bool,X363) = s(bool,X366) | s(bool,'T_0') = s(bool,X364) | s(bool,'F_0') = s(bool,X366)) ) | $spl62),
  inference(superposition,[],[f1999,f2782])).
fof(f6780,plain,(
  ( ! [X356,X358,X357,X355] : (p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) = s(bool,X357) | s(bool,X355) = s(bool,X358) | s(bool,'F_0') = s(bool,X358)) )),
  inference(subsumption_resolution,[],[f5666,f267])).
fof(f5666,plain,(
  ( ! [X356,X358,X357,X355] : (p(s(bool,X356)) | p(s(bool,X357)) | s(bool,X355) = s(bool,X357) | s(bool,X355) = s(bool,X358) | s(bool,'T_0') = s(bool,X356) | s(bool,'F_0') = s(bool,X358)) )),
  inference(superposition,[],[f1796,f2782])).
fof(f6778,plain,(
  ( ! [X352,X354,X351,X353] : (p(s(bool,X352)) | p(s(bool,X353)) | s(bool,X351) = s(bool,X353) | s(bool,X351) = s(bool,X354) | s(bool,'F_0') = s(bool,X354)) )),
  inference(subsumption_resolution,[],[f5665,f267])).
fof(f5665,plain,(
  ( ! [X352,X354,X351,X353] : (p(s(bool,X352)) | p(s(bool,X353)) | s(bool,X351) = s(bool,X353) | s(bool,X351) = s(bool,X354) | s(bool,'T_0') = s(bool,X352) | s(bool,'F_0') = s(bool,X354)) )),
  inference(superposition,[],[f1796,f2782])).
fof(f6777,plain,(
  ( ! [X350,X349,X348] : (s(bool,'F_0') != s(bool,X349) | s(bool,'T_0') != s(bool,X348) | s(bool,X348) = s(bool,X350) | s(bool,'F_0') = s(bool,X350)) )),
  inference(subsumption_resolution,[],[f5664,f1729])).
fof(f5664,plain,(
  ( ! [X350,X349,X348] : (s(bool,'F_0') != s(bool,X349) | s(bool,'T_0') != s(bool,X348) | s(bool,X348) = s(bool,X350) | s(bool,'T_0') = s(bool,X349) | s(bool,'F_0') = s(bool,X350)) )),
  inference(superposition,[],[f1729,f2782])).
fof(f6776,plain,(
  ( ! [X341,X339,X340,X342] : (~p(s(bool,X341)) | s(bool,X339) != s(bool,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X340) | s(bool,'F_0') = s(bool,X342)) ) | $spl48),
  inference(subsumption_resolution,[],[f5662,f306])).
fof(f5662,plain,(
  ( ! [X341,X339,X340,X342] : (s(bool,'F_0') != s(bool,X340) | ~p(s(bool,X341)) | s(bool,X339) != s(bool,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X340) | s(bool,'F_0') = s(bool,X342)) ) | $spl48),
  inference(superposition,[],[f1356,f2782])).
fof(f6775,plain,(
  ( ! [X333,X331,X332,X334,X330] : (s(bool,X331) != s(bool,X332) | s(bool,'F_0') = s(bool,X330) | s(bool,'T_0') = s(bool,X333) | s(bool,X330) = s(bool,X334) | s(bool,'T_0') = s(bool,X331) | s(bool,'F_0') = s(bool,X334)) )),
  inference(subsumption_resolution,[],[f5660,f3625])).
fof(f5660,plain,(
  ( ! [X333,X331,X332,X334,X330] : (s(bool,X331) != s(bool,X332) | s(bool,'F_0') = s(bool,X330) | s(bool,X332) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,X330) = s(bool,X334) | s(bool,'T_0') = s(bool,X331) | s(bool,'F_0') = s(bool,X334)) )),
  inference(superposition,[],[f906,f2782])).
fof(f6774,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,'F_0') = s(bool,X317) | s(bool,'T_0') = s(bool,X319) | s(bool,X317) = s(bool,X320) | s(bool,'T_0') = s(bool,X318) | s(bool,'F_0') = s(bool,X320)) )),
  inference(subsumption_resolution,[],[f5657,f2328])).
fof(f5657,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,X318) != s(bool,X319) | s(bool,'F_0') = s(bool,X317) | s(bool,'T_0') = s(bool,X319) | s(bool,X317) = s(bool,X320) | s(bool,'T_0') = s(bool,X318) | s(bool,'F_0') = s(bool,X320)) )),
  inference(superposition,[],[f826,f2782])).
fof(f6772,plain,(
  ( ! [X313,X315,X316,X314] : (s(bool,'F_0') = s(bool,X315) | s(bool,X313) != s(bool,X315) | s(bool,X313) = s(bool,X316) | s(bool,'T_0') = s(bool,X314) | s(bool,'F_0') = s(bool,X316)) )),
  inference(subsumption_resolution,[],[f5656,f306])).
fof(f5656,plain,(
  ( ! [X313,X315,X316,X314] : (s(bool,'F_0') != s(bool,X314) | s(bool,'F_0') = s(bool,X315) | s(bool,X313) != s(bool,X315) | s(bool,X313) = s(bool,X316) | s(bool,'T_0') = s(bool,X314) | s(bool,'F_0') = s(bool,X316)) )),
  inference(superposition,[],[f824,f2782])).
fof(f6771,plain,(
  ( ! [X310,X312,X309,X311] : (~p(s(bool,X309)) | s(bool,'T_0') = s(bool,X311) | s(bool,X309) = s(bool,X312) | s(bool,'T_0') = s(bool,X310) | s(bool,'F_0') = s(bool,X312)) )),
  inference(subsumption_resolution,[],[f5655,f2328])).
fof(f5655,plain,(
  ( ! [X310,X312,X309,X311] : (s(bool,X310) != s(bool,X311) | ~p(s(bool,X309)) | s(bool,'T_0') = s(bool,X311) | s(bool,X309) = s(bool,X312) | s(bool,'T_0') = s(bool,X310) | s(bool,'F_0') = s(bool,X312)) )),
  inference(superposition,[],[f806,f2782])).
fof(f6770,plain,(
  ( ! [X300,X302,X304,X303] : (s(bool,X300) != s(bool,X302) | s(bool,'F_0') = s(bool,X304) | s(bool,X300) = s(bool,X304) | s(bool,'F_0') = s(bool,X303) | s(bool,X302) != s(bool,X303)) ) | $spl162),
  inference(cnf_transformation,[],[f6770_D])).
fof(f6770_D,plain,(
  ( ! [X300,X302,X304,X303] : (s(bool,X300) != s(bool,X302) | s(bool,'F_0') = s(bool,X304) | s(bool,X300) = s(bool,X304) | s(bool,'F_0') = s(bool,X303) | s(bool,X302) != s(bool,X303)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6769,plain,(
  ( ! [X300,X302,X304,X301,X303] : (p(s(bool,X301)) | s(bool,X300) != s(bool,X302) | s(bool,X302) != s(bool,X303) | s(bool,'F_0') = s(bool,X303) | s(bool,X300) = s(bool,X304) | s(bool,'F_0') = s(bool,X304)) ) | $spl30),
  inference(subsumption_resolution,[],[f5653,f267])).
fof(f5653,plain,(
  ( ! [X300,X302,X304,X301,X303] : (p(s(bool,X301)) | s(bool,X300) != s(bool,X302) | s(bool,X302) != s(bool,X303) | s(bool,'F_0') = s(bool,X303) | s(bool,X300) = s(bool,X304) | s(bool,'T_0') = s(bool,X301) | s(bool,'F_0') = s(bool,X304)) ) | $spl30),
  inference(superposition,[],[f755,f2782])).
fof(f6768,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X293) != s(bool,X294) | p(s(bool,X294)) | ~p(s(bool,X292)) | s(bool,X292) = s(bool,X295) | s(bool,'F_0') = s(bool,X295)) )),
  inference(subsumption_resolution,[],[f5651,f2028])).
fof(f5651,plain,(
  ( ! [X292,X294,X293,X295] : (s(bool,X293) != s(bool,X294) | p(s(bool,X294)) | ~p(s(bool,X292)) | s(bool,X292) = s(bool,X295) | s(bool,'T_0') = s(bool,X293) | s(bool,'F_0') = s(bool,X295)) )),
  inference(superposition,[],[f368,f2782])).
fof(f6766,plain,(
  ( ! [X280,X277,X279,X281,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | ~p(s(bool,X280)) | s(bool,X277) != s(bool,X280) | s(bool,X277) = s(bool,X281) | s(bool,'F_0') = s(bool,X281)) )),
  inference(subsumption_resolution,[],[f5648,f267])).
fof(f5648,plain,(
  ( ! [X280,X277,X279,X281,X278] : (p(s(bool,X278)) | p(s(bool,X279)) | ~p(s(bool,X280)) | s(bool,X277) != s(bool,X280) | s(bool,X277) = s(bool,X281) | s(bool,'T_0') = s(bool,X278) | s(bool,'F_0') = s(bool,X281)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6765,plain,(
  ( ! [X269,X271,X267,X268,X270] : (s(bool,X268) != s(bool,X269) | p(s(bool,X270)) | ~p(s(bool,X267)) | p(s(bool,X269)) | s(bool,X267) = s(bool,X271) | s(bool,'F_0') = s(bool,X271)) )),
  inference(subsumption_resolution,[],[f5646,f2028])).
fof(f5646,plain,(
  ( ! [X269,X271,X267,X268,X270] : (s(bool,X268) != s(bool,X269) | p(s(bool,X270)) | ~p(s(bool,X267)) | p(s(bool,X269)) | s(bool,X267) = s(bool,X271) | s(bool,'T_0') = s(bool,X268) | s(bool,'F_0') = s(bool,X271)) )),
  inference(superposition,[],[f307,f2782])).
fof(f6458,plain,(
  ( ! [X261,X263,X262] : (s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X261) | s(bool,X261) = s(bool,X263) | s(bool,'F_0') = s(bool,X263)) )),
  inference(duplicate_literal_removal,[],[f5644])).
fof(f5644,plain,(
  ( ! [X261,X263,X262] : (s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X261) | s(bool,X261) = s(bool,X263) | s(bool,'T_0') = s(bool,X262) | s(bool,'F_0') = s(bool,X263)) )),
  inference(superposition,[],[f306,f2782])).
fof(f6459,plain,(
  ( ! [X255,X257,X256] : (s(bool,'T_0') = s(bool,X256) | ~p(s(bool,X255)) | s(bool,X255) = s(bool,X257) | s(bool,'F_0') = s(bool,X257)) )),
  inference(duplicate_literal_removal,[],[f5642])).
fof(f5642,plain,(
  ( ! [X255,X257,X256] : (s(bool,'T_0') = s(bool,X256) | ~p(s(bool,X255)) | s(bool,X255) = s(bool,X257) | s(bool,'T_0') = s(bool,X256) | s(bool,'F_0') = s(bool,X257)) )),
  inference(superposition,[],[f270,f2782])).
fof(f6764,plain,(
  ( ! [X253,X254,X252] : (p(s(bool,X253)) | s(bool,'T_0') != s(bool,X252) | s(bool,X252) = s(bool,X254) | s(bool,'F_0') = s(bool,X254)) )),
  inference(subsumption_resolution,[],[f5641,f267])).
fof(f5641,plain,(
  ( ! [X253,X254,X252] : (p(s(bool,X253)) | s(bool,'T_0') != s(bool,X252) | s(bool,X252) = s(bool,X254) | s(bool,'T_0') = s(bool,X253) | s(bool,'F_0') = s(bool,X254)) )),
  inference(superposition,[],[f269,f2782])).
fof(f6460,plain,(
  ( ! [X243,X244,X245] : (s(bool,'T_0') = s(bool,X244) | ~p(s(bool,X243)) | s(bool,X243) = s(bool,X245) | s(bool,'F_0') = s(bool,X245)) )),
  inference(duplicate_literal_removal,[],[f5638])).
fof(f5638,plain,(
  ( ! [X243,X244,X245] : (s(bool,'T_0') = s(bool,X244) | ~p(s(bool,X243)) | s(bool,X243) = s(bool,X245) | s(bool,'T_0') = s(bool,X244) | s(bool,'F_0') = s(bool,X245)) )),
  inference(superposition,[],[f268,f2782])).
fof(f6762,plain,(
  ( ! [X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,X240) | s(bool,X240) = s(bool,X242) | s(bool,'F_0') = s(bool,X242)) )),
  inference(subsumption_resolution,[],[f5637,f267])).
fof(f5637,plain,(
  ( ! [X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,X240) | s(bool,X240) = s(bool,X242) | s(bool,'T_0') = s(bool,X241) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f267,f2782])).
fof(f6761,plain,(
  ( ! [X235,X236,X234] : (p(s(bool,X235)) | s(bool,'F_0') = s(bool,X234) | s(bool,X234) = s(bool,X236) | s(bool,'F_0') = s(bool,X236)) )),
  inference(subsumption_resolution,[],[f5635,f267])).
fof(f5635,plain,(
  ( ! [X235,X236,X234] : (p(s(bool,X235)) | s(bool,'F_0') = s(bool,X234) | s(bool,X234) = s(bool,X236) | s(bool,'T_0') = s(bool,X235) | s(bool,'F_0') = s(bool,X236)) )),
  inference(superposition,[],[f264,f2782])).
fof(f6759,plain,(
  ( ! [X227,X225,X226] : (~p(s(bool,X225)) | s(bool,X225) = s(bool,X227) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,X227)) )),
  inference(subsumption_resolution,[],[f5632,f306])).
fof(f5632,plain,(
  ( ! [X227,X225,X226] : (s(bool,'F_0') != s(bool,X226) | ~p(s(bool,X225)) | s(bool,X225) = s(bool,X227) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,X227)) )),
  inference(superposition,[],[f263,f2782])).
fof(f6758,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~p(s(bool,X189)) | s(bool,X186) != s(bool,X189) | s(bool,X188) != s(bool,X190) | s(bool,X186) = s(bool,X191) | s(bool,'F_0') = s(bool,X191)) )),
  inference(subsumption_resolution,[],[f5622,f267])).
fof(f5622,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,X187)) | ~p(s(bool,X188)) | ~p(s(bool,X189)) | s(bool,X186) != s(bool,X189) | s(bool,X188) != s(bool,X190) | s(bool,X186) = s(bool,X191) | s(bool,'T_0') = s(bool,X187) | s(bool,'F_0') = s(bool,X191)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6757,plain,(
  ( ! [X185,X182,X180,X184,X183,X181] : (s(bool,X181) != s(bool,X182) | p(s(bool,X182)) | ~p(s(bool,X183)) | ~p(s(bool,X180)) | s(bool,X183) != s(bool,X184) | s(bool,X180) = s(bool,X185) | s(bool,'F_0') = s(bool,X185)) )),
  inference(subsumption_resolution,[],[f5621,f2028])).
fof(f5621,plain,(
  ( ! [X185,X182,X180,X184,X183,X181] : (s(bool,X181) != s(bool,X182) | p(s(bool,X182)) | ~p(s(bool,X183)) | ~p(s(bool,X180)) | s(bool,X183) != s(bool,X184) | s(bool,X180) = s(bool,X185) | s(bool,'T_0') = s(bool,X181) | s(bool,'F_0') = s(bool,X185)) )),
  inference(superposition,[],[f224,f2782])).
fof(f6753,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X138) != s(bool,X141) | s(bool,X138) = s(bool,X143) | s(bool,'F_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f5614,f267])).
fof(f5614,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X139)) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X138) != s(bool,X141) | s(bool,X138) = s(bool,X143) | s(bool,'T_0') = s(bool,X139) | s(bool,'F_0') = s(bool,X143)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6752,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,X127) != s(bool,X128) | p(s(bool,X129)) | ~p(s(bool,X126)) | ~p(s(bool,X130)) | p(s(bool,X128)) | s(bool,X126) = s(bool,X131) | s(bool,'F_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f5612,f2028])).
fof(f5612,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,X127) != s(bool,X128) | p(s(bool,X129)) | ~p(s(bool,X126)) | ~p(s(bool,X130)) | p(s(bool,X128)) | s(bool,X126) = s(bool,X131) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X131)) )),
  inference(superposition,[],[f223,f2782])).
fof(f6750,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X108) != s(bool,X111) | s(bool,X108) = s(bool,X113) | s(bool,'F_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f5609,f267])).
fof(f5609,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | s(bool,X110) != s(bool,X112) | s(bool,X108) != s(bool,X111) | s(bool,X108) = s(bool,X113) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X113)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6747,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X84)) | p(s(bool,X86)) | s(bool,X87) != s(bool,X88) | s(bool,X84) = s(bool,X89) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f5605,f2028])).
fof(f5605,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X84)) | p(s(bool,X86)) | s(bool,X87) != s(bool,X88) | s(bool,X84) = s(bool,X89) | s(bool,'T_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f222,f2782])).
fof(f6746,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X66) != s(bool,X69) | s(bool,X68) != s(bool,X70) | s(bool,X66) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f5602,f267])).
fof(f5602,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X66) != s(bool,X69) | s(bool,X68) != s(bool,X70) | s(bool,X66) = s(bool,X71) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6745,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X54)) | s(bool,X57) != s(bool,X58) | s(bool,X54) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f5600,f2028])).
fof(f5600,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X54)) | s(bool,X57) != s(bool,X58) | s(bool,X54) = s(bool,X59) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f220,f2782])).
fof(f6740,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X24) != s(bool,X27) | s(bool,X24) = s(bool,X29) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5595,f267])).
fof(f5595,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X24) != s(bool,X27) | s(bool,X24) = s(bool,X29) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6737,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,X7) != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X6)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f5592,f2028])).
fof(f5592,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,X7) != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X6)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f218,f2782])).
fof(f6736,plain,(
  ( ! [X629,X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'T_0') = s(bool,X629) | s(bool,X629) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) | $spl156),
  inference(cnf_transformation,[],[f6736_D])).
fof(f6736_D,plain,(
  ( ! [X629,X625,X627,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s(bool,'T_0') = s(bool,X629) | s(bool,X629) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6735,plain,(
  ( ! [X629,X625,X627,X628,X624,X626] : (~p(s(bool,X628)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | s(bool,X629) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s(bool,'T_0') = s(bool,X629)) )),
  inference(subsumption_resolution,[],[f5586,f263])).
fof(f5586,plain,(
  ( ! [X629,X625,X627,X628,X624,X626] : (~p(s(bool,X628)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | s(bool,X629) = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627))))) | s(bool,'T_0') = s(bool,X629) | s(bool,'F_0') = s(bool,X628)) )),
  inference(superposition,[],[f278,f2782])).
fof(f6734,plain,(
  ( ! [X617,X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'T_0') = s(bool,X617) | s(bool,X617) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) | $spl154),
  inference(cnf_transformation,[],[f6734_D])).
fof(f6734_D,plain,(
  ( ! [X617,X613,X615,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | s(bool,'T_0') = s(bool,X617) | s(bool,X617) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6733,plain,(
  ( ! [X617,X613,X615,X616,X612,X614] : (~p(s(bool,X616)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615)))))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s(bool,X617) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s(bool,'T_0') = s(bool,X617)) )),
  inference(subsumption_resolution,[],[f5584,f263])).
fof(f5584,plain,(
  ( ! [X617,X613,X615,X616,X612,X614] : (~p(s(bool,X616)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615)))))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s(bool,X617) = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615))))) | s(bool,'T_0') = s(bool,X617) | s(bool,'F_0') = s(bool,X616)) )),
  inference(superposition,[],[f278,f2782])).
fof(f5117,plain,(
  ( ! [X61,X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,X61) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59))) | s(bool,X61) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59)))) ) | $spl116),
  inference(cnf_transformation,[],[f5117_D])).
fof(f5117_D,plain,(
  ( ! [X61,X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,X61) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59))) | s(bool,X61) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6732,plain,(
  ( ! [X606,X602,X605,X601,X603,X604] : (s(bool,'T_0') != s(bool,X605) | s('type.ffi.ffi_state'(X602),'const.decSem.state_ffi_1'(s('type.decSem.state'(X602),X603))) = s('type.ffi.ffi_state'(X602),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X602),X604))) | s(bool,X606) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X601),s('type.decSem.state'(X602),X603),s('type.exhSem.state'(X602),X604))) | s(bool,'T_0') = s(bool,X606)) )),
  inference(subsumption_resolution,[],[f5582,f1729])).
fof(f5582,plain,(
  ( ! [X606,X602,X605,X601,X603,X604] : (s(bool,'T_0') != s(bool,X605) | s('type.ffi.ffi_state'(X602),'const.decSem.state_ffi_1'(s('type.decSem.state'(X602),X603))) = s('type.ffi.ffi_state'(X602),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X602),X604))) | s(bool,X606) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X601),s('type.decSem.state'(X602),X603),s('type.exhSem.state'(X602),X604))) | s(bool,'T_0') = s(bool,X606) | s(bool,'F_0') = s(bool,X605)) )),
  inference(superposition,[],[f2266,f2782])).
fof(f6461,plain,(
  ( ! [X600,X596,X598,X597,X599,X595] : (s(bool,'F_0') = s(bool,X599) | s('type.ffi.ffi_state'(X596),'const.decSem.state_ffi_1'(s('type.decSem.state'(X596),X597))) = s('type.ffi.ffi_state'(X596),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X596),X598))) | s(bool,X600) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X595),s('type.decSem.state'(X596),X597),s('type.exhSem.state'(X596),X598))) | s(bool,'T_0') = s(bool,X600)) )),
  inference(duplicate_literal_removal,[],[f5581])).
fof(f5581,plain,(
  ( ! [X600,X596,X598,X597,X599,X595] : (s(bool,'F_0') = s(bool,X599) | s('type.ffi.ffi_state'(X596),'const.decSem.state_ffi_1'(s('type.decSem.state'(X596),X597))) = s('type.ffi.ffi_state'(X596),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X596),X598))) | s(bool,X600) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X595),s('type.decSem.state'(X596),X597),s('type.exhSem.state'(X596),X598))) | s(bool,'T_0') = s(bool,X600) | s(bool,'F_0') = s(bool,X599)) )),
  inference(superposition,[],[f2265,f2782])).
fof(f5119,plain,(
  ( ! [X64,X62,X67,X65,X63] : (s(bool,'T_0') = s(bool,X67) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X63),X64))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X63),X65))) | s(bool,X67) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X62),s('type.decSem.state'(X63),X64),s('type.exhSem.state'(X63),X65)))) ) | $spl118),
  inference(cnf_transformation,[],[f5119_D])).
fof(f5119_D,plain,(
  ( ! [X64,X62,X67,X65,X63] : (s(bool,'T_0') = s(bool,X67) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X63),X64))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X63),X65))) | s(bool,X67) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X62),s('type.decSem.state'(X63),X64),s('type.exhSem.state'(X63),X65)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6731,plain,(
  ( ! [X592,X594,X590,X593,X589,X591] : (s(bool,'T_0') != s(bool,X593) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X590),X591))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X590),X592))) | s(bool,X594) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X589),s('type.decSem.state'(X590),X591),s('type.exhSem.state'(X590),X592))) | s(bool,'T_0') = s(bool,X594)) )),
  inference(subsumption_resolution,[],[f5580,f1729])).
fof(f5580,plain,(
  ( ! [X592,X594,X590,X593,X589,X591] : (s(bool,'T_0') != s(bool,X593) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X590),X591))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X590),X592))) | s(bool,X594) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X589),s('type.decSem.state'(X590),X591),s('type.exhSem.state'(X590),X592))) | s(bool,'T_0') = s(bool,X594) | s(bool,'F_0') = s(bool,X593)) )),
  inference(superposition,[],[f2062,f2782])).
fof(f6462,plain,(
  ( ! [X588,X584,X586,X585,X587,X583] : (s(bool,'F_0') = s(bool,X587) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X584),X585))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X584),X586))) | s(bool,X588) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X583),s('type.decSem.state'(X584),X585),s('type.exhSem.state'(X584),X586))) | s(bool,'T_0') = s(bool,X588)) )),
  inference(duplicate_literal_removal,[],[f5579])).
fof(f5579,plain,(
  ( ! [X588,X584,X586,X585,X587,X583] : (s(bool,'F_0') = s(bool,X587) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X584),X585))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X584),X586))) | s(bool,X588) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X583),s('type.decSem.state'(X584),X585),s('type.exhSem.state'(X584),X586))) | s(bool,'T_0') = s(bool,X588) | s(bool,'F_0') = s(bool,X587)) )),
  inference(superposition,[],[f2061,f2782])).
fof(f6730,plain,(
  ( ! [X573,X571,X576,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'T_0') = s(bool,X576) | s(bool,X576) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) | $spl152),
  inference(cnf_transformation,[],[f6730_D])).
fof(f6730_D,plain,(
  ( ! [X573,X571,X576,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,'T_0') = s(bool,X576) | s(bool,X576) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6729,plain,(
  ( ! [X573,X575,X571,X576,X572,X574] : (~p(s(bool,X575)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,X576) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574))) | s(bool,'T_0') = s(bool,X576)) )),
  inference(subsumption_resolution,[],[f5577,f263])).
fof(f5577,plain,(
  ( ! [X573,X575,X571,X576,X572,X574] : (~p(s(bool,X575)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,X576) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574))) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,X575)) )),
  inference(superposition,[],[f277,f2782])).
fof(f6728,plain,(
  ( ! [X569,X565,X567,X568,X570,X566] : (~p(s(bool,X569)) | s('type.ffi.ffi_state'(X566),'const.decSem.state_ffi_1'(s('type.decSem.state'(X566),X567))) = s('type.ffi.ffi_state'(X566),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X566),X568))) | s(bool,X570) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X565),s('type.decSem.state'(X566),X567),s('type.exhSem.state'(X566),X568))) | s(bool,'T_0') = s(bool,X570)) )),
  inference(subsumption_resolution,[],[f5576,f263])).
fof(f5576,plain,(
  ( ! [X569,X565,X567,X568,X570,X566] : (~p(s(bool,X569)) | s('type.ffi.ffi_state'(X566),'const.decSem.state_ffi_1'(s('type.decSem.state'(X566),X567))) = s('type.ffi.ffi_state'(X566),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X566),X568))) | s(bool,X570) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X565),s('type.decSem.state'(X566),X567),s('type.exhSem.state'(X566),X568))) | s(bool,'T_0') = s(bool,X570) | s(bool,'F_0') = s(bool,X569)) )),
  inference(superposition,[],[f276,f2782])).
fof(f6727,plain,(
  ( ! [X561,X559,X564,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'T_0') = s(bool,X564) | s(bool,X564) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) | $spl150),
  inference(cnf_transformation,[],[f6727_D])).
fof(f6727_D,plain,(
  ( ! [X561,X559,X564,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,'T_0') = s(bool,X564) | s(bool,X564) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6726,plain,(
  ( ! [X561,X563,X559,X564,X560,X562] : (~p(s(bool,X563)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,X564) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562))) | s(bool,'T_0') = s(bool,X564)) )),
  inference(subsumption_resolution,[],[f5575,f263])).
fof(f5575,plain,(
  ( ! [X561,X563,X559,X564,X560,X562] : (~p(s(bool,X563)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,X564) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562))) | s(bool,'T_0') = s(bool,X564) | s(bool,'F_0') = s(bool,X563)) )),
  inference(superposition,[],[f275,f2782])).
fof(f6725,plain,(
  ( ! [X557,X553,X555,X556,X558,X554] : (~p(s(bool,X557)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X554),X555))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X554),X556))) | s(bool,X558) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X553),s('type.decSem.state'(X554),X555),s('type.exhSem.state'(X554),X556))) | s(bool,'T_0') = s(bool,X558)) )),
  inference(subsumption_resolution,[],[f5574,f263])).
fof(f5574,plain,(
  ( ! [X557,X553,X555,X556,X558,X554] : (~p(s(bool,X557)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X554),X555))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X554),X556))) | s(bool,X558) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X553),s('type.decSem.state'(X554),X555),s('type.exhSem.state'(X554),X556))) | s(bool,'T_0') = s(bool,X558) | s(bool,'F_0') = s(bool,X557)) )),
  inference(superposition,[],[f274,f2782])).
fof(f6724,plain,(
  ( ! [X552] : (s(bool,X552) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X552)) ) | $spl148),
  inference(cnf_transformation,[],[f6724_D])).
fof(f6724_D,plain,(
  ( ! [X552] : (s(bool,X552) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X552)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6723,plain,(
  ( ! [X551,X552] : (~p(s(bool,X551)) | s(bool,X552) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X552)) )),
  inference(subsumption_resolution,[],[f5573,f263])).
fof(f5573,plain,(
  ( ! [X551,X552] : (~p(s(bool,X551)) | s(bool,X552) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X552) | s(bool,'F_0') = s(bool,X551)) )),
  inference(superposition,[],[f206,f2782])).
fof(f6468,plain,(
  ( ! [X482,X481,X483] : (s(bool,X481) = s(bool,X482) | s(bool,'F_0') = s(bool,X482) | s(bool,'T_0') = s(bool,X483) | s(bool,'F_0') = s(bool,X481)) )),
  inference(duplicate_literal_removal,[],[f5555])).
fof(f5555,plain,(
  ( ! [X482,X481,X483] : (s(bool,X481) = s(bool,X482) | s(bool,'F_0') = s(bool,X482) | s(bool,'T_0') = s(bool,X483) | s(bool,'T_0') = s(bool,X483) | s(bool,'F_0') = s(bool,X481)) )),
  inference(superposition,[],[f306,f2782])).
fof(f6471,plain,(
  ( ! [X476,X474,X475] : (s(bool,X474) = s(bool,X475) | ~p(s(bool,X475)) | s(bool,'T_0') = s(bool,X476) | s(bool,'F_0') = s(bool,X474)) )),
  inference(duplicate_literal_removal,[],[f5552])).
fof(f5552,plain,(
  ( ! [X476,X474,X475] : (s(bool,X474) = s(bool,X475) | ~p(s(bool,X475)) | s(bool,'T_0') = s(bool,X476) | s(bool,'T_0') = s(bool,X476) | s(bool,'F_0') = s(bool,X474)) )),
  inference(superposition,[],[f270,f2782])).
fof(f6473,plain,(
  ( ! [X469,X468,X470] : (s(bool,X468) = s(bool,X469) | ~p(s(bool,X469)) | s(bool,'T_0') = s(bool,X470) | s(bool,'F_0') = s(bool,X468)) )),
  inference(duplicate_literal_removal,[],[f5550])).
fof(f5550,plain,(
  ( ! [X469,X468,X470] : (s(bool,X468) = s(bool,X469) | ~p(s(bool,X469)) | s(bool,'T_0') = s(bool,X470) | s(bool,'T_0') = s(bool,X470) | s(bool,'F_0') = s(bool,X468)) )),
  inference(superposition,[],[f268,f2782])).
fof(f6713,plain,(
  ( ! [X416,X418,X415,X417,X419] : (~p(s(bool,X416)) | s(bool,X417) = s(bool,X418) | s(bool,'T_0') = s(bool,X418) | s(bool,X415) = s(bool,X417) | s(bool,X415) = s(bool,X419) | s(bool,'T_0') = s(bool,X419)) )),
  inference(subsumption_resolution,[],[f5536,f263])).
fof(f5536,plain,(
  ( ! [X416,X418,X415,X417,X419] : (~p(s(bool,X416)) | s(bool,X417) = s(bool,X418) | s(bool,'T_0') = s(bool,X418) | s(bool,X415) = s(bool,X417) | s(bool,X415) = s(bool,X419) | s(bool,'T_0') = s(bool,X419) | s(bool,'F_0') = s(bool,X416)) )),
  inference(superposition,[],[f2783,f2782])).
fof(f4352,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl98),
  inference(cnf_transformation,[],[f4352_D])).
fof(f4352_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6712,plain,(
  ( ! [X412,X414,X410,X413,X411] : (~p(s(bool,X411)) | s(bool,X412) != s(bool,X413) | p(s(bool,X413)) | s(bool,X410) != s(bool,X412) | s(bool,X410) = s(bool,X414) | s(bool,'T_0') = s(bool,X414)) )),
  inference(subsumption_resolution,[],[f5535,f263])).
fof(f5535,plain,(
  ( ! [X412,X414,X410,X413,X411] : (~p(s(bool,X411)) | s(bool,X412) != s(bool,X413) | p(s(bool,X413)) | s(bool,X410) != s(bool,X412) | s(bool,X410) = s(bool,X414) | s(bool,'T_0') = s(bool,X414) | s(bool,'F_0') = s(bool,X411)) )),
  inference(superposition,[],[f2388,f2782])).
fof(f6632,plain,(
  ( ! [X333,X331,X332,X334] : (s(bool,X331) != s(bool,X332) | s(bool,'T_0') = s(bool,X334) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X333) | s(bool,X332) != s(bool,X333)) ) | $spl136),
  inference(cnf_transformation,[],[f6632_D])).
fof(f6632_D,plain,(
  ( ! [X333,X331,X332,X334] : (s(bool,X331) != s(bool,X332) | s(bool,'T_0') = s(bool,X334) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X333) | s(bool,X332) != s(bool,X333)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6706,plain,(
  ( ! [X344,X346,X343,X345,X347] : (~p(s(bool,X344)) | s(bool,X343) != s(bool,X345) | s(bool,'T_0') = s(bool,X346) | s(bool,X345) != s(bool,X346) | s(bool,X343) = s(bool,X347) | s(bool,'T_0') = s(bool,X347)) ) | $spl52),
  inference(subsumption_resolution,[],[f5519,f263])).
fof(f5519,plain,(
  ( ! [X344,X346,X343,X345,X347] : (~p(s(bool,X344)) | s(bool,X343) != s(bool,X345) | s(bool,'T_0') = s(bool,X346) | s(bool,X345) != s(bool,X346) | s(bool,X343) = s(bool,X347) | s(bool,'T_0') = s(bool,X347) | s(bool,'F_0') = s(bool,X344)) ) | $spl52),
  inference(superposition,[],[f1367,f2782])).
fof(f6660,plain,(
  ( ! [X629,X625,X627,X628,X624,X626] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | s(bool,X628) = s(bool,X629) | s(bool,'T_0') = s(bool,X629) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627)))))) )),
  inference(subsumption_resolution,[],[f5442,f1797])).
fof(f5442,plain,(
  ( ! [X629,X625,X627,X628,X624,X626] : (~p(s(bool,X628)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624),s('type.decSem.state'(X625),X626),s('type.exhSem.state'(X625),X627)))) | s('type.ffi.ffi_state'(X625),'const.decSem.state_ffi_1'(s('type.decSem.state'(X625),X626))) != s('type.ffi.ffi_state'(X625),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X625),X627))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X625),X627)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X625),X626))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X625),X627))) | s(bool,X628) = s(bool,X629) | s(bool,'T_0') = s(bool,X629) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X624))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X625),X626))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X625),X627)))))) )),
  inference(superposition,[],[f278,f2782])).
fof(f6658,plain,(
  ( ! [X617,X613,X615,X616,X612,X614] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615)))))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s(bool,X616) = s(bool,X617) | s(bool,'T_0') = s(bool,X617) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615)))))) )),
  inference(subsumption_resolution,[],[f5440,f1797])).
fof(f5440,plain,(
  ( ! [X617,X613,X615,X616,X612,X614] : (~p(s(bool,X616)) | p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612),s('type.decSem.state'(X613),X614),s('type.exhSem.state'(X613),X615)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X613),X615)))))) | s('type.ffi.ffi_state'(X613),'const.decSem.state_ffi_1'(s('type.decSem.state'(X613),X614))) != s('type.ffi.ffi_state'(X613),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X613),X615))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X613),X614))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X613),X615))) | s(bool,X616) = s(bool,X617) | s(bool,'T_0') = s(bool,X617) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X612))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X613),X614))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X613),X615)))))) )),
  inference(superposition,[],[f278,f2782])).
fof(f6656,plain,(
  ( ! [X573,X575,X571,X576,X572,X574] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,X575) = s(bool,X576) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) )),
  inference(subsumption_resolution,[],[f5433,f1797])).
fof(f5433,plain,(
  ( ! [X573,X575,X571,X576,X572,X574] : (~p(s(bool,X575)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X572),X573))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X572),X574)))))) | s(bool,X575) = s(bool,X576) | s(bool,'T_0') = s(bool,X576) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X571),s('type.decSem.state'(X572),X573),s('type.exhSem.state'(X572),X574)))) )),
  inference(superposition,[],[f277,f2782])).
fof(f6654,plain,(
  ( ! [X561,X563,X559,X564,X560,X562] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,X563) = s(bool,X564) | s(bool,'T_0') = s(bool,X564) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) )),
  inference(subsumption_resolution,[],[f5431,f1797])).
fof(f5431,plain,(
  ( ! [X561,X563,X559,X564,X560,X562] : (~p(s(bool,X563)) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X560),X561))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X560),X562)))))) | s(bool,X563) = s(bool,X564) | s(bool,'T_0') = s(bool,X564) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X559),s('type.decSem.state'(X560),X561),s('type.exhSem.state'(X560),X562)))) )),
  inference(superposition,[],[f275,f2782])).
fof(f6651,plain,(
  ( ! [X551,X552] : (s(bool,X551) = s(bool,X552) | s(bool,'T_0') = s(bool,X552) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f5429,f1797])).
fof(f5429,plain,(
  ( ! [X551,X552] : (~p(s(bool,X551)) | s(bool,X551) = s(bool,X552) | s(bool,'T_0') = s(bool,X552) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f206,f2782])).
fof(f6484,plain,(
  ( ! [X333,X331,X332,X334,X330] : (s(bool,X331) != s(bool,X332) | s(bool,'F_0') = s(bool,X330) | s(bool,X332) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334)) )),
  inference(duplicate_literal_removal,[],[f5372])).
fof(f5372,plain,(
  ( ! [X333,X331,X332,X334,X330] : (s(bool,X331) != s(bool,X332) | s(bool,'F_0') = s(bool,X330) | s(bool,X332) != s(bool,X333) | s(bool,'T_0') = s(bool,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334) | s(bool,'F_0') = s(bool,X330)) )),
  inference(superposition,[],[f906,f2782])).
fof(f6589,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f5299,f3625])).
fof(f5299,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6488,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'F_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(duplicate_literal_removal,[],[f5297])).
fof(f5297,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'F_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6489,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X65)) )),
  inference(duplicate_literal_removal,[],[f5286])).
fof(f5286,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X65) = s(bool,X66) | s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6584,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f5282,f3625])).
fof(f5282,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,X45) | s(bool,X45) = s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6490,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f5265])).
fof(f5265,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f2782,f306])).
fof(f6579,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | p(s(bool,X56))) )),
  inference(subsumption_resolution,[],[f5264,f267])).
fof(f5264,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | p(s(bool,X56))) )),
  inference(superposition,[],[f2782,f264])).
fof(f6578,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X52) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f5263,f3625])).
fof(f5263,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6545,plain,(
  ( ! [X61,X64,X62,X65] : (s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X62)) ) | $spl126),
  inference(cnf_transformation,[],[f6545_D])).
fof(f6545_D,plain,(
  ( ! [X61,X64,X62,X65] : (s(bool,X61) = s(bool,X62) | s(bool,'F_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X62)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6491,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X35) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(duplicate_literal_removal,[],[f5261])).
fof(f5261,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X35) = s(bool,X39) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6494,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X64) | s(bool,X65) = s(bool,X66) | s(bool,X64) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X63)) )),
  inference(duplicate_literal_removal,[],[f5251])).
fof(f5251,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X64) | s(bool,X65) = s(bool,X66) | s(bool,X64) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6571,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,X102) | s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X99),X100))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X99),X101)))) )),
  inference(subsumption_resolution,[],[f5233,f2062])).
fof(f5233,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (s(bool,'F_0') = s(bool,X102) | s(bool,X102) = s(bool,X103) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X98),s('type.decSem.state'(X99),X100),s('type.exhSem.state'(X99),X101))) | s(bool,'F_0') = s(bool,X103) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X99),X100))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X99),X101)))) )),
  inference(superposition,[],[f2782,f2061])).
fof(f6569,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,X96) | s(bool,X96) = s(bool,X97) | s(bool,'F_0') = s(bool,X97) | s('type.ffi.ffi_state'(X93),'const.decSem.state_ffi_1'(s('type.decSem.state'(X93),X94))) = s('type.ffi.ffi_state'(X93),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X93),X95)))) )),
  inference(subsumption_resolution,[],[f5232,f2266])).
fof(f5232,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,X96) | s(bool,X96) = s(bool,X97) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X92),s('type.decSem.state'(X93),X94),s('type.exhSem.state'(X93),X95))) | s(bool,'F_0') = s(bool,X97) | s('type.ffi.ffi_state'(X93),'const.decSem.state_ffi_1'(s('type.decSem.state'(X93),X94))) = s('type.ffi.ffi_state'(X93),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X93),X95)))) )),
  inference(superposition,[],[f2782,f2265])).
fof(f6497,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,'F_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f5221])).
fof(f5221,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6564,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f5217,f3625])).
fof(f5217,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6500,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X14)) )),
  inference(duplicate_literal_removal,[],[f5211])).
fof(f5211,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2782,f306])).
fof(f6557,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f5207,f267])).
fof(f5207,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X0))) )),
  inference(superposition,[],[f2782,f264])).
fof(f6501,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'F_0') = s(bool,X95)) )),
  inference(duplicate_literal_removal,[],[f5202])).
fof(f5202,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f2782,f329])).
fof(f6504,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X84) | s(bool,X83) = s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(duplicate_literal_removal,[],[f5199])).
fof(f5199,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X83) = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X85) | s(bool,X83) = s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6550,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5189,f3625])).
fof(f5189,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6505,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,X31) = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(duplicate_literal_removal,[],[f5187])).
fof(f5187,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,X31) = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6508,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X9)) )),
  inference(duplicate_literal_removal,[],[f5181])).
fof(f5181,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f2782,f306])).
fof(f6509,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(duplicate_literal_removal,[],[f5180])).
fof(f5180,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f2782,f270])).
fof(f6510,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) )),
  inference(duplicate_literal_removal,[],[f5179])).
fof(f5179,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f2782,f268])).
fof(f6511,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(bool,'F_0') = s(bool,X106) | s(bool,X107) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X102),s('type.decSem.state'(X103),X104),s('type.exhSem.state'(X103),X105))) | s(bool,'T_0') = s(bool,X107) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X103),X104))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X103),X105)))) )),
  inference(duplicate_literal_removal,[],[f5175])).
fof(f5175,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(bool,'F_0') = s(bool,X106) | s(bool,X107) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X102),s('type.decSem.state'(X103),X104),s('type.exhSem.state'(X103),X105))) | s(bool,'T_0') = s(bool,X107) | s(bool,'F_0') = s(bool,X106) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X103),X104))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X103),X105)))) )),
  inference(superposition,[],[f2782,f2061])).
fof(f6512,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X96),s('type.decSem.state'(X97),X98),s('type.exhSem.state'(X97),X99))) | s(bool,'T_0') = s(bool,X101) | s('type.ffi.ffi_state'(X97),'const.decSem.state_ffi_1'(s('type.decSem.state'(X97),X98))) = s('type.ffi.ffi_state'(X97),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X97),X99)))) )),
  inference(duplicate_literal_removal,[],[f5174])).
fof(f5174,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X96),s('type.decSem.state'(X97),X98),s('type.exhSem.state'(X97),X99))) | s(bool,'T_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X100) | s('type.ffi.ffi_state'(X97),'const.decSem.state_ffi_1'(s('type.decSem.state'(X97),X98))) = s('type.ffi.ffi_state'(X97),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X97),X99)))) )),
  inference(superposition,[],[f2782,f2265])).
fof(f6513,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,X81) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81) | ~p(s(bool,X80))) )),
  inference(duplicate_literal_removal,[],[f5169])).
fof(f5169,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,X81) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81) | ~p(s(bool,X80))) )),
  inference(superposition,[],[f2782,f268])).
fof(f6514,plain,(
  ( ! [X78,X79,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78) | ~p(s(bool,X77))) )),
  inference(duplicate_literal_removal,[],[f5168])).
fof(f5168,plain,(
  ( ! [X78,X79,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X78) | ~p(s(bool,X77))) )),
  inference(superposition,[],[f2782,f270])).
fof(f6515,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(duplicate_literal_removal,[],[f5167])).
fof(f5167,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X74)) )),
  inference(superposition,[],[f2782,f306])).
fof(f6518,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X62) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(duplicate_literal_removal,[],[f5164])).
fof(f5164,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X62) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6543,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f5161,f3625])).
fof(f5161,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X47) = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6519,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f5159])).
fof(f5159,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6529,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f5133,f3625])).
fof(f5133,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X55) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f6522,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X41) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(duplicate_literal_removal,[],[f5131])).
fof(f5131,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X41) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f2782,f2782])).
fof(f5118,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (s(bool,'T_0') = s(bool,X67) | s(bool,X67) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X62),s('type.decSem.state'(X63),X64),s('type.exhSem.state'(X63),X65))) | s(bool,'T_0') = s(bool,X66) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X63),X64))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X63),X65)))) )),
  inference(subsumption_resolution,[],[f5096,f306])).
fof(f5096,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (s(bool,'F_0') != s(bool,X66) | s(bool,'T_0') = s(bool,X67) | s(bool,X67) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X62),s('type.decSem.state'(X63),X64),s('type.exhSem.state'(X63),X65))) | s(bool,'T_0') = s(bool,X66) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X63),X64))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X63),X65)))) )),
  inference(superposition,[],[f3625,f2061])).
fof(f5116,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X61) | s(bool,X61) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59))) | s(bool,'T_0') = s(bool,X60) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59)))) )),
  inference(subsumption_resolution,[],[f5095,f306])).
fof(f5095,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,'F_0') != s(bool,X60) | s(bool,'T_0') = s(bool,X61) | s(bool,X61) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59))) | s(bool,'T_0') = s(bool,X60) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59)))) )),
  inference(superposition,[],[f3625,f2265])).
fof(f5048,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl114),
  inference(cnf_transformation,[],[f5048_D])).
fof(f5048_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f5002,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) != s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X66) != s(bool,X70)) ) | $spl52),
  inference(resolution,[],[f1367,f307])).
fof(f4995,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) != s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X42) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl52),
  inference(resolution,[],[f1367,f224])).
fof(f4993,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X30) != s(bool,X34)) ) | $spl52),
  inference(resolution,[],[f1367,f223])).
fof(f4991,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X18) != s(bool,X22)) ) | $spl52),
  inference(resolution,[],[f1367,f222])).
fof(f4990,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X12) != s(bool,X16) | s(bool,X15) != s(bool,X17)) ) | $spl52),
  inference(resolution,[],[f1367,f220])).
fof(f4988,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X0) != s(bool,X4)) ) | $spl52),
  inference(resolution,[],[f1367,f218])).
fof(f4915,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X50)) ) | $spl112),
  inference(cnf_transformation,[],[f4915_D])).
fof(f4915_D,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X50)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4916,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'T_0') != s(bool,X53)) )),
  inference(subsumption_resolution,[],[f4901,f267])).
fof(f4901,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'T_0') != s(bool,X53) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f3958,f306])).
fof(f4907,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,X50)) )),
  inference(duplicate_literal_removal,[],[f4900])).
fof(f4900,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | s(bool,'T_0') != s(bool,X50) | p(s(bool,X48))) )),
  inference(superposition,[],[f3958,f264])).
fof(f4912,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f4892,f267])).
fof(f4892,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f3958,f2328])).
fof(f4910,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f4891,f267])).
fof(f4891,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f3958,f2328])).
fof(f4908,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl108),
  inference(cnf_transformation,[],[f4908_D])).
fof(f4908_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4875,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'T_0') != s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) )),
  inference(resolution,[],[f3958,f307])).
fof(f4867,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f3958,f224])).
fof(f4864,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) )),
  inference(resolution,[],[f3958,f223])).
fof(f4863,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) )),
  inference(resolution,[],[f3958,f222])).
fof(f4862,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3958,f220])).
fof(f4860,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) )),
  inference(resolution,[],[f3958,f218])).
fof(f4855,plain,(
  ( ! [X478,X480,X482,X477,X479] : (p(s(bool,X482)) | s('type.ffi.ffi_state'(X478),'const.decSem.state_ffi_1'(s('type.decSem.state'(X478),X479))) = s('type.ffi.ffi_state'(X478),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X478),X480))) | s(bool,X482) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X477),s('type.decSem.state'(X478),X479),s('type.exhSem.state'(X478),X480)))) ) | $spl106),
  inference(cnf_transformation,[],[f4855_D])).
fof(f4855_D,plain,(
  ( ! [X478,X480,X482,X477,X479] : (p(s(bool,X482)) | s('type.ffi.ffi_state'(X478),'const.decSem.state_ffi_1'(s('type.decSem.state'(X478),X479))) = s('type.ffi.ffi_state'(X478),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X478),X480))) | s(bool,X482) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X477),s('type.decSem.state'(X478),X479),s('type.exhSem.state'(X478),X480)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4854,plain,(
  ( ! [X478,X480,X482,X477,X479,X481] : (p(s(bool,X482)) | p(s(bool,X481)) | s(bool,X482) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X477),s('type.decSem.state'(X478),X479),s('type.exhSem.state'(X478),X480))) | s('type.ffi.ffi_state'(X478),'const.decSem.state_ffi_1'(s('type.decSem.state'(X478),X479))) = s('type.ffi.ffi_state'(X478),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X478),X480)))) ) | $spl66),
  inference(subsumption_resolution,[],[f4803,f264])).
fof(f4803,plain,(
  ( ! [X478,X480,X482,X477,X479,X481] : (s(bool,'F_0') != s(bool,X481) | p(s(bool,X482)) | p(s(bool,X481)) | s(bool,X482) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X477),s('type.decSem.state'(X478),X479),s('type.exhSem.state'(X478),X480))) | s('type.ffi.ffi_state'(X478),'const.decSem.state_ffi_1'(s('type.decSem.state'(X478),X479))) = s('type.ffi.ffi_state'(X478),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X478),X480)))) ) | $spl66),
  inference(superposition,[],[f2249,f2265])).
fof(f2309,plain,(
  ( ! [X6,X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X2)))) ) | $spl68),
  inference(cnf_transformation,[],[f2309_D])).
fof(f2309_D,plain,(
  ( ! [X6,X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2))) | s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X2)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4844,plain,(
  ( ! [X364,X366,X362,X365,X367,X363] : (p(s(bool,X366)) | ~p(s(bool,X367)) | s(bool,X367) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X362),s('type.decSem.state'(X363),X364),s('type.exhSem.state'(X363),X365))) | s('type.ffi.ffi_state'(X363),'const.decSem.state_ffi_1'(s('type.decSem.state'(X363),X364))) = s('type.ffi.ffi_state'(X363),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X363),X365)))) )),
  inference(subsumption_resolution,[],[f4781,f271])).
fof(f4781,plain,(
  ( ! [X364,X366,X362,X365,X367,X363] : (p(s(bool,'F_0')) | p(s(bool,X366)) | ~p(s(bool,X367)) | s(bool,X367) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X362),s('type.decSem.state'(X363),X364),s('type.exhSem.state'(X363),X365))) | s('type.ffi.ffi_state'(X363),'const.decSem.state_ffi_1'(s('type.decSem.state'(X363),X364))) = s('type.ffi.ffi_state'(X363),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X363),X365)))) )),
  inference(superposition,[],[f307,f2265])).
fof(f4840,plain,(
  ( ! [X251,X249,X246,X250,X248,X247,X245] : (~p(s(bool,X249)) | ~p(s(bool,X250)) | s(bool,X250) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X245),s('type.decSem.state'(X246),X247),s('type.exhSem.state'(X246),X248))) | s(bool,X249) != s(bool,X251) | s('type.ffi.ffi_state'(X246),'const.decSem.state_ffi_1'(s('type.decSem.state'(X246),X247))) = s('type.ffi.ffi_state'(X246),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X246),X248)))) )),
  inference(subsumption_resolution,[],[f4757,f271])).
fof(f4757,plain,(
  ( ! [X251,X249,X246,X250,X248,X247,X245] : (p(s(bool,'F_0')) | ~p(s(bool,X249)) | ~p(s(bool,X250)) | s(bool,X250) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X245),s('type.decSem.state'(X246),X247),s('type.exhSem.state'(X246),X248))) | s(bool,X249) != s(bool,X251) | s('type.ffi.ffi_state'(X246),'const.decSem.state_ffi_1'(s('type.decSem.state'(X246),X247))) = s('type.ffi.ffi_state'(X246),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X246),X248)))) )),
  inference(superposition,[],[f224,f2265])).
fof(f4837,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (p(s(bool,X193)) | ~p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X194) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X189),s('type.decSem.state'(X190),X191),s('type.exhSem.state'(X190),X192))) | s('type.ffi.ffi_state'(X190),'const.decSem.state_ffi_1'(s('type.decSem.state'(X190),X191))) = s('type.ffi.ffi_state'(X190),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X190),X192)))) )),
  inference(subsumption_resolution,[],[f4749,f271])).
fof(f4749,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (p(s(bool,'F_0')) | p(s(bool,X193)) | ~p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X194) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X189),s('type.decSem.state'(X190),X191),s('type.exhSem.state'(X190),X192))) | s('type.ffi.ffi_state'(X190),'const.decSem.state_ffi_1'(s('type.decSem.state'(X190),X191))) = s('type.ffi.ffi_state'(X190),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X190),X192)))) )),
  inference(superposition,[],[f223,f2265])).
fof(f4835,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (p(s(bool,X158)) | ~p(s(bool,X159)) | s(bool,X158) != s(bool,X160) | s(bool,X159) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X154),s('type.decSem.state'(X155),X156),s('type.exhSem.state'(X155),X157))) | s('type.ffi.ffi_state'(X155),'const.decSem.state_ffi_1'(s('type.decSem.state'(X155),X156))) = s('type.ffi.ffi_state'(X155),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X155),X157)))) )),
  inference(subsumption_resolution,[],[f4744,f271])).
fof(f4744,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X160] : (p(s(bool,'F_0')) | p(s(bool,X158)) | ~p(s(bool,X159)) | s(bool,X158) != s(bool,X160) | s(bool,X159) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X154),s('type.decSem.state'(X155),X156),s('type.exhSem.state'(X155),X157))) | s('type.ffi.ffi_state'(X155),'const.decSem.state_ffi_1'(s('type.decSem.state'(X155),X156))) = s('type.ffi.ffi_state'(X155),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X155),X157)))) )),
  inference(superposition,[],[f222,f2265])).
fof(f4832,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (~p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X110) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X105),s('type.decSem.state'(X106),X107),s('type.exhSem.state'(X106),X108))) | s(bool,X109) != s(bool,X111) | s('type.ffi.ffi_state'(X106),'const.decSem.state_ffi_1'(s('type.decSem.state'(X106),X107))) = s('type.ffi.ffi_state'(X106),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X106),X108)))) )),
  inference(subsumption_resolution,[],[f4737,f271])).
fof(f4737,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (p(s(bool,'F_0')) | ~p(s(bool,X109)) | ~p(s(bool,X110)) | s(bool,X110) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X105),s('type.decSem.state'(X106),X107),s('type.exhSem.state'(X106),X108))) | s(bool,X109) != s(bool,X111) | s('type.ffi.ffi_state'(X106),'const.decSem.state_ffi_1'(s('type.decSem.state'(X106),X107))) = s('type.ffi.ffi_state'(X106),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X106),X108)))) )),
  inference(superposition,[],[f220,f2265])).
fof(f4829,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,X62) | s(bool,X61) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59))) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59)))) )),
  inference(subsumption_resolution,[],[f4730,f271])).
fof(f4730,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (p(s(bool,'F_0')) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,X62) | s(bool,X61) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X56),s('type.decSem.state'(X57),X58),s('type.exhSem.state'(X57),X59))) | s('type.ffi.ffi_state'(X57),'const.decSem.state_ffi_1'(s('type.decSem.state'(X57),X58))) = s('type.ffi.ffi_state'(X57),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X57),X59)))) )),
  inference(superposition,[],[f218,f2265])).
fof(f4653,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,X66) = s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f2783,f307])).
fof(f4646,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X42) = s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f2783,f224])).
fof(f4644,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f2783,f223])).
fof(f4642,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f2783,f222])).
fof(f4641,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f2783,f220])).
fof(f4639,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f2783,f218])).
fof(f4611,plain,(
  ( ! [X472,X474,X469,X471,X470] : (p(s(bool,X474)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X470),X471))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X470),X472))) | s(bool,X474) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469),s('type.decSem.state'(X470),X471),s('type.exhSem.state'(X470),X472)))) ) | $spl102),
  inference(cnf_transformation,[],[f4611_D])).
fof(f4611_D,plain,(
  ( ! [X472,X474,X469,X471,X470] : (p(s(bool,X474)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X470),X471))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X470),X472))) | s(bool,X474) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469),s('type.decSem.state'(X470),X471),s('type.exhSem.state'(X470),X472)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4610,plain,(
  ( ! [X472,X474,X469,X471,X473,X470] : (p(s(bool,X474)) | p(s(bool,X473)) | s(bool,X474) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469),s('type.decSem.state'(X470),X471),s('type.exhSem.state'(X470),X472))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X470),X471))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X470),X472)))) ) | $spl66),
  inference(subsumption_resolution,[],[f4560,f264])).
fof(f4560,plain,(
  ( ! [X472,X474,X469,X471,X473,X470] : (s(bool,'F_0') != s(bool,X473) | p(s(bool,X474)) | p(s(bool,X473)) | s(bool,X474) = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X469),s('type.decSem.state'(X470),X471),s('type.exhSem.state'(X470),X472))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X470),X471))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X470),X472)))) ) | $spl66),
  inference(superposition,[],[f2249,f2061])).
fof(f2101,plain,(
  ( ! [X6,X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X2)))) ) | $spl64),
  inference(cnf_transformation,[],[f2101_D])).
fof(f2101_D,plain,(
  ( ! [X6,X4,X2,X0,X1] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X2)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4600,plain,(
  ( ! [X356,X358,X354,X357,X359,X355] : (p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X359) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X354),s('type.decSem.state'(X355),X356),s('type.exhSem.state'(X355),X357))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X355),X356))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X355),X357)))) )),
  inference(subsumption_resolution,[],[f4538,f271])).
fof(f4538,plain,(
  ( ! [X356,X358,X354,X357,X359,X355] : (p(s(bool,'F_0')) | p(s(bool,X358)) | ~p(s(bool,X359)) | s(bool,X359) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X354),s('type.decSem.state'(X355),X356),s('type.exhSem.state'(X355),X357))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X355),X356))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X355),X357)))) )),
  inference(superposition,[],[f307,f2061])).
fof(f4596,plain,(
  ( ! [X243,X241,X239,X237,X242,X240,X238] : (~p(s(bool,X241)) | ~p(s(bool,X242)) | s(bool,X242) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X237),s('type.decSem.state'(X238),X239),s('type.exhSem.state'(X238),X240))) | s(bool,X241) != s(bool,X243) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X238),X239))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X238),X240)))) )),
  inference(subsumption_resolution,[],[f4514,f271])).
fof(f4514,plain,(
  ( ! [X243,X241,X239,X237,X242,X240,X238] : (p(s(bool,'F_0')) | ~p(s(bool,X241)) | ~p(s(bool,X242)) | s(bool,X242) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X237),s('type.decSem.state'(X238),X239),s('type.exhSem.state'(X238),X240))) | s(bool,X241) != s(bool,X243) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X238),X239))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X238),X240)))) )),
  inference(superposition,[],[f224,f2061])).
fof(f4593,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (p(s(bool,X185)) | ~p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X186) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X181),s('type.decSem.state'(X182),X183),s('type.exhSem.state'(X182),X184))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X182),X183))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X182),X184)))) )),
  inference(subsumption_resolution,[],[f4506,f271])).
fof(f4506,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (p(s(bool,'F_0')) | p(s(bool,X185)) | ~p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X186) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X181),s('type.decSem.state'(X182),X183),s('type.exhSem.state'(X182),X184))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X182),X183))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X182),X184)))) )),
  inference(superposition,[],[f223,f2061])).
fof(f4591,plain,(
  ( ! [X146,X152,X151,X149,X147,X150,X148] : (p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,X152) | s(bool,X151) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X146),s('type.decSem.state'(X147),X148),s('type.exhSem.state'(X147),X149))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X147),X148))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X147),X149)))) )),
  inference(subsumption_resolution,[],[f4501,f271])).
fof(f4501,plain,(
  ( ! [X146,X152,X151,X149,X147,X150,X148] : (p(s(bool,'F_0')) | p(s(bool,X150)) | ~p(s(bool,X151)) | s(bool,X150) != s(bool,X152) | s(bool,X151) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X146),s('type.decSem.state'(X147),X148),s('type.exhSem.state'(X147),X149))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X147),X148))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X147),X149)))) )),
  inference(superposition,[],[f222,f2061])).
fof(f4588,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98] : (~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X97),s('type.decSem.state'(X98),X99),s('type.exhSem.state'(X98),X100))) | s(bool,X101) != s(bool,X103) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X98),X99))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X98),X100)))) )),
  inference(subsumption_resolution,[],[f4494,f271])).
fof(f4494,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98] : (p(s(bool,'F_0')) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X97),s('type.decSem.state'(X98),X99),s('type.exhSem.state'(X98),X100))) | s(bool,X101) != s(bool,X103) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X98),X99))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X98),X100)))) )),
  inference(superposition,[],[f220,f2061])).
fof(f4585,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X48),s('type.decSem.state'(X49),X50),s('type.exhSem.state'(X49),X51))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X49),X50))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X49),X51)))) )),
  inference(subsumption_resolution,[],[f4487,f271])).
fof(f4487,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (p(s(bool,'F_0')) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X48),s('type.decSem.state'(X49),X50),s('type.exhSem.state'(X49),X51))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X49),X50))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X49),X51)))) )),
  inference(superposition,[],[f218,f2061])).
fof(f3719,plain,(
  ( ! [X253,X252,X250] : (s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X253) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X252)) ) | $spl78),
  inference(cnf_transformation,[],[f3719_D])).
fof(f3719_D,plain,(
  ( ! [X253,X252,X250] : (s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X253) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X252)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4461,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f4444,f267])).
fof(f4444,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f3624,f2328])).
fof(f4460,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f4443,f267])).
fof(f4443,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f3624,f2328])).
fof(f4458,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl100),
  inference(cnf_transformation,[],[f4458_D])).
fof(f4458_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4428,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) = s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'T_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) )),
  inference(resolution,[],[f3624,f307])).
fof(f4420,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f3624,f224])).
fof(f4417,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) )),
  inference(resolution,[],[f3624,f223])).
fof(f4416,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) )),
  inference(resolution,[],[f3624,f222])).
fof(f4415,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) = s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f3624,f220])).
fof(f4413,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) )),
  inference(resolution,[],[f3624,f218])).
fof(f1201,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f1201_D])).
fof(f1201_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X4) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4376,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) != s(bool,X71) | s(bool,'F_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) ) | $spl30),
  inference(resolution,[],[f755,f307])).
fof(f4368,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl30),
  inference(resolution,[],[f755,f224])).
fof(f4365,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl30),
  inference(resolution,[],[f755,f223])).
fof(f4364,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl30),
  inference(resolution,[],[f755,f222])).
fof(f4363,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl30),
  inference(resolution,[],[f755,f220])).
fof(f4361,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl30),
  inference(resolution,[],[f755,f218])).
fof(f4355,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) ) | $spl70),
  inference(subsumption_resolution,[],[f4338,f267])).
fof(f4338,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl70),
  inference(superposition,[],[f2350,f2328])).
fof(f4354,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15)) ) | $spl70),
  inference(subsumption_resolution,[],[f4337,f267])).
fof(f4337,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) ) | $spl70),
  inference(superposition,[],[f2350,f2328])).
fof(f4322,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X69) != s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) ) | $spl70),
  inference(resolution,[],[f2350,f307])).
fof(f4314,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl70),
  inference(resolution,[],[f2350,f224])).
fof(f4311,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl70),
  inference(resolution,[],[f2350,f223])).
fof(f4310,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl70),
  inference(resolution,[],[f2350,f222])).
fof(f4309,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl70),
  inference(resolution,[],[f2350,f220])).
fof(f4307,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl70),
  inference(resolution,[],[f2350,f218])).
fof(f4184,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X29)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.ffi.ffi_state'(X26),'const.decSem.state_ffi_1'(s('type.decSem.state'(X26),X27))) != s('type.ffi.ffi_state'(X26),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X26),X28))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X26),X27))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X26),X28))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) )),
  inference(subsumption_resolution,[],[f4086,f267])).
fof(f4086,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X29)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.ffi.ffi_state'(X26),'const.decSem.state_ffi_1'(s('type.decSem.state'(X26),X27))) != s('type.ffi.ffi_state'(X26),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X26),X28))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X26),X27))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X26),X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) )),
  inference(superposition,[],[f278,f2328])).
fof(f4182,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X24)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.decSem.state_ffi_1'(s('type.decSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X21),X23))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X21),X22))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X21),X23))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) )),
  inference(subsumption_resolution,[],[f4085,f267])).
fof(f4085,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X24)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.ffi.ffi_state'(X21),'const.decSem.state_ffi_1'(s('type.decSem.state'(X21),X22))) != s('type.ffi.ffi_state'(X21),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X21),X23))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X21),X23)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X21),X22))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X21),X23))) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f278,f2328])).
fof(f4170,plain,(
  ( ! [X6,X4,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7))))) | ~p(s(bool,X9)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7))))))) ) | $spl94),
  inference(cnf_transformation,[],[f4170_D])).
fof(f4170_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7))))) | ~p(s(bool,X9)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7))))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4070,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84),s('type.decSem.state'(X85),X86),s('type.exhSem.state'(X85),X87)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X85),X87)))))) | s('type.ffi.ffi_state'(X85),'const.decSem.state_ffi_1'(s('type.decSem.state'(X85),X86))) != s('type.ffi.ffi_state'(X85),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X85),X87))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X85),X86))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X85),X87))) | p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X85),X87)))))) )),
  inference(resolution,[],[f278,f307])).
fof(f4063,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53),s('type.decSem.state'(X54),X55),s('type.exhSem.state'(X54),X56)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X54),X55))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X54),X56)))))) | s('type.ffi.ffi_state'(X54),'const.decSem.state_ffi_1'(s('type.decSem.state'(X54),X55))) != s('type.ffi.ffi_state'(X54),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X54),X56))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X54),X55))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X54),X56))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X54),X55))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X54),X56))))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f278,f224])).
fof(f4061,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39),s('type.decSem.state'(X40),X41),s('type.exhSem.state'(X40),X42)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X40),X41))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X40),X42)))))) | s('type.ffi.ffi_state'(X40),'const.decSem.state_ffi_1'(s('type.decSem.state'(X40),X41))) != s('type.ffi.ffi_state'(X40),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X40),X42))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X40),X41))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X40),X42))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X40),X41))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X40),X42)))))) )),
  inference(resolution,[],[f278,f223])).
fof(f4059,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.ffi.ffi_state'(X26),'const.decSem.state_ffi_1'(s('type.decSem.state'(X26),X27))) != s('type.ffi.ffi_state'(X26),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X26),X28))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X26),X27))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X26),X28))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X26),X28)))))) )),
  inference(resolution,[],[f278,f222])).
fof(f4058,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18),s('type.decSem.state'(X19),X20),s('type.exhSem.state'(X19),X21)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X19),X21)))))) | s('type.ffi.ffi_state'(X19),'const.decSem.state_ffi_1'(s('type.decSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X19),X21))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X19),X20))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X19),X21))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X19),X20))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X19),X21))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f278,f220])).
fof(f4056,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7)))))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7)))))) )),
  inference(resolution,[],[f278,f218])).
fof(f4169,plain,(
  ( ! [X6,X4,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7))))) | ~p(s(bool,X9)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7)))))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7)))) ) | $spl92),
  inference(cnf_transformation,[],[f4169_D])).
fof(f4169_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7))))) | ~p(s(bool,X9)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7)))))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4049,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84),s('type.decSem.state'(X85),X86),s('type.exhSem.state'(X85),X87)))) | s('type.ffi.ffi_state'(X85),'const.decSem.state_ffi_1'(s('type.decSem.state'(X85),X86))) != s('type.ffi.ffi_state'(X85),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X85),X87))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X85),X87)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X85),X86))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X85),X87))) | p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X85),X87)))))) )),
  inference(resolution,[],[f278,f307])).
fof(f4042,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53),s('type.decSem.state'(X54),X55),s('type.exhSem.state'(X54),X56)))) | s('type.ffi.ffi_state'(X54),'const.decSem.state_ffi_1'(s('type.decSem.state'(X54),X55))) != s('type.ffi.ffi_state'(X54),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X54),X56))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X54),X55))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X54),X56)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X54),X55))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X54),X56))) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X54),X55))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X54),X56))))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f278,f224])).
fof(f4040,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39),s('type.decSem.state'(X40),X41),s('type.exhSem.state'(X40),X42)))) | s('type.ffi.ffi_state'(X40),'const.decSem.state_ffi_1'(s('type.decSem.state'(X40),X41))) != s('type.ffi.ffi_state'(X40),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X40),X42))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X40),X41))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X40),X42)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X40),X41))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X40),X42))) | p(s(bool,X43)) | ~p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X40),X41))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X40),X42)))))) )),
  inference(resolution,[],[f278,f223])).
fof(f4038,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) | s('type.ffi.ffi_state'(X26),'const.decSem.state_ffi_1'(s('type.decSem.state'(X26),X27))) != s('type.ffi.ffi_state'(X26),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X26),X28))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X26),X28)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X26),X27))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X26),X28))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) )),
  inference(resolution,[],[f278,f222])).
fof(f4037,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18),s('type.decSem.state'(X19),X20),s('type.exhSem.state'(X19),X21)))) | s('type.ffi.ffi_state'(X19),'const.decSem.state_ffi_1'(s('type.decSem.state'(X19),X20))) != s('type.ffi.ffi_state'(X19),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X19),X21))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X19),X20))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X19),X21)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X19),X20))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X19),X21))) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X19),X20))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X19),X21))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f278,f220])).
fof(f4035,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.decSem.state'(X5),X6),s('type.exhSem.state'(X5),X7)))) | s('type.ffi.ffi_state'(X5),'const.decSem.state_ffi_1'(s('type.decSem.state'(X5),X6))) != s('type.ffi.ffi_state'(X5),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X5),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X5),X7)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X5),X6))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X5),X7))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X5),X6))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X5),X7)))))) )),
  inference(resolution,[],[f278,f218])).
fof(f4168,plain,(
  ( ! [X21,X19,X17,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X17),X19)))))) | s(bool,X21) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16),s('type.decSem.state'(X17),X18),s('type.exhSem.state'(X17),X19))) | p(s(bool,X21)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X17),X18))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X17),X19))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X17),X19)))))) | s('type.ffi.ffi_state'(X17),'const.decSem.state_ffi_1'(s('type.decSem.state'(X17),X18))) != s('type.ffi.ffi_state'(X17),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X17),X19)))) ) | $spl90),
  inference(cnf_transformation,[],[f4168_D])).
fof(f4168_D,plain,(
  ( ! [X21,X19,X17,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X17),X19)))))) | s(bool,X21) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16),s('type.decSem.state'(X17),X18),s('type.exhSem.state'(X17),X19))) | p(s(bool,X21)) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X17),X18))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X17),X19))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X17),X19)))))) | s('type.ffi.ffi_state'(X17),'const.decSem.state_ffi_1'(s('type.decSem.state'(X17),X18))) != s('type.ffi.ffi_state'(X17),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X17),X19)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4032,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X100))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X101),X102))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X101),X103)))))) | s('type.ffi.ffi_state'(X101),'const.decSem.state_ffi_1'(s('type.decSem.state'(X101),X102))) != s('type.ffi.ffi_state'(X101),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X101),X103))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X100))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X101),X102))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X101),X103)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X101),X102))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X101),X103))) | p(s(bool,X104)) | p(s(bool,X105)) | s(bool,X105) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X100),s('type.decSem.state'(X101),X102),s('type.exhSem.state'(X101),X103)))) )),
  inference(resolution,[],[f278,f307])).
fof(f4024,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X66),X67))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X66),X68)))))) | s('type.ffi.ffi_state'(X66),'const.decSem.state_ffi_1'(s('type.decSem.state'(X66),X67))) != s('type.ffi.ffi_state'(X66),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X66),X68))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X66),X67))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X66),X68)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X66),X67))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X66),X68))) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X65),s('type.decSem.state'(X66),X67),s('type.exhSem.state'(X66),X68))) | s(bool,X70) != s(bool,X71)) )),
  inference(resolution,[],[f278,f224])).
fof(f4021,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X44))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X45),X46))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X45),X47)))))) | s('type.ffi.ffi_state'(X45),'const.decSem.state_ffi_1'(s('type.decSem.state'(X45),X46))) != s('type.ffi.ffi_state'(X45),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X45),X47))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X44))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X45),X46))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X45),X47)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X45),X46))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X45),X47))) | p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X50) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X44),s('type.decSem.state'(X45),X46),s('type.exhSem.state'(X45),X47)))) )),
  inference(resolution,[],[f278,f223])).
fof(f4020,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X37))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X38),X39))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X38),X40)))))) | s('type.ffi.ffi_state'(X38),'const.decSem.state_ffi_1'(s('type.decSem.state'(X38),X39))) != s('type.ffi.ffi_state'(X38),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X38),X40))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X37))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X38),X39))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X38),X40)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X38),X39))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X38),X40))) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X42) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X37),s('type.decSem.state'(X38),X39),s('type.exhSem.state'(X38),X40)))) )),
  inference(resolution,[],[f278,f222])).
fof(f4019,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X31),X32))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X31),X33)))))) | s('type.ffi.ffi_state'(X31),'const.decSem.state_ffi_1'(s('type.decSem.state'(X31),X32))) != s('type.ffi.ffi_state'(X31),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X31),X33))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X31),X32))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X31),X33)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X31),X32))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X31),X33))) | p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30),s('type.decSem.state'(X31),X32),s('type.exhSem.state'(X31),X33))) | s(bool,X35) != s(bool,X36)) )),
  inference(resolution,[],[f278,f220])).
fof(f4017,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X17),X19)))))) | s('type.ffi.ffi_state'(X17),'const.decSem.state_ffi_1'(s('type.decSem.state'(X17),X18))) != s('type.ffi.ffi_state'(X17),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X17),X19))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X17),X18))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X17),X19)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X17),X18))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X17),X19))) | p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X16),s('type.decSem.state'(X17),X18),s('type.exhSem.state'(X17),X19)))) )),
  inference(resolution,[],[f278,f218])).
fof(f386,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | $spl15),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f390,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | $spl19),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f4167,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(trivial_inequality_removal,[],[f4166])).
fof(f4166,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f4165,f310])).
fof(f310,plain,(
  s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1)))),
  inference(resolution,[],[f205,f274])).
fof(f4165,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f4164,f215])).
fof(f215,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6,X11] : ! [X24,X25] : s(X6,'const.combin.K_2'(s(X6,X24),s(X11,X25))) = s(X6,happ(s(fun(X11,X6),'const.combin.K_1'(s(X6,X24))),s(X11,X25)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','ARITY_1_2_const.combin.K')).
fof(f4164,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f4163,f215])).
fof(f4163,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f4162,f211])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))))) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))))) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))))) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X20] : ! [X21,X22] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X22),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X20),X21))))) = s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X20),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.decSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.decSem.decSem_28_prove_ltb')).
fof(f4162,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) != s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f4161,f207])).
fof(f207,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X1))))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X1))))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X1))))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X20] : ! [X21,X22] : s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X22),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X20),X21))))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X20),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.exhSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.exhSem.exhSem_94_prove_ltb')).
fof(f4161,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4160,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X20] : ! [X21,X22] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X20),X21))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X20),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.decSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.decSem.decSem_|49|9_prove_ltb')).
fof(f4160,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4159,f208])).
fof(f208,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X20] : ! [X21,X22] : s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X20),X21))) = s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X20),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.exhSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.exhSem.exhSem_85_prove_ltb')).
fof(f4159,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(subsumption_resolution,[],[f4158,f312])).
fof(f312,plain,(
  s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) = s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1)))),
  inference(resolution,[],[f205,f276])).
fof(f4158,plain,(
  s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4157,f214])).
fof(f214,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X20] : ! [X21,X22] : s('type.ffi.ffi_state'(X20),'const.decSem.state_ffi_1'(s('type.decSem.state'(X20),X21))) = s('type.ffi.ffi_state'(X20),'const.decSem.state_ffi_1'(s('type.decSem.state'(X20),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.decSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.decSem.decSem_22_prove_ltb')).
fof(f4157,plain,(
  s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) != s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4156,f210])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X20] : ! [X21,X22] : s('type.ffi.ffi_state'(X20),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X20),X21))) = s('type.ffi.ffi_state'(X20),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X20),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.exhSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.exhSem.exhSem_88_prove_ltb')).
fof(f4156,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4155,f213])).
fof(f213,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.decSem.state'(X0),X1)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X20] : ! [X21,X22] : s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X20),X21))) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X20),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.decSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.decSem.decSem_25_prove_ltb')).
fof(f4155,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f4016,f209])).
fof(f209,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.exhSem.state'(X0),X1)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X20] : ! [X21,X22] : s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X20),X21))) = s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X20),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X22),s('type.exhSem.state'(X20),X21)))))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.exhSem.exhSem_9|49|_prove_ltb')).
fof(f4016,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(resolution,[],[f278,f206])).
fof(f3999,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) != s(bool,X2) | p(s(bool,X1))) ) | $spl88),
  inference(cnf_transformation,[],[f3999_D])).
fof(f3999_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X0) != s(bool,X2) | p(s(bool,X1))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3960,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X66) != s(bool,X67) | p(s(bool,X67)) | s(bool,X66) != s(bool,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X68) != s(bool,X70)) )),
  inference(resolution,[],[f2388,f307])).
fof(f3953,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f2388,f224])).
fof(f3951,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f2388,f223])).
fof(f3949,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f2388,f222])).
fof(f3948,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f2388,f220])).
fof(f3946,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f2388,f218])).
fof(f3937,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,X18) | p(s(bool,X18))) ) | $spl86),
  inference(cnf_transformation,[],[f3937_D])).
fof(f3937_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,X18) | p(s(bool,X18))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3938,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X22)) | p(s(bool,X19)) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | $spl66),
  inference(subsumption_resolution,[],[f3906,f267])).
fof(f3906,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X22)) | p(s(bool,X19)) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl66),
  inference(superposition,[],[f2249,f2328])).
fof(f3936,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | p(s(bool,X15)) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) ) | $spl66),
  inference(subsumption_resolution,[],[f3905,f267])).
fof(f3905,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | p(s(bool,X15)) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) ) | $spl66),
  inference(superposition,[],[f2249,f2328])).
fof(f2350,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl70),
  inference(cnf_transformation,[],[f2350_D])).
fof(f2350_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3928,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl66),
  inference(subsumption_resolution,[],[f3891,f267])).
fof(f3891,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl66),
  inference(superposition,[],[f2249,f2328])).
fof(f3927,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X33)) | p(s(bool,X32)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl66),
  inference(subsumption_resolution,[],[f3890,f267])).
fof(f3890,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X33)) | p(s(bool,X32)) | s(bool,'T_0') = s(bool,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl66),
  inference(superposition,[],[f2249,f2328])).
fof(f3873,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X29)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) )),
  inference(subsumption_resolution,[],[f3849,f267])).
fof(f3849,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X29)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.decSem.state'(X26),X27),s('type.exhSem.state'(X26),X28)))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X26),X27))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X26),X28)))))) )),
  inference(superposition,[],[f277,f2328])).
fof(f3871,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X24)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23)))))) )),
  inference(subsumption_resolution,[],[f3848,f267])).
fof(f3848,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X24)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X20))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X21),X22))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X21),X23))))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f277,f2328])).
fof(f3866,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3)))))) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3))) | ~p(s(bool,X5))) ) | $spl82),
  inference(cnf_transformation,[],[f3866_D])).
fof(f3866_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3)))))) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3))) | ~p(s(bool,X5))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3826,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X80))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X81),X82))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X81),X83)))))) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X80),s('type.decSem.state'(X81),X82),s('type.exhSem.state'(X81),X83)))) )),
  inference(resolution,[],[f277,f307])).
fof(f3819,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X50),X51))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X50),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49),s('type.decSem.state'(X50),X51),s('type.exhSem.state'(X50),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f277,f224])).
fof(f3817,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X35))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X36),X37))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X36),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X35),s('type.decSem.state'(X36),X37),s('type.exhSem.state'(X36),X38)))) )),
  inference(resolution,[],[f277,f223])).
fof(f3815,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X22),X23))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X22),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.decSem.state'(X22),X23),s('type.exhSem.state'(X22),X24)))) )),
  inference(resolution,[],[f277,f222])).
fof(f3814,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X15),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14),s('type.decSem.state'(X15),X16),s('type.exhSem.state'(X15),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f277,f220])).
fof(f3812,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) )),
  inference(resolution,[],[f277,f218])).
fof(f3865,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3))))) | p(s(bool,X5))) ) | $spl80),
  inference(cnf_transformation,[],[f3865_D])).
fof(f3865_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3))))) | p(s(bool,X5))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3811,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84),s('type.decSem.state'(X85),X86),s('type.exhSem.state'(X85),X87)))) | p(s(bool,X88)) | p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X85),X87)))))) )),
  inference(resolution,[],[f277,f307])).
fof(f3803,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49),s('type.decSem.state'(X50),X51),s('type.exhSem.state'(X50),X52)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X50),X51))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X50),X52))))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f277,f224])).
fof(f3800,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X28),s('type.decSem.state'(X29),X30),s('type.exhSem.state'(X29),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X28))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X29),X30))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X29),X31)))))) )),
  inference(resolution,[],[f277,f223])).
fof(f3799,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.decSem.state'(X22),X23),s('type.exhSem.state'(X22),X24)))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X22),X23))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X22),X24)))))) )),
  inference(resolution,[],[f277,f222])).
fof(f3798,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14),s('type.decSem.state'(X15),X16),s('type.exhSem.state'(X15),X17)))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X15),X16))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X15),X17))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f277,f220])).
fof(f3796,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(X1),X3)))))) )),
  inference(resolution,[],[f277,f218])).
fof(f3785,plain,(
  ( ! [X253,X251,X252,X250] : (s(bool,'T_0') != s(bool,X251) | s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X253)) )),
  inference(subsumption_resolution,[],[f3402,f1729])).
fof(f3402,plain,(
  ( ! [X253,X251,X252,X250] : (s(bool,'T_0') != s(bool,X251) | s(bool,'F_0') = s(bool,X251) | s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X253)) )),
  inference(superposition,[],[f906,f2328])).
fof(f3780,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f3394,f308])).
fof(f3394,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f307,f2328])).
fof(f3776,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f3371,f308])).
fof(f3371,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f224,f2328])).
fof(f3771,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f3363,f308])).
fof(f3363,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f223,f2328])).
fof(f3768,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f3358,f308])).
fof(f3358,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f222,f2328])).
fof(f3763,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3351,f308])).
fof(f3351,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f220,f2328])).
fof(f3758,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3344,f308])).
fof(f3344,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f218,f2328])).
fof(f3745,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X255) != s(bool,X256) | p(s(bool,X257)) | p(s(bool,X256)) | s(bool,X255) = s(bool,X258) | s(bool,'T_0') = s(bool,X258)) )),
  inference(subsumption_resolution,[],[f3317,f308])).
fof(f3317,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X255) != s(bool,X256) | p(s(bool,X257)) | ~p(s(bool,'T_0')) | p(s(bool,X256)) | s(bool,X255) = s(bool,X258) | s(bool,'T_0') = s(bool,X258)) )),
  inference(superposition,[],[f307,f2328])).
fof(f3742,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f3298,f308])).
fof(f3298,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X202)) | ~p(s(bool,X203)) | ~p(s(bool,'T_0')) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f224,f2328])).
fof(f3739,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,X156) != s(bool,X157) | p(s(bool,X158)) | ~p(s(bool,X159)) | p(s(bool,X157)) | s(bool,X156) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(subsumption_resolution,[],[f3289,f308])).
fof(f3289,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,X156) != s(bool,X157) | p(s(bool,X158)) | ~p(s(bool,'T_0')) | ~p(s(bool,X159)) | p(s(bool,X157)) | s(bool,X156) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f223,f2328])).
fof(f3736,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,X121) != s(bool,X122) | p(s(bool,X123)) | p(s(bool,X122)) | s(bool,X123) != s(bool,X124) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f3282,f308])).
fof(f3282,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,X121) != s(bool,X122) | p(s(bool,X123)) | ~p(s(bool,'T_0')) | p(s(bool,X122)) | s(bool,X123) != s(bool,X124) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f222,f2328])).
fof(f3734,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f3277,f308])).
fof(f3277,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | ~p(s(bool,'T_0')) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f220,f2328])).
fof(f3730,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3269,f308])).
fof(f3269,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,'T_0')) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f218,f2328])).
fof(f3718,plain,(
  ( ! [X253,X251,X252,X250] : (s(bool,'T_0') != s(bool,X251) | s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X253)) )),
  inference(subsumption_resolution,[],[f3203,f1729])).
fof(f3203,plain,(
  ( ! [X253,X251,X252,X250] : (s(bool,'T_0') != s(bool,X251) | s(bool,'F_0') = s(bool,X251) | s(bool,X250) != s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,X250) = s(bool,X253) | s(bool,'T_0') = s(bool,X253)) )),
  inference(superposition,[],[f906,f2328])).
fof(f3713,plain,(
  ( ! [X227,X225,X228,X226] : (p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f3195,f308])).
fof(f3195,plain,(
  ( ! [X227,X225,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,X226)) | p(s(bool,X227)) | s(bool,X225) != s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f307,f2328])).
fof(f3709,plain,(
  ( ! [X167,X165,X169,X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(subsumption_resolution,[],[f3172,f308])).
fof(f3172,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X165) != s(bool,X166) | s(bool,X167) != s(bool,X168) | s(bool,X165) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f224,f2328])).
fof(f3704,plain,(
  ( ! [X127,X125,X128,X126,X129] : (p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f3164,f308])).
fof(f3164,plain,(
  ( ! [X127,X125,X128,X126,X129] : (~p(s(bool,'T_0')) | p(s(bool,X126)) | ~p(s(bool,X127)) | p(s(bool,X128)) | s(bool,X125) != s(bool,X128) | s(bool,X125) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f223,f2328])).
fof(f3701,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f3159,f308])).
fof(f3159,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f222,f2328])).
fof(f3696,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3152,f308])).
fof(f3152,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X65) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f220,f2328])).
fof(f3691,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3145,f308])).
fof(f3145,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,X30) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f218,f2328])).
fof(f3678,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X255) != s(bool,X256) | p(s(bool,X257)) | p(s(bool,X256)) | s(bool,X255) = s(bool,X258) | s(bool,'T_0') = s(bool,X258)) )),
  inference(subsumption_resolution,[],[f3118,f308])).
fof(f3118,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X255) != s(bool,X256) | p(s(bool,X257)) | ~p(s(bool,'T_0')) | p(s(bool,X256)) | s(bool,X255) = s(bool,X258) | s(bool,'T_0') = s(bool,X258)) )),
  inference(superposition,[],[f307,f2328])).
fof(f3675,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f3099,f308])).
fof(f3099,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(bool,X201) != s(bool,X202) | p(s(bool,X202)) | ~p(s(bool,X203)) | ~p(s(bool,'T_0')) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f224,f2328])).
fof(f3672,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,X156) != s(bool,X157) | p(s(bool,X158)) | ~p(s(bool,X159)) | p(s(bool,X157)) | s(bool,X156) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(subsumption_resolution,[],[f3090,f308])).
fof(f3090,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,X156) != s(bool,X157) | p(s(bool,X158)) | ~p(s(bool,'T_0')) | ~p(s(bool,X159)) | p(s(bool,X157)) | s(bool,X156) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f223,f2328])).
fof(f3669,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,X121) != s(bool,X122) | p(s(bool,X123)) | p(s(bool,X122)) | s(bool,X123) != s(bool,X124) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f3083,f308])).
fof(f3083,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,X121) != s(bool,X122) | p(s(bool,X123)) | ~p(s(bool,'T_0')) | p(s(bool,X122)) | s(bool,X123) != s(bool,X124) | s(bool,X121) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f222,f2328])).
fof(f3667,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f3078,f308])).
fof(f3078,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X96) != s(bool,X97) | p(s(bool,X97)) | ~p(s(bool,X98)) | ~p(s(bool,'T_0')) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f220,f2328])).
fof(f3663,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3070,f308])).
fof(f3070,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,'T_0')) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f218,f2328])).
fof(f3662,plain,(
  ( ! [X368,X370,X367,X369,X371] : (p(s(bool,X371)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370)))))) )),
  inference(subsumption_resolution,[],[f3048,f267])).
fof(f3048,plain,(
  ( ! [X368,X370,X367,X369,X371] : (p(s(bool,X371)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,X371) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370)))))) )),
  inference(superposition,[],[f275,f2328])).
fof(f3626,plain,(
  ( ! [X368,X370,X367,X369,X371] : (p(s(bool,X371)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370)))))) )),
  inference(subsumption_resolution,[],[f2932,f267])).
fof(f2932,plain,(
  ( ! [X368,X370,X367,X369,X371] : (p(s(bool,X371)) | ~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367),s('type.decSem.state'(X368),X369),s('type.exhSem.state'(X368),X370)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X367))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X368),X369))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X368),X370))))) | s(bool,'T_0') = s(bool,X371)) )),
  inference(superposition,[],[f275,f2328])).
fof(f2680,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3)))))) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3))) | ~p(s(bool,X5))) ) | $spl74),
  inference(cnf_transformation,[],[f2680_D])).
fof(f2680_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3)))))) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3))) | ~p(s(bool,X5))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2648,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X80))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X81),X82))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X81),X83)))))) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X80),s('type.decSem.state'(X81),X82),s('type.exhSem.state'(X81),X83)))) )),
  inference(resolution,[],[f275,f307])).
fof(f2641,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X50),X51))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X50),X52)))))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49),s('type.decSem.state'(X50),X51),s('type.exhSem.state'(X50),X52))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f275,f224])).
fof(f2639,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X35))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X36),X37))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X36),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X35),s('type.decSem.state'(X36),X37),s('type.exhSem.state'(X36),X38)))) )),
  inference(resolution,[],[f275,f223])).
fof(f2637,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X22),X23))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X22),X24)))))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.decSem.state'(X22),X23),s('type.exhSem.state'(X22),X24)))) )),
  inference(resolution,[],[f275,f222])).
fof(f2636,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X15),X16))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X15),X17)))))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14),s('type.decSem.state'(X15),X16),s('type.exhSem.state'(X15),X17))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f275,f220])).
fof(f2634,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3)))))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) )),
  inference(resolution,[],[f275,f218])).
fof(f2679,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3))))) | p(s(bool,X5))) ) | $spl72),
  inference(cnf_transformation,[],[f2679_D])).
fof(f2679_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3))))) | p(s(bool,X5))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2633,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84),s('type.decSem.state'(X85),X86),s('type.exhSem.state'(X85),X87)))) | p(s(bool,X88)) | p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X84))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X85),X86))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X85),X87)))))) )),
  inference(resolution,[],[f275,f307])).
fof(f2625,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49),s('type.decSem.state'(X50),X51),s('type.exhSem.state'(X50),X52)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X49))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X50),X51))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X50),X52))))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f275,f224])).
fof(f2622,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X28),s('type.decSem.state'(X29),X30),s('type.exhSem.state'(X29),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X28))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X29),X30))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X29),X31)))))) )),
  inference(resolution,[],[f275,f223])).
fof(f2621,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.decSem.state'(X22),X23),s('type.exhSem.state'(X22),X24)))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X22),X23))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X22),X24)))))) )),
  inference(resolution,[],[f275,f222])).
fof(f2620,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14),s('type.decSem.state'(X15),X16),s('type.exhSem.state'(X15),X17)))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X15),X16))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X15),X17))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f275,f220])).
fof(f2618,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.decSem.state'(X1),X2),s('type.exhSem.state'(X1),X3)))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(X1),X2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(X1),X3)))))) )),
  inference(resolution,[],[f275,f218])).
fof(f393,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) ) | $spl22),
  inference(cnf_transformation,[],[f393_D])).
fof(f393_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2431,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f307])).
fof(f2424,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f206,f224])).
fof(f2422,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f223])).
fof(f2420,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f222])).
fof(f2419,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f206,f220])).
fof(f2417,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f218])).
fof(f755,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2369,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) ) | $spl62),
  inference(resolution,[],[f1999,f307])).
fof(f2361,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl62),
  inference(resolution,[],[f1999,f224])).
fof(f2358,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl62),
  inference(resolution,[],[f1999,f223])).
fof(f2357,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl62),
  inference(resolution,[],[f1999,f222])).
fof(f2356,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl62),
  inference(resolution,[],[f1999,f220])).
fof(f2354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl62),
  inference(resolution,[],[f1999,f218])).
fof(f2330,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X54) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f1797,f307])).
fof(f2322,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1797,f224])).
fof(f2319,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1797,f223])).
fof(f2318,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1797,f222])).
fof(f2317,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1797,f220])).
fof(f2315,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1797,f218])).
fof(f2268,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s('type.ffi.ffi_state'(X76),'const.decSem.state_ffi_1'(s('type.decSem.state'(X76),X77))) = s('type.ffi.ffi_state'(X76),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X76),X78))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X81),s('type.decSem.state'(X76),X77),s('type.exhSem.state'(X76),X78)))) )),
  inference(resolution,[],[f276,f307])).
fof(f2261,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s('type.ffi.ffi_state'(X48),'const.decSem.state_ffi_1'(s('type.decSem.state'(X48),X49))) = s('type.ffi.ffi_state'(X48),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X48),X50))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53),s('type.decSem.state'(X48),X49),s('type.exhSem.state'(X48),X50))) | s(bool,X51) != s(bool,X54)) )),
  inference(resolution,[],[f276,f224])).
fof(f2259,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s('type.ffi.ffi_state'(X35),'const.decSem.state_ffi_1'(s('type.decSem.state'(X35),X36))) = s('type.ffi.ffi_state'(X35),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X35),X37))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X41),s('type.decSem.state'(X35),X36),s('type.exhSem.state'(X35),X37)))) )),
  inference(resolution,[],[f276,f223])).
fof(f2257,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.ffi.ffi_state'(X21),'const.decSem.state_ffi_1'(s('type.decSem.state'(X21),X22))) = s('type.ffi.ffi_state'(X21),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X21),X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X27),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) )),
  inference(resolution,[],[f276,f222])).
fof(f2256,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.ffi.ffi_state'(X14),'const.decSem.state_ffi_1'(s('type.decSem.state'(X14),X15))) = s('type.ffi.ffi_state'(X14),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X14),X16))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X19),s('type.decSem.state'(X14),X15),s('type.exhSem.state'(X14),X16))) | s(bool,X17) != s(bool,X20)) )),
  inference(resolution,[],[f276,f220])).
fof(f2254,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.decSem.state_ffi_1'(s('type.decSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(X0),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2)))) )),
  inference(resolution,[],[f276,f218])).
fof(f2249,plain,(
  ( ! [X0,X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) = s(bool,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f2249_D])).
fof(f2249_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2219,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | s(bool,X54) = s(bool,X55) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f1796,f307])).
fof(f2211,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1796,f224])).
fof(f2208,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1796,f223])).
fof(f2207,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1796,f222])).
fof(f2206,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1796,f220])).
fof(f2204,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1796,f218])).
fof(f2201,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | s(bool,X54) = s(bool,X55) | p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X55) != s(bool,X57)) )),
  inference(resolution,[],[f1796,f307])).
fof(f2193,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) = s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1796,f224])).
fof(f2190,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X20)) | s(bool,X20) = s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1796,f223])).
fof(f2189,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,X15) = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1796,f222])).
fof(f2188,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X10)) | s(bool,X10) = s(bool,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1796,f220])).
fof(f2186,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1796,f218])).
fof(f2179,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f2176])).
fof(f2176,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl28),
  inference(superposition,[],[f536,f329])).
fof(f2178,plain,(
  $false | $spl28),
  inference(equality_resolution,[],[f536])).
fof(f2155,plain,(
  $false | $spl26),
  inference(equality_resolution,[],[f496])).
fof(f2135,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl32),
  inference(subsumption_resolution,[],[f2111,f263])).
fof(f2111,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') = s(bool,X8) | ~p(s(bool,X8))) ) | $spl32),
  inference(superposition,[],[f931,f268])).
fof(f2134,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl32),
  inference(subsumption_resolution,[],[f2110,f263])).
fof(f2110,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X6))) ) | $spl32),
  inference(superposition,[],[f931,f270])).
fof(f2133,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl32),
  inference(subsumption_resolution,[],[f2109,f931])).
fof(f2109,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl32),
  inference(superposition,[],[f931,f306])).
fof(f2132,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | $spl32),
  inference(duplicate_literal_removal,[],[f2108])).
fof(f2108,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl32),
  inference(superposition,[],[f931,f306])).
fof(f2064,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X76),X77))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X76),X78))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X81),s('type.decSem.state'(X76),X77),s('type.exhSem.state'(X76),X78)))) )),
  inference(resolution,[],[f274,f307])).
fof(f2057,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X48),X49))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X48),X50))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X53),s('type.decSem.state'(X48),X49),s('type.exhSem.state'(X48),X50))) | s(bool,X51) != s(bool,X54)) )),
  inference(resolution,[],[f274,f224])).
fof(f2055,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X35),X36))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X35),X37))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X41),s('type.decSem.state'(X35),X36),s('type.exhSem.state'(X35),X37)))) )),
  inference(resolution,[],[f274,f223])).
fof(f2053,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X21),X22))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X21),X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X27),s('type.decSem.state'(X21),X22),s('type.exhSem.state'(X21),X23)))) )),
  inference(resolution,[],[f274,f222])).
fof(f2052,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X14),X15))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X14),X16))) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X18) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X19),s('type.decSem.state'(X14),X15),s('type.exhSem.state'(X14),X16))) | s(bool,X17) != s(bool,X20)) )),
  inference(resolution,[],[f274,f220])).
fof(f2050,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(X0),X1))) = s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(X0),X2))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.decSem.state'(X0),X1),s('type.exhSem.state'(X0),X2)))) )),
  inference(resolution,[],[f274,f218])).
fof(f1999,plain,(
  ( ! [X2,X0] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl62),
  inference(cnf_transformation,[],[f1999_D])).
fof(f1999_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2015,plain,(
  ( ! [X194,X195,X193] : (s(bool,X193) != s(bool,X194) | s(bool,'F_0') = s(bool,X194) | s(bool,'T_0') = s(bool,X195) | p(s(bool,X193))) )),
  inference(subsumption_resolution,[],[f1889,f306])).
fof(f1889,plain,(
  ( ! [X194,X195,X193] : (s(bool,X193) != s(bool,X194) | s(bool,'F_0') = s(bool,X194) | s(bool,'F_0') != s(bool,X195) | s(bool,'T_0') = s(bool,X195) | p(s(bool,X193))) )),
  inference(superposition,[],[f906,f264])).
fof(f1794,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') = s(bool,X39) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f264,f307])).
fof(f1786,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f264,f224])).
fof(f1783,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f264,f223])).
fof(f1782,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f264,f222])).
fof(f1781,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f264,f220])).
fof(f1779,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f264,f218])).
fof(f1356,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5)) ) | $spl48),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1731,plain,(
  ( ! [X39,X37,X38] : (s(bool,'F_0') != s(bool,X37) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X37) != s(bool,X39)) )),
  inference(resolution,[],[f263,f307])).
fof(f1725,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') != s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f263,f224])).
fof(f1723,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f263,f223])).
fof(f1721,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f263,f222])).
fof(f1720,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f263,f220])).
fof(f1718,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f263,f218])).
fof(f1717,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) | $spl60),
  inference(cnf_transformation,[],[f1717_D])).
fof(f1717_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1710,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f271,f307])).
fof(f1704,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f271,f224])).
fof(f1702,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f271,f223])).
fof(f1700,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f271,f222])).
fof(f1699,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f271,f220])).
fof(f1697,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f271,f218])).
fof(f1695,plain,(
  ~p(s(bool,'F_0')) | $spl59),
  inference(cnf_transformation,[],[f1695_D])).
fof(f1695_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f234,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.IMP_CLAUSES|split|4')).
fof(f1696,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f247,f308])).
fof(f247,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f187])).
fof(f253,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) & p(s(bool,X12))) <=> p(s(bool,'F_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.AND_CLAUSES|split|2')).
fof(f255,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : ((p(s(bool,X12)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.AND_CLAUSES|split|3')).
fof(f1646,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f1439,f273])).
fof(f273,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.state_rel_with_clock',aHL_TRUTH)).
fof(f1439,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(backward_demodulation,[],[f1355,f206])).
fof(f1639,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f1431,f273])).
fof(f1431,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(backward_demodulation,[],[f1355,f271])).
fof(f1618,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f1614])).
fof(f1614,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl46),
  inference(superposition,[],[f992,f1355])).
fof(f1407,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f307])).
fof(f1397,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f313,f224])).
fof(f1394,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f223])).
fof(f1393,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f222])).
fof(f1392,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f313,f220])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f313,f218])).
fof(f1368,plain,(
  ( ! [X117] : (s(bool,'F_0') = s(bool,X117) | s(bool,X117) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) | $spl54),
  inference(cnf_transformation,[],[f1368_D])).
fof(f1368_D,plain,(
  ( ! [X117] : (s(bool,'F_0') = s(bool,X117) | s(bool,X117) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f933,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl35),
  inference(cnf_transformation,[],[f933_D])).
fof(f933_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1072,plain,(
  ( ! [X117] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X117) | s(bool,X117) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(superposition,[],[f824,f329])).
fof(f932,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl34),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1367,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) != s(bool,X26)) ) | $spl52),
  inference(cnf_transformation,[],[f1367_D])).
fof(f1367_D,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) != s(bool,X26)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1265,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | ~p(s(bool,X21))) )),
  inference(superposition,[],[f906,f270])).
fof(f1266,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | ~p(s(bool,X24))) )),
  inference(superposition,[],[f906,f268])).
fof(f1366,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | $spl44),
  inference(subsumption_resolution,[],[f1294,f1278])).
fof(f1294,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(superposition,[],[f1278,f306])).
fof(f1313,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X9) | s(bool,X8) != s(bool,X9)) ) | $spl44),
  inference(subsumption_resolution,[],[f1297,f806])).
fof(f1297,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X8))) ) | $spl44),
  inference(superposition,[],[f1278,f266])).
fof(f266,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.EQ_CLAUSES|split|3')).
fof(f1314,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl44),
  inference(subsumption_resolution,[],[f1298,f806])).
fof(f1298,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11) | p(s(bool,X10))) ) | $spl44),
  inference(superposition,[],[f1278,f264])).
fof(f1360,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X17))) ) | $spl35),
  inference(subsumption_resolution,[],[f1337,f268])).
fof(f1337,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X17)) | p(s(bool,X15))) ) | $spl35),
  inference(superposition,[],[f1284,f264])).
fof(f1284,plain,(
  ( ! [X23,X21,X22] : (s(bool,X22) != s(bool,X23) | s(bool,X21) != s(bool,X22) | s(bool,'T_0') = s(bool,X23) | ~p(s(bool,X21))) ) | $spl35),
  inference(subsumption_resolution,[],[f1265,f933])).
fof(f1359,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X14))) ) | $spl35),
  inference(subsumption_resolution,[],[f1336,f268])).
fof(f1336,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X14)) | p(s(bool,X12))) ) | $spl35),
  inference(superposition,[],[f1284,f266])).
fof(f1348,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X5))) ) | $spl35),
  inference(duplicate_literal_removal,[],[f1333])).
fof(f1333,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X3)) ) | $spl35),
  inference(superposition,[],[f1284,f306])).
fof(f1278,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1281,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f1261,f263])).
fof(f1261,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X9) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X9))) )),
  inference(superposition,[],[f906,f268])).
fof(f1280,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1260,f263])).
fof(f1260,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X6) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f906,f270])).
fof(f1274,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1258])).
fof(f1258,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f906,f306])).
fof(f1276,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X37)) ) | $spl40),
  inference(cnf_transformation,[],[f1276_D])).
fof(f1276_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1275,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1255,f306])).
fof(f1255,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f906,f329])).
fof(f1229,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f307])).
fof(f1219,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f311,f224])).
fof(f1216,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f223])).
fof(f1215,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f222])).
fof(f1214,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f311,f220])).
fof(f1212,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f311,f218])).
fof(f1210,plain,(
  $false | $spl21),
  inference(trivial_inequality_removal,[],[f1209])).
fof(f1209,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) | $spl21),
  inference(forward_demodulation,[],[f392,f310])).
fof(f1170,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X75) != s(bool,X76) | s(bool,X75) != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X76) != s(bool,X79)) ) | $spl30),
  inference(resolution,[],[f755,f307])).
fof(f1160,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl30),
  inference(resolution,[],[f755,f224])).
fof(f1157,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl30),
  inference(resolution,[],[f755,f223])).
fof(f1156,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl30),
  inference(resolution,[],[f755,f222])).
fof(f1155,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl30),
  inference(resolution,[],[f755,f220])).
fof(f1153,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl30),
  inference(resolution,[],[f755,f218])).
fof(f1144,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f1133])).
fof(f1133,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f340,f329])).
fof(f1143,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f340])).
fof(f935,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7)) )),
  inference(subsumption_resolution,[],[f918,f796])).
fof(f918,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X7) | s(bool,X6) != s(bool,X7) | ~p(s(bool,X6))) )),
  inference(superposition,[],[f824,f268])).
fof(f934,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f917,f796])).
fof(f917,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f824,f270])).
fof(f930,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f915,f824])).
fof(f915,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f824,f306])).
fof(f828,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'F_0') = s(bool,X58) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) )),
  inference(resolution,[],[f796,f307])).
fof(f818,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f796,f224])).
fof(f815,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f796,f223])).
fof(f814,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f796,f222])).
fof(f813,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f796,f220])).
fof(f811,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f796,f218])).
fof(f725,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'F_0') = s(bool,X58) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X59) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f649,f307])).
fof(f649,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | $spl6),
  inference(subsumption_resolution,[],[f628,f308])).
fof(f628,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f341,f306])).
fof(f715,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl6),
  inference(resolution,[],[f649,f224])).
fof(f712,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f649,f223])).
fof(f711,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl6),
  inference(resolution,[],[f649,f222])).
fof(f710,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl6),
  inference(resolution,[],[f649,f220])).
fof(f708,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl6),
  inference(resolution,[],[f649,f218])).
fof(f600,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f569,f343])).
fof(f569,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f343,f245])).
fof(f245,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f187])).
fof(f599,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f567,f343])).
fof(f567,plain,(
  ( ! [X42] : (p(s(bool,X42))) ) | $spl10),
  inference(resolution,[],[f343,f242])).
fof(f598,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f566,f343])).
fof(f566,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl10),
  inference(resolution,[],[f343,f239])).
fof(f597,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f565,f343])).
fof(f565,plain,(
  ( ! [X40] : (p(s(bool,X40))) ) | $spl10),
  inference(resolution,[],[f343,f236])).
fof(f573,plain,(
  $false | $spl10),
  inference(resolution,[],[f343,f272])).
fof(f272,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.state_rel_with_clock','thm.bool.TRUTH')).
fof(f574,plain,(
  $false | $spl10),
  inference(resolution,[],[f343,f273])).
fof(f575,plain,(
  $false | $spl10),
  inference(resolution,[],[f343,f308])).
fof(f578,plain,(
  $false | $spl10),
  inference(resolution,[],[f343,f205])).
fof(f533,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f507])).
fof(f507,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | ~p(s(bool,X6))) ) | $spl8),
  inference(superposition,[],[f342,f268])).
fof(f534,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f506])).
fof(f506,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X4))) ) | $spl8),
  inference(superposition,[],[f342,f270])).
fof(f537,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl8),
  inference(subsumption_resolution,[],[f505,f342])).
fof(f505,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f342,f306])).
fof(f535,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f504,f342])).
fof(f504,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f342,f306])).
fof(f493,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f490])).
fof(f490,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23) | p(s(bool,X22))) ) | $spl2),
  inference(superposition,[],[f339,f264])).
fof(f494,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f489])).
fof(f489,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21) | p(s(bool,X20))) ) | $spl2),
  inference(superposition,[],[f339,f266])).
fof(f501,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | $spl2),
  inference(subsumption_resolution,[],[f488,f267])).
fof(f488,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl2),
  inference(superposition,[],[f339,f306])).
fof(f471,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,X71) | ~p(s(bool,X72))) ) | $spl2),
  inference(resolution,[],[f339,f234])).
fof(f497,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,X69) | p(s(bool,X70))) ) | $spl2),
  inference(subsumption_resolution,[],[f470,f308])).
fof(f470,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,'T_0'))) ) | $spl2),
  inference(resolution,[],[f339,f247])).
fof(f469,plain,(
  ( ! [X68,X67] : (s(bool,'F_0') != s(bool,X67) | p(s(bool,X68))) ) | $spl2),
  inference(resolution,[],[f339,f253])).
fof(f468,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X66))) ) | $spl2),
  inference(resolution,[],[f339,f255])).
fof(f447,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f422,f344])).
fof(f422,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f344,f234])).
fof(f418,plain,(
  $false | $spl12),
  inference(resolution,[],[f344,f271])).
fof(f427,plain,(
  $false | $spl12),
  inference(resolution,[],[f344,f206])).
fof(f367,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f307])).
fof(f356,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f206,f224])).
fof(f354,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f223])).
fof(f394,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) ) | $spl24),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f341,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f341_D])).
fof(f341_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f353,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f206,f222])).
fof(f352,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f222])).
fof(f351,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f206,f220])).
fof(f350,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f206,f218])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))),s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))) )),
  inference(resolution,[],[f206,f218])).
fof(f388,plain,(
  s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | $spl17),
  inference(cnf_transformation,[],[f388_D])).
fof(f388_D,plain,(
  s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f392,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) | $spl21),
  inference(cnf_transformation,[],[f392_D])).
fof(f392_D,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f384,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f383,f215])).
fof(f383,plain,(
  s('type.num.num','const.combin.K_2'(s('type.num.num',sK4),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f382,f215])).
fof(f382,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),sK2))))) != s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f381,f211])).
fof(f381,plain,(
  s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),sK1))))) != s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f380,f207])).
fof(f380,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f379,f212])).
fof(f379,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f378,f208])).
fof(f378,plain,(
  s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),sK2))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f377,f214])).
fof(f377,plain,(
  s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),sK1))) != s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f376,f210])).
fof(f376,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),sK2))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f375,f213])).
fof(f375,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),sK1)))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(forward_demodulation,[],[f348,f209])).
fof(f348,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.option.option'('type.conSem.v'),fun('type.option.option'('type.exhSem.v'),bool)),'const.option.OPTREL_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.option.option'('type.conSem.v')),'const.decSem.state_globals_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.option.option'('type.exhSem.v')),'const.exhSem.state_globals_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.ffi.ffi_state'(sK0),'const.decSem.state_ffi_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.ffi.ffi_state'(sK0),'const.exhSem.state_ffi_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.semanticPrimitives.store_v'('type.conSem.v'),fun('type.semanticPrimitives.store_v'('type.exhSem.v'),bool)),'const.evalProps.sv_rel_1'(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),'const.dec_to_exhProof.v_rel_1'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.decSem.state_refs_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))),s('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'const.exhSem.state_refs_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))))) | s('type.num.num','const.decSem.state_clock_1'(s('type.decSem.state'(sK0),'const.decSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.decSem.state'(sK0),sK2))))) != s('type.num.num','const.exhSem.state_clock_1'(s('type.exhSem.state'(sK0),'const.exhSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',sK4))),s('type.exhSem.state'(sK0),sK1)))))),
  inference(resolution,[],[f206,f278])).
fof(f338,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) | $spl0),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f331,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f307])).
fof(f321,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f205,f224])).
fof(f340,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) | $spl4),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f320,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f224])).
fof(f318,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f223])).
fof(f317,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f222])).
fof(f316,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f205,f220])).
fof(f315,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f220])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.dec_to_exhProof.state_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK3),s('type.decSem.state'(sK0),sK2),s('type.exhSem.state'(sK0),sK1)))) )),
  inference(resolution,[],[f205,f218])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 143
% Passive clauses: 1091
% Generated clauses: 9393
% Final active clauses: 105
% Final passive clauses: 816
% Input formulas: 63
% Initial clauses: 102
% 
% Selected by SInE selection: 59
% SInE iterations: 1
% 
% Duplicate literals: 463
% Trivial inequalities: 203
% Fw subsumption resolutions: 1782
% Fw demodulations: 344
% Bw demodulations: 23
% 
% Simple tautologies: 365
% Equational tautologies: 40
% Forward subsumptions: 4361
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 956
% Factoring: 1
% Forward superposition: 1658
% Backward superposition: 3579
% Self superposition: 113
% Equality factoring: 100
% Equality resolution: 24
% 
% Split clauses: 764
% Split components: 1703
% SAT solver clauses: 824
% SAT solver unit clauses: 22
% SAT solver binary clauses: 395
% 
% Sat splits: 764
% Sat splitting refutations: 22
% 
% Memory used [KB]: 4733
% Time elapsed: 0.203 s
% ------------------------------
% ------------------------------
% Success in time 11.95 s
