Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (28563)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 2714
% Generated clauses: 23692
% Final active clauses: 128
% Final passive clauses: 2351
% Input formulas: 37
% Initial clauses: 51
% 
% Duplicate literals: 1017
% Trivial inequalities: 520
% Fw subsumption resolutions: 3513
% Fw demodulations: 3480
% Bw demodulations: 7
% 
% Simple tautologies: 521
% Equational tautologies: 49
% Forward subsumptions: 11240
% 
% Binary resolution: 147
% Factoring: 2
% Forward superposition: 7767
% Backward superposition: 7053
% Self superposition: 203
% Equality factoring: 120
% Equality resolution: 40
% 
% Split clauses: 456
% Split components: 952
% SAT solver clauses: 843
% SAT solver unit clauses: 23
% SAT solver binary clauses: 146
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 456
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 836
% TWLsolver calls for satisfiability: 134
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (28584)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 729
% Generated clauses: 16436
% Final active clauses: 126
% Final passive clauses: 588
% Input formulas: 37
% Initial clauses: 51
% 
% Duplicate literals: 440
% Trivial inequalities: 258
% Fw subsumption resolutions: 1650
% Fw demodulations: 2715
% Bw demodulations: 10
% Condensations: 1290
% Global subsumptions: 5
% 
% Simple tautologies: 196
% Equational tautologies: 23
% Forward subsumptions: 7779
% 
% Binary resolution: 174
% Factoring: 2
% Forward superposition: 4321
% Backward superposition: 4611
% Self superposition: 115
% Equality factoring: 46
% Equality resolution: 38
% 
% Split clauses: 53
% Split components: 113
% SAT solver clauses: 1233
% SAT solver unit clauses: 78
% SAT solver binary clauses: 33
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 53
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1212
% TWLsolver calls for satisfiability: 279
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f45707,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f116,f170,f168,f115,f169,f167,f175,f501,f504,f5816,f5817,f5834,f5840,f5842,f5846,f5848,f5850,f5852,f5856,f5858,f5860,f5862,f5864,f5866,f5869,f6006,f6007,f6008,f6009,f6519,f6545,f6549,f6550,f6551,f6552,f6512,f6554,f6553,f13192,f13208,f13440,f13442,f13445,f13450,f13452,f13454,f13456,f13458,f13460,f13466,f13468,f13470,f13473,f13476,f13481,f13489,f13891,f13892,f13893,f13894,f13895,f13896,f13897,f13899,f16084,f16054,f16045,f16175,f15946,f16652,f17146,f17147,f17151,f17164,f17209,f17212,f17215,f17217,f17224,f17226,f17227,f17229,f17231,f17242,f17245,f17247,f27515,f27488,f27477,f27616,f27376,f28076,f27373,f28081,f32073,f32042,f32032,f32163,f31925,f32666,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33600,f33570,f33601,f33571,f33602,f33603,f33573,f33604,f33574,f33605,f33606,f33591,f33593,f33607,f33608,f33594,f33587,f33609,f33595,f33610,f33611,f33598,f33588,f33597,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33857,f33862,f33866,f33868,f33869,f33870,f33871,f33872,f33873,f33874,f33848,f33850,f33875,f33876,f33877,f33878,f33879,f33880,f33854,f33853,f33881,f33648,f33885,f33649,f33650,f33884,f33651,f33886,f33895,f33852,f33896,f33897,f33898,f33900,f33901,f33903,f33845,f33906,f33907,f33909,f33910,f33883,f33851,f33682,f33847,f33911,f36918,f37022,f37026,f37029,f37036,f36917,f37064,f37078,f37030,f37079,f37027,f16176,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37215,f37151,f37216,f37217,f37218,f37219,f37220,f37214,f37221,f37222,f37223,f37225,f37226,f37227,f37208,f37210,f37229,f37230,f37211,f33590,f37231,f37232,f37213,f37212,f37233,f37236,f37234,f41605,f41626,f42830,f42831,f42833,f42715,f42834,f42838,f42712,f42861,f42878,f42832,f6544,f43296,f43298,f43300,f43297,f502,f44737,f45447,f45448,f45452,f45464,f45504,f45507,f45510,f45512,f45519,f45525,f45529,f45547,f45553,f45557,f45559,f45560,f45565,f45569,f45574,f45580,f45587,f45592,f45594,f45596,f45598,f45600,f45629,f45632,f45634,f45636,f45638,f45640,f45642,f45644])).
fof(f45644,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f45643,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633,f45635,f45637])).
fof(f45637,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X40,happ(s(fun(X34,X40),X41),s(X34,X42))) = s(X40,happ(s(fun(X34,X40),happ(s(fun(fun(X34,X40),fun(X34,X40)),X43),s(fun(X34,X40),'const.relation.RESTRICT_3'(s(fun(X34,X40),X41),s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(X34,X42))))),s(X34,X42))) | s(fun(X34,X40),X41) != s(fun(X34,X40),'const.relation.WFREC_2'(s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(fun(fun(X34,X40),fun(X34,X40)),X43)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45405,f117])).
fof(f45405,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (~p(s(bool,'T_0')) | s(X40,happ(s(fun(X34,X40),X41),s(X34,X42))) = s(X40,happ(s(fun(X34,X40),happ(s(fun(fun(X34,X40),fun(X34,X40)),X43),s(fun(X34,X40),'const.relation.RESTRICT_3'(s(fun(X34,X40),X41),s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(X34,X42))))),s(X34,X42))) | s(fun(X34,X40),X41) != s(fun(X34,X40),'const.relation.WFREC_2'(s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(fun(fun(X34,X40),fun(X34,X40)),X43)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f5784])).
fof(f45635,plain,(
  ( ! [X819,X817,X810,X814,X812,X818,X816,X811,X815,X813] : (s(X816,happ(s(fun(X810,X816),X817),s(X810,X818))) = s(X816,happ(s(fun(X810,X816),happ(s(fun(fun(X810,X816),fun(X810,X816)),X819),s(fun(X810,X816),'const.relation.RESTRICT_3'(s(fun(X810,X816),X817),s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(X810,X818))))),s(X810,X818))) | s(fun(X810,X816),X817) != s(fun(X810,X816),'const.relation.WFREC_2'(s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(fun(fun(X810,X816),fun(X810,X816)),X819))) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X811),X812) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X811),X812)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45404,f117])).
fof(f45404,plain,(
  ( ! [X819,X817,X810,X814,X812,X818,X816,X811,X815,X813] : (~p(s(bool,'T_0')) | s(X816,happ(s(fun(X810,X816),X817),s(X810,X818))) = s(X816,happ(s(fun(X810,X816),happ(s(fun(fun(X810,X816),fun(X810,X816)),X819),s(fun(X810,X816),'const.relation.RESTRICT_3'(s(fun(X810,X816),X817),s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(X810,X818))))),s(X810,X818))) | s(fun(X810,X816),X817) != s(fun(X810,X816),'const.relation.WFREC_2'(s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(fun(fun(X810,X816),fun(X810,X816)),X819))) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X811),X812) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X811),X812)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39963])).
fof(f45633,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (s(X59,happ(s(fun(X53,X59),X60),s(X53,X61))) = s(X59,happ(s(fun(X53,X59),happ(s(fun(fun(X53,X59),fun(X53,X59)),X62),s(fun(X53,X59),'const.relation.RESTRICT_3'(s(fun(X53,X59),X60),s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(X53,X61))))),s(X53,X61))) | s(fun(X53,X59),X60) != s(fun(X53,X59),'const.relation.WFREC_2'(s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(fun(fun(X53,X59),fun(X53,X59)),X62)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45403,f117])).
fof(f45403,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (~p(s(bool,'T_0')) | s(X59,happ(s(fun(X53,X59),X60),s(X53,X61))) = s(X59,happ(s(fun(X53,X59),happ(s(fun(fun(X53,X59),fun(X53,X59)),X62),s(fun(X53,X59),'const.relation.RESTRICT_3'(s(fun(X53,X59),X60),s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(X53,X61))))),s(X53,X61))) | s(fun(X53,X59),X60) != s(fun(X53,X59),'const.relation.WFREC_2'(s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(fun(fun(X53,X59),fun(X53,X59)),X62)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f5786])).
fof(f45631,plain,(
  ( ! [X825,X831,X833,X826,X824,X830,X832,X827] : (s(X830,happ(s(fun(X824,X830),X831),s(X824,X832))) = s(X830,happ(s(fun(X824,X830),happ(s(fun(fun(X824,X830),fun(X824,X830)),X833),s(fun(X824,X830),'const.relation.RESTRICT_3'(s(fun(X824,X830),X831),s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(X824,X832))))),s(X824,X832))) | s(fun(X824,X830),X831) != s(fun(X824,X830),'const.relation.WFREC_2'(s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(fun(fun(X824,X830),fun(X824,X830)),X833))) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X825),X827) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X825),X827)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45402,f117])).
fof(f45402,plain,(
  ( ! [X825,X831,X833,X826,X824,X830,X832,X827] : (~p(s(bool,'T_0')) | s(X830,happ(s(fun(X824,X830),X831),s(X824,X832))) = s(X830,happ(s(fun(X824,X830),happ(s(fun(fun(X824,X830),fun(X824,X830)),X833),s(fun(X824,X830),'const.relation.RESTRICT_3'(s(fun(X824,X830),X831),s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(X824,X832))))),s(X824,X832))) | s(fun(X824,X830),X831) != s(fun(X824,X830),'const.relation.WFREC_2'(s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(fun(fun(X824,X830),fun(X824,X830)),X833))) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X825),X827) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X825),X827)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39810])).
fof(f45628,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45396,f117])).
fof(f45396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f162])).
fof(f45599,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45337,f117])).
fof(f45337,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f4082])).
fof(f45597,plain,(
  ( ! [X947,X945,X946,X944,X943] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X943),X944),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X945),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X946),s(fun(X943,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X947))) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X943),X944) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X943),X944)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45334,f117])).
fof(f45334,plain,(
  ( ! [X947,X945,X946,X944,X943] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X943),X944),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X945),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X946),s(fun(X943,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X947))) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X943),X944) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X943),X944)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39982])).
fof(f45595,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45333,f117])).
fof(f45333,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f4083])).
fof(f45593,plain,(
  ( ! [X963,X965,X964] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X963,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X964),s(X963,sK2(X965,X963)))) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X963),X965) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X963),X965)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45331,f117])).
fof(f45331,plain,(
  ( ! [X963,X965,X964] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X963,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X964),s(X963,sK2(X965,X963)))) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X963),X965) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X963),X965)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39829])).
fof(f45591,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45328,f117])).
fof(f45328,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f157])).
fof(f45579,plain,(
  ( ! [X177,X176] : (s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X176),X177) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X176),X177)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45578,f117])).
fof(f45578,plain,(
  ( ! [X177,X176] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X176),X177) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X176),X177)) ) | ($spl0 | $spl3 | $spl12)),
  inference(forward_demodulation,[],[f45577,f44738])).
fof(f45577,plain,(
  ( ! [X177,X180,X176] : (~p(s(bool,happ(s(fun(X176,bool),X180),s(X176,sK2(X177,X176))))) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X176),X177) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X176),X177)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45165,f44738])).
fof(f45165,plain,(
  ( ! [X177,X180,X178,X176,X179] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X176),X177),s(bool,X178),s(bool,X179),s(fun(X176,bool),X180))) | ~p(s(bool,happ(s(fun(X176,bool),X180),s(X176,sK2(X177,X176))))) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X176),X177) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X176),X177)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39856])).
fof(f45573,plain,(
  ( ! [X182,X184] : (s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X182),X184) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X182),X184)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45572,f117])).
fof(f45572,plain,(
  ( ! [X182,X184] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X182),X184) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X182),X184)) ) | ($spl0 | $spl3 | $spl12)),
  inference(forward_demodulation,[],[f45571,f44738])).
fof(f45571,plain,(
  ( ! [X185,X182,X186,X184,X183] : (~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X182),X184),s(bool,X185),s(bool,X186),s(fun(X182,bool),X183)))) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X182),X184) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X182),X184)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45141,f44738])).
fof(f45141,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') != s(bool,happ(s(fun(X182,bool),X183),s(X182,sK2(X184,X182)))) | ~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X182),X184),s(bool,X185),s(bool,X186),s(fun(X182,bool),X183)))) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X182),X184) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X182),X184)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f39702])).
fof(f45552,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(condensation,[],[f45551])).
fof(f45551,plain,(
  ( ! [X198,X197] : (s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,X197)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45053,f44738])).
fof(f45053,plain,(
  ( ! [X198,X197] : (s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X197) | s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,X197)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f29190])).
fof(f45546,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(condensation,[],[f45545])).
fof(f45545,plain,(
  ( ! [X191,X192] : (s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,X191)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45044,f44738])).
fof(f45044,plain,(
  ( ! [X191,X192] : (s(bool,'T_0') != s(bool,X191) | s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,X191)))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f29017])).
fof(f45528,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(condensation,[],[f45527])).
fof(f45527,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X25) = s(bool,X26)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45023,f16251])).
fof(f45023,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X25))) | s(bool,X25) = s(bool,X26)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f27068])).
fof(f45524,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(condensation,[],[f45523])).
fof(f45523,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X5) = s(bool,X6)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45020,f44738])).
fof(f45020,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f27062])).
fof(f45518,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(condensation,[],[f45517])).
fof(f45517,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f45010,f44738])).
fof(f45010,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f16731])).
fof(f45511,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44982,f44738])).
fof(f44982,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.Match_type_error_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f13006])).
fof(f45509,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44973,f44738])).
fof(f44973,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.No_match_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f12347])).
fof(f45506,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44956,f44738])).
fof(f44956,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f11638])).
fof(f45503,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(forward_demodulation,[],[f45502,f44738])).
fof(f45502,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44952,f44738])).
fof(f44952,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f11207])).
fof(f45451,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44750,f44738])).
fof(f44750,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f387])).
fof(f45446,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44741,f44738])).
fof(f44741,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f134])).
fof(f44738,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44735,f1303])).
fof(f44735,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44728,f6545])).
fof(f44734,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X16) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X14),'const.semanticPrimitives.Match_type_error_0'),s(bool,X15),s(bool,X16),s(fun(X14,bool),X17)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f611])).
fof(f44733,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X11) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X10),'const.semanticPrimitives.No_match_0'),s(bool,X11),s(bool,X12),s(fun(X10,bool),X13)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f610])).
fof(f44732,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X9)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f609])).
fof(f44731,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X8)))) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f608])).
fof(f44730,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f594])).
fof(f44728,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f44727,f170])).
fof(f44727,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | $spl0),
  inference(subsumption_resolution,[],[f44550,f177])).
fof(f44550,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | $spl0),
  inference(superposition,[],[f3250,f200])).
fof(f44548,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X55),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X51,X52),X53),s(fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X54))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X51,X52),fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))))),'const.pair.,_2'(s('type.pair.prod'(X51,X52),X53),s(fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X54))),s(fun('type.pair.prod'(X51,X52),fun(fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'const.pair.pair_CASE_0'))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X51,X52),fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))))),'const.pair.,_2'(s('type.pair.prod'(X51,X52),X53),s(fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X54))),s(fun('type.pair.prod'(X51,X52),fun(fun(X51,fun(X52,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f3250,f4906])).
fof(f44547,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X50),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X46,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),happ(s(fun(X45,fun(X46,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X49),s(X45,X47))),s(X46,X48))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X45,X46),'const.pair.,_2'(s(X45,X47),s(X46,X48))),s(fun(X45,fun(X46,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X49))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X45,X46),'const.pair.,_2'(s(X45,X47),s(X46,X48))),s(fun(X45,fun(X46,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X49)))) )),
  inference(superposition,[],[f3250,f165])).
fof(f44546,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X44),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X40,X41),X42),s(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X43))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),happ(s(fun('type.pair.prod'(X40,X41),fun(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X40,X41),X42))),s(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X43))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),happ(s(fun('type.pair.prod'(X40,X41),fun(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X40,X41),X42))),s(fun(X40,fun(X41,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X43)))) )),
  inference(superposition,[],[f3250,f164])).
fof(f44545,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X39),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X34),X36),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X37),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X38),s(fun(X34,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X35))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X34,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X35),s(X34,sK2(X36,X34)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X34,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X35),s(X34,sK2(X36,X34)))) | s('type.semanticPrimitives.match_result'(X34),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X34),X36) | s('type.semanticPrimitives.match_result'(X34),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X34),X36)) )),
  inference(superposition,[],[f3250,f2917])).
fof(f44544,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X33),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X29,X28),'const.pair.,_2'(s(X29,X31),s(X28,X32))),s(fun(X29,fun(X28,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X30))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X28,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),happ(s(fun(X29,fun(X28,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X30),s(X29,X31))),s(X28,X32))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X28,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),happ(s(fun(X29,fun(X28,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))))),X30),s(X29,X31))),s(X28,X32)))) )),
  inference(superposition,[],[f3250,f165])).
fof(f44543,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X27),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X26))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),'const.combin.I_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X26))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),'const.combin.I_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X26)))) )),
  inference(superposition,[],[f3250,f172])).
fof(f44542,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X24),s(X20,sK2(X21,X20)))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X20),X21),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X22),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X23),s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X24))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X20),X21),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X22),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X23),s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X24))) | s('type.semanticPrimitives.match_result'(X20),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X20),X21) | s('type.semanticPrimitives.match_result'(X20),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X20),X21)) )),
  inference(superposition,[],[f3250,f2917])).
fof(f44541,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X19),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),happ(s(fun(X14,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X18),s(X14,X15))))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X14),'const.semanticPrimitives.Match_1'(s(X14,X15))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X16),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X17),s(fun(X14,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X18))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X14),'const.semanticPrimitives.Match_1'(s(X14,X15))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X16),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X17),s(fun(X14,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X18)))) )),
  inference(superposition,[],[f3250,f163])).
fof(f44540,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X11))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X10),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X11),s(fun(X9,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X12))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X10),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X11),s(fun(X9,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X12)))) )),
  inference(superposition,[],[f3250,f161])).
fof(f44539,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X5))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X4),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X5),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X6),s(fun(X4,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X7))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X4),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X5),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X6),s(fun(X4,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')))),X7)))) )),
  inference(superposition,[],[f3250,f160])).
fof(f44153,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(fun(X16,fun(X17,bool)),bool),happ(s(fun('type.pair.prod'(X16,X17),fun(fun(X16,fun(X17,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X16,X17),X18))),s(fun(X16,fun(X17,bool)),X19))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X16,X17),X18),s(fun(X16,fun(X17,bool)),X19)))) )),
  inference(subsumption_resolution,[],[f43975,f117])).
fof(f43975,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(fun(X16,fun(X17,bool)),bool),happ(s(fun('type.pair.prod'(X16,X17),fun(fun(X16,fun(X17,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X16,X17),X18))),s(fun(X16,fun(X17,bool)),X19))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X16,X17),X18),s(fun(X16,fun(X17,bool)),X19)))) )),
  inference(superposition,[],[f134,f4907])).
fof(f43970,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(fun(X157,fun(X158,bool))),X160),s(bool,X161),s(bool,X162),s(fun(fun(X157,fun(X158,bool)),bool),happ(s(fun('type.pair.prod'(X157,X158),fun(fun(X157,fun(X158,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X157,X158),X159))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X157,X158),X159),s(fun(X157,fun(X158,bool)),sK2(X160,fun(X157,fun(X158,bool)))))) | s('type.semanticPrimitives.match_result'(fun(X157,fun(X158,bool))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(fun(X157,fun(X158,bool))),X160) | s('type.semanticPrimitives.match_result'(fun(X157,fun(X158,bool))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(fun(X157,fun(X158,bool))),X160)) )),
  inference(superposition,[],[f4907,f2917])).
fof(f43939,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X57,fun(X58,bool)),bool),happ(s(fun('type.pair.prod'(X57,X58),fun(fun(X57,fun(X58,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X57,X58),X64))),s(fun(X57,fun(X58,bool)),happ(s(fun(X60,fun(X57,fun(X58,bool))),happ(s(fun(X59,fun(X60,fun(X57,fun(X58,bool)))),X63),s(X59,X61))),s(X60,X62))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X57,X58),X64),s(fun(X57,fun(X58,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,X61),s(X60,X62))),s(fun(X59,fun(X60,fun(X57,fun(X58,bool)))),X63)))))) )),
  inference(superposition,[],[f4907,f165])).
fof(f43937,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X42,fun(X43,bool)),bool),happ(s(fun('type.pair.prod'(X42,X43),fun(fun(X42,fun(X43,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X42,X43),X49))),s(fun(X42,fun(X43,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X44),X46),s(fun(X42,fun(X43,bool)),X47),s(fun(X42,fun(X43,bool)),X48),s(fun(X44,fun(X42,fun(X43,bool))),X45))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X43),X49),s(fun(X42,fun(X43,bool)),happ(s(fun(X44,fun(X42,fun(X43,bool))),X45),s(X44,sK2(X46,X44)))))) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X44),X46) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X44),X46)) )),
  inference(superposition,[],[f4907,f2917])).
fof(f43936,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),X41))),s(fun(X34,fun(X35,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X37,X36),'const.pair.,_2'(s(X37,X39),s(X36,X40))),s(fun(X37,fun(X36,fun(X34,fun(X35,bool)))),X38))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X41),s(fun(X34,fun(X35,bool)),happ(s(fun(X36,fun(X34,fun(X35,bool))),happ(s(fun(X37,fun(X36,fun(X34,fun(X35,bool)))),X38),s(X37,X39))),s(X36,X40)))))) )),
  inference(superposition,[],[f4907,f165])).
fof(f43934,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X22,fun(X23,bool)),bool),happ(s(fun('type.pair.prod'(X22,X23),fun(fun(X22,fun(X23,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X22,X23),X29))),s(fun(X22,fun(X23,bool)),happ(s(fun(X24,fun(X22,fun(X23,bool))),X28),s(X24,sK2(X25,X24)))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),X29),s(fun(X22,fun(X23,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X24),X25),s(fun(X22,fun(X23,bool)),X26),s(fun(X22,fun(X23,bool)),X27),s(fun(X24,fun(X22,fun(X23,bool))),X28))))) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X24),X25) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X24),X25)) )),
  inference(superposition,[],[f4907,f2917])).
fof(f43928,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X57,fun(X58,bool)),bool),happ(s(fun('type.pair.prod'(X57,X58),fun(fun(X57,fun(X58,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X57,X58),happ(s(fun(X60,'type.pair.prod'(X57,X58)),happ(s(fun(X59,fun(X60,'type.pair.prod'(X57,X58))),X63),s(X59,X61))),s(X60,X62))))),s(fun(X57,fun(X58,bool)),X64))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X57,X58),'const.pair.pair_CASE_2'(s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,X61),s(X60,X62))),s(fun(X59,fun(X60,'type.pair.prod'(X57,X58))),X63))),s(fun(X57,fun(X58,bool)),X64)))) )),
  inference(superposition,[],[f4907,f165])).
fof(f43926,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X42,fun(X43,bool)),bool),happ(s(fun('type.pair.prod'(X42,X43),fun(fun(X42,fun(X43,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X42,X43),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X44),X46),s('type.pair.prod'(X42,X43),X47),s('type.pair.prod'(X42,X43),X48),s(fun(X44,'type.pair.prod'(X42,X43)),X45))))),s(fun(X42,fun(X43,bool)),X49))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X43),happ(s(fun(X44,'type.pair.prod'(X42,X43)),X45),s(X44,sK2(X46,X44)))),s(fun(X42,fun(X43,bool)),X49))) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X44),X46) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X44),X46)) )),
  inference(superposition,[],[f4907,f2917])).
fof(f43925,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),'const.pair.pair_CASE_2'(s('type.pair.prod'(X37,X36),'const.pair.,_2'(s(X37,X39),s(X36,X40))),s(fun(X37,fun(X36,'type.pair.prod'(X34,X35))),X38))))),s(fun(X34,fun(X35,bool)),X41))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),happ(s(fun(X36,'type.pair.prod'(X34,X35)),happ(s(fun(X37,fun(X36,'type.pair.prod'(X34,X35))),X38),s(X37,X39))),s(X36,X40))),s(fun(X34,fun(X35,bool)),X41)))) )),
  inference(superposition,[],[f4907,f165])).
fof(f43923,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X22,fun(X23,bool)),bool),happ(s(fun('type.pair.prod'(X22,X23),fun(fun(X22,fun(X23,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X22,X23),happ(s(fun(X24,'type.pair.prod'(X22,X23)),X28),s(X24,sK2(X25,X24)))))),s(fun(X22,fun(X23,bool)),X29))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X24),X25),s('type.pair.prod'(X22,X23),X26),s('type.pair.prod'(X22,X23),X27),s(fun(X24,'type.pair.prod'(X22,X23)),X28))),s(fun(X22,fun(X23,bool)),X29))) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X24),X25) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X24),X25)) )),
  inference(superposition,[],[f4907,f2917])).
fof(f43326,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,happ(s(fun(fun(X159,fun(X160,bool)),bool),happ(s(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),'const.min.@_1'(s(fun(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),bool),happ(s(fun('type.pair.prod'('type.pair.prod'(X159,X160),fun(X159,fun(X160,bool))),fun(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'('type.pair.prod'(X159,X160),fun(X159,fun(X160,bool))),'const.pair.,_2'(s('type.pair.prod'(X159,X160),X161),s(fun(X159,fun(X160,bool)),X162))))))),s('type.pair.prod'(X159,X160),X161))),s(fun(X159,fun(X160,bool)),X162)))) | ~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X159,X160),X161),s(fun(X159,fun(X160,bool)),X162))))) )),
  inference(forward_demodulation,[],[f43285,f165])).
fof(f43285,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X159,X160),X161),s(fun(X159,fun(X160,bool)),X162)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X159,X160),fun(X159,fun(X160,bool))),'const.pair.,_2'(s('type.pair.prod'(X159,X160),X161),s(fun(X159,fun(X160,bool)),X162))),s(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),'const.min.@_1'(s(fun(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),bool),happ(s(fun('type.pair.prod'('type.pair.prod'(X159,X160),fun(X159,fun(X160,bool))),fun(fun('type.pair.prod'(X159,X160),fun(fun(X159,fun(X160,bool)),bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'('type.pair.prod'(X159,X160),fun(X159,fun(X160,bool))),'const.pair.,_2'(s('type.pair.prod'(X159,X160),X161),s(fun(X159,fun(X160,bool)),X162))))))))))) )),
  inference(superposition,[],[f4996,f4906])).
fof(f43325,plain,(
  ( ! [X158,X156,X154,X157,X155] : (p(s(bool,happ(s(fun(X155,bool),happ(s(fun(X154,fun(X155,bool)),'const.min.@_1'(s(fun(fun(X154,fun(X155,bool)),bool),happ(s(fun('type.pair.prod'(X154,X155),fun(fun(X154,fun(X155,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X154,X155),'const.pair.,_2'(s(X154,X156),s(X155,X157))))))),s(X154,X156))),s(X155,X157)))) | ~p(s(bool,happ(s(fun(X155,bool),happ(s(fun(X154,fun(X155,bool)),X158),s(X154,X156))),s(X155,X157))))) )),
  inference(forward_demodulation,[],[f43284,f165])).
fof(f43284,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),happ(s(fun(X154,fun(X155,bool)),X158),s(X154,X156))),s(X155,X157)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X154,X155),'const.pair.,_2'(s(X154,X156),s(X155,X157))),s(fun(X154,fun(X155,bool)),'const.min.@_1'(s(fun(fun(X154,fun(X155,bool)),bool),happ(s(fun('type.pair.prod'(X154,X155),fun(fun(X154,fun(X155,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X154,X155),'const.pair.,_2'(s(X154,X156),s(X155,X157))))))))))) )),
  inference(superposition,[],[f4996,f165])).
fof(f43303,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X2),s(fun(X0,fun(X1,bool)),X3))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X2),s(fun(X0,fun(X1,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X2))))))))) ) | $spl0),
  inference(condensation,[],[f43302])).
fof(f43302,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X58) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),'const.min.@_1'(s(fun(fun(X54,fun(X55,bool)),bool),happ(s(fun('type.pair.prod'(X54,X55),fun(fun(X54,fun(X55,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X54,X55),X56)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),X57)))) ) | $spl0),
  inference(subsumption_resolution,[],[f43301,f135])).
fof(f43301,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X58) | ~p(s(bool,X58)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),'const.min.@_1'(s(fun(fun(X54,fun(X55,bool)),bool),happ(s(fun('type.pair.prod'(X54,X55),fun(fun(X54,fun(X55,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X54,X55),X56)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),X57)))) ) | $spl0),
  inference(forward_demodulation,[],[f43263,f141])).
fof(f43263,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X58)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),'const.min.@_1'(s(fun(fun(X54,fun(X55,bool)),bool),happ(s(fun('type.pair.prod'(X54,X55),fun(fun(X54,fun(X55,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X54,X55),X56)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),X56),s(fun(X54,fun(X55,bool)),X57))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X58)))) ) | $spl0),
  inference(superposition,[],[f4996,f755])).
fof(f43295,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X9,X10),X11),s(fun(X9,fun(X10,bool)),'const.min.@_1'(s(fun(fun(X9,fun(X10,bool)),bool),happ(s(fun('type.pair.prod'(X9,X10),fun(fun(X9,fun(X10,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X9,X10),X11)))))))) | s(bool,X13) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X9,X10),X11),s(fun(X9,fun(X10,bool)),X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f43254,f117])).
fof(f43254,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X9,X10),X11),s(fun(X9,fun(X10,bool)),'const.min.@_1'(s(fun(fun(X9,fun(X10,bool)),bool),happ(s(fun('type.pair.prod'(X9,X10),fun(fun(X9,fun(X10,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X9,X10),X11)))))))) | s(bool,X13) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X9,X10),X11),s(fun(X9,fun(X10,bool)),X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f4996,f205])).
fof(f43294,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))),s(X1,X3)))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f43293,f165])).
fof(f43293,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))),s(fun(X0,fun(X1,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X3)))) )),
  inference(subsumption_resolution,[],[f43252,f117])).
fof(f43252,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))),s(fun(X0,fun(X1,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X1,bool)),bool),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X2))),s(X1,X3)))) )),
  inference(superposition,[],[f4996,f3570])).
fof(f43238,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X57,X58),happ(s(fun(X60,'type.pair.prod'(X57,X58)),happ(s(fun(X59,fun(X60,'type.pair.prod'(X57,X58))),X63),s(X59,X61))),s(X60,X62))),s(fun(X57,fun(X58,bool)),X64)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X57,X58),'const.pair.pair_CASE_2'(s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,X61),s(X60,X62))),s(fun(X59,fun(X60,'type.pair.prod'(X57,X58))),X63))),s(fun(X57,fun(X58,bool)),'const.min.@_1'(s(fun(fun(X57,fun(X58,bool)),bool),happ(s(fun('type.pair.prod'(X57,X58),fun(fun(X57,fun(X58,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X57,X58),'const.pair.pair_CASE_2'(s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,X61),s(X60,X62))),s(fun(X59,fun(X60,'type.pair.prod'(X57,X58))),X63))))))))))) )),
  inference(superposition,[],[f4996,f165])).
fof(f43236,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X43),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X44),X46),s('type.pair.prod'(X42,X43),X47),s('type.pair.prod'(X42,X43),X48),s(fun(X44,'type.pair.prod'(X42,X43)),X45))),s(fun(X42,fun(X43,bool)),X49)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X43),happ(s(fun(X44,'type.pair.prod'(X42,X43)),X45),s(X44,sK2(X46,X44)))),s(fun(X42,fun(X43,bool)),'const.min.@_1'(s(fun(fun(X42,fun(X43,bool)),bool),happ(s(fun('type.pair.prod'(X42,X43),fun(fun(X42,fun(X43,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X42,X43),happ(s(fun(X44,'type.pair.prod'(X42,X43)),X45),s(X44,sK2(X46,X44))))))))))) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X44),X46) | s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X44),X46)) )),
  inference(superposition,[],[f4996,f2917])).
fof(f43235,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),'const.pair.pair_CASE_2'(s('type.pair.prod'(X37,X36),'const.pair.,_2'(s(X37,X39),s(X36,X40))),s(fun(X37,fun(X36,'type.pair.prod'(X34,X35))),X38))),s(fun(X34,fun(X35,bool)),X41)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),happ(s(fun(X36,'type.pair.prod'(X34,X35)),happ(s(fun(X37,fun(X36,'type.pair.prod'(X34,X35))),X38),s(X37,X39))),s(X36,X40))),s(fun(X34,fun(X35,bool)),'const.min.@_1'(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),happ(s(fun(X36,'type.pair.prod'(X34,X35)),happ(s(fun(X37,fun(X36,'type.pair.prod'(X34,X35))),X38),s(X37,X39))),s(X36,X40))))))))))) )),
  inference(superposition,[],[f4996,f165])).
fof(f43233,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),happ(s(fun(X24,'type.pair.prod'(X22,X23)),X28),s(X24,sK2(X25,X24)))),s(fun(X22,fun(X23,bool)),X29)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X24),X25),s('type.pair.prod'(X22,X23),X26),s('type.pair.prod'(X22,X23),X27),s(fun(X24,'type.pair.prod'(X22,X23)),X28))),s(fun(X22,fun(X23,bool)),'const.min.@_1'(s(fun(fun(X22,fun(X23,bool)),bool),happ(s(fun('type.pair.prod'(X22,X23),fun(fun(X22,fun(X23,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X22,X23),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X24),X25),s('type.pair.prod'(X22,X23),X26),s('type.pair.prod'(X22,X23),X27),s(fun(X24,'type.pair.prod'(X22,X23)),X28)))))))))) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X24),X25) | s('type.semanticPrimitives.match_result'(X24),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X24),X25)) )),
  inference(superposition,[],[f4996,f2917])).
fof(f42237,plain,(
  ( ! [X850,X848,X843,X847,X845,X849,X842,X846,X844] : (s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))) | s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))) | s(X847,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(X847,X849),s(X847,X850),s(fun(X842,X847),X848))) = s(X847,happ(s(fun(X842,X847),X848),s(X842,sK2('const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0')),X842))))) )),
  inference(forward_demodulation,[],[f42236,f4906])).
fof(f42236,plain,(
  ( ! [X850,X848,X843,X847,X845,X849,X842,X846,X844] : (s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))) | s(X847,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(X847,X849),s(X847,X850),s(fun(X842,X847),X848))) = s(X847,happ(s(fun(X842,X847),X848),s(X842,sK2('const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0')),X842)))) | s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0')))) )),
  inference(forward_demodulation,[],[f42073,f4906])).
fof(f42073,plain,(
  ( ! [X850,X848,X843,X847,X845,X849,X842,X846,X844] : (s(X847,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(X847,X849),s(X847,X850),s(fun(X842,X847),X848))) = s(X847,happ(s(fun(X842,X847),X848),s(X842,sK2('const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0')),X842)))) | s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0'))) | s('type.semanticPrimitives.match_result'(X842),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X842),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X843,X844),fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842)))),'const.pair.,_2'(s('type.pair.prod'(X843,X844),X845),s(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),X846))),s(fun('type.pair.prod'(X843,X844),fun(fun(X843,fun(X844,'type.semanticPrimitives.match_result'(X842))),'type.semanticPrimitives.match_result'(X842))),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f2917,f4906])).
fof(f42222,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'T_0') = s(bool,X635) | s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X631,X632),X633),s(fun(X631,fun(X632,bool)),X634))) | s(bool,X635) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X631,X632),fun(X631,fun(X632,bool))),'const.pair.,_2'(s('type.pair.prod'(X631,X632),X633),s(fun(X631,fun(X632,bool)),X634))),s(fun('type.pair.prod'(X631,X632),fun(fun(X631,fun(X632,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f42037,f141])).
fof(f42037,plain,(
  ( ! [X633,X635,X631,X632,X634] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X631,X632),X633),s(fun(X631,fun(X632,bool)),X634))) | s(bool,X635) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X635) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X631,X632),fun(X631,fun(X632,bool))),'const.pair.,_2'(s('type.pair.prod'(X631,X632),X633),s(fun(X631,fun(X632,bool)),X634))),s(fun('type.pair.prod'(X631,X632),fun(fun(X631,fun(X632,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f16731,f4906])).
fof(f41986,plain,(
  ( ! [X389,X391,X393,X390,X392] : (s(bool,X393) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X389,X390),X391),s(fun(X389,fun(X390,bool)),X392))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X393))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X389,X390),fun(X389,fun(X390,bool))),'const.pair.,_2'(s('type.pair.prod'(X389,X390),X391),s(fun(X389,fun(X390,bool)),X392))),s(fun('type.pair.prod'(X389,X390),fun(fun(X389,fun(X390,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f2234,f4906])).
fof(f42157,plain,(
  ( ! [X377,X376,X378,X375] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X375,X376),X377),s(fun(X375,fun(X376,bool)),X378))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X375,X376),fun(X375,fun(X376,bool))),'const.pair.,_2'(s('type.pair.prod'(X375,X376),X377),s(fun(X375,fun(X376,bool)),X378))),s(fun('type.pair.prod'(X375,X376),fun(fun(X375,fun(X376,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41983,f141])).
fof(f41983,plain,(
  ( ! [X377,X376,X378,X375] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X375,X376),X377),s(fun(X375,fun(X376,bool)),X378))))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X375,X376),fun(X375,fun(X376,bool))),'const.pair.,_2'(s('type.pair.prod'(X375,X376),X377),s(fun(X375,fun(X376,bool)),X378))),s(fun('type.pair.prod'(X375,X376),fun(fun(X375,fun(X376,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f2047,f4906])).
fof(f41981,plain,(
  ( ! [X366,X368,X370,X367,X369] : (s(bool,'const.combin.I_1'(s(bool,X370))) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X366,X367),X368),s(fun(X366,fun(X367,bool)),X369))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X366,X367),fun(X366,fun(X367,bool))),'const.pair.,_2'(s('type.pair.prod'(X366,X367),X368),s(fun(X366,fun(X367,bool)),X369))),s(fun('type.pair.prod'(X366,X367),fun(fun(X366,fun(X367,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'T_0') = s(bool,X370)) ) | $spl0),
  inference(superposition,[],[f1916,f4906])).
fof(f42153,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,X365) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X361,X362),X363),s(fun(X361,fun(X362,bool)),X364))) | s(bool,'F_0') = s(bool,X365) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X361,X362),fun(X361,fun(X362,bool))),'const.pair.,_2'(s('type.pair.prod'(X361,X362),X363),s(fun(X361,fun(X362,bool)),X364))),s(fun('type.pair.prod'(X361,X362),fun(fun(X361,fun(X362,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41980,f141])).
fof(f41980,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,X365) != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X361,X362),X363),s(fun(X361,fun(X362,bool)),X364))))) | s(bool,'F_0') = s(bool,X365) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X361,X362),fun(X361,fun(X362,bool))),'const.pair.,_2'(s('type.pair.prod'(X361,X362),X363),s(fun(X361,fun(X362,bool)),X364))),s(fun('type.pair.prod'(X361,X362),fun(fun(X361,fun(X362,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f1916,f4906])).
fof(f42138,plain,(
  ( ! [X327,X329,X326,X328] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X326,X327),X328),s(fun(X326,fun(X327,bool)),X329))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X326,X327),fun(X326,fun(X327,bool))),'const.pair.,_2'(s('type.pair.prod'(X326,X327),X328),s(fun(X326,fun(X327,bool)),X329))),s(fun('type.pair.prod'(X326,X327),fun(fun(X326,fun(X327,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41972,f141])).
fof(f41972,plain,(
  ( ! [X327,X329,X326,X328] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X326,X327),X328),s(fun(X326,fun(X327,bool)),X329))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X326,X327),fun(X326,fun(X327,bool))),'const.pair.,_2'(s('type.pair.prod'(X326,X327),X328),s(fun(X326,fun(X327,bool)),X329))),s(fun('type.pair.prod'(X326,X327),fun(fun(X326,fun(X327,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f1629,f4906])).
fof(f42137,plain,(
  ( ! [X325,X321,X323,X324,X322] : (s(bool,X325) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X321,X322),X323),s(fun(X321,fun(X322,bool)),X324))) | s(bool,X325) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X321,X322),fun(X321,fun(X322,bool))),'const.pair.,_2'(s('type.pair.prod'(X321,X322),X323),s(fun(X321,fun(X322,bool)),X324))),s(fun('type.pair.prod'(X321,X322),fun(fun(X321,fun(X322,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,X325)) ) | $spl0),
  inference(forward_demodulation,[],[f41971,f141])).
fof(f41971,plain,(
  ( ! [X325,X321,X323,X324,X322] : (s(bool,X325) != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X321,X322),X323),s(fun(X321,fun(X322,bool)),X324))))) | s(bool,X325) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X321,X322),fun(X321,fun(X322,bool))),'const.pair.,_2'(s('type.pair.prod'(X321,X322),X323),s(fun(X321,fun(X322,bool)),X324))),s(fun('type.pair.prod'(X321,X322),fun(fun(X321,fun(X322,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,X325)) ) | $spl0),
  inference(superposition,[],[f1617,f4906])).
fof(f41970,plain,(
  ( ! [X316,X318,X320,X317,X319] : (s(bool,'const.combin.I_1'(s(bool,X320))) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X316,X317),X318),s(fun(X316,fun(X317,bool)),X319))) | s(bool,X320) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X316,X317),fun(X316,fun(X317,bool))),'const.pair.,_2'(s('type.pair.prod'(X316,X317),X318),s(fun(X316,fun(X317,bool)),X319))),s(fun('type.pair.prod'(X316,X317),fun(fun(X316,fun(X317,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X316,X317),fun(X316,fun(X317,bool))),'const.pair.,_2'(s('type.pair.prod'(X316,X317),X318),s(fun(X316,fun(X317,bool)),X319))),s(fun('type.pair.prod'(X316,X317),fun(fun(X316,fun(X317,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f1617,f4906])).
fof(f41969,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X311,X312),X313),s(fun(X311,fun(X312,bool)),X314))) | s(bool,X315) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X311,X312),fun(X311,fun(X312,bool))),'const.pair.,_2'(s('type.pair.prod'(X311,X312),X313),s(fun(X311,fun(X312,bool)),X314))),s(fun('type.pair.prod'(X311,X312),fun(fun(X311,fun(X312,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X315)))) ) | $spl0),
  inference(superposition,[],[f1515,f4906])).
fof(f41966,plain,(
  ( ! [X296,X298,X295,X297,X299] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X295,X296),X297),s(fun(X295,fun(X296,bool)),X298))) | s(bool,X299) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X295,X296),fun(X295,fun(X296,bool))),'const.pair.,_2'(s('type.pair.prod'(X295,X296),X297),s(fun(X295,fun(X296,bool)),X298))),s(fun('type.pair.prod'(X295,X296),fun(fun(X295,fun(X296,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,X299)) ) | $spl0),
  inference(superposition,[],[f900,f4906])).
fof(f41957,plain,(
  ( ! [X253,X251,X249,X254,X252,X250,X248] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X252),'const.semanticPrimitives.Match_type_error_0'),s(bool,X253),s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X248,X249),X250),s(fun(X248,fun(X249,bool)),X251))),s(fun(X252,bool),X254))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X248,X249),fun(X248,fun(X249,bool))),'const.pair.,_2'(s('type.pair.prod'(X248,X249),X250),s(fun(X248,fun(X249,bool)),X251))),s(fun('type.pair.prod'(X248,X249),fun(fun(X248,fun(X249,bool)),bool)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f414,f4906])).
fof(f41954,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X232] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X231),'const.semanticPrimitives.No_match_0'),s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X227,X228),X229),s(fun(X227,fun(X228,bool)),X230))),s(bool,X232),s(fun(X231,bool),X233))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X227,X228),fun(X227,fun(X228,bool))),'const.pair.,_2'(s('type.pair.prod'(X227,X228),X229),s(fun(X227,fun(X228,bool)),X230))),s(fun('type.pair.prod'(X227,X228),fun(fun(X227,fun(X228,bool)),bool)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f316,f4906])).
fof(f42083,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X87,X88),X89),s(fun(X87,fun(X88,bool)),X90))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X87,X88),fun(X87,fun(X88,bool))),'const.pair.,_2'(s('type.pair.prod'(X87,X88),X89),s(fun(X87,fun(X88,bool)),X90))),s(fun('type.pair.prod'(X87,X88),fun(fun(X87,fun(X88,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41900,f141])).
fof(f41900,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X87,X88),X89),s(fun(X87,fun(X88,bool)),X90))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X87,X88),fun(X87,fun(X88,bool))),'const.pair.,_2'(s('type.pair.prod'(X87,X88),X89),s(fun(X87,fun(X88,bool)),X90))),s(fun('type.pair.prod'(X87,X88),fun(fun(X87,fun(X88,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f4906,f3762])).
fof(f41899,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'const.combin.I_1'(s(bool,X86))) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X82,X83),X84),s(fun(X82,fun(X83,bool)),X85))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X82,X83),fun(X82,fun(X83,bool))),'const.pair.,_2'(s('type.pair.prod'(X82,X83),X84),s(fun(X82,fun(X83,bool)),X85))),s(fun('type.pair.prod'(X82,X83),fun(fun(X82,fun(X83,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'T_0') = s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f4906,f2195])).
fof(f41896,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X68,X69),X70),s(fun(X68,fun(X69,bool)),X71))) | s(bool,X72) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X68,X69),fun(X68,fun(X69,bool))),'const.pair.,_2'(s('type.pair.prod'(X68,X69),X70),s(fun(X68,fun(X69,bool)),X71))),s(fun('type.pair.prod'(X68,X69),fun(fun(X68,fun(X69,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X72)))) ) | $spl0),
  inference(superposition,[],[f4906,f755])).
fof(f41895,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,X67) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X63,X64),X65),s(fun(X63,fun(X64,bool)),X66))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X63,X64),fun(X63,fun(X64,bool))),'const.pair.,_2'(s('type.pair.prod'(X63,X64),X65),s(fun(X63,fun(X64,bool)),X66))),s(fun('type.pair.prod'(X63,X64),fun(fun(X63,fun(X64,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X67)))) ) | $spl0),
  inference(superposition,[],[f4906,f755])).
fof(f41891,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X43,X44),X45),s(fun(X43,fun(X44,bool)),X46))) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X43,X44),fun(X43,fun(X44,bool))),'const.pair.,_2'(s('type.pair.prod'(X43,X44),X45),s(fun(X43,fun(X44,bool)),X46))),s(fun('type.pair.prod'(X43,X44),fun(fun(X43,fun(X44,bool)),bool)),'const.pair.pair_CASE_0')))) ) | $spl0),
  inference(superposition,[],[f4906,f612])).
fof(f41890,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X38,X39),X40),s(fun(X38,fun(X39,bool)),X41))) | s(bool,X42) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X38,X39),fun(X38,fun(X39,bool))),'const.pair.,_2'(s('type.pair.prod'(X38,X39),X40),s(fun(X38,fun(X39,bool)),X41))),s(fun('type.pair.prod'(X38,X39),fun(fun(X38,fun(X39,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'F_0') = s(bool,X42)) ) | $spl0),
  inference(superposition,[],[f4906,f612])).
fof(f41887,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,X27) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X24),X25),s(fun(X23,fun(X24,bool)),X26))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X23,X24),fun(X23,fun(X24,bool))),'const.pair.,_2'(s('type.pair.prod'(X23,X24),X25),s(fun(X23,fun(X24,bool)),X26))),s(fun('type.pair.prod'(X23,X24),fun(fun(X23,fun(X24,bool)),bool)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f4906,f205])).
fof(f41886,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X19),X20),s(fun(X18,fun(X19,bool)),X21))) | s(bool,X22) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X18,X19),fun(X18,fun(X19,bool))),'const.pair.,_2'(s('type.pair.prod'(X18,X19),X20),s(fun(X18,fun(X19,bool)),X21))),s(fun('type.pair.prod'(X18,X19),fun(fun(X18,fun(X19,bool)),bool)),'const.pair.pair_CASE_0'))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f4906,f205])).
fof(f41885,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X14,X15),X16),s(fun(X14,fun(X15,bool)),X17))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X14,X15),fun(X14,fun(X15,bool))),'const.pair.,_2'(s('type.pair.prod'(X14,X15),X16),s(fun(X14,fun(X15,bool)),X17))),s(fun('type.pair.prod'(X14,X15),fun(fun(X14,fun(X15,bool)),bool)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f4906,f128])).
fof(f41424,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X15,X16),'const.pair.,_2'(s(X15,X17),s(X16,X18))),s(fun(X15,fun(X16,bool)),X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X15,fun(X16,bool)),X19),s(X15,X17))),s(X16,X18)))) )),
  inference(subsumption_resolution,[],[f41066,f117])).
fof(f41066,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X15,X16),'const.pair.,_2'(s(X15,X17),s(X16,X18))),s(fun(X15,fun(X16,bool)),X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X15,fun(X16,bool)),X19),s(X15,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f134,f3570])).
fof(f41033,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65,X73] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X65,X66),'const.pair.,_2'(s(X65,X72),s(X66,X73))),s(fun(X65,fun(X66,bool)),happ(s(fun(X68,fun(X65,fun(X66,bool))),happ(s(fun(X67,fun(X68,fun(X65,fun(X66,bool)))),X71),s(X67,X69))),s(X68,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun(X66,bool),happ(s(fun(X65,fun(X66,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X67,X68),'const.pair.,_2'(s(X67,X69),s(X68,X70))),s(fun(X67,fun(X68,fun(X65,fun(X66,bool)))),X71))),s(X65,X72))),s(X66,X73)))) )),
  inference(superposition,[],[f3570,f165])).
fof(f41031,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,X55),s(X49,X56))),s(fun(X48,fun(X49,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X50),X52),s(fun(X48,fun(X49,bool)),X53),s(fun(X48,fun(X49,bool)),X54),s(fun(X50,fun(X48,fun(X49,bool))),X51))))) | s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),happ(s(fun(X48,fun(X49,bool)),happ(s(fun(X50,fun(X48,fun(X49,bool))),X51),s(X50,sK2(X52,X50)))),s(X48,X55))),s(X49,X56))) | s('type.semanticPrimitives.match_result'(X50),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X50),X52) | s('type.semanticPrimitives.match_result'(X50),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X50),X52)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f41030,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X40),'const.pair.,_2'(s(X39,X46),s(X40,X47))),s(fun(X39,fun(X40,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X41),'const.pair.,_2'(s(X42,X44),s(X41,X45))),s(fun(X42,fun(X41,fun(X39,fun(X40,bool)))),X43))))) | s(bool,'F_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X39,fun(X40,bool)),happ(s(fun(X41,fun(X39,fun(X40,bool))),happ(s(fun(X42,fun(X41,fun(X39,fun(X40,bool)))),X43),s(X42,X44))),s(X41,X45))),s(X39,X46))),s(X40,X47)))) )),
  inference(superposition,[],[f3570,f165])).
fof(f41028,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X25,X26),'const.pair.,_2'(s(X25,X32),s(X26,X33))),s(fun(X25,fun(X26,bool)),happ(s(fun(X27,fun(X25,fun(X26,bool))),X31),s(X27,sK2(X28,X27)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),happ(s(fun(X25,fun(X26,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X27),X28),s(fun(X25,fun(X26,bool)),X29),s(fun(X25,fun(X26,bool)),X30),s(fun(X27,fun(X25,fun(X26,bool))),X31))),s(X25,X32))),s(X26,X33))) | s('type.semanticPrimitives.match_result'(X27),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X27),X28) | s('type.semanticPrimitives.match_result'(X27),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X27),X28)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f41022,plain,(
  ( ! [X472,X469,X471,X473,X470] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X471,'type.semanticPrimitives.match_result'(X469)),'const.pair.,_2'(s(X471,X472),s('type.semanticPrimitives.match_result'(X469),X470))),s(fun(X471,fun('type.semanticPrimitives.match_result'(X469),bool)),X473))) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.match_result'(X469),bool),happ(s(fun(X471,fun('type.semanticPrimitives.match_result'(X469),bool)),X473),s(X471,X472))),s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_1'(s(X469,sK2(X470,X469)))))) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X469),X470) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X469),X470)) )),
  inference(superposition,[],[f3570,f144])).
fof(f41378,plain,(
  ( ! [X389,X385,X387,X388,X386] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X386,fun(bool,bool)),X388),s(X386,X387))),s(bool,X385))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X386,bool),'const.pair.,_2'(s(X386,X387),s(bool,'F_0'))),s(fun(X386,fun(bool,bool)),X388))) | s(bool,X385) = s(bool,'const.relation.WF_1'(s(fun(X389,fun(X389,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41002,f141])).
fof(f41002,plain,(
  ( ! [X389,X385,X387,X388,X386] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X386,bool),'const.pair.,_2'(s(X386,X387),s(bool,'F_0'))),s(fun(X386,fun(bool,bool)),X388))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X386,fun(bool,bool)),X388),s(X386,X387))),s(bool,'const.combin.I_1'(s(bool,X385))))) | s(bool,X385) = s(bool,'const.relation.WF_1'(s(fun(X389,fun(X389,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f41368,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'F_0') = s(bool,X348) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X349,fun(bool,bool)),X351),s(X349,X350))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,bool),'const.pair.,_2'(s(X349,X350),s(bool,X348))),s(fun(X349,fun(bool,bool)),X351)))) ) | $spl0),
  inference(forward_demodulation,[],[f41367,f141])).
fof(f41367,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X349,fun(bool,bool)),X351),s(X349,X350))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,bool),'const.pair.,_2'(s(X349,X350),s(bool,X348))),s(fun(X349,fun(bool,bool)),X351))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X348)))) ) | $spl0),
  inference(forward_demodulation,[],[f40993,f183])).
fof(f40993,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,bool),'const.pair.,_2'(s(X349,X350),s(bool,X348))),s(fun(X349,fun(bool,bool)),X351))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X349,fun(bool,bool)),X351),s(X349,X350))),s(bool,'const.relation.WF_1'(s(fun(X347,fun(X347,bool)),'const.relation.EMPTY_REL_0'))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X348)))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f40977,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X269,bool),'const.pair.,_2'(s(X269,X270),s(bool,X268))),s(fun(X269,fun(bool,bool)),X271))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X269,fun(bool,bool)),X271),s(X269,X270))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f3570,f128])).
fof(f40974,plain,(
  ( ! [X255,X253,X257,X254,X256] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X254,bool),'const.pair.,_2'(s(X254,X255),s(bool,X253))),s(fun(X254,fun(bool,bool)),X256))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X254,fun(bool,bool)),X256),s(X254,X255))),s(bool,'F_0'))) | s(bool,X253) = s(bool,X257) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X257)))) ) | $spl0),
  inference(superposition,[],[f3570,f755])).
fof(f41351,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'F_0') = s(bool,X252) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X250),s(bool,X248))),s(fun(X249,fun(bool,bool)),X251))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X249,fun(bool,bool)),X251),s(X249,X250))),s(bool,'F_0'))) | s(bool,X248) = s(bool,X252)) ) | $spl0),
  inference(forward_demodulation,[],[f40973,f141])).
fof(f40973,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X250),s(bool,X248))),s(fun(X249,fun(bool,bool)),X251))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X249,fun(bool,bool)),X251),s(X249,X250))),s(bool,'F_0'))) | s(bool,X252) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X248) = s(bool,X252)) ) | $spl0),
  inference(superposition,[],[f3570,f11021])).
fof(f41327,plain,(
  ( ! [X142,X140,X138,X141] : (s(bool,'F_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,'T_0'))),s(fun(X140,fun(bool,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X140,fun(bool,bool)),X142),s(X140,X141))),s(bool,X138)))) ) | $spl0),
  inference(forward_demodulation,[],[f41326,f141])).
fof(f41326,plain,(
  ( ! [X142,X140,X138,X141] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,'T_0'))),s(fun(X140,fun(bool,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X140,fun(bool,bool)),X142),s(X140,X141))),s(bool,X138))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X138)))) ) | $spl0),
  inference(forward_demodulation,[],[f40950,f183])).
fof(f40950,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,'const.relation.WF_1'(s(fun(X139,fun(X139,bool)),'const.relation.EMPTY_REL_0'))))),s(fun(X140,fun(bool,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X140,fun(bool,bool)),X142),s(X140,X141))),s(bool,X138))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X138)))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f40949,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X134,bool),'const.pair.,_2'(s(X134,X135),s(bool,'F_0'))),s(fun(X134,fun(bool,bool)),X136))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X134,fun(bool,bool)),X136),s(X134,X135))),s(bool,X133))) | s(bool,X133) = s(bool,X137) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X137)))) ) | $spl0),
  inference(superposition,[],[f3570,f755])).
fof(f40944,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X109))),s(fun(X110,fun(bool,bool)),X112))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X110,fun(bool,bool)),X112),s(X110,X111))),s(bool,X108))) | s(bool,'F_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X108)) ) | $spl0),
  inference(superposition,[],[f3570,f612])).
fof(f40943,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X104,bool),'const.pair.,_2'(s(X104,X105),s(bool,'F_0'))),s(fun(X104,fun(bool,bool)),X106))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X104,fun(bool,bool)),X106),s(X104,X105))),s(bool,X103))) | s(bool,X103) = s(bool,X107) | s(bool,'F_0') = s(bool,X107)) ) | $spl0),
  inference(superposition,[],[f3570,f612])).
fof(f40940,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X89))),s(fun(X90,fun(bool,bool)),X92))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X90,fun(bool,bool)),X92),s(X90,X91))),s(bool,X88))) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f3570,f205])).
fof(f40939,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X84,bool),'const.pair.,_2'(s(X84,X85),s(bool,'T_0'))),s(fun(X84,fun(bool,bool)),X86))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X84,fun(bool,bool)),X86),s(X84,X85))),s(bool,X83))) | s(bool,X83) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f3570,f205])).
fof(f40938,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X80,bool),'const.pair.,_2'(s(X80,X81),s(bool,'F_0'))),s(fun(X80,fun(bool,bool)),X82))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(X80,fun(bool,bool)),X82),s(X80,X81))),s(bool,X79))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f3570,f128])).
fof(f40936,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65,X73] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X71,X65),'const.pair.,_2'(s(X71,X72),s(X65,happ(s(fun(X67,X65),happ(s(fun(X66,fun(X67,X65)),X70),s(X66,X68))),s(X67,X69))))),s(fun(X71,fun(X65,bool)),X73))) | s(bool,'F_0') = s(bool,happ(s(fun(X65,bool),happ(s(fun(X71,fun(X65,bool)),X73),s(X71,X72))),s(X65,'const.pair.pair_CASE_2'(s('type.pair.prod'(X66,X67),'const.pair.,_2'(s(X66,X68),s(X67,X69))),s(fun(X66,fun(X67,X65)),X70)))))) )),
  inference(superposition,[],[f3570,f165])).
fof(f40934,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X48),'const.pair.,_2'(s(X54,X55),s(X48,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),X51),s(X48,X52),s(X48,X53),s(fun(X49,X48),X50))))),s(fun(X54,fun(X48,bool)),X56))) | s(bool,'F_0') = s(bool,happ(s(fun(X48,bool),happ(s(fun(X54,fun(X48,bool)),X56),s(X54,X55))),s(X48,happ(s(fun(X49,X48),X50),s(X49,sK2(X51,X49)))))) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X49),X51) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X49),X51)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f40933,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X45,X39),'const.pair.,_2'(s(X45,X46),s(X39,'const.pair.pair_CASE_2'(s('type.pair.prod'(X41,X40),'const.pair.,_2'(s(X41,X43),s(X40,X44))),s(fun(X41,fun(X40,X39)),X42))))),s(fun(X45,fun(X39,bool)),X47))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X45,fun(X39,bool)),X47),s(X45,X46))),s(X39,happ(s(fun(X40,X39),happ(s(fun(X41,fun(X40,X39)),X42),s(X41,X43))),s(X40,X44)))))) )),
  inference(superposition,[],[f3570,f165])).
fof(f40931,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X31,X25),'const.pair.,_2'(s(X31,X32),s(X25,happ(s(fun(X26,X25),X30),s(X26,sK2(X27,X26)))))),s(fun(X31,fun(X25,bool)),X33))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),happ(s(fun(X31,fun(X25,bool)),X33),s(X31,X32))),s(X25,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X26),X27),s(X25,X28),s(X25,X29),s(fun(X26,X25),X30))))) | s('type.semanticPrimitives.match_result'(X26),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X26),X27) | s('type.semanticPrimitives.match_result'(X26),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X26),X27)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f41314,plain,(
  ( ! [X472,X469,X471,X473,X470] : (s(bool,'T_0') = s(bool,happ(s(fun(X471,bool),happ(s(fun('type.semanticPrimitives.match_result'(X469),fun(X471,bool)),X473),s('type.semanticPrimitives.match_result'(X469),X470))),s(X471,X472))) | s(bool,'F_0') = s(bool,happ(s(fun(X471,bool),happ(s(fun('type.semanticPrimitives.match_result'(X469),fun(X471,bool)),X473),s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_1'(s(X469,sK2(X470,X469)))))),s(X471,X472))) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X469),X470) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X469),X470)) )),
  inference(forward_demodulation,[],[f40925,f165])).
fof(f40925,plain,(
  ( ! [X472,X469,X471,X473,X470] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.match_result'(X469),X471),'const.pair.,_2'(s('type.semanticPrimitives.match_result'(X469),X470),s(X471,X472))),s(fun('type.semanticPrimitives.match_result'(X469),fun(X471,bool)),X473))) | s(bool,'F_0') = s(bool,happ(s(fun(X471,bool),happ(s(fun('type.semanticPrimitives.match_result'(X469),fun(X471,bool)),X473),s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_1'(s(X469,sK2(X470,X469)))))),s(X471,X472))) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X469),X470) | s('type.semanticPrimitives.match_result'(X469),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X469),X470)) )),
  inference(superposition,[],[f3570,f144])).
fof(f41292,plain,(
  ( ! [X389,X385,X387,X388,X386] : (s(bool,'F_0') = s(bool,happ(s(fun(X386,bool),happ(s(fun(bool,fun(X386,bool)),X388),s(bool,X385))),s(X386,X387))) | s(bool,'T_0') = s(bool,happ(s(fun(X386,bool),happ(s(fun(bool,fun(X386,bool)),X388),s(bool,'F_0'))),s(X386,X387))) | s(bool,X385) = s(bool,'const.relation.WF_1'(s(fun(X389,fun(X389,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f41291,f141])).
fof(f41291,plain,(
  ( ! [X389,X385,X387,X388,X386] : (s(bool,'T_0') = s(bool,happ(s(fun(X386,bool),happ(s(fun(bool,fun(X386,bool)),X388),s(bool,'F_0'))),s(X386,X387))) | s(bool,'F_0') = s(bool,happ(s(fun(X386,bool),happ(s(fun(bool,fun(X386,bool)),X388),s(bool,'const.combin.I_1'(s(bool,X385))))),s(X386,X387))) | s(bool,X385) = s(bool,'const.relation.WF_1'(s(fun(X389,fun(X389,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f40905,f165])).
fof(f40905,plain,(
  ( ! [X389,X385,X387,X388,X386] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X386),'const.pair.,_2'(s(bool,'F_0'),s(X386,X387))),s(fun(bool,fun(X386,bool)),X388))) | s(bool,'F_0') = s(bool,happ(s(fun(X386,bool),happ(s(fun(bool,fun(X386,bool)),X388),s(bool,'const.combin.I_1'(s(bool,X385))))),s(X386,X387))) | s(bool,X385) = s(bool,'const.relation.WF_1'(s(fun(X389,fun(X389,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f41290,plain,(
  ( ! [X383,X382,X384,X381] : (s(bool,'F_0') = s(bool,happ(s(fun(X382,bool),happ(s(fun(bool,fun(X382,bool)),X384),s(bool,X381))),s(X382,X383))) | s(bool,'T_0') = s(bool,happ(s(fun(X382,bool),happ(s(fun(bool,fun(X382,bool)),X384),s(bool,'F_0'))),s(X382,X383))) | s(bool,'T_0') = s(bool,X381)) ) | $spl0),
  inference(forward_demodulation,[],[f41289,f141])).
fof(f41289,plain,(
  ( ! [X383,X382,X384,X381] : (s(bool,'T_0') = s(bool,happ(s(fun(X382,bool),happ(s(fun(bool,fun(X382,bool)),X384),s(bool,'F_0'))),s(X382,X383))) | s(bool,'F_0') = s(bool,happ(s(fun(X382,bool),happ(s(fun(bool,fun(X382,bool)),X384),s(bool,'const.combin.I_1'(s(bool,X381))))),s(X382,X383))) | s(bool,'T_0') = s(bool,X381)) ) | $spl0),
  inference(forward_demodulation,[],[f40904,f165])).
fof(f40904,plain,(
  ( ! [X383,X382,X384,X381] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X382),'const.pair.,_2'(s(bool,'F_0'),s(X382,X383))),s(fun(bool,fun(X382,bool)),X384))) | s(bool,'F_0') = s(bool,happ(s(fun(X382,bool),happ(s(fun(bool,fun(X382,bool)),X384),s(bool,'const.combin.I_1'(s(bool,X381))))),s(X382,X383))) | s(bool,'T_0') = s(bool,X381)) ) | $spl0),
  inference(superposition,[],[f3570,f1764])).
fof(f41288,plain,(
  ( ! [X377,X379,X380,X376,X378] : (s(bool,'F_0') = s(bool,happ(s(fun(X377,bool),happ(s(fun(bool,fun(X377,bool)),X379),s(bool,X376))),s(X377,X378))) | s(bool,'T_0') = s(bool,happ(s(fun(X377,bool),happ(s(fun(bool,fun(X377,bool)),X379),s(bool,'F_0'))),s(X377,X378))) | s(bool,X376) = s(bool,X380) | s(bool,'F_0') = s(bool,X380)) ) | $spl0),
  inference(forward_demodulation,[],[f41287,f141])).
fof(f41287,plain,(
  ( ! [X377,X379,X380,X376,X378] : (s(bool,'T_0') = s(bool,happ(s(fun(X377,bool),happ(s(fun(bool,fun(X377,bool)),X379),s(bool,'F_0'))),s(X377,X378))) | s(bool,'F_0') = s(bool,happ(s(fun(X377,bool),happ(s(fun(bool,fun(X377,bool)),X379),s(bool,'const.combin.I_1'(s(bool,X376))))),s(X377,X378))) | s(bool,X376) = s(bool,X380) | s(bool,'F_0') = s(bool,X380)) ) | $spl0),
  inference(forward_demodulation,[],[f40903,f165])).
fof(f40903,plain,(
  ( ! [X377,X379,X380,X376,X378] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X377),'const.pair.,_2'(s(bool,'F_0'),s(X377,X378))),s(fun(bool,fun(X377,bool)),X379))) | s(bool,'F_0') = s(bool,happ(s(fun(X377,bool),happ(s(fun(bool,fun(X377,bool)),X379),s(bool,'const.combin.I_1'(s(bool,X376))))),s(X377,X378))) | s(bool,X376) = s(bool,X380) | s(bool,'F_0') = s(bool,X380)) ) | $spl0),
  inference(superposition,[],[f3570,f755])).
fof(f41275,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'F_0') = s(bool,X348) | s(bool,'F_0') = s(bool,happ(s(fun(X349,bool),happ(s(fun(bool,fun(X349,bool)),X351),s(bool,'T_0'))),s(X349,X350))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X349),'const.pair.,_2'(s(bool,X348),s(X349,X350))),s(fun(bool,fun(X349,bool)),X351)))) ) | $spl0),
  inference(forward_demodulation,[],[f41274,f141])).
fof(f41274,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'F_0') = s(bool,happ(s(fun(X349,bool),happ(s(fun(bool,fun(X349,bool)),X351),s(bool,'T_0'))),s(X349,X350))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X349),'const.pair.,_2'(s(bool,X348),s(X349,X350))),s(fun(bool,fun(X349,bool)),X351))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X348)))) ) | $spl0),
  inference(forward_demodulation,[],[f40896,f183])).
fof(f40896,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X349),'const.pair.,_2'(s(bool,X348),s(X349,X350))),s(fun(bool,fun(X349,bool)),X351))) | s(bool,'F_0') = s(bool,happ(s(fun(X349,bool),happ(s(fun(bool,fun(X349,bool)),X351),s(bool,'const.relation.WF_1'(s(fun(X347,fun(X347,bool)),'const.relation.EMPTY_REL_0'))))),s(X349,X350))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X348)))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f41259,plain,(
  ( ! [X300,X304,X303,X305] : (s(bool,'F_0') = s(bool,happ(s(fun(X303,bool),happ(s(fun(bool,fun(X303,bool)),X305),s(bool,X300))),s(X303,X304))) | s(bool,'T_0') = s(bool,happ(s(fun(X303,bool),happ(s(fun(bool,fun(X303,bool)),X305),s(bool,'T_0'))),s(X303,X304))) | s(bool,'F_0') = s(bool,X300)) )),
  inference(forward_demodulation,[],[f41258,f160])).
fof(f41258,plain,(
  ( ! [X300,X302,X304,X301,X303,X299,X305] : (s(bool,'T_0') = s(bool,happ(s(fun(X303,bool),happ(s(fun(bool,fun(X303,bool)),X305),s(bool,'T_0'))),s(X303,X304))) | s(bool,'F_0') = s(bool,happ(s(fun(X303,bool),happ(s(fun(bool,fun(X303,bool)),X305),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X299),'const.semanticPrimitives.No_match_0'),s(bool,X300),s(bool,X301),s(fun(X299,bool),X302))))),s(X303,X304))) | s(bool,'F_0') = s(bool,X300)) )),
  inference(forward_demodulation,[],[f40888,f165])).
fof(f40888,plain,(
  ( ! [X300,X302,X304,X301,X303,X299,X305] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X303),'const.pair.,_2'(s(bool,'T_0'),s(X303,X304))),s(fun(bool,fun(X303,bool)),X305))) | s(bool,'F_0') = s(bool,happ(s(fun(X303,bool),happ(s(fun(bool,fun(X303,bool)),X305),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X299),'const.semanticPrimitives.No_match_0'),s(bool,X300),s(bool,X301),s(fun(X299,bool),X302))))),s(X303,X304))) | s(bool,'F_0') = s(bool,X300)) )),
  inference(superposition,[],[f3570,f316])).
fof(f40880,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X269),'const.pair.,_2'(s(bool,X268),s(X269,X270))),s(fun(bool,fun(X269,bool)),X271))) | s(bool,'F_0') = s(bool,happ(s(fun(X269,bool),happ(s(fun(bool,fun(X269,bool)),X271),s(bool,'F_0'))),s(X269,X270))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f3570,f128])).
fof(f40877,plain,(
  ( ! [X255,X253,X257,X254,X256] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X254),'const.pair.,_2'(s(bool,X253),s(X254,X255))),s(fun(bool,fun(X254,bool)),X256))) | s(bool,'F_0') = s(bool,happ(s(fun(X254,bool),happ(s(fun(bool,fun(X254,bool)),X256),s(bool,'F_0'))),s(X254,X255))) | s(bool,X253) = s(bool,X257) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X257)))) ) | $spl0),
  inference(superposition,[],[f3570,f755])).
fof(f41243,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'F_0') = s(bool,X252) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X249),'const.pair.,_2'(s(bool,X248),s(X249,X250))),s(fun(bool,fun(X249,bool)),X251))) | s(bool,'F_0') = s(bool,happ(s(fun(X249,bool),happ(s(fun(bool,fun(X249,bool)),X251),s(bool,'F_0'))),s(X249,X250))) | s(bool,X248) = s(bool,X252)) ) | $spl0),
  inference(forward_demodulation,[],[f40876,f141])).
fof(f40876,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X249),'const.pair.,_2'(s(bool,X248),s(X249,X250))),s(fun(bool,fun(X249,bool)),X251))) | s(bool,'F_0') = s(bool,happ(s(fun(X249,bool),happ(s(fun(bool,fun(X249,bool)),X251),s(bool,'F_0'))),s(X249,X250))) | s(bool,X252) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X248) = s(bool,X252)) ) | $spl0),
  inference(superposition,[],[f3570,f11021])).
fof(f41219,plain,(
  ( ! [X142,X140,X138,X141] : (s(bool,'F_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,'T_0'),s(X140,X141))),s(fun(bool,fun(X140,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(X140,bool),happ(s(fun(bool,fun(X140,bool)),X142),s(bool,X138))),s(X140,X141)))) ) | $spl0),
  inference(forward_demodulation,[],[f41218,f141])).
fof(f41218,plain,(
  ( ! [X142,X140,X138,X141] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,'T_0'),s(X140,X141))),s(fun(bool,fun(X140,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(X140,bool),happ(s(fun(bool,fun(X140,bool)),X142),s(bool,X138))),s(X140,X141))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X138)))) ) | $spl0),
  inference(forward_demodulation,[],[f40853,f183])).
fof(f40853,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,'const.relation.WF_1'(s(fun(X139,fun(X139,bool)),'const.relation.EMPTY_REL_0'))),s(X140,X141))),s(fun(bool,fun(X140,bool)),X142))) | s(bool,'F_0') = s(bool,happ(s(fun(X140,bool),happ(s(fun(bool,fun(X140,bool)),X142),s(bool,X138))),s(X140,X141))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X138)))) ) | $spl0),
  inference(superposition,[],[f3570,f1796])).
fof(f40852,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X134),'const.pair.,_2'(s(bool,'F_0'),s(X134,X135))),s(fun(bool,fun(X134,bool)),X136))) | s(bool,'F_0') = s(bool,happ(s(fun(X134,bool),happ(s(fun(bool,fun(X134,bool)),X136),s(bool,X133))),s(X134,X135))) | s(bool,X133) = s(bool,X137) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X137)))) ) | $spl0),
  inference(superposition,[],[f3570,f755])).
fof(f40847,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X110),'const.pair.,_2'(s(bool,X109),s(X110,X111))),s(fun(bool,fun(X110,bool)),X112))) | s(bool,'F_0') = s(bool,happ(s(fun(X110,bool),happ(s(fun(bool,fun(X110,bool)),X112),s(bool,X108))),s(X110,X111))) | s(bool,'F_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X108)) ) | $spl0),
  inference(superposition,[],[f3570,f612])).
fof(f40846,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X104),'const.pair.,_2'(s(bool,'F_0'),s(X104,X105))),s(fun(bool,fun(X104,bool)),X106))) | s(bool,'F_0') = s(bool,happ(s(fun(X104,bool),happ(s(fun(bool,fun(X104,bool)),X106),s(bool,X103))),s(X104,X105))) | s(bool,X103) = s(bool,X107) | s(bool,'F_0') = s(bool,X107)) ) | $spl0),
  inference(superposition,[],[f3570,f612])).
fof(f40843,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X90),'const.pair.,_2'(s(bool,X89),s(X90,X91))),s(fun(bool,fun(X90,bool)),X92))) | s(bool,'F_0') = s(bool,happ(s(fun(X90,bool),happ(s(fun(bool,fun(X90,bool)),X92),s(bool,X88))),s(X90,X91))) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f3570,f205])).
fof(f40842,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X84),'const.pair.,_2'(s(bool,'T_0'),s(X84,X85))),s(fun(bool,fun(X84,bool)),X86))) | s(bool,'F_0') = s(bool,happ(s(fun(X84,bool),happ(s(fun(bool,fun(X84,bool)),X86),s(bool,X83))),s(X84,X85))) | s(bool,X83) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f3570,f205])).
fof(f40841,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(bool,X80),'const.pair.,_2'(s(bool,'F_0'),s(X80,X81))),s(fun(bool,fun(X80,bool)),X82))) | s(bool,'F_0') = s(bool,happ(s(fun(X80,bool),happ(s(fun(bool,fun(X80,bool)),X82),s(bool,X79))),s(X80,X81))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f3570,f128])).
fof(f40839,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65,X73] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X65,X71),'const.pair.,_2'(s(X65,happ(s(fun(X67,X65),happ(s(fun(X66,fun(X67,X65)),X70),s(X66,X68))),s(X67,X69))),s(X71,X72))),s(fun(X65,fun(X71,bool)),X73))) | s(bool,'F_0') = s(bool,happ(s(fun(X71,bool),happ(s(fun(X65,fun(X71,bool)),X73),s(X65,'const.pair.pair_CASE_2'(s('type.pair.prod'(X66,X67),'const.pair.,_2'(s(X66,X68),s(X67,X69))),s(fun(X66,fun(X67,X65)),X70))))),s(X71,X72)))) )),
  inference(superposition,[],[f3570,f165])).
fof(f40837,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X48,X54),'const.pair.,_2'(s(X48,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),X51),s(X48,X52),s(X48,X53),s(fun(X49,X48),X50))),s(X54,X55))),s(fun(X48,fun(X54,bool)),X56))) | s(bool,'F_0') = s(bool,happ(s(fun(X54,bool),happ(s(fun(X48,fun(X54,bool)),X56),s(X48,happ(s(fun(X49,X48),X50),s(X49,sK2(X51,X49)))))),s(X54,X55))) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X49),X51) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X49),X51)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f40836,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X45),'const.pair.,_2'(s(X39,'const.pair.pair_CASE_2'(s('type.pair.prod'(X41,X40),'const.pair.,_2'(s(X41,X43),s(X40,X44))),s(fun(X41,fun(X40,X39)),X42))),s(X45,X46))),s(fun(X39,fun(X45,bool)),X47))) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),happ(s(fun(X39,fun(X45,bool)),X47),s(X39,happ(s(fun(X40,X39),happ(s(fun(X41,fun(X40,X39)),X42),s(X41,X43))),s(X40,X44))))),s(X45,X46)))) )),
  inference(superposition,[],[f3570,f165])).
fof(f40834,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X25,X31),'const.pair.,_2'(s(X25,happ(s(fun(X26,X25),X30),s(X26,sK2(X27,X26)))),s(X31,X32))),s(fun(X25,fun(X31,bool)),X33))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),happ(s(fun(X25,fun(X31,bool)),X33),s(X25,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X26),X27),s(X25,X28),s(X25,X29),s(fun(X26,X25),X30))))),s(X31,X32))) | s('type.semanticPrimitives.match_result'(X26),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X26),X27) | s('type.semanticPrimitives.match_result'(X26),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X26),X27)) )),
  inference(superposition,[],[f3570,f2917])).
fof(f39982,plain,(
  ( ! [X947,X945,X946,X944,X948,X943] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X948),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X943,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X947),s(X943,sK2(X944,X943)))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X943),X944),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X945),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X946),s(fun(X943,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X947))) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X943),X944) | s('type.semanticPrimitives.match_result'(X943),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X943),X944)) )),
  inference(superposition,[],[f157,f2917])).
fof(f39981,plain,(
  ( ! [X938,X942,X940,X939,X937,X941] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X937,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X941),s(X937,sK2(X938,X937)))) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X942),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X937),X938),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X939),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X940),s(fun(X937,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X941))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'(X937),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X937),X938) | s('type.semanticPrimitives.match_result'(X937),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X937),X938)) )),
  inference(superposition,[],[f158,f2917])).
fof(f39963,plain,(
  ( ! [X819,X817,X810,X814,X812,X818,X816,X811,X815,X813] : (~p(s(bool,'const.relation.WF_1'(s(fun(X810,fun(X810,bool)),happ(s(fun(X811,fun(X810,fun(X810,bool))),X815),s(X811,sK2(X812,X811))))))) | s(X816,happ(s(fun(X810,X816),X817),s(X810,X818))) = s(X816,happ(s(fun(X810,X816),happ(s(fun(fun(X810,X816),fun(X810,X816)),X819),s(fun(X810,X816),'const.relation.RESTRICT_3'(s(fun(X810,X816),X817),s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(X810,X818))))),s(X810,X818))) | s(fun(X810,X816),X817) != s(fun(X810,X816),'const.relation.WFREC_2'(s(fun(X810,fun(X810,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X811),X812),s(fun(X810,fun(X810,bool)),X813),s(fun(X810,fun(X810,bool)),X814),s(fun(X811,fun(X810,fun(X810,bool))),X815))),s(fun(fun(X810,X816),fun(X810,X816)),X819))) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X811),X812) | s('type.semanticPrimitives.match_result'(X811),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X811),X812)) )),
  inference(superposition,[],[f162,f2917])).
fof(f39962,plain,(
  ( ! [X808,X806,X804,X809,X807,X805] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X804,fun(X804,bool)),bool),'const.relation.WF_0'),s(fun(X804,fun(X804,bool)),happ(s(fun(X805,fun(X804,fun(X804,bool))),X809),s(X805,sK2(X806,X805)))))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X804,fun(X804,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X805),X806),s(fun(X804,fun(X804,bool)),X807),s(fun(X804,fun(X804,bool)),X808),s(fun(X805,fun(X804,fun(X804,bool))),X809))))) | s('type.semanticPrimitives.match_result'(X805),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X805),X806) | s('type.semanticPrimitives.match_result'(X805),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X805),X806)) )),
  inference(superposition,[],[f546,f2917])).
fof(f39961,plain,(
  ( ! [X798,X802,X800,X799,X797,X803,X801] : (~p(s(bool,happ(s(fun(X797,bool),happ(s(fun(X798,fun(X797,bool)),X802),s(X798,sK2(X799,X798)))),s(X797,X803)))) | p(s(bool,happ(s(fun(X797,bool),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X798),X799),s(fun(X797,bool),X800),s(fun(X797,bool),X801),s(fun(X798,fun(X797,bool)),X802))),s(X797,'const.min.@_1'(s(fun(X797,bool),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X798),X799),s(fun(X797,bool),X800),s(fun(X797,bool),X801),s(fun(X798,fun(X797,bool)),X802)))))))) | s('type.semanticPrimitives.match_result'(X798),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X798),X799) | s('type.semanticPrimitives.match_result'(X798),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X798),X799)) )),
  inference(superposition,[],[f149,f2917])).
fof(f39958,plain,(
  ( ! [X779,X777,X775,X773,X778,X776,X780,X774,X772] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X772),'const.semanticPrimitives.Match_1'(s(X772,X778))),s(bool,X779),s(bool,X780),s(fun(X772,bool),happ(s(fun(X773,fun(X772,bool)),X777),s(X773,sK2(X774,X773)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X772,bool),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X773),X774),s(fun(X772,bool),X775),s(fun(X772,bool),X776),s(fun(X773,fun(X772,bool)),X777))),s(X772,X778))) | s('type.semanticPrimitives.match_result'(X773),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X773),X774) | s('type.semanticPrimitives.match_result'(X773),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X773),X774)) )),
  inference(superposition,[],[f2989,f2917])).
fof(f40176,plain,(
  ( ! [X736,X732,X734,X733,X735,X731] : (s(bool,'T_0') = s(bool,X736) | s(bool,'F_0') != s(bool,happ(s(fun(X731,bool),X735),s(X731,sK2(X732,X731)))) | s(bool,X736) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X731),X732),s(bool,X733),s(bool,X734),s(fun(X731,bool),X735))) | s('type.semanticPrimitives.match_result'(X731),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X731),X732) | s('type.semanticPrimitives.match_result'(X731),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X731),X732)) ) | $spl0),
  inference(forward_demodulation,[],[f39951,f141])).
fof(f39951,plain,(
  ( ! [X736,X732,X734,X733,X735,X731] : (s(bool,'F_0') != s(bool,happ(s(fun(X731,bool),X735),s(X731,sK2(X732,X731)))) | s(bool,X736) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X736) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X731),X732),s(bool,X733),s(bool,X734),s(fun(X731,bool),X735))) | s('type.semanticPrimitives.match_result'(X731),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X731),X732) | s('type.semanticPrimitives.match_result'(X731),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X731),X732)) ) | $spl0),
  inference(superposition,[],[f16731,f2917])).
fof(f40171,plain,(
  ( ! [X705,X707,X704,X706,X703] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X703),X704),s(bool,X705),s(bool,X706),s(fun(X703,bool),X707))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X703,bool),X707),s(X703,sK2(X704,X703)))))) | s('type.semanticPrimitives.match_result'(X703),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X703),X704) | s('type.semanticPrimitives.match_result'(X703),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X703),X704)) ) | $spl0),
  inference(forward_demodulation,[],[f39946,f141])).
fof(f39946,plain,(
  ( ! [X707,X705,X704,X706,X703] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X703,bool),X707),s(X703,sK2(X704,X703)))))) | s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X703),X704),s(bool,X705),s(bool,X706),s(fun(X703,bool),X707))) | s('type.semanticPrimitives.match_result'(X703),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X703),X704) | s('type.semanticPrimitives.match_result'(X703),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X703),X704)) ) | $spl0),
  inference(superposition,[],[f11207,f2917])).
fof(f39902,plain,(
  ( ! [X453,X451,X452,X454,X450] : (s(bool,'T_0') != s(bool,happ(s(fun(X450,bool),X454),s(X450,sK2(X451,X450)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X450),X451),s(bool,X452),s(bool,X453),s(fun(X450,bool),X454))))) | s('type.semanticPrimitives.match_result'(X450),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X450),X451) | s('type.semanticPrimitives.match_result'(X450),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X450),X451)) ) | $spl0),
  inference(superposition,[],[f2368,f2917])).
fof(f39901,plain,(
  ( ! [X445,X447,X449,X444,X446,X448] : (s(bool,X449) != s(bool,happ(s(fun(X444,bool),X448),s(X444,sK2(X445,X444)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X444),X445),s(bool,X446),s(bool,X447),s(fun(X444,bool),X448))))) | s(bool,'T_0') = s(bool,X449) | s('type.semanticPrimitives.match_result'(X444),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X444),X445) | s('type.semanticPrimitives.match_result'(X444),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X444),X445)) ) | $spl0),
  inference(superposition,[],[f2234,f2917])).
fof(f39900,plain,(
  ( ! [X441,X443,X438,X440,X442,X439] : (s(bool,X443) != s(bool,happ(s(fun(X438,bool),X442),s(X438,sK2(X439,X438)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X443))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X438),X439),s(bool,X440),s(bool,X441),s(fun(X438,bool),X442))) | s('type.semanticPrimitives.match_result'(X438),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X438),X439) | s('type.semanticPrimitives.match_result'(X438),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X438),X439)) ) | $spl0),
  inference(superposition,[],[f2234,f2917])).
fof(f39898,plain,(
  ( ! [X428,X430,X426,X429,X431,X427] : (s(bool,X431) = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X426,bool),X430),s(X426,sK2(X427,X426)))))) | s(bool,'T_0') = s(bool,X431) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X426),X427),s(bool,X428),s(bool,X429),s(fun(X426,bool),X430))) | s('type.semanticPrimitives.match_result'(X426),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X426),X427) | s('type.semanticPrimitives.match_result'(X426),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X426),X427)) ) | $spl0),
  inference(superposition,[],[f2195,f2917])).
fof(f39897,plain,(
  ( ! [X422,X424,X421,X423,X425] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X421,bool),X425),s(X421,sK2(X422,X421)))))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X421),X422),s(bool,X423),s(bool,X424),s(fun(X421,bool),X425))) | s('type.semanticPrimitives.match_result'(X421),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X421),X422) | s('type.semanticPrimitives.match_result'(X421),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X421),X422)) ) | $spl0),
  inference(superposition,[],[f2047,f2917])).
fof(f39895,plain,(
  ( ! [X412,X414,X410,X413,X415,X411] : (s(bool,'const.combin.I_1'(s(bool,X415))) != s(bool,happ(s(fun(X410,bool),X414),s(X410,sK2(X411,X410)))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X410),X411),s(bool,X412),s(bool,X413),s(fun(X410,bool),X414))) | s(bool,'T_0') = s(bool,X415) | s('type.semanticPrimitives.match_result'(X410),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X410),X411) | s('type.semanticPrimitives.match_result'(X410),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X410),X411)) ) | $spl0),
  inference(superposition,[],[f1916,f2917])).
fof(f39894,plain,(
  ( ! [X408,X405,X407,X409,X404,X406] : (s(bool,X409) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X404,bool),X408),s(X404,sK2(X405,X404)))))) | s(bool,'F_0') = s(bool,X409) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X404),X405),s(bool,X406),s(bool,X407),s(fun(X404,bool),X408))) | s('type.semanticPrimitives.match_result'(X404),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X404),X405) | s('type.semanticPrimitives.match_result'(X404),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X404),X405)) ) | $spl0),
  inference(superposition,[],[f1916,f2917])).
fof(f39893,plain,(
  ( ! [X399,X401,X403,X398,X400,X402] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X398,bool),X402),s(X398,sK2(X399,X398)))))) | s(bool,X403) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X398),X399),s(bool,X400),s(bool,X401),s(fun(X398,bool),X402))) | s(bool,'F_0') = s(bool,X403) | s('type.semanticPrimitives.match_result'(X398),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X398),X399) | s('type.semanticPrimitives.match_result'(X398),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X398),X399)) ) | $spl0),
  inference(superposition,[],[f1875,f2917])).
fof(f39892,plain,(
  ( ! [X397,X393,X395,X396,X394] : (s(bool,'F_0') != s(bool,happ(s(fun(X393,bool),X397),s(X393,sK2(X394,X393)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X393),X394),s(bool,X395),s(bool,X396),s(fun(X393,bool),X397))))) | s('type.semanticPrimitives.match_result'(X393),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X393),X394) | s('type.semanticPrimitives.match_result'(X393),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X393),X394)) ) | $spl0),
  inference(superposition,[],[f1858,f2917])).
fof(f40132,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,'F_0') != s(bool,happ(s(fun(X361,bool),X365),s(X361,sK2(X362,X361)))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X361),X362),s(bool,X363),s(bool,X364),s(fun(X361,bool),X365))) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X361),X362) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X361),X362)) ) | $spl0),
  inference(forward_demodulation,[],[f39886,f141])).
fof(f39886,plain,(
  ( ! [X364,X362,X365,X361,X363] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,happ(s(fun(X361,bool),X365),s(X361,sK2(X362,X361)))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X361),X362),s(bool,X363),s(bool,X364),s(fun(X361,bool),X365))) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X361),X362) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X361),X362)) ) | $spl0),
  inference(superposition,[],[f1629,f2917])).
fof(f39885,plain,(
  ( ! [X356,X358,X360,X357,X359,X355] : (s(bool,X360) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X355,bool),X359),s(X355,sK2(X356,X355)))))) | s(bool,X360) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X355),X356),s(bool,X357),s(bool,X358),s(fun(X355,bool),X359))) | s(bool,'F_0') = s(bool,X360) | s('type.semanticPrimitives.match_result'(X355),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X355),X356) | s('type.semanticPrimitives.match_result'(X355),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X355),X356)) ) | $spl0),
  inference(superposition,[],[f1617,f2917])).
fof(f39884,plain,(
  ( ! [X350,X352,X354,X349,X351,X353] : (s(bool,'const.combin.I_1'(s(bool,X354))) != s(bool,happ(s(fun(X349,bool),X353),s(X349,sK2(X350,X349)))) | s(bool,X354) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X349),X350),s(bool,X351),s(bool,X352),s(fun(X349,bool),X353))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X349),X350),s(bool,X351),s(bool,X352),s(fun(X349,bool),X353))) | s('type.semanticPrimitives.match_result'(X349),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X349),X350) | s('type.semanticPrimitives.match_result'(X349),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X349),X350)) ) | $spl0),
  inference(superposition,[],[f1617,f2917])).
fof(f39883,plain,(
  ( ! [X344,X346,X343,X345,X347,X348] : (s(bool,'T_0') != s(bool,happ(s(fun(X343,bool),X347),s(X343,sK2(X344,X343)))) | s(bool,X348) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X343),X344),s(bool,X345),s(bool,X346),s(fun(X343,bool),X347))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X348))) | s('type.semanticPrimitives.match_result'(X343),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X343),X344) | s('type.semanticPrimitives.match_result'(X343),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X343),X344)) ) | $spl0),
  inference(superposition,[],[f1515,f2917])).
fof(f39882,plain,(
  ( ! [X341,X337,X339,X340,X342,X338] : (s(bool,'T_0') != s(bool,happ(s(fun(X337,bool),X341),s(X337,sK2(X338,X337)))) | s(bool,X342) = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X337),X338),s(bool,X339),s(bool,X340),s(fun(X337,bool),X341))))) | s(bool,'F_0') = s(bool,X342) | s('type.semanticPrimitives.match_result'(X337),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X337),X338) | s('type.semanticPrimitives.match_result'(X337),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X337),X338)) ) | $spl0),
  inference(superposition,[],[f1236,f2917])).
fof(f39881,plain,(
  ( ! [X333,X335,X331,X332,X334,X330,X336] : (p(s(bool,happ(s(fun(X330,bool),X334),s(X330,sK2(X331,X330))))) | s(bool,X335) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X330),X331),s(bool,X332),s(bool,X333),s(fun(X330,bool),X334))) | s(bool,X335) = s(bool,X336) | s(bool,'F_0') = s(bool,X336) | s('type.semanticPrimitives.match_result'(X330),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X330),X331) | s('type.semanticPrimitives.match_result'(X330),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X330),X331)) ) | $spl0),
  inference(superposition,[],[f1204,f2917])).
fof(f39880,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (s(bool,'T_0') != s(bool,happ(s(fun(X324,bool),X328),s(X324,sK2(X325,X324)))) | s(bool,X329) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X324),X325),s(bool,X326),s(bool,X327),s(fun(X324,bool),X328))) | s(bool,'F_0') = s(bool,X329) | s('type.semanticPrimitives.match_result'(X324),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X324),X325) | s('type.semanticPrimitives.match_result'(X324),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X324),X325)) ) | $spl0),
  inference(superposition,[],[f900,f2917])).
fof(f39871,plain,(
  ( ! [X269,X271,X273,X275,X268,X270,X272,X274] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X273),'const.semanticPrimitives.Match_type_error_0'),s(bool,X274),s(bool,happ(s(fun(X268,bool),X272),s(X268,sK2(X269,X268)))),s(fun(X273,bool),X275))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X268),X269),s(bool,X270),s(bool,X271),s(fun(X268,bool),X272))) | s('type.semanticPrimitives.match_result'(X268),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X268),X269) | s('type.semanticPrimitives.match_result'(X268),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X268),X269)) )),
  inference(superposition,[],[f414,f2917])).
fof(f39868,plain,(
  ( ! [X251,X249,X246,X244,X250,X248,X247,X245] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X249),'const.semanticPrimitives.No_match_0'),s(bool,happ(s(fun(X244,bool),X248),s(X244,sK2(X245,X244)))),s(bool,X250),s(fun(X249,bool),X251))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X244),X245),s(bool,X246),s(bool,X247),s(fun(X244,bool),X248))) | s('type.semanticPrimitives.match_result'(X244),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X244),X245) | s('type.semanticPrimitives.match_result'(X244),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X244),X245)) )),
  inference(superposition,[],[f316,f2917])).
fof(f39867,plain,(
  ( ! [X243,X241,X239,X242,X240] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun(X239,bool),X243),s(X239,sK2(X240,X239)))))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X239),X240),s(bool,X241),s(bool,X242),s(fun(X239,bool),X243))) | s('type.semanticPrimitives.match_result'(X239),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X239),X240) | s('type.semanticPrimitives.match_result'(X239),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X239),X240)) )),
  inference(superposition,[],[f278,f2917])).
fof(f39866,plain,(
  ( ! [X237,X235,X233,X238,X236,X234] : (p(s(bool,happ(s(fun(X233,bool),X237),s(X233,sK2(X234,X233))))) | s(bool,X238) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X233),X234),s(bool,X235),s(bool,X236),s(fun(X233,bool),X237))) | s(bool,'T_0') = s(bool,X238) | s('type.semanticPrimitives.match_result'(X233),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X233),X234) | s('type.semanticPrimitives.match_result'(X233),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X233),X234)) )),
  inference(superposition,[],[f256,f2917])).
fof(f39860,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (~p(s(bool,happ(s(fun(X198,bool),X202),s(X198,sK2(X199,X198))))) | ~p(s(bool,X203)) | s(bool,X203) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X198),X199),s(bool,X200),s(bool,X201),s(fun(X198,bool),X202))) | s('type.semanticPrimitives.match_result'(X198),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X198),X199) | s('type.semanticPrimitives.match_result'(X198),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X198),X199)) )),
  inference(superposition,[],[f156,f2917])).
fof(f39858,plain,(
  ( ! [X191,X189,X187,X190,X188,X186] : (p(s(bool,happ(s(fun(X186,bool),X190),s(X186,sK2(X187,X186))))) | p(s(bool,X191)) | s(bool,X191) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X186),X187),s(bool,X188),s(bool,X189),s(fun(X186,bool),X190))) | s('type.semanticPrimitives.match_result'(X186),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X186),X187) | s('type.semanticPrimitives.match_result'(X186),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X186),X187)) )),
  inference(superposition,[],[f153,f2917])).
fof(f39857,plain,(
  ( ! [X185,X182,X184,X183,X181] : (p(s(bool,happ(s(fun(X181,bool),X185),s(X181,sK2(X182,X181))))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X181),X182),s(bool,X183),s(bool,X184),s(fun(X181,bool),X185))) | s('type.semanticPrimitives.match_result'(X181),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X181),X182) | s('type.semanticPrimitives.match_result'(X181),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X181),X182)) )),
  inference(superposition,[],[f135,f2917])).
fof(f39856,plain,(
  ( ! [X177,X180,X178,X176,X179] : (~p(s(bool,happ(s(fun(X176,bool),X180),s(X176,sK2(X177,X176))))) | s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X176),X177),s(bool,X178),s(bool,X179),s(fun(X176,bool),X180))) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X176),X177) | s('type.semanticPrimitives.match_result'(X176),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X176),X177)) )),
  inference(superposition,[],[f134,f2917])).
fof(f39855,plain,(
  ( ! [X175,X173,X171,X174,X172] : (~p(s(bool,happ(s(fun(X171,bool),X175),s(X171,sK2(X172,X171))))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X171),X172),s(bool,X173),s(bool,X174),s(fun(X171,bool),X175))) | s('type.semanticPrimitives.match_result'(X171),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X171),X172) | s('type.semanticPrimitives.match_result'(X171),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X171),X172)) )),
  inference(superposition,[],[f131,f2917])).
fof(f39854,plain,(
  ( ! [X167,X169,X166,X170,X168] : (p(s(bool,happ(s(fun(X166,bool),X170),s(X166,sK2(X167,X166))))) | s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X166),X167),s(bool,X168),s(bool,X169),s(fun(X166,bool),X170))) | s('type.semanticPrimitives.match_result'(X166),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X166),X167) | s('type.semanticPrimitives.match_result'(X166),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X166),X167)) )),
  inference(superposition,[],[f130,f2917])).
fof(f39849,plain,(
  ( ! [X146,X144,X142,X140,X147,X145,X143,X141,X139] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X139),'const.semanticPrimitives.Match_1'(s(X139,happ(s(fun(X140,X139),X144),s(X140,sK2(X141,X140)))))),s(bool,X145),s(bool,X146),s(fun(X139,bool),X147))) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),X147),s(X139,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X140),X141),s(X139,X142),s(X139,X143),s(fun(X140,X139),X144))))) | s('type.semanticPrimitives.match_result'(X140),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X140),X141) | s('type.semanticPrimitives.match_result'(X140),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X140),X141)) )),
  inference(superposition,[],[f2989,f2917])).
fof(f39829,plain,(
  ( ! [X963,X967,X965,X968,X966,X964] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X968),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X963),X965),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X966),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X967),s(fun(X963,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X964))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X963,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X964),s(X963,sK2(X965,X963)))) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X963),X965) | s('type.semanticPrimitives.match_result'(X963),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X963),X965)) )),
  inference(superposition,[],[f157,f2917])).
fof(f39828,plain,(
  ( ! [X957,X961,X958,X962,X960,X959] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X957),X959),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X960),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X961),s(fun(X957,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X958))) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X962),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X957,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),X958),s(X957,sK2(X959,X957)))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'(X957),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X957),X959) | s('type.semanticPrimitives.match_result'(X957),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X957),X959)) )),
  inference(superposition,[],[f158,f2917])).
fof(f39810,plain,(
  ( ! [X825,X831,X829,X833,X826,X824,X830,X828,X832,X827] : (~p(s(bool,'const.relation.WF_1'(s(fun(X824,fun(X824,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X825),X827),s(fun(X824,fun(X824,bool)),X828),s(fun(X824,fun(X824,bool)),X829),s(fun(X825,fun(X824,fun(X824,bool))),X826)))))) | s(X830,happ(s(fun(X824,X830),X831),s(X824,X832))) = s(X830,happ(s(fun(X824,X830),happ(s(fun(fun(X824,X830),fun(X824,X830)),X833),s(fun(X824,X830),'const.relation.RESTRICT_3'(s(fun(X824,X830),X831),s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(X824,X832))))),s(X824,X832))) | s(fun(X824,X830),X831) != s(fun(X824,X830),'const.relation.WFREC_2'(s(fun(X824,fun(X824,bool)),happ(s(fun(X825,fun(X824,fun(X824,bool))),X826),s(X825,sK2(X827,X825)))),s(fun(fun(X824,X830),fun(X824,X830)),X833))) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X825),X827) | s('type.semanticPrimitives.match_result'(X825),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X825),X827)) )),
  inference(superposition,[],[f162,f2917])).
fof(f39809,plain,(
  ( ! [X819,X823,X821,X818,X822,X820] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X818,fun(X818,bool)),bool),'const.relation.WF_0'),s(fun(X818,fun(X818,bool)),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X819),X821),s(fun(X818,fun(X818,bool)),X822),s(fun(X818,fun(X818,bool)),X823),s(fun(X819,fun(X818,fun(X818,bool))),X820))))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X818,fun(X818,bool)),happ(s(fun(X819,fun(X818,fun(X818,bool))),X820),s(X819,sK2(X821,X819)))))) | s('type.semanticPrimitives.match_result'(X819),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X819),X821) | s('type.semanticPrimitives.match_result'(X819),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X819),X821)) )),
  inference(superposition,[],[f546,f2917])).
fof(f39808,plain,(
  ( ! [X817,X814,X812,X816,X811,X815,X813] : (~p(s(bool,happ(s(fun(X811,bool),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X812),X814),s(fun(X811,bool),X815),s(fun(X811,bool),X816),s(fun(X812,fun(X811,bool)),X813))),s(X811,X817)))) | p(s(bool,happ(s(fun(X811,bool),happ(s(fun(X812,fun(X811,bool)),X813),s(X812,sK2(X814,X812)))),s(X811,'const.min.@_1'(s(fun(X811,bool),happ(s(fun(X812,fun(X811,bool)),X813),s(X812,sK2(X814,X812))))))))) | s('type.semanticPrimitives.match_result'(X812),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X812),X814) | s('type.semanticPrimitives.match_result'(X812),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X812),X814)) )),
  inference(superposition,[],[f149,f2917])).
fof(f39805,plain,(
  ( ! [X786,X790,X788,X793,X787,X791,X789,X794,X792] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X786),'const.semanticPrimitives.Match_1'(s(X786,X792))),s(bool,X793),s(bool,X794),s(fun(X786,bool),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X787),X789),s(fun(X786,bool),X790),s(fun(X786,bool),X791),s(fun(X787,fun(X786,bool)),X788))))) | s(bool,'F_0') = s(bool,happ(s(fun(X786,bool),happ(s(fun(X787,fun(X786,bool)),X788),s(X787,sK2(X789,X787)))),s(X786,X792))) | s('type.semanticPrimitives.match_result'(X787),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X787),X789) | s('type.semanticPrimitives.match_result'(X787),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X787),X789)) )),
  inference(superposition,[],[f2989,f2917])).
fof(f39804,plain,(
  ( ! [X784,X779,X783,X781,X785,X778,X782,X780] : (s(X779,'const.pair.pair_CASE_2'(s('type.pair.prod'(X780,X778),'const.pair.,_2'(s(X780,sK2(X782,X780)),s(X778,X785))),s(fun(X780,fun(X778,X779)),X781))) = s(X779,happ(s(fun(X778,X779),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X780),X782),s(fun(X778,X779),X783),s(fun(X778,X779),X784),s(fun(X780,fun(X778,X779)),X781))),s(X778,X785))) | s('type.semanticPrimitives.match_result'(X780),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X780),X782) | s('type.semanticPrimitives.match_result'(X780),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X780),X782)) )),
  inference(superposition,[],[f165,f2917])).
fof(f40127,plain,(
  ( ! [X738,X741,X737,X739,X740,X742] : (s(bool,'T_0') = s(bool,X742) | s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X737),X739),s(bool,X740),s(bool,X741),s(fun(X737,bool),X738))) | s(bool,X742) = s(bool,happ(s(fun(X737,bool),X738),s(X737,sK2(X739,X737)))) | s('type.semanticPrimitives.match_result'(X737),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X737),X739) | s('type.semanticPrimitives.match_result'(X737),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X737),X739)) ) | $spl0),
  inference(forward_demodulation,[],[f39797,f141])).
fof(f39797,plain,(
  ( ! [X738,X741,X737,X739,X740,X742] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X737),X739),s(bool,X740),s(bool,X741),s(fun(X737,bool),X738))) | s(bool,X742) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X742) = s(bool,happ(s(fun(X737,bool),X738),s(X737,sK2(X739,X737)))) | s('type.semanticPrimitives.match_result'(X737),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X737),X739) | s('type.semanticPrimitives.match_result'(X737),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X737),X739)) ) | $spl0),
  inference(superposition,[],[f16731,f2917])).
fof(f40122,plain,(
  ( ! [X712,X710,X713,X709,X711] : (s(bool,'F_0') = s(bool,happ(s(fun(X709,bool),X710),s(X709,sK2(X711,X709)))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X709),X711),s(bool,X712),s(bool,X713),s(fun(X709,bool),X710))))) | s('type.semanticPrimitives.match_result'(X709),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X709),X711) | s('type.semanticPrimitives.match_result'(X709),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X709),X711)) ) | $spl0),
  inference(forward_demodulation,[],[f39792,f141])).
fof(f39792,plain,(
  ( ! [X712,X710,X713,X709,X711] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X709),X711),s(bool,X712),s(bool,X713),s(fun(X709,bool),X710))))) | s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,happ(s(fun(X709,bool),X710),s(X709,sK2(X711,X709)))) | s('type.semanticPrimitives.match_result'(X709),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X709),X711) | s('type.semanticPrimitives.match_result'(X709),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X709),X711)) ) | $spl0),
  inference(superposition,[],[f11207,f2917])).
fof(f39748,plain,(
  ( ! [X457,X459,X460,X456,X458] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X456),X458),s(bool,X459),s(bool,X460),s(fun(X456,bool),X457))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun(X456,bool),X457),s(X456,sK2(X458,X456)))))) | s('type.semanticPrimitives.match_result'(X456),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X456),X458) | s('type.semanticPrimitives.match_result'(X456),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X456),X458)) ) | $spl0),
  inference(superposition,[],[f2368,f2917])).
fof(f39747,plain,(
  ( ! [X453,X455,X451,X452,X454,X450] : (s(bool,X455) != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X450),X452),s(bool,X453),s(bool,X454),s(fun(X450,bool),X451))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X450,bool),X451),s(X450,sK2(X452,X450)))))) | s(bool,'T_0') = s(bool,X455) | s('type.semanticPrimitives.match_result'(X450),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X450),X452) | s('type.semanticPrimitives.match_result'(X450),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X450),X452)) ) | $spl0),
  inference(superposition,[],[f2234,f2917])).
fof(f39746,plain,(
  ( ! [X445,X447,X449,X444,X446,X448] : (s(bool,X449) != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X444),X446),s(bool,X447),s(bool,X448),s(fun(X444,bool),X445))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X449))) | s(bool,'T_0') = s(bool,happ(s(fun(X444,bool),X445),s(X444,sK2(X446,X444)))) | s('type.semanticPrimitives.match_result'(X444),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X444),X446) | s('type.semanticPrimitives.match_result'(X444),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X444),X446)) ) | $spl0),
  inference(superposition,[],[f2234,f2917])).
fof(f39744,plain,(
  ( ! [X436,X432,X434,X437,X433,X435] : (s(bool,X437) = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X432),X434),s(bool,X435),s(bool,X436),s(fun(X432,bool),X433))))) | s(bool,'T_0') = s(bool,X437) | s(bool,'T_0') = s(bool,happ(s(fun(X432,bool),X433),s(X432,sK2(X434,X432)))) | s('type.semanticPrimitives.match_result'(X432),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X432),X434) | s('type.semanticPrimitives.match_result'(X432),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X432),X434)) ) | $spl0),
  inference(superposition,[],[f2195,f2917])).
fof(f39743,plain,(
  ( ! [X428,X430,X429,X431,X427] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X427),X429),s(bool,X430),s(bool,X431),s(fun(X427,bool),X428))))) | s(bool,'T_0') = s(bool,happ(s(fun(X427,bool),X428),s(X427,sK2(X429,X427)))) | s('type.semanticPrimitives.match_result'(X427),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X427),X429) | s('type.semanticPrimitives.match_result'(X427),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X427),X429)) ) | $spl0),
  inference(superposition,[],[f2047,f2917])).
fof(f39741,plain,(
  ( ! [X420,X416,X418,X421,X417,X419] : (s(bool,'const.combin.I_1'(s(bool,X421))) != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X416),X418),s(bool,X419),s(bool,X420),s(fun(X416,bool),X417))) | s(bool,'F_0') = s(bool,happ(s(fun(X416,bool),X417),s(X416,sK2(X418,X416)))) | s(bool,'T_0') = s(bool,X421) | s('type.semanticPrimitives.match_result'(X416),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X416),X418) | s('type.semanticPrimitives.match_result'(X416),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X416),X418)) ) | $spl0),
  inference(superposition,[],[f1916,f2917])).
fof(f39740,plain,(
  ( ! [X412,X414,X410,X413,X415,X411] : (s(bool,X415) != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X410),X412),s(bool,X413),s(bool,X414),s(fun(X410,bool),X411))))) | s(bool,'F_0') = s(bool,X415) | s(bool,'T_0') = s(bool,happ(s(fun(X410,bool),X411),s(X410,sK2(X412,X410)))) | s('type.semanticPrimitives.match_result'(X410),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X410),X412) | s('type.semanticPrimitives.match_result'(X410),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X410),X412)) ) | $spl0),
  inference(superposition,[],[f1916,f2917])).
fof(f39739,plain,(
  ( ! [X408,X405,X407,X409,X404,X406] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X404),X406),s(bool,X407),s(bool,X408),s(fun(X404,bool),X405))))) | s(bool,X409) = s(bool,happ(s(fun(X404,bool),X405),s(X404,sK2(X406,X404)))) | s(bool,'F_0') = s(bool,X409) | s('type.semanticPrimitives.match_result'(X404),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X404),X406) | s('type.semanticPrimitives.match_result'(X404),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X404),X406)) ) | $spl0),
  inference(superposition,[],[f1875,f2917])).
fof(f39738,plain,(
  ( ! [X399,X401,X403,X400,X402] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X399),X401),s(bool,X402),s(bool,X403),s(fun(X399,bool),X400))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X399,bool),X400),s(X399,sK2(X401,X399)))))) | s('type.semanticPrimitives.match_result'(X399),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X399),X401) | s('type.semanticPrimitives.match_result'(X399),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X399),X401)) ) | $spl0),
  inference(superposition,[],[f1858,f2917])).
fof(f40083,plain,(
  ( ! [X368,X370,X367,X369,X371] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X367),X369),s(bool,X370),s(bool,X371),s(fun(X367,bool),X368))) | s(bool,'F_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK2(X369,X367)))) | s('type.semanticPrimitives.match_result'(X367),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X367),X369) | s('type.semanticPrimitives.match_result'(X367),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X367),X369)) ) | $spl0),
  inference(forward_demodulation,[],[f39732,f141])).
fof(f39732,plain,(
  ( ! [X368,X370,X367,X369,X371] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X367),X369),s(bool,X370),s(bool,X371),s(fun(X367,bool),X368))) | s(bool,'F_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK2(X369,X367)))) | s('type.semanticPrimitives.match_result'(X367),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X367),X369) | s('type.semanticPrimitives.match_result'(X367),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X367),X369)) ) | $spl0),
  inference(superposition,[],[f1629,f2917])).
fof(f39731,plain,(
  ( ! [X364,X366,X362,X365,X361,X363] : (s(bool,X366) != s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X361),X363),s(bool,X364),s(bool,X365),s(fun(X361,bool),X362))))) | s(bool,X366) = s(bool,happ(s(fun(X361,bool),X362),s(X361,sK2(X363,X361)))) | s(bool,'F_0') = s(bool,X366) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X361),X363) | s('type.semanticPrimitives.match_result'(X361),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X361),X363)) ) | $spl0),
  inference(superposition,[],[f1617,f2917])).
fof(f39730,plain,(
  ( ! [X356,X358,X360,X357,X359,X355] : (s(bool,'const.combin.I_1'(s(bool,X360))) != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X355),X357),s(bool,X358),s(bool,X359),s(fun(X355,bool),X356))) | s(bool,X360) = s(bool,happ(s(fun(X355,bool),X356),s(X355,sK2(X357,X355)))) | s(bool,'F_0') = s(bool,happ(s(fun(X355,bool),X356),s(X355,sK2(X357,X355)))) | s('type.semanticPrimitives.match_result'(X355),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X355),X357) | s('type.semanticPrimitives.match_result'(X355),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X355),X357)) ) | $spl0),
  inference(superposition,[],[f1617,f2917])).
fof(f39729,plain,(
  ( ! [X350,X352,X354,X349,X351,X353] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X349),X351),s(bool,X352),s(bool,X353),s(fun(X349,bool),X350))) | s(bool,X354) = s(bool,happ(s(fun(X349,bool),X350),s(X349,sK2(X351,X349)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X354))) | s('type.semanticPrimitives.match_result'(X349),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X349),X351) | s('type.semanticPrimitives.match_result'(X349),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X349),X351)) ) | $spl0),
  inference(superposition,[],[f1515,f2917])).
fof(f39728,plain,(
  ( ! [X344,X346,X343,X345,X347,X348] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X343),X345),s(bool,X346),s(bool,X347),s(fun(X343,bool),X344))) | s(bool,X348) = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X343,bool),X344),s(X343,sK2(X345,X343)))))) | s(bool,'F_0') = s(bool,X348) | s('type.semanticPrimitives.match_result'(X343),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X343),X345) | s('type.semanticPrimitives.match_result'(X343),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X343),X345)) ) | $spl0),
  inference(superposition,[],[f1236,f2917])).
fof(f39727,plain,(
  ( ! [X341,X337,X339,X340,X342,X336,X338] : (p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X336),X338),s(bool,X339),s(bool,X340),s(fun(X336,bool),X337)))) | s(bool,X341) = s(bool,happ(s(fun(X336,bool),X337),s(X336,sK2(X338,X336)))) | s(bool,X341) = s(bool,X342) | s(bool,'F_0') = s(bool,X342) | s('type.semanticPrimitives.match_result'(X336),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X336),X338) | s('type.semanticPrimitives.match_result'(X336),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X336),X338)) ) | $spl0),
  inference(superposition,[],[f1204,f2917])).
fof(f39726,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X330),X332),s(bool,X333),s(bool,X334),s(fun(X330,bool),X331))) | s(bool,X335) = s(bool,happ(s(fun(X330,bool),X331),s(X330,sK2(X332,X330)))) | s(bool,'F_0') = s(bool,X335) | s('type.semanticPrimitives.match_result'(X330),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X330),X332) | s('type.semanticPrimitives.match_result'(X330),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X330),X332)) ) | $spl0),
  inference(superposition,[],[f900,f2917])).
fof(f39717,plain,(
  ( ! [X280,X277,X279,X275,X281,X276,X278,X274] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X279),'const.semanticPrimitives.Match_type_error_0'),s(bool,X280),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X274),X276),s(bool,X277),s(bool,X278),s(fun(X274,bool),X275))),s(fun(X279,bool),X281))) | s(bool,'F_0') = s(bool,happ(s(fun(X274,bool),X275),s(X274,sK2(X276,X274)))) | s('type.semanticPrimitives.match_result'(X274),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X274),X276) | s('type.semanticPrimitives.match_result'(X274),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X274),X276)) )),
  inference(superposition,[],[f414,f2917])).
fof(f39714,plain,(
  ( ! [X255,X253,X251,X257,X254,X252,X250,X256] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X255),'const.semanticPrimitives.No_match_0'),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X250),X252),s(bool,X253),s(bool,X254),s(fun(X250,bool),X251))),s(bool,X256),s(fun(X255,bool),X257))) | s(bool,'F_0') = s(bool,happ(s(fun(X250,bool),X251),s(X250,sK2(X252,X250)))) | s('type.semanticPrimitives.match_result'(X250),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X250),X252) | s('type.semanticPrimitives.match_result'(X250),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X250),X252)) )),
  inference(superposition,[],[f316,f2917])).
fof(f39713,plain,(
  ( ! [X249,X246,X248,X247,X245] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X245),X247),s(bool,X248),s(bool,X249),s(fun(X245,bool),X246))))) | s(bool,'F_0') = s(bool,happ(s(fun(X245,bool),X246),s(X245,sK2(X247,X245)))) | s('type.semanticPrimitives.match_result'(X245),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X245),X247) | s('type.semanticPrimitives.match_result'(X245),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X245),X247)) )),
  inference(superposition,[],[f278,f2917])).
fof(f39712,plain,(
  ( ! [X243,X241,X239,X244,X242,X240] : (p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X239),X241),s(bool,X242),s(bool,X243),s(fun(X239,bool),X240)))) | s(bool,X244) = s(bool,happ(s(fun(X239,bool),X240),s(X239,sK2(X241,X239)))) | s(bool,'T_0') = s(bool,X244) | s('type.semanticPrimitives.match_result'(X239),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X239),X241) | s('type.semanticPrimitives.match_result'(X239),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X239),X241)) )),
  inference(superposition,[],[f256,f2917])).
fof(f39706,plain,(
  ( ! [X208,X206,X204,X209,X207,X205] : (~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X204),X206),s(bool,X207),s(bool,X208),s(fun(X204,bool),X205)))) | ~p(s(bool,X209)) | s(bool,X209) = s(bool,happ(s(fun(X204,bool),X205),s(X204,sK2(X206,X204)))) | s('type.semanticPrimitives.match_result'(X204),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X204),X206) | s('type.semanticPrimitives.match_result'(X204),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X204),X206)) )),
  inference(superposition,[],[f156,f2917])).
fof(f39704,plain,(
  ( ! [X196,X194,X192,X197,X195,X193] : (p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X192),X194),s(bool,X195),s(bool,X196),s(fun(X192,bool),X193)))) | p(s(bool,X197)) | s(bool,X197) = s(bool,happ(s(fun(X192,bool),X193),s(X192,sK2(X194,X192)))) | s('type.semanticPrimitives.match_result'(X192),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X192),X194) | s('type.semanticPrimitives.match_result'(X192),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X192),X194)) )),
  inference(superposition,[],[f153,f2917])).
fof(f39703,plain,(
  ( ! [X191,X189,X187,X190,X188] : (p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X187),X189),s(bool,X190),s(bool,X191),s(fun(X187,bool),X188)))) | s(bool,'F_0') = s(bool,happ(s(fun(X187,bool),X188),s(X187,sK2(X189,X187)))) | s('type.semanticPrimitives.match_result'(X187),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X187),X189) | s('type.semanticPrimitives.match_result'(X187),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X187),X189)) )),
  inference(superposition,[],[f135,f2917])).
fof(f39702,plain,(
  ( ! [X185,X182,X186,X184,X183] : (~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X182),X184),s(bool,X185),s(bool,X186),s(fun(X182,bool),X183)))) | s(bool,'F_0') != s(bool,happ(s(fun(X182,bool),X183),s(X182,sK2(X184,X182)))) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X182),X184) | s('type.semanticPrimitives.match_result'(X182),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X182),X184)) )),
  inference(superposition,[],[f134,f2917])).
fof(f39701,plain,(
  ( ! [X177,X180,X178,X181,X179] : (~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X177),X179),s(bool,X180),s(bool,X181),s(fun(X177,bool),X178)))) | s(bool,'T_0') = s(bool,happ(s(fun(X177,bool),X178),s(X177,sK2(X179,X177)))) | s('type.semanticPrimitives.match_result'(X177),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X177),X179) | s('type.semanticPrimitives.match_result'(X177),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X177),X179)) )),
  inference(superposition,[],[f131,f2917])).
fof(f39700,plain,(
  ( ! [X175,X173,X176,X174,X172] : (p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X172),X174),s(bool,X175),s(bool,X176),s(fun(X172,bool),X173)))) | s(bool,'T_0') != s(bool,happ(s(fun(X172,bool),X173),s(X172,sK2(X174,X172)))) | s('type.semanticPrimitives.match_result'(X172),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X172),X174) | s('type.semanticPrimitives.match_result'(X172),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X172),X174)) )),
  inference(superposition,[],[f130,f2917])).
fof(f39698,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X159),X161),s(bool,X162),s(bool,X163),s(fun(X159,bool),X160)))) | p(s(bool,happ(s(fun(X159,bool),X160),s(X159,'const.min.@_1'(s(fun(X159,bool),X160)))))) | s('type.semanticPrimitives.match_result'(X159),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X159),X161) | s('type.semanticPrimitives.match_result'(X159),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X159),X161)) )),
  inference(superposition,[],[f149,f2917])).
fof(f39695,plain,(
  ( ! [X146,X144,X142,X149,X147,X145,X143,X150,X148] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X142),'const.semanticPrimitives.Match_1'(s(X142,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X143),X145),s(X142,X146),s(X142,X147),s(fun(X143,X142),X144))))),s(bool,X148),s(bool,X149),s(fun(X142,bool),X150))) | s(bool,'F_0') = s(bool,happ(s(fun(X142,bool),X150),s(X142,happ(s(fun(X143,X142),X144),s(X143,sK2(X145,X143)))))) | s('type.semanticPrimitives.match_result'(X143),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X143),X145) | s('type.semanticPrimitives.match_result'(X143),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X143),X145)) )),
  inference(superposition,[],[f2989,f2917])).
fof(f40069,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'F_0') = s(bool,happ(s(fun(X138,bool),X142),s(X138,sK2(X139,X138)))) | s('type.semanticPrimitives.match_result'(X138),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X138),X139) | s('type.semanticPrimitives.match_result'(X138),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X138),X139) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X138),X139),s(bool,X140),s(bool,X141),s(fun(X138,bool),X142)))))) ) | $spl0),
  inference(forward_demodulation,[],[f39664,f141])).
fof(f39664,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,happ(s(fun(X138,bool),X142),s(X138,sK2(X139,X138)))) | s('type.semanticPrimitives.match_result'(X138),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X138),X139) | s('type.semanticPrimitives.match_result'(X138),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X138),X139) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X138),X139),s(bool,X140),s(bool,X141),s(fun(X138,bool),X142)))))) ) | $spl0),
  inference(superposition,[],[f2917,f8450])).
fof(f39660,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (s(bool,'const.combin.I_1'(s(bool,X122))) = s(bool,happ(s(fun(X117,bool),X121),s(X117,sK2(X118,X117)))) | s('type.semanticPrimitives.match_result'(X117),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X117),X118) | s('type.semanticPrimitives.match_result'(X117),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X117),X118) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X117),X118),s(bool,X119),s(bool,X120),s(fun(X117,bool),X121))) | s(bool,'T_0') = s(bool,X122)) ) | $spl0),
  inference(superposition,[],[f2917,f2195])).
fof(f40064,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s(bool,'T_0') = s(bool,happ(s(fun(X106,bool),X110),s(X106,sK2(X107,X106)))) | s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X106),X107) | s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X106),X107) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X106),X107),s(bool,X108),s(bool,X109),s(fun(X106,bool),X110)))))) ) | $spl0),
  inference(forward_demodulation,[],[f39658,f183])).
fof(f39658,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (s(bool,'const.relation.WF_1'(s(fun(X111,fun(X111,bool)),'const.relation.EMPTY_REL_0'))) = s(bool,happ(s(fun(X106,bool),X110),s(X106,sK2(X107,X106)))) | s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X106),X107) | s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X106),X107) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X106),X107),s(bool,X108),s(bool,X109),s(fun(X106,bool),X110)))))) ) | $spl0),
  inference(superposition,[],[f2917,f1796])).
fof(f39655,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (s(bool,X93) = s(bool,happ(s(fun(X88,bool),X92),s(X88,sK2(X89,X88)))) | s('type.semanticPrimitives.match_result'(X88),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X88),X89) | s('type.semanticPrimitives.match_result'(X88),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X88),X89) | s(bool,'F_0') = s(bool,X93) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X88),X89),s(bool,X90),s(bool,X91),s(fun(X88,bool),X92)))))) ) | $spl0),
  inference(superposition,[],[f2917,f755])).
fof(f40063,plain,(
  ( ! [X30,X31,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,sK2('const.semanticPrimitives.Match_1'(s(X30,X31)),X30)))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,X31)))) )),
  inference(subsumption_resolution,[],[f40062,f148])).
fof(f40062,plain,(
  ( ! [X30,X31,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,sK2('const.semanticPrimitives.Match_1'(s(X30,X31)),X30)))) | s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,X31))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,X31)))) )),
  inference(subsumption_resolution,[],[f39645,f147])).
fof(f39645,plain,(
  ( ! [X30,X31,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,sK2('const.semanticPrimitives.Match_1'(s(X30,X31)),X30)))) | s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,X31))) | s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,X31))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,X31)))) )),
  inference(superposition,[],[f2917,f2989])).
fof(f40061,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X131),X132) | s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X131),X132) | s(X133,happ(s(fun(X131,X133),X134),s(X131,sK2('const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132)),X131)))) = s(X133,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X131),X132),s(X133,X135),s(X133,X136),s(fun(X131,X133),X134)))) )),
  inference(forward_demodulation,[],[f40060,f141])).
fof(f40060,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X131),X132) | s(X133,happ(s(fun(X131,X133),X134),s(X131,sK2('const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132)),X131)))) = s(X133,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X131),X132),s(X133,X135),s(X133,X136),s(fun(X131,X133),X134))) | s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X131),'const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132)))) )),
  inference(forward_demodulation,[],[f39441,f141])).
fof(f39441,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s(X133,happ(s(fun(X131,X133),X134),s(X131,sK2('const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132)),X131)))) = s(X133,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X131),X132),s(X133,X135),s(X133,X136),s(fun(X131,X133),X134))) | s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X131),'const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132))) | s('type.semanticPrimitives.match_result'(X131),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X131),'const.combin.I_1'(s('type.semanticPrimitives.match_result'(X131),X132)))) )),
  inference(superposition,[],[f2917,f141])).
fof(f39440,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X129] : (s(X127,happ(s(fun(X121,X127),X128),s(X121,sK2('const.pair.pair_CASE_2'(s('type.pair.prod'(X122,X123),'const.pair.,_2'(s(X122,X124),s(X123,X125))),s(fun(X122,fun(X123,'type.semanticPrimitives.match_result'(X121))),X126)),X121)))) = s(X127,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X121),happ(s(fun(X123,'type.semanticPrimitives.match_result'(X121)),happ(s(fun(X122,fun(X123,'type.semanticPrimitives.match_result'(X121))),X126),s(X122,X124))),s(X123,X125))),s(X127,X129),s(X127,X130),s(fun(X121,X127),X128))) | s('type.semanticPrimitives.match_result'(X121),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X121),'const.pair.pair_CASE_2'(s('type.pair.prod'(X122,X123),'const.pair.,_2'(s(X122,X124),s(X123,X125))),s(fun(X122,fun(X123,'type.semanticPrimitives.match_result'(X121))),X126))) | s('type.semanticPrimitives.match_result'(X121),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X121),'const.pair.pair_CASE_2'(s('type.pair.prod'(X122,X123),'const.pair.,_2'(s(X122,X124),s(X123,X125))),s(fun(X122,fun(X123,'type.semanticPrimitives.match_result'(X121))),X126)))) )),
  inference(superposition,[],[f2917,f165])).
fof(f40059,plain,(
  ( ! [X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))) | s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))) | s(X117,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))),s(X117,X119),s(X117,X120),s(fun(X112,X117),X118))) = s(X117,happ(s(fun(X112,X117),X118),s(X112,sK2(happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116)),X112))))) )),
  inference(forward_demodulation,[],[f40058,f164])).
fof(f40058,plain,(
  ( ! [X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))) | s(X117,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))),s(X117,X119),s(X117,X120),s(fun(X112,X117),X118))) = s(X117,happ(s(fun(X112,X117),X118),s(X112,sK2(happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116)),X112)))) | s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X112),happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116)))) )),
  inference(forward_demodulation,[],[f39439,f164])).
fof(f39439,plain,(
  ( ! [X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s(X117,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X112),'const.pair.pair_CASE_2'(s('type.pair.prod'(X113,X114),X115),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))),s(X117,X119),s(X117,X120),s(fun(X112,X117),X118))) = s(X117,happ(s(fun(X112,X117),X118),s(X112,sK2(happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116)),X112)))) | s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X112),happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116))) | s('type.semanticPrimitives.match_result'(X112),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X112),happ(s(fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112)),happ(s(fun('type.pair.prod'(X113,X114),fun(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),'type.semanticPrimitives.match_result'(X112))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X113,X114),X115))),s(fun(X113,fun(X114,'type.semanticPrimitives.match_result'(X112))),X116)))) )),
  inference(superposition,[],[f2917,f164])).
fof(f39438,plain,(
  ( ! [X103,X111,X109,X107,X105,X102,X110,X108,X106,X104] : (s(X108,happ(s(fun(X102,X108),X109),s(X102,sK2(happ(s(fun(X103,'type.semanticPrimitives.match_result'(X102)),X104),s(X103,sK2(X105,X103))),X102)))) = s(X108,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X102),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X103),X105),s('type.semanticPrimitives.match_result'(X102),X106),s('type.semanticPrimitives.match_result'(X102),X107),s(fun(X103,'type.semanticPrimitives.match_result'(X102)),X104))),s(X108,X110),s(X108,X111),s(fun(X102,X108),X109))) | s('type.semanticPrimitives.match_result'(X102),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X102),happ(s(fun(X103,'type.semanticPrimitives.match_result'(X102)),X104),s(X103,sK2(X105,X103)))) | s('type.semanticPrimitives.match_result'(X102),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X102),happ(s(fun(X103,'type.semanticPrimitives.match_result'(X102)),X104),s(X103,sK2(X105,X103)))) | s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X103),X105) | s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X103),X105)) )),
  inference(superposition,[],[f2917,f2917])).
fof(f39437,plain,(
  ( ! [X94,X92,X101,X99,X97,X95,X93,X100,X98,X96] : (s(X98,happ(s(fun(X92,X98),X99),s(X92,sK2(happ(s(fun(X93,'type.semanticPrimitives.match_result'(X92)),happ(s(fun(X94,fun(X93,'type.semanticPrimitives.match_result'(X92))),X95),s(X94,X96))),s(X93,X97)),X92)))) = s(X98,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X92),'const.pair.pair_CASE_2'(s('type.pair.prod'(X94,X93),'const.pair.,_2'(s(X94,X96),s(X93,X97))),s(fun(X94,fun(X93,'type.semanticPrimitives.match_result'(X92))),X95))),s(X98,X100),s(X98,X101),s(fun(X92,X98),X99))) | s('type.semanticPrimitives.match_result'(X92),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X92),happ(s(fun(X93,'type.semanticPrimitives.match_result'(X92)),happ(s(fun(X94,fun(X93,'type.semanticPrimitives.match_result'(X92))),X95),s(X94,X96))),s(X93,X97))) | s('type.semanticPrimitives.match_result'(X92),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X92),happ(s(fun(X93,'type.semanticPrimitives.match_result'(X92)),happ(s(fun(X94,fun(X93,'type.semanticPrimitives.match_result'(X92))),X95),s(X94,X96))),s(X93,X97)))) )),
  inference(superposition,[],[f2917,f165])).
fof(f40057,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X86),X87) | s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X86),X87) | s(X88,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X86),X87),s(X88,X90),s(X88,X91),s(fun(X86,X88),X89))) = s(X88,happ(s(fun(X86,X88),X89),s(X86,sK2(happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87)),X86))))) )),
  inference(forward_demodulation,[],[f40056,f172])).
fof(f40056,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X86),X87) | s(X88,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X86),X87),s(X88,X90),s(X88,X91),s(fun(X86,X88),X89))) = s(X88,happ(s(fun(X86,X88),X89),s(X86,sK2(happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87)),X86)))) | s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X86),happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87)))) )),
  inference(forward_demodulation,[],[f39436,f172])).
fof(f39436,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(X88,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X86),X87),s(X88,X90),s(X88,X91),s(fun(X86,X88),X89))) = s(X88,happ(s(fun(X86,X88),X89),s(X86,sK2(happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87)),X86)))) | s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X86),happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87))) | s('type.semanticPrimitives.match_result'(X86),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X86),happ(s(fun('type.semanticPrimitives.match_result'(X86),'type.semanticPrimitives.match_result'(X86)),'const.combin.I_0'),s('type.semanticPrimitives.match_result'(X86),X87)))) )),
  inference(superposition,[],[f2917,f172])).
fof(f39435,plain,(
  ( ! [X80,X78,X76,X85,X83,X81,X79,X77,X84,X82] : (s(X82,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X76),happ(s(fun(X77,'type.semanticPrimitives.match_result'(X76)),X81),s(X77,sK2(X78,X77)))),s(X82,X84),s(X82,X85),s(fun(X76,X82),X83))) = s(X82,happ(s(fun(X76,X82),X83),s(X76,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),X78),s('type.semanticPrimitives.match_result'(X76),X79),s('type.semanticPrimitives.match_result'(X76),X80),s(fun(X77,'type.semanticPrimitives.match_result'(X76)),X81)),X76)))) | s('type.semanticPrimitives.match_result'(X76),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X76),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),X78),s('type.semanticPrimitives.match_result'(X76),X79),s('type.semanticPrimitives.match_result'(X76),X80),s(fun(X77,'type.semanticPrimitives.match_result'(X76)),X81))) | s('type.semanticPrimitives.match_result'(X76),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X76),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),X78),s('type.semanticPrimitives.match_result'(X76),X79),s('type.semanticPrimitives.match_result'(X76),X80),s(fun(X77,'type.semanticPrimitives.match_result'(X76)),X81))) | s('type.semanticPrimitives.match_result'(X77),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X77),X78) | s('type.semanticPrimitives.match_result'(X77),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X77),X78)) )),
  inference(superposition,[],[f2917,f2917])).
fof(f40055,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X75,X73] : (s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))) | s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))) | s(X72,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))),s(X72,X74),s(X72,X75),s(fun(X66,X72),X73))) = s(X72,happ(s(fun(X66,X72),X73),s(X66,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71)),X66))))) )),
  inference(forward_demodulation,[],[f40054,f163])).
fof(f40054,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X75,X73] : (s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))) | s(X72,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))),s(X72,X74),s(X72,X75),s(fun(X66,X72),X73))) = s(X72,happ(s(fun(X66,X72),X73),s(X66,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71)),X66)))) | s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71)))) )),
  inference(forward_demodulation,[],[f39434,f163])).
fof(f39434,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X75,X73] : (s(X72,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X66),happ(s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71),s(X67,X68))),s(X72,X74),s(X72,X75),s(fun(X66,X72),X73))) = s(X72,happ(s(fun(X66,X72),X73),s(X66,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71)),X66)))) | s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71))) | s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X66),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_1'(s(X67,X68))),s('type.semanticPrimitives.match_result'(X66),X69),s('type.semanticPrimitives.match_result'(X66),X70),s(fun(X67,'type.semanticPrimitives.match_result'(X66)),X71)))) )),
  inference(superposition,[],[f2917,f163])).
fof(f40053,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X57),X60) | s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X57),X60) | s(X62,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X57),X60),s(X62,X64),s(X62,X65),s(fun(X57,X62),X63))) = s(X62,happ(s(fun(X57,X62),X63),s(X57,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61)),X57))))) )),
  inference(forward_demodulation,[],[f40052,f161])).
fof(f40052,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X57),X60) | s(X62,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X57),X60),s(X62,X64),s(X62,X65),s(fun(X57,X62),X63))) = s(X62,happ(s(fun(X57,X62),X63),s(X57,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61)),X57)))) | s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61)))) )),
  inference(forward_demodulation,[],[f39433,f161])).
fof(f39433,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (s(X62,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X57),X60),s(X62,X64),s(X62,X65),s(fun(X57,X62),X63))) = s(X62,happ(s(fun(X57,X62),X63),s(X57,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61)),X57)))) | s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61))) | s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X57),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X58),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'(X57),X59),s('type.semanticPrimitives.match_result'(X57),X60),s(fun(X58,'type.semanticPrimitives.match_result'(X57)),X61)))) )),
  inference(superposition,[],[f2917,f161])).
fof(f40051,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X48),X50) | s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X48),X50) | s(X53,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X48),X50),s(X53,X55),s(X53,X56),s(fun(X48,X53),X54))) = s(X53,happ(s(fun(X48,X53),X54),s(X48,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52)),X48))))) )),
  inference(forward_demodulation,[],[f40050,f160])).
fof(f40050,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X48),X50) | s(X53,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X48),X50),s(X53,X55),s(X53,X56),s(fun(X48,X53),X54))) = s(X53,happ(s(fun(X48,X53),X54),s(X48,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52)),X48)))) | s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52)))) )),
  inference(forward_demodulation,[],[f39432,f160])).
fof(f39432,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(X53,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X48),X50),s(X53,X55),s(X53,X56),s(fun(X48,X53),X54))) = s(X53,happ(s(fun(X48,X53),X54),s(X48,sK2('const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52)),X48)))) | s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52))) | s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X48),'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'(X48),X50),s('type.semanticPrimitives.match_result'(X48),X51),s(fun(X49,'type.semanticPrimitives.match_result'(X48)),X52)))) )),
  inference(superposition,[],[f2917,f160])).
fof(f40049,plain,(
  ( ! [X39,X37,X38,X40] : (s(X39,happ(s(fun(X37,X39),X40),s(X37,X38))) = s(X39,happ(s(fun(X37,X39),X40),s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))),X37))))) )),
  inference(forward_demodulation,[],[f40048,f163])).
fof(f40048,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(X39,happ(s(fun(X37,X39),X40),s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))),X37)))) = s(X39,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,X38))),s(X39,X41),s(X39,X42),s(fun(X37,X39),X40)))) )),
  inference(subsumption_resolution,[],[f40047,f148])).
fof(f40047,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(X39,happ(s(fun(X37,X39),X40),s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))),X37)))) = s(X39,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,X38))),s(X39,X41),s(X39,X42),s(fun(X37,X39),X40))) | s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))))) )),
  inference(subsumption_resolution,[],[f39430,f147])).
fof(f39430,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(X39,happ(s(fun(X37,X39),X40),s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))),X37)))) = s(X39,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,X38))),s(X39,X41),s(X39,X42),s(fun(X37,X39),X40))) | s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37)))) | s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X37),'const.semanticPrimitives.Match_1'(s(X37,sK2('const.semanticPrimitives.Match_1'(s(X37,X38)),X37))))) )),
  inference(superposition,[],[f2917,f2200])).
fof(f40046,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X33,happ(s(fun(X31,X33),X34),s(X31,sK2('const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31))),X31)))) = s(X33,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X31),X32),s(X33,X35),s(X33,X36),s(fun(X31,X33),X34))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X31),X32) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X31),X32)) )),
  inference(subsumption_resolution,[],[f40045,f148])).
fof(f40045,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X33,happ(s(fun(X31,X33),X34),s(X31,sK2('const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31))),X31)))) = s(X33,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X31),X32),s(X33,X35),s(X33,X36),s(fun(X31,X33),X34))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31)))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X31),X32) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X31),X32)) )),
  inference(subsumption_resolution,[],[f39429,f147])).
fof(f39429,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X33,happ(s(fun(X31,X33),X34),s(X31,sK2('const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31))),X31)))) = s(X33,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X31),X32),s(X33,X35),s(X33,X36),s(fun(X31,X33),X34))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31)))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_1'(s(X31,sK2(X32,X31)))) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X31),X32) | s('type.semanticPrimitives.match_result'(X31),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X31),X32)) )),
  inference(superposition,[],[f2917,f144])).
fof(f40028,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X141),X143),s(bool,X144),s(bool,X145),s(fun(X141,bool),X142))) | s('type.semanticPrimitives.match_result'(X141),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X141),X143) | s('type.semanticPrimitives.match_result'(X141),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X141),X143) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X141,bool),X142),s(X141,sK2(X143,X141))))))) ) | $spl0),
  inference(forward_demodulation,[],[f39409,f141])).
fof(f39409,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X141),X143),s(bool,X144),s(bool,X145),s(fun(X141,bool),X142))) | s('type.semanticPrimitives.match_result'(X141),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X141),X143) | s('type.semanticPrimitives.match_result'(X141),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X141),X143) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X141,bool),X142),s(X141,sK2(X143,X141))))))) ) | $spl0),
  inference(superposition,[],[f2917,f8450])).
fof(f39405,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (s(bool,'const.combin.I_1'(s(bool,X123))) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X120),X122),s(bool,X124),s(bool,X125),s(fun(X120,bool),X121))) | s('type.semanticPrimitives.match_result'(X120),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X120),X122) | s('type.semanticPrimitives.match_result'(X120),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X120),X122) | s(bool,'T_0') = s(bool,happ(s(fun(X120,bool),X121),s(X120,sK2(X122,X120)))) | s(bool,'T_0') = s(bool,X123)) ) | $spl0),
  inference(superposition,[],[f2917,f2195])).
fof(f40023,plain,(
  ( ! [X111,X109,X114,X110,X113] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X109),X111),s(bool,X113),s(bool,X114),s(fun(X109,bool),X110))) | s('type.semanticPrimitives.match_result'(X109),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X109),X111) | s('type.semanticPrimitives.match_result'(X109),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X109),X111) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X109,bool),X110),s(X109,sK2(X111,X109))))))) ) | $spl0),
  inference(forward_demodulation,[],[f39403,f183])).
fof(f39403,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (s(bool,'const.relation.WF_1'(s(fun(X112,fun(X112,bool)),'const.relation.EMPTY_REL_0'))) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X109),X111),s(bool,X113),s(bool,X114),s(fun(X109,bool),X110))) | s('type.semanticPrimitives.match_result'(X109),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X109),X111) | s('type.semanticPrimitives.match_result'(X109),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X109),X111) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X109,bool),X110),s(X109,sK2(X111,X109))))))) ) | $spl0),
  inference(superposition,[],[f2917,f1796])).
fof(f39402,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X103),X105),s(bool,X106),s(bool,X107),s(fun(X103,bool),X104))) | s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X103),X105) | s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X103),X105) | s(bool,X108) = s(bool,happ(s(fun(X103,bool),X104),s(X103,sK2(X105,X103)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X108)))) ) | $spl0),
  inference(superposition,[],[f2917,f755])).
fof(f39401,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (s(bool,X100) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X97),X99),s(bool,X101),s(bool,X102),s(fun(X97,bool),X98))) | s('type.semanticPrimitives.match_result'(X97),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X97),X99) | s('type.semanticPrimitives.match_result'(X97),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X97),X99) | s(bool,'F_0') = s(bool,happ(s(fun(X97,bool),X98),s(X97,sK2(X99,X97)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X100)))) ) | $spl0),
  inference(superposition,[],[f2917,f755])).
fof(f39400,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (s(bool,X94) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X91),X93),s(bool,X95),s(bool,X96),s(fun(X91,bool),X92))) | s('type.semanticPrimitives.match_result'(X91),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X91),X93) | s('type.semanticPrimitives.match_result'(X91),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X91),X93) | s(bool,'F_0') = s(bool,X94) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X91,bool),X92),s(X91,sK2(X93,X91))))))) ) | $spl0),
  inference(superposition,[],[f2917,f755])).
fof(f39397,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,X76) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X73),X75),s(bool,X77),s(bool,X78),s(fun(X73,bool),X74))) | s('type.semanticPrimitives.match_result'(X73),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X73),X75) | s('type.semanticPrimitives.match_result'(X73),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X73),X75) | s(bool,'F_0') = s(bool,X76) | s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),X74),s(X73,sK2(X75,X73))))) ) | $spl0),
  inference(superposition,[],[f2917,f612])).
fof(f39396,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X67),X69),s(bool,X70),s(bool,X71),s(fun(X67,bool),X68))) | s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X67),X69) | s('type.semanticPrimitives.match_result'(X67),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X67),X69) | s(bool,X72) = s(bool,happ(s(fun(X67,bool),X68),s(X67,sK2(X69,X67)))) | s(bool,'F_0') = s(bool,X72)) ) | $spl0),
  inference(superposition,[],[f2917,f612])).
fof(f39393,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,X52) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X49),X51),s(bool,X53),s(bool,X54),s(fun(X49,bool),X50))) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X49),X51) | s('type.semanticPrimitives.match_result'(X49),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X49),X51) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,happ(s(fun(X49,bool),X50),s(X49,sK2(X51,X49))))) )),
  inference(superposition,[],[f2917,f205])).
fof(f39392,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X43),X45),s(bool,X46),s(bool,X47),s(fun(X43,bool),X44))) | s('type.semanticPrimitives.match_result'(X43),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X43),X45) | s('type.semanticPrimitives.match_result'(X43),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X43),X45) | s(bool,X48) = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK2(X45,X43)))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f2917,f205])).
fof(f39391,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X38),X40),s(bool,X41),s(bool,X42),s(fun(X38,bool),X39))) | s('type.semanticPrimitives.match_result'(X38),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X38),X40) | s('type.semanticPrimitives.match_result'(X38),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X38),X40) | s(bool,'T_0') = s(bool,happ(s(fun(X38,bool),X39),s(X38,sK2(X40,X38))))) )),
  inference(superposition,[],[f2917,f128])).
fof(f39390,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'const.relation.WF_1'(s(fun(X31,fun(X31,bool)),sK2(X32,fun(X31,fun(X31,bool)))))) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(fun(X31,fun(X31,bool))),X32),s(bool,X33),s(bool,X34),s(fun(fun(X31,fun(X31,bool)),bool),'const.relation.WF_0'))) | s('type.semanticPrimitives.match_result'(fun(X31,fun(X31,bool))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(fun(X31,fun(X31,bool))),X32) | s('type.semanticPrimitives.match_result'(fun(X31,fun(X31,bool))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(fun(X31,fun(X31,bool))),X32)) )),
  inference(superposition,[],[f2917,f143])).
fof(f39389,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(fun(X27,fun(X27,bool))),X28),s(bool,X29),s(bool,X30),s(fun(fun(X27,fun(X27,bool)),bool),'const.relation.WF_0'))) | s('type.semanticPrimitives.match_result'(fun(X27,fun(X27,bool))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(fun(X27,fun(X27,bool))),X28) | s('type.semanticPrimitives.match_result'(fun(X27,fun(X27,bool))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(fun(X27,fun(X27,bool))),X28) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X27,fun(X27,bool)),sK2(X28,fun(X27,fun(X27,bool))))))) )),
  inference(superposition,[],[f2917,f546])).
fof(f39388,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(X20,'const.pair.pair_CASE_2'(s('type.pair.prod'(X21,X22),X23),s(fun(X21,fun(X22,X20)),sK2(X24,fun(X21,fun(X22,X20)))))) = s(X20,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(fun(X21,fun(X22,X20))),X24),s(X20,X25),s(X20,X26),s(fun(fun(X21,fun(X22,X20)),X20),happ(s(fun('type.pair.prod'(X21,X22),fun(fun(X21,fun(X22,X20)),X20)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X21,X22),X23))))) | s('type.semanticPrimitives.match_result'(fun(X21,fun(X22,X20))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(fun(X21,fun(X22,X20))),X24) | s('type.semanticPrimitives.match_result'(fun(X21,fun(X22,X20))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(fun(X21,fun(X22,X20))),X24)) )),
  inference(superposition,[],[f2917,f164])).
fof(f40021,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X12,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X13),X15),s(X12,X16),s(X12,X17),s(fun(X13,X12),X14))) = s(X12,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X13),X15),s(X12,X18),s(X12,X19),s(fun(X13,X12),X14))) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X13),X15) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X13),X15)) )),
  inference(duplicate_literal_removal,[],[f39387])).
fof(f39387,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X12,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X13),X15),s(X12,X16),s(X12,X17),s(fun(X13,X12),X14))) = s(X12,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X13),X15),s(X12,X18),s(X12,X19),s(fun(X13,X12),X14))) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X13),X15) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X13),X15) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X13),X15) | s('type.semanticPrimitives.match_result'(X13),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X13),X15)) )),
  inference(superposition,[],[f2917,f2917])).
fof(f39386,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X4,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X5),'const.pair.,_2'(s(X6,X8),s(X5,sK2(X9,X5)))),s(fun(X6,fun(X5,X4)),X7))) = s(X4,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X5),X9),s(X4,X10),s(X4,X11),s(fun(X5,X4),happ(s(fun(X6,fun(X5,X4)),X7),s(X6,X8))))) | s('type.semanticPrimitives.match_result'(X5),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X5),X9) | s('type.semanticPrimitives.match_result'(X5),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X5),X9)) )),
  inference(superposition,[],[f2917,f165])).
fof(f38480,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X15),'const.semanticPrimitives.Match_1'(s(X15,X16))),s(bool,X17),s(bool,X18),s(fun(X15,bool),X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X15,bool),X19),s(X15,X16)))) )),
  inference(subsumption_resolution,[],[f38163,f117])).
fof(f38163,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X15),'const.semanticPrimitives.Match_1'(s(X15,X16))),s(bool,X17),s(bool,X18),s(fun(X15,bool),X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X15,bool),X19),s(X15,X16)))) )),
  inference(superposition,[],[f134,f2989])).
fof(f38130,plain,(
  ( ! [X47,X54,X52,X50,X48,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X47),'const.semanticPrimitives.Match_1'(s(X47,X53))),s(bool,X54),s(bool,X55),s(fun(X47,bool),happ(s(fun(X49,fun(X47,bool)),happ(s(fun(X48,fun(X49,fun(X47,bool))),X52),s(X48,X50))),s(X49,X51))))) | s(bool,'F_0') = s(bool,happ(s(fun(X47,bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,X50),s(X49,X51))),s(fun(X48,fun(X49,fun(X47,bool))),X52))),s(X47,X53)))) )),
  inference(superposition,[],[f2989,f165])).
fof(f38128,plain,(
  ( ! [X30,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,X36))),s(bool,X37),s(bool,X38),s(fun(X30,bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X32,X31),'const.pair.,_2'(s(X32,X34),s(X31,X35))),s(fun(X32,fun(X31,fun(X30,bool))),X33))))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),happ(s(fun(X31,fun(X30,bool)),happ(s(fun(X32,fun(X31,fun(X30,bool))),X33),s(X32,X34))),s(X31,X35))),s(X30,X36)))) )),
  inference(superposition,[],[f2989,f165])).
fof(f37958,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,X31))),s(bool,X32),s(bool,X33),s(fun(X30,bool),X34))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X34),s(X30,sK2('const.semanticPrimitives.Match_1'(s(X30,X31)),X30))))) )),
  inference(superposition,[],[f2989,f2200])).
fof(f37957,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X25),X26),s(bool,X27),s(bool,X28),s(fun(X25,bool),X29))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),X29),s(X25,sK2(X26,X25)))) | s('type.semanticPrimitives.match_result'(X25),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X25),X26) | s('type.semanticPrimitives.match_result'(X25),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X25),X26)) )),
  inference(superposition,[],[f2989,f144])).
fof(f37955,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),X18),s(bool,X19),s(bool,X20),s(fun(bool,bool),X21))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X21),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK2(X18,bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X18) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X18)) )),
  inference(superposition,[],[f2989,f2203])).
fof(f37948,plain,(
  ( ! [X436,X438,X440,X437,X439] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'('type.semanticPrimitives.match_result'(X436)),'const.semanticPrimitives.Match_1'(s('type.semanticPrimitives.match_result'(X436),X437))),s(bool,X438),s(bool,X439),s(fun('type.semanticPrimitives.match_result'(X436),bool),X440))) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.match_result'(X436),bool),X440),s('type.semanticPrimitives.match_result'(X436),'const.semanticPrimitives.Match_1'(s(X436,sK2(X437,X436)))))) | s('type.semanticPrimitives.match_result'(X436),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X436),X437) | s('type.semanticPrimitives.match_result'(X436),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X436),X437)) )),
  inference(superposition,[],[f2989,f144])).
fof(f38426,plain,(
  ( ! [X315,X316,X318,X317] : (s(bool,'F_0') = s(bool,X315) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X315))),s(bool,X316),s(bool,X317),s(fun(bool,bool),X318)))) ) | $spl0),
  inference(forward_demodulation,[],[f38425,f141])).
fof(f38425,plain,(
  ( ! [X315,X316,X318,X317] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X315))),s(bool,X316),s(bool,X317),s(fun(bool,bool),X318))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X315)))) ) | $spl0),
  inference(forward_demodulation,[],[f37919,f183])).
fof(f37919,plain,(
  ( ! [X315,X316,X318,X314,X317] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X315))),s(bool,X316),s(bool,X317),s(fun(bool,bool),X318))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'const.relation.WF_1'(s(fun(X314,fun(X314,bool)),'const.relation.EMPTY_REL_0'))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X315)))) ) | $spl0),
  inference(superposition,[],[f2989,f1796])).
fof(f37906,plain,(
  ( ! [X253,X251,X252,X250] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X250))),s(bool,X251),s(bool,X252),s(fun(bool,bool),X253))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X253),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X250)) )),
  inference(superposition,[],[f2989,f128])).
fof(f37903,plain,(
  ( ! [X239,X237,X235,X238,X236] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X235))),s(bool,X236),s(bool,X237),s(fun(bool,bool),X238))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X238),s(bool,'F_0'))) | s(bool,X235) = s(bool,X239) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X239)))) ) | $spl0),
  inference(superposition,[],[f2989,f755])).
fof(f38408,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(bool,'F_0') = s(bool,X234) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X230))),s(bool,X231),s(bool,X232),s(fun(bool,bool),X233))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X233),s(bool,'F_0'))) | s(bool,X230) = s(bool,X234)) ) | $spl0),
  inference(forward_demodulation,[],[f37902,f141])).
fof(f37902,plain,(
  ( ! [X231,X233,X230,X234,X232] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X230))),s(bool,X231),s(bool,X232),s(fun(bool,bool),X233))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X233),s(bool,'F_0'))) | s(bool,X234) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X230) = s(bool,X234)) ) | $spl0),
  inference(superposition,[],[f2989,f11021])).
fof(f37873,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X91))),s(bool,X92),s(bool,X93),s(fun(bool,bool),X94))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X94),s(bool,X90))) | s(bool,'F_0') = s(bool,X91) | s(bool,'F_0') = s(bool,X90)) ) | $spl0),
  inference(superposition,[],[f2989,f612])).
fof(f37869,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X71))),s(bool,X72),s(bool,X73),s(fun(bool,bool),X74))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X74),s(bool,X70))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f2989,f205])).
fof(f37868,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))),s(bool,X66),s(bool,X67),s(fun(bool,bool),X68))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X68),s(bool,X65))) | s(bool,X65) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f2989,f205])).
fof(f37865,plain,(
  ( ! [X47,X54,X52,X50,X48,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X47),'const.semanticPrimitives.Match_1'(s(X47,happ(s(fun(X49,X47),happ(s(fun(X48,fun(X49,X47)),X52),s(X48,X50))),s(X49,X51))))),s(bool,X53),s(bool,X54),s(fun(X47,bool),X55))) | s(bool,'F_0') = s(bool,happ(s(fun(X47,bool),X55),s(X47,'const.pair.pair_CASE_2'(s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,X50),s(X49,X51))),s(fun(X48,fun(X49,X47)),X52)))))) )),
  inference(superposition,[],[f2989,f165])).
fof(f37863,plain,(
  ( ! [X30,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.Match_1'(s(X30,'const.pair.pair_CASE_2'(s('type.pair.prod'(X32,X31),'const.pair.,_2'(s(X32,X34),s(X31,X35))),s(fun(X32,fun(X31,X30)),X33))))),s(bool,X36),s(bool,X37),s(fun(X30,bool),X38))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X38),s(X30,happ(s(fun(X31,X30),happ(s(fun(X32,fun(X31,X30)),X33),s(X32,X34))),s(X31,X35)))))) )),
  inference(superposition,[],[f2989,f165])).
fof(f37243,plain,(
  ( ! [X68,X69,X67] : (s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,'T_0') = s(bool,g(s(X67,sK3(X68,X69,bool,X67)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,f(s(X67,sK3(X68,X69,bool,X67))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f37190,f141])).
fof(f37190,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,'T_0') = s(bool,g(s(X67,sK3(X68,X69,bool,X67)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,f(s(X67,sK3(X68,X69,bool,X67))))))) ) | $spl0),
  inference(superposition,[],[f2387,f8450])).
fof(f37204,plain,(
  ( ! [X45,X43,X46,X44] : (s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,g(s(X43,sK3(X44,X45,bool,X43)))) | s(bool,X46) = s(bool,f(s(X43,sK3(X44,X45,bool,X43)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X46)))) ) | $spl0),
  inference(trivial_inequality_removal,[],[f37183])).
fof(f37183,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'T_0') = s(bool,g(s(X43,sK3(X44,X45,bool,X43)))) | s(bool,X46) = s(bool,f(s(X43,sK3(X44,X45,bool,X43)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X46)))) ) | $spl0),
  inference(superposition,[],[f2387,f755])).
fof(f37206,plain,(
  ( ! [X21,X19,X22,X20] : (s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,X22) = s(bool,f(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f37177])).
fof(f37177,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,X22) = s(bool,f(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f2387,f612])).
fof(f37207,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f37172])).
fof(f37172,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f2387,f128])).
fof(f36087,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X7,happ(s(fun(bool,X7),X8),s(bool,'F_0'))) = s(X7,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(bool),X6),s(X7,X9),s(X7,X10),s(fun(bool,X7),X8))) | s(bool,'T_0') = s(bool,sK2(X6,bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X6) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X6)) )),
  inference(superposition,[],[f163,f2203])).
fof(f36011,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'(bool),X0) = s('type.semanticPrimitives.match_result'(bool),X1) | s(bool,'T_0') = s(bool,sK2(X1,bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X0) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X0)) )),
  inference(superposition,[],[f2203,f2203])).
fof(f35367,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0)) )),
  inference(superposition,[],[f3317,f144])).
fof(f34908,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,happ(s(fun(fun(X11,fun(X11,bool)),bool),'const.relation.WF_0'),s(fun(X11,fun(X11,bool)),X12))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X11,fun(X11,bool)),X12)))) )),
  inference(subsumption_resolution,[],[f34751,f117])).
fof(f34751,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(fun(X11,fun(X11,bool)),bool),'const.relation.WF_0'),s(fun(X11,fun(X11,bool)),X12))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X11,fun(X11,bool)),X12)))) )),
  inference(superposition,[],[f134,f546])).
fof(f34716,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X29,fun(X29,bool)),bool),'const.relation.WF_0'),s(fun(X29,fun(X29,bool)),happ(s(fun(X31,fun(X29,fun(X29,bool))),happ(s(fun(X30,fun(X31,fun(X29,fun(X29,bool)))),X34),s(X30,X32))),s(X31,X33))))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X29,fun(X29,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X30,X31),'const.pair.,_2'(s(X30,X32),s(X31,X33))),s(fun(X30,fun(X31,fun(X29,fun(X29,bool)))),X34)))))) )),
  inference(superposition,[],[f546,f165])).
fof(f34714,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X18,fun(X18,bool)),bool),'const.relation.WF_0'),s(fun(X18,fun(X18,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X20,X19),'const.pair.,_2'(s(X20,X22),s(X19,X23))),s(fun(X20,fun(X19,fun(X18,fun(X18,bool)))),X21))))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X18,fun(X18,bool)),happ(s(fun(X19,fun(X18,fun(X18,bool))),happ(s(fun(X20,fun(X19,fun(X18,fun(X18,bool)))),X21),s(X20,X22))),s(X19,X23)))))) )),
  inference(superposition,[],[f546,f165])).
fof(f34306,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f34250,f117])).
fof(f34250,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4069,f205])).
fof(f34228,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X1))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(unit_resulting_resolution,[],[f171,f4069,f153])).
fof(f34229,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X1))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X3))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(unit_resulting_resolution,[],[f4069,f4069,f153])).
fof(f34230,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X1))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(unit_resulting_resolution,[],[f4069,f130])).
fof(f33810,plain,(
  ( ! [X171] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X171))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X171)))) ) | $spl0),
  inference(superposition,[],[f2200,f7810])).
fof(f33933,plain,(
  ( ! [X54,X56,X55] : (s(bool,X55) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X54)),bool)) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X54)))) ) | $spl0),
  inference(subsumption_resolution,[],[f33740,f171])).
fof(f33740,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'F_0')) | s(bool,X55) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X54)),bool)) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X54)))) ) | $spl0),
  inference(superposition,[],[f1204,f7810])).
fof(f33930,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X6) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X5)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X5)))) ) | $spl0),
  inference(subsumption_resolution,[],[f33720,f171])).
fof(f33720,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X6) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X5)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X5)))) ) | $spl0),
  inference(superposition,[],[f153,f7810])).
fof(f33929,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1)))) ) | $spl0),
  inference(subsumption_resolution,[],[f33716,f171])).
fof(f33716,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1)))) ) | $spl0),
  inference(superposition,[],[f130,f7810])).
fof(f33535,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,f(s(bool,X5))) != s(X6,g(s(bool,X5))) | s(fun(bool,X6),X7) = s(fun(bool,X6),X8) | s(bool,'T_0') = s(bool,sK3(X7,X8,X6,bool)) | s(bool,X5) = s(bool,X9) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) ) | $spl0),
  inference(superposition,[],[f2378,f755])).
fof(f33583,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X4) | s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,X1,bool)) | s(bool,X0) = s(bool,X4)) ) | $spl0),
  inference(forward_demodulation,[],[f33534,f141])).
fof(f33534,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,X1,bool)) | s(bool,X4) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f2378,f11021])).
fof(f30513,plain,(
  ( ! [X198,X196,X197,X195] : (s(bool,X196) = s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,X196) = s(bool,X197) | s(bool,'F_0') = s(bool,X197) | s(bool,X195) = s(bool,'const.relation.WF_1'(s(fun(X198,fun(X198,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f30436,f171])).
fof(f30436,plain,(
  ( ! [X198,X196,X197,X195] : (p(s(bool,'F_0')) | s(bool,X196) = s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,X196) = s(bool,X197) | s(bool,'F_0') = s(bool,X197) | s(bool,X195) = s(bool,'const.relation.WF_1'(s(fun(X198,fun(X198,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f1204,f1796])).
fof(f30505,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X146) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X144),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X145))))) | s(bool,X146) = s(bool,X147) | s(bool,'F_0') = s(bool,X147)) ) | $spl0),
  inference(subsumption_resolution,[],[f30422,f171])).
fof(f30422,plain,(
  ( ! [X146,X144,X147,X145] : (p(s(bool,'F_0')) | s(bool,X146) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X144),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X145))))) | s(bool,X146) = s(bool,X147) | s(bool,'F_0') = s(bool,X147)) ) | $spl0),
  inference(superposition,[],[f1204,f152])).
fof(f30504,plain,(
  ( ! [X142,X143,X141] : (s(bool,X142) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X141),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,X142) = s(bool,X143) | s(bool,'F_0') = s(bool,X143)) ) | $spl0),
  inference(subsumption_resolution,[],[f30421,f171])).
fof(f30421,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,'F_0')) | s(bool,X142) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X141),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,X142) = s(bool,X143) | s(bool,'F_0') = s(bool,X143)) ) | $spl0),
  inference(superposition,[],[f1204,f126])).
fof(f30503,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X138),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,X139) = s(bool,X140) | s(bool,'F_0') = s(bool,X140)) ) | $spl0),
  inference(subsumption_resolution,[],[f30420,f171])).
fof(f30420,plain,(
  ( ! [X140,X138,X139] : (p(s(bool,'F_0')) | s(bool,X139) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X138),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,X139) = s(bool,X140) | s(bool,'F_0') = s(bool,X140)) ) | $spl0),
  inference(superposition,[],[f1204,f125])).
fof(f29190,plain,(
  ( ! [X198,X197] : (s(bool,'T_0') != s(bool,X197) | s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X198) = s(bool,'const.combin.I_1'(s(bool,X197)))) ) | $spl0),
  inference(superposition,[],[f27062,f141])).
fof(f29017,plain,(
  ( ! [X191,X192] : (s(bool,'F_0') != s(bool,X191) | s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X192) = s(bool,'const.combin.I_1'(s(bool,X191)))) ) | $spl0),
  inference(superposition,[],[f16731,f141])).
fof(f29092,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') != s(bool,happ(s(fun(X186,bool),happ(s(fun(X185,fun(X186,bool)),X189),s(X185,X187))),s(X186,X188))) | s(bool,X190) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X185,X186),'const.pair.,_2'(s(X185,X187),s(X186,X188))),s(fun(X185,fun(X186,bool)),X189)))) ) | $spl0),
  inference(forward_demodulation,[],[f29016,f141])).
fof(f29016,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'F_0') != s(bool,happ(s(fun(X186,bool),happ(s(fun(X185,fun(X186,bool)),X189),s(X185,X187))),s(X186,X188))) | s(bool,X190) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X190) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X185,X186),'const.pair.,_2'(s(X185,X187),s(X186,X188))),s(fun(X185,fun(X186,bool)),X189)))) ) | $spl0),
  inference(superposition,[],[f16731,f165])).
fof(f29091,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,'T_0') = s(bool,X184) | s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X180,X181),X182),s(fun(X180,fun(X181,bool)),X183))) | s(bool,X184) = s(bool,happ(s(fun(fun(X180,fun(X181,bool)),bool),happ(s(fun('type.pair.prod'(X180,X181),fun(fun(X180,fun(X181,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X180,X181),X182))),s(fun(X180,fun(X181,bool)),X183)))) ) | $spl0),
  inference(forward_demodulation,[],[f29015,f141])).
fof(f29015,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X180,X181),X182),s(fun(X180,fun(X181,bool)),X183))) | s(bool,X184) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X184) = s(bool,happ(s(fun(fun(X180,fun(X181,bool)),bool),happ(s(fun('type.pair.prod'(X180,X181),fun(fun(X180,fun(X181,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X180,X181),X182))),s(fun(X180,fun(X181,bool)),X183)))) ) | $spl0),
  inference(superposition,[],[f16731,f164])).
fof(f29090,plain,(
  ( ! [X173,X174] : (s(bool,'T_0') = s(bool,X174) | s(bool,'F_0') != s(bool,X173) | s(bool,X174) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X173)))) ) | $spl0),
  inference(forward_demodulation,[],[f29014,f141])).
fof(f29014,plain,(
  ( ! [X173,X174] : (s(bool,'F_0') != s(bool,X173) | s(bool,X174) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X174) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X173)))) ) | $spl0),
  inference(superposition,[],[f16731,f172])).
fof(f29089,plain,(
  ( ! [X167,X171,X169,X172,X170,X168] : (s(bool,'T_0') = s(bool,X172) | s(bool,'F_0') != s(bool,happ(s(fun(X167,bool),X171),s(X167,X168))) | s(bool,X172) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X167),'const.semanticPrimitives.Match_1'(s(X167,X168))),s(bool,X169),s(bool,X170),s(fun(X167,bool),X171)))) ) | $spl0),
  inference(forward_demodulation,[],[f29013,f141])).
fof(f29013,plain,(
  ( ! [X167,X171,X169,X172,X170,X168] : (s(bool,'F_0') != s(bool,happ(s(fun(X167,bool),X171),s(X167,X168))) | s(bool,X172) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X172) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X167),'const.semanticPrimitives.Match_1'(s(X167,X168))),s(bool,X169),s(bool,X170),s(fun(X167,bool),X171)))) ) | $spl0),
  inference(superposition,[],[f16731,f163])).
fof(f29088,plain,(
  ( ! [X165,X163,X166,X164,X162] : (s(bool,'T_0') = s(bool,X166) | s(bool,'F_0') != s(bool,X164) | s(bool,X166) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X162),'const.semanticPrimitives.Match_type_error_0'),s(bool,X163),s(bool,X164),s(fun(X162,bool),X165)))) ) | $spl0),
  inference(forward_demodulation,[],[f29012,f141])).
fof(f29012,plain,(
  ( ! [X165,X163,X166,X164,X162] : (s(bool,'F_0') != s(bool,X164) | s(bool,X166) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X166) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X162),'const.semanticPrimitives.Match_type_error_0'),s(bool,X163),s(bool,X164),s(fun(X162,bool),X165)))) ) | $spl0),
  inference(superposition,[],[f16731,f161])).
fof(f29087,plain,(
  ( ! [X158,X161,X159,X157,X160] : (s(bool,'T_0') = s(bool,X161) | s(bool,'F_0') != s(bool,X158) | s(bool,X161) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X157),'const.semanticPrimitives.No_match_0'),s(bool,X158),s(bool,X159),s(fun(X157,bool),X160)))) ) | $spl0),
  inference(forward_demodulation,[],[f29011,f141])).
fof(f29011,plain,(
  ( ! [X158,X161,X159,X157,X160] : (s(bool,'F_0') != s(bool,X158) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X161) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X157),'const.semanticPrimitives.No_match_0'),s(bool,X158),s(bool,X159),s(fun(X157,bool),X160)))) ) | $spl0),
  inference(superposition,[],[f16731,f160])).
fof(f29076,plain,(
  ( ! [X111,X109,X110] : (s(bool,'T_0') = s(bool,X111) | s(bool,'F_0') != s(bool,'const.relation.WF_1'(s(fun(X109,fun(X109,bool)),X110))) | s(bool,X111) = s(bool,happ(s(fun(fun(X109,fun(X109,bool)),bool),'const.relation.WF_0'),s(fun(X109,fun(X109,bool)),X110)))) ) | $spl0),
  inference(forward_demodulation,[],[f28990,f141])).
fof(f28990,plain,(
  ( ! [X111,X109,X110] : (s(bool,'F_0') != s(bool,'const.relation.WF_1'(s(fun(X109,fun(X109,bool)),X110))) | s(bool,X111) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X111) = s(bool,happ(s(fun(fun(X109,fun(X109,bool)),bool),'const.relation.WF_0'),s(fun(X109,fun(X109,bool)),X110)))) ) | $spl0),
  inference(superposition,[],[f16731,f143])).
fof(f28932,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X18) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f16731,f128])).
fof(f29050,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(forward_demodulation,[],[f28928,f141])).
fof(f28928,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f16731,f11021])).
fof(f28192,plain,(
  ( ! [X182,X183] : (s(bool,X182) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X183) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X182) = s(bool,X183)) ) | $spl0),
  inference(subsumption_resolution,[],[f27214,f141])).
fof(f27214,plain,(
  ( ! [X182,X183] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X182) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X183) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X182) = s(bool,X183)) ) | $spl0),
  inference(superposition,[],[f11207,f11021])).
fof(f27072,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,X39) | s(bool,X39) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X38) = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f11330,f11021])).
fof(f27071,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,'const.combin.I_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X33) = s(bool,X35)) ) | $spl0),
  inference(superposition,[],[f8450,f11021])).
fof(f27070,plain,(
  ( ! [X30,X31,X32] : (s(bool,X31) = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X30) = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f8211,f11021])).
fof(f27069,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) != s(bool,'const.combin.I_1'(s(bool,X27))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,X29) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X27) = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f1943,f11021])).
fof(f27068,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X25))) | s(bool,X26) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X25) = s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f1881,f11021])).
fof(f27067,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X23))) | s(bool,X24) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,X24)) ) | $spl0),
  inference(superposition,[],[f1858,f11021])).
fof(f27066,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,'const.combin.I_1'(s(bool,X19))) | s(bool,'F_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X19) = s(bool,X21)) ) | $spl0),
  inference(superposition,[],[f1629,f11021])).
fof(f27062,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X6) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X5) = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f177,f11021])).
fof(f28103,plain,(
  ( ! [X241,X242,X240] : (s(bool,X241) = s(bool,'const.combin.I_1'(s(bool,X240))) | s(bool,X241) = s(bool,X242) | s(bool,'T_0') = s(bool,X242) | s(bool,X240) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f26971,f177])).
fof(f26971,plain,(
  ( ! [X241,X242,X240] : (s(bool,X241) = s(bool,'const.combin.I_1'(s(bool,X240))) | s(bool,X241) = s(bool,X242) | s(bool,'T_0') = s(bool,X242) | s(bool,X240) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f10324,f11021])).
fof(f28094,plain,(
  ( ! [X208,X209,X207] : (s(bool,X207) != s(bool,X208) | s(bool,X209) = s(bool,'const.combin.I_1'(s(bool,X208))) | s(bool,'F_0') = s(bool,X209) | s(bool,X207) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f26961,f177])).
fof(f26961,plain,(
  ( ! [X208,X209,X207] : (s(bool,X207) != s(bool,X208) | s(bool,X209) = s(bool,'const.combin.I_1'(s(bool,X208))) | s(bool,'F_0') = s(bool,X209) | s(bool,X207) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1236,f11021])).
fof(f27963,plain,(
  ( ! [X246,X248,X247] : (s(bool,X246) != s(bool,X247) | s(bool,X247) = s(bool,'const.relation.WF_1'(s(fun(X248,fun(X248,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X246)) ) | $spl0),
  inference(subsumption_resolution,[],[f26794,f1881])).
fof(f26794,plain,(
  ( ! [X246,X248,X247] : (s(bool,X246) != s(bool,X247) | s(bool,X247) = s(bool,'const.relation.WF_1'(s(fun(X248,fun(X248,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X246)) ) | $spl0),
  inference(superposition,[],[f12476,f11021])).
fof(f27962,plain,(
  ( ! [X241,X242,X240] : (s(bool,X241) = s(bool,'const.combin.I_1'(s(bool,X240))) | s(bool,X241) = s(bool,X242) | s(bool,'T_0') = s(bool,X242) | s(bool,'F_0') = s(bool,X240)) ) | $spl0),
  inference(subsumption_resolution,[],[f26792,f1881])).
fof(f26792,plain,(
  ( ! [X241,X242,X240] : (s(bool,X241) = s(bool,'const.combin.I_1'(s(bool,X240))) | s(bool,X241) = s(bool,X242) | s(bool,'T_0') = s(bool,X242) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X240)) ) | $spl0),
  inference(superposition,[],[f10324,f11021])).
fof(f27870,plain,(
  ( ! [X130,X128,X129] : (s(bool,X129) != s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,X129) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X128) = s(bool,X130) | s(bool,'F_0') = s(bool,X130)) ) | $spl0),
  inference(subsumption_resolution,[],[f26645,f11330])).
fof(f26645,plain,(
  ( ! [X130,X128,X129] : (s(bool,X129) != s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,X129) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X129) | s(bool,X128) = s(bool,X130) | s(bool,'F_0') = s(bool,X130)) ) | $spl0),
  inference(superposition,[],[f1617,f11021])).
fof(f27727,plain,(
  ( ! [X109,X108] : (s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,X108))) | s(bool,X108) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f26440,f177])).
fof(f26440,plain,(
  ( ! [X109,X108] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,X108))) | s(bool,X108) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f11021,f8450])).
fof(f27724,plain,(
  ( ! [X99,X98] : (s(bool,X99) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X99) = s(bool,'const.combin.I_1'(s(bool,X98))) | s(bool,'F_0') = s(bool,X98)) ) | $spl0),
  inference(subsumption_resolution,[],[f26438,f177])).
fof(f26438,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X99) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X99) = s(bool,'const.combin.I_1'(s(bool,X98))) | s(bool,'F_0') = s(bool,X98)) ) | $spl0),
  inference(superposition,[],[f11021,f207])).
fof(f26385,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f11021,f612])).
fof(f26384,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X7)))) ) | $spl0),
  inference(superposition,[],[f11021,f755])).
fof(f26383,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f11021,f11021])).
fof(f27671,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X76) | s(bool,X76) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X75) = s(bool,X77)) ) | $spl0),
  inference(forward_demodulation,[],[f27670,f141])).
fof(f27670,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,X76) | s(bool,X76) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X77) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X75) = s(bool,X77)) ) | $spl0),
  inference(subsumption_resolution,[],[f26341,f177])).
fof(f26341,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,X76) | s(bool,X76) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X77) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X75) = s(bool,X77)) ) | $spl0),
  inference(superposition,[],[f11021,f10324])).
fof(f26209,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X4)))) ) | $spl0),
  inference(superposition,[],[f11021,f8450])).
fof(f27586,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) = s(bool,'const.combin.I_1'(s(bool,X62))) | s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f26204,f177])).
fof(f26204,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) = s(bool,'const.combin.I_1'(s(bool,X62))) | s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f11021,f10324])).
fof(f26189,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f11021,f128])).
fof(f26187,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) = s(bool,'const.combin.I_1'(s(bool,X8))) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f11021,f612])).
fof(f26186,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X7)))) ) | $spl0),
  inference(superposition,[],[f11021,f755])).
fof(f26185,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f11021,f11021])).
fof(f23419,plain,(
  ( ! [X125,X126] : (s(bool,X126) != s(bool,'const.combin.I_1'(s(bool,X125))) | s(bool,X126) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(subsumption_resolution,[],[f23321,f11330])).
fof(f23321,plain,(
  ( ! [X125,X126] : (s(bool,X126) != s(bool,'const.combin.I_1'(s(bool,X125))) | s(bool,X126) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(superposition,[],[f1617,f8211])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.match_result_rel_def|split|3')).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0')))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0')))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.match_result_rel_def|split|5')).
fof(f134,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f135,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f152,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X6] : s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.match_result_rel_def|split|6')).
fof(f171,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f120,f117])).
fof(f120,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f207,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f128,f141])).
fof(f278,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f172,f128])).
fof(f316,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X5),'const.semanticPrimitives.No_match_0'),s(bool,X6),s(bool,X7),s(fun(X5,bool),X8))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f160,f128])).
fof(f385,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f366,f141])).
fof(f366,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(superposition,[],[f207,f172])).
fof(f386,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X6),'const.semanticPrimitives.No_match_0'),s(bool,X7),s(bool,X8),s(fun(X6,bool),X9)))) )),
  inference(forward_demodulation,[],[f367,f141])).
fof(f367,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X6),'const.semanticPrimitives.No_match_0'),s(bool,X7),s(bool,X8),s(fun(X6,bool),X9)))) )),
  inference(superposition,[],[f207,f160])).
fof(f387,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f373,f117])).
fof(f373,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f134,f207])).
fof(f389,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f388,f141])).
fof(f388,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f377,f117])).
fof(f377,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f156,f207])).
fof(f414,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X5),'const.semanticPrimitives.Match_type_error_0'),s(bool,X6),s(bool,X7),s(fun(X5,bool),X8))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f161,f128])).
fof(f440,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,X73) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X71),'const.semanticPrimitives.Match_type_error_0'),s(bool,X72),s(bool,X73),s(fun(X71,bool),X74)))) )),
  inference(forward_demodulation,[],[f429,f141])).
fof(f429,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X73))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X71),'const.semanticPrimitives.Match_type_error_0'),s(bool,X72),s(bool,X73),s(fun(X71,bool),X74)))) )),
  inference(superposition,[],[f207,f161])).
fof(f462,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f449,f141])).
fof(f449,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(superposition,[],[f387,f172])).
fof(f463,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X6),'const.semanticPrimitives.No_match_0'),s(bool,X7),s(bool,X8),s(fun(X6,bool),X9)))) )),
  inference(forward_demodulation,[],[f450,f141])).
fof(f450,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X6),'const.semanticPrimitives.No_match_0'),s(bool,X7),s(bool,X8),s(fun(X6,bool),X9)))) )),
  inference(superposition,[],[f387,f160])).
fof(f464,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') != s(bool,X12) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X10),'const.semanticPrimitives.Match_type_error_0'),s(bool,X11),s(bool,X12),s(fun(X10,bool),X13)))) )),
  inference(forward_demodulation,[],[f451,f141])).
fof(f451,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X12))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X10),'const.semanticPrimitives.Match_type_error_0'),s(bool,X11),s(bool,X12),s(fun(X10,bool),X13)))) )),
  inference(superposition,[],[f387,f161])).
fof(f546,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X2,fun(X2,bool)),bool),'const.relation.WF_0'),s(fun(X2,fun(X2,bool)),X3))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),X3)))) )),
  inference(superposition,[],[f143,f128])).
fof(f576,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X21,fun(X21,bool)),X22))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X21,fun(X21,bool)),bool),'const.relation.WF_0'),s(fun(X21,fun(X21,bool)),X22)))) )),
  inference(forward_demodulation,[],[f556,f141])).
fof(f556,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X21,fun(X21,bool)),X22))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X21,fun(X21,bool)),bool),'const.relation.WF_0'),s(fun(X21,fun(X21,bool)),X22)))) )),
  inference(superposition,[],[f207,f143])).
fof(f578,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),X27))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X26,fun(X26,bool)),bool),'const.relation.WF_0'),s(fun(X26,fun(X26,bool)),X27)))) )),
  inference(forward_demodulation,[],[f558,f141])).
fof(f558,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),X27))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X26,fun(X26,bool)),bool),'const.relation.WF_0'),s(fun(X26,fun(X26,bool)),X27)))) )),
  inference(superposition,[],[f387,f143])).
fof(f594,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f584,f117])).
fof(f584,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f389,f200])).
fof(f596,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.relation.WF_1'(s(fun(X8,fun(X8,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f586,f117])).
fof(f586,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.relation.WF_1'(s(fun(X8,fun(X8,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f389,f183])).
fof(f598,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f588,f117])).
fof(f588,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X12) = s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f389,f207])).
fof(f608,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X8)))) ) | $spl0),
  inference(superposition,[],[f594,f141])).
fof(f609,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X9)))) ) | $spl0),
  inference(superposition,[],[f594,f172])).
fof(f610,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X11) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X10),'const.semanticPrimitives.No_match_0'),s(bool,X11),s(bool,X12),s(fun(X10,bool),X13)))) ) | $spl0),
  inference(superposition,[],[f594,f160])).
fof(f611,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X16) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X14),'const.semanticPrimitives.Match_type_error_0'),s(bool,X15),s(bool,X16),s(fun(X14,bool),X17)))) ) | $spl0),
  inference(superposition,[],[f594,f161])).
fof(f691,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,happ(s(fun(X17,bool),X18),s(X17,'const.min.@_1'(s(fun(X17,bool),X18)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) ) | $spl0),
  inference(subsumption_resolution,[],[f690,f117])).
fof(f690,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,'const.min.@_1'(s(fun(X17,bool),X18)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) ) | $spl0),
  inference(forward_demodulation,[],[f629,f200])).
fof(f629,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,'const.min.@_1'(s(fun(X17,bool),X18)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) ) | $spl0),
  inference(superposition,[],[f149,f594])).
fof(f751,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.relation.WF_1'(s(fun(X27,fun(X27,bool)),X28))) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X27,fun(X27,bool)),bool),'const.relation.WF_0'),s(fun(X27,fun(X27,bool)),X28)))) ) | $spl0),
  inference(superposition,[],[f612,f143])).
fof(f755,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X35)))) ) | $spl0),
  inference(superposition,[],[f612,f141])).
fof(f756,plain,(
  ( ! [X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X37)))) ) | $spl0),
  inference(superposition,[],[f612,f172])).
fof(f757,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) = s(bool,X43) | s(bool,'F_0') = s(bool,X43) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X39),'const.semanticPrimitives.No_match_0'),s(bool,X40),s(bool,X41),s(fun(X39,bool),X42)))) ) | $spl0),
  inference(superposition,[],[f612,f160])).
fof(f758,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X46) = s(bool,X48) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X44),'const.semanticPrimitives.Match_type_error_0'),s(bool,X45),s(bool,X46),s(fun(X44,bool),X47)))) ) | $spl0),
  inference(superposition,[],[f612,f161])).
fof(f792,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X18,fun(X18,bool)),X19))) | s(bool,X20) = s(bool,happ(s(fun(fun(X18,fun(X18,bool)),bool),'const.relation.WF_0'),s(fun(X18,fun(X18,bool)),X19))) | s(bool,'F_0') = s(bool,X20)) ) | $spl0),
  inference(superposition,[],[f612,f143])).
fof(f797,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,X28) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f612,f172])).
fof(f798,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,X31) | s(bool,X34) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X30),'const.semanticPrimitives.No_match_0'),s(bool,X31),s(bool,X32),s(fun(X30,bool),X33))) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f612,f160])).
fof(f799,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,'F_0') = s(bool,X37) | s(bool,X39) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X35),'const.semanticPrimitives.Match_type_error_0'),s(bool,X36),s(bool,X37),s(fun(X35,bool),X38))) | s(bool,'F_0') = s(bool,X39)) ) | $spl0),
  inference(superposition,[],[f612,f161])).
fof(f1204,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f926,f171])).
fof(f926,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f153,f612])).
fof(f1232,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,'const.relation.WF_1'(s(fun(X29,fun(X29,bool)),X30))) | s(bool,X31) = s(bool,happ(s(fun(fun(X29,fun(X29,bool)),bool),'const.relation.WF_0'),s(fun(X29,fun(X29,bool)),X30))) | s(bool,'F_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f900,f143])).
fof(f1236,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X37) | s(bool,X38) = s(bool,'const.combin.I_1'(s(bool,X37))) | s(bool,'F_0') = s(bool,X38)) ) | $spl0),
  inference(superposition,[],[f900,f141])).
fof(f1237,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,X40) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X39))) | s(bool,'F_0') = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f900,f172])).
fof(f1238,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,X45) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X41),'const.semanticPrimitives.No_match_0'),s(bool,X42),s(bool,X43),s(fun(X41,bool),X44))) | s(bool,'F_0') = s(bool,X45)) ) | $spl0),
  inference(superposition,[],[f900,f160])).
fof(f1239,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') != s(bool,X48) | s(bool,X50) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X46),'const.semanticPrimitives.Match_type_error_0'),s(bool,X47),s(bool,X48),s(fun(X46,bool),X49))) | s(bool,'F_0') = s(bool,X50)) ) | $spl0),
  inference(superposition,[],[f900,f161])).
fof(f1348,plain,(
  ( ! [X41,X42,X40] : (s(bool,X42) = s(bool,'const.relation.WF_1'(s(fun(X40,fun(X40,bool)),X41))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X40,fun(X40,bool)),bool),'const.relation.WF_0'),s(fun(X40,fun(X40,bool)),X41))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X42)))) ) | $spl0),
  inference(superposition,[],[f755,f143])).
fof(f1354,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X54)))) ) | $spl0),
  inference(superposition,[],[f755,f172])).
fof(f1355,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X56) = s(bool,X59) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X55),'const.semanticPrimitives.No_match_0'),s(bool,X56),s(bool,X57),s(fun(X55,bool),X58))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X59)))) ) | $spl0),
  inference(superposition,[],[f755,f160])).
fof(f1356,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X60),'const.semanticPrimitives.Match_type_error_0'),s(bool,X61),s(bool,X62),s(fun(X60,bool),X63))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X64)))) ) | $spl0),
  inference(superposition,[],[f755,f161])).
fof(f1357,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X3)))) ) | $spl0),
  inference(superposition,[],[f755,f755])).
fof(f1358,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6))) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f755,f612])).
fof(f1360,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X14))) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f755,f128])).
fof(f1376,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X27,fun(X27,bool)),X28))) | s(bool,X29) = s(bool,happ(s(fun(fun(X27,fun(X27,bool)),bool),'const.relation.WF_0'),s(fun(X27,fun(X27,bool)),X28))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f755,f143])).
fof(f1381,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X39) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X39)))) ) | $spl0),
  inference(superposition,[],[f755,f172])).
fof(f1382,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X44) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X40),'const.semanticPrimitives.No_match_0'),s(bool,X41),s(bool,X42),s(fun(X40,bool),X43))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X44)))) ) | $spl0),
  inference(superposition,[],[f755,f160])).
fof(f1383,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'F_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X45),'const.semanticPrimitives.Match_type_error_0'),s(bool,X46),s(bool,X47),s(fun(X45,bool),X48))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X49)))) ) | $spl0),
  inference(superposition,[],[f755,f161])).
fof(f1515,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f177,f755])).
fof(f1617,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f755])).
fof(f1629,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1611])).
fof(f1611,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(equality_factoring,[],[f755])).
fof(f1757,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X1)))) ) | $spl0),
  inference(condensation,[],[f1756])).
fof(f1756,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X24)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1313])).
fof(f1313,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X24)))) ) | $spl0),
  inference(superposition,[],[f755,f755])).
fof(f1764,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(condensation,[],[f1763])).
fof(f1763,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1305])).
fof(f1305,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f755,f128])).
fof(f1796,plain,(
  ( ! [X30,X31] : (s(bool,X31) = s(bool,'const.relation.WF_1'(s(fun(X30,fun(X30,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1377,f177])).
fof(f1377,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X31) = s(bool,'const.relation.WF_1'(s(fun(X30,fun(X30,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl0),
  inference(superposition,[],[f755,f183])).
fof(f1798,plain,(
  ( ! [X35,X34] : (s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(subsumption_resolution,[],[f1379,f177])).
fof(f1379,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f755,f207])).
fof(f1850,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) != s(bool,X47) | s(bool,X47) = s(bool,X48) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X46)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1483,f177])).
fof(f1483,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) != s(bool,X47) | s(bool,X47) = s(bool,X48) | s(bool,'F_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X46)))) ) | $spl0),
  inference(superposition,[],[f900,f755])).
fof(f1858,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl0),
  inference(condensation,[],[f1857])).
fof(f1857,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') != s(bool,X81) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1498,f387])).
fof(f1498,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') != s(bool,X81) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X80))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl0),
  inference(superposition,[],[f387,f755])).
fof(f1875,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f1584,f171])).
fof(f1584,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f130,f755])).
fof(f1880,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'F_0') = s(bool,X26)) ) | $spl0),
  inference(subsumption_resolution,[],[f1592,f171])).
fof(f1592,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X25) = s(bool,'const.combin.I_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'F_0') = s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f256,f755])).
fof(f1881,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f177,f1629])).
fof(f1913,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f1629,f755])).
fof(f1914,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f1629,f612])).
fof(f1916,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f1629,f128])).
fof(f1943,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl0),
  inference(condensation,[],[f1942])).
fof(f1942,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f1894])).
fof(f1894,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | $spl0),
  inference(superposition,[],[f1629,f755])).
fof(f2023,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f1982,f177])).
fof(f1982,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X4) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f127,f612])).
fof(f2024,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1988,f177])).
fof(f1988,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X16) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X16)))) ) | $spl0),
  inference(superposition,[],[f127,f755])).
fof(f2044,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X1)))) ) | $spl0),
  inference(superposition,[],[f1881,f755])).
fof(f2203,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),X0) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK2(X0,bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X0) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X0)) )),
  inference(superposition,[],[f144,f128])).
fof(f2205,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.match_result'(bool),X2) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X2) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X2) | s(bool,X3) = s(bool,sK2(X2,bool)) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f144,f612])).
fof(f2206,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.match_result'(bool),X4) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X5))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X4) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X4) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,sK2(X4,bool))) ) | $spl0),
  inference(superposition,[],[f144,f612])).
fof(f2209,plain,(
  ( ! [X10,X11] : (s('type.semanticPrimitives.match_result'(bool),X10) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X11))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X10) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X10) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK2(X10,bool))))) ) | $spl0),
  inference(superposition,[],[f144,f755])).
fof(f2211,plain,(
  ( ! [X14,X15] : (s('type.semanticPrimitives.match_result'(bool),X14) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X14) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X14) | s(bool,X15) = s(bool,sK2(X14,bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl0),
  inference(superposition,[],[f144,f755])).
fof(f2223,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.match_result'(bool),X1) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s(bool,'F_0') = s(bool,sK2(X1,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f2204,f200])).
fof(f2204,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.match_result'(bool),X1) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X1) | s(bool,'F_0') = s(bool,sK2(X1,bool))) ) | $spl0),
  inference(superposition,[],[f144,f594])).
fof(f2231,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f1858,f755])).
fof(f2234,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f1858,f128])).
fof(f2378,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f159,f128])).
fof(f2380,plain,(
  ( ! [X6,X8,X7,X9] : (s(X8,f(s(bool,'F_0'))) != s(X8,g(s(bool,'F_0'))) | s(fun(bool,X8),X6) = s(fun(bool,X8),X7) | s(bool,X9) = s(bool,sK3(X6,X7,X8,bool)) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f159,f612])).
fof(f2381,plain,(
  ( ! [X12,X10,X13,X11] : (s(X12,f(s(bool,X13))) != s(X12,g(s(bool,X13))) | s(fun(bool,X12),X10) = s(fun(bool,X12),X11) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,sK3(X10,X11,X12,bool))) ) | $spl0),
  inference(superposition,[],[f159,f612])).
fof(f2384,plain,(
  ( ! [X24,X23,X25,X22] : (s(X24,f(s(bool,X25))) != s(X24,g(s(bool,X25))) | s(fun(bool,X24),X22) = s(fun(bool,X24),X23) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X22,X23,X24,bool))))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2385,plain,(
  ( ! [X28,X26,X29,X27] : (s(X28,f(s(bool,X29))) != s(X28,g(s(bool,X29))) | s(fun(bool,X28),X26) = s(fun(bool,X28),X27) | s(bool,'F_0') = s(bool,sK3(X26,X27,X28,bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2386,plain,(
  ( ! [X30,X33,X31,X32] : (s(X32,f(s(bool,'F_0'))) != s(X32,g(s(bool,'F_0'))) | s(fun(bool,X32),X30) = s(fun(bool,X32),X31) | s(bool,X33) = s(bool,sK3(X30,X31,X32,bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X33)))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2387,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f159,f128])).
fof(f2389,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(X6,sK3(X7,X8,bool,X6)))) | s(fun(X6,bool),X7) = s(fun(X6,bool),X8) | s(bool,X9) = s(bool,g(s(X6,sK3(X7,X8,bool,X6)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f159,f612])).
fof(f2390,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) != s(bool,f(s(X10,sK3(X11,X12,bool,X10)))) | s(fun(X10,bool),X11) = s(fun(X10,bool),X12) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,g(s(X10,sK3(X11,X12,bool,X10))))) ) | $spl0),
  inference(superposition,[],[f159,f612])).
fof(f2393,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X25) != s(bool,f(s(X22,sK3(X23,X24,bool,X22)))) | s(fun(X22,bool),X23) = s(fun(X22,bool),X24) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X22,sK3(X23,X24,bool,X22))))))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2394,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) != s(bool,f(s(X26,sK3(X27,X28,bool,X26)))) | s(fun(X26,bool),X27) = s(fun(X26,bool),X28) | s(bool,'F_0') = s(bool,g(s(X26,sK3(X27,X28,bool,X26)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2395,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,f(s(X30,sK3(X31,X32,bool,X30)))) | s(fun(X30,bool),X31) = s(fun(X30,bool),X32) | s(bool,X33) = s(bool,g(s(X30,sK3(X31,X32,bool,X30)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X33)))) ) | $spl0),
  inference(superposition,[],[f159,f755])).
fof(f2396,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'F_0') = s(bool,sK3(X3,X4,X5,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f2379,f200])).
fof(f2379,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) != s(X5,g(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'F_0') = s(bool,sK3(X3,X4,X5,bool))) ) | $spl0),
  inference(superposition,[],[f159,f594])).
fof(f2397,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK3(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'F_0') = s(bool,g(s(X3,sK3(X4,X5,bool,X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2388,f200])).
fof(f2388,plain,(
  ( ! [X4,X5,X3] : (s(bool,f(s(X3,sK3(X4,X5,bool,X3)))) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'F_0') = s(bool,g(s(X3,sK3(X4,X5,bool,X3))))) ) | $spl0),
  inference(superposition,[],[f159,f594])).
fof(f2431,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f1943,f755])).
fof(f2432,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f1943,f612])).
fof(f2434,plain,(
  ( ! [X10,X9] : (s(bool,X10) != s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f1943,f128])).
fof(f2989,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X6),'const.semanticPrimitives.Match_1'(s(X6,X7))),s(bool,X8),s(bool,X9),s(fun(X6,bool),X10))) | s(bool,'F_0') = s(bool,happ(s(fun(X6,bool),X10),s(X6,X7)))) )),
  inference(superposition,[],[f163,f128])).
fof(f2991,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),X20),s(X16,X17))) | s(bool,X21) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X16),'const.semanticPrimitives.Match_1'(s(X16,X17))),s(bool,X18),s(bool,X19),s(fun(X16,bool),X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl0),
  inference(superposition,[],[f163,f612])).
fof(f2992,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X27) = s(bool,happ(s(fun(X22,bool),X26),s(X22,X23))) | s(bool,'F_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X22),'const.semanticPrimitives.Match_1'(s(X22,X23))),s(bool,X24),s(bool,X25),s(fun(X22,bool),X26)))) ) | $spl0),
  inference(superposition,[],[f163,f612])).
fof(f2996,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X51) = s(bool,happ(s(fun(X46,bool),X50),s(X46,X47))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X46),'const.semanticPrimitives.Match_1'(s(X46,X47))),s(bool,X48),s(bool,X49),s(fun(X46,bool),X50))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X51)))) ) | $spl0),
  inference(superposition,[],[f163,f755])).
fof(f2997,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),X56),s(X52,X53))) | s(bool,X57) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X52),'const.semanticPrimitives.Match_1'(s(X52,X53))),s(bool,X54),s(bool,X55),s(fun(X52,bool),X56))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X57)))) ) | $spl0),
  inference(superposition,[],[f163,f755])).
fof(f3029,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') != s(bool,happ(s(fun(X185,bool),X189),s(X185,X186))) | s(bool,X190) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X185),'const.semanticPrimitives.Match_1'(s(X185,X186))),s(bool,X187),s(bool,X188),s(fun(X185,bool),X189))) | s(bool,'F_0') = s(bool,X190)) ) | $spl0),
  inference(superposition,[],[f900,f163])).
fof(f3062,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X15),s(X11,X12))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X11),'const.semanticPrimitives.Match_1'(s(X11,X12))),s(bool,X13),s(bool,X14),s(fun(X11,bool),X15)))) ) | $spl0),
  inference(forward_demodulation,[],[f2990,f200])).
fof(f2990,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,happ(s(fun(X11,bool),X15),s(X11,X12))) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X11),'const.semanticPrimitives.Match_1'(s(X11,X12))),s(bool,X13),s(bool,X14),s(fun(X11,bool),X15)))) ) | $spl0),
  inference(superposition,[],[f163,f594])).
fof(f3075,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'F_0') != s(bool,happ(s(fun(X121,bool),X125),s(X121,X122))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X121),'const.semanticPrimitives.Match_1'(s(X121,X122))),s(bool,X123),s(bool,X124),s(fun(X121,bool),X125)))) )),
  inference(forward_demodulation,[],[f3018,f141])).
fof(f3018,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X121,bool),X125),s(X121,X122))))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X121),'const.semanticPrimitives.Match_1'(s(X121,X122))),s(bool,X123),s(bool,X124),s(fun(X121,bool),X125)))) )),
  inference(superposition,[],[f387,f163])).
fof(f3250,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X0)) )),
  inference(superposition,[],[f152,f144])).
fof(f3468,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f3417,f177])).
fof(f3417,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f2368,f612])).
fof(f3469,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X7))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3420,f177])).
fof(f3420,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X7))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f2368,f755])).
fof(f3570,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X8),s(X7,X9))),s(fun(X6,fun(X7,bool)),X10))) | s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X8))),s(X7,X9)))) )),
  inference(superposition,[],[f165,f128])).
fof(f3572,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),happ(s(fun(X16,fun(X17,bool)),X20),s(X16,X18))),s(X17,X19))) | s(bool,X21) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X16,X17),'const.pair.,_2'(s(X16,X18),s(X17,X19))),s(fun(X16,fun(X17,bool)),X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl0),
  inference(superposition,[],[f165,f612])).
fof(f3573,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X27) = s(bool,happ(s(fun(X23,bool),happ(s(fun(X22,fun(X23,bool)),X26),s(X22,X24))),s(X23,X25))) | s(bool,'F_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),'const.pair.,_2'(s(X22,X24),s(X23,X25))),s(fun(X22,fun(X23,bool)),X26)))) ) | $spl0),
  inference(superposition,[],[f165,f612])).
fof(f3577,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X51) = s(bool,happ(s(fun(X47,bool),happ(s(fun(X46,fun(X47,bool)),X50),s(X46,X48))),s(X47,X49))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X46,X47),'const.pair.,_2'(s(X46,X48),s(X47,X49))),s(fun(X46,fun(X47,bool)),X50))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X51)))) ) | $spl0),
  inference(superposition,[],[f165,f755])).
fof(f3578,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),happ(s(fun(X52,fun(X53,bool)),X56),s(X52,X54))),s(X53,X55))) | s(bool,X57) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X54),s(X53,X55))),s(fun(X52,fun(X53,bool)),X56))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X57)))) ) | $spl0),
  inference(superposition,[],[f165,f755])).
fof(f3622,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X141,bool),happ(s(fun(X140,fun(X141,bool)),X144),s(X140,X142))),s(X141,X143))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X140,X141),'const.pair.,_2'(s(X140,X142),s(X141,X143))),s(fun(X140,fun(X141,bool)),X144)))) )),
  inference(superposition,[],[f207,f165])).
fof(f3624,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X152,bool),happ(s(fun(X151,fun(X152,bool)),X155),s(X151,X153))),s(X152,X154))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X151,X152),'const.pair.,_2'(s(X151,X153),s(X152,X154))),s(fun(X151,fun(X152,bool)),X155)))) )),
  inference(superposition,[],[f387,f165])).
fof(f3631,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X192,bool),happ(s(fun(X191,fun(X192,bool)),X195),s(X191,X193))),s(X192,X194))))) | s(bool,X196) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X191,X192),'const.pair.,_2'(s(X191,X193),s(X192,X194))),s(fun(X191,fun(X192,bool)),X195))) | s(bool,'F_0') = s(bool,X196)) ) | $spl0),
  inference(superposition,[],[f755,f165])).
fof(f3635,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (s(bool,'T_0') != s(bool,happ(s(fun(X216,bool),happ(s(fun(X215,fun(X216,bool)),X219),s(X215,X217))),s(X216,X218))) | s(bool,X220) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X215,X216),'const.pair.,_2'(s(X215,X217),s(X216,X218))),s(fun(X215,fun(X216,bool)),X219))) | s(bool,'F_0') = s(bool,X220)) ) | $spl0),
  inference(superposition,[],[f900,f165])).
fof(f3637,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X227,bool),happ(s(fun(X226,fun(X227,bool)),X230),s(X226,X228))),s(X227,X229))))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X226,X227),'const.pair.,_2'(s(X226,X228),s(X227,X229))),s(fun(X226,fun(X227,bool)),X230)))) ) | $spl0),
  inference(superposition,[],[f1764,f165])).
fof(f3669,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X15),s(X11,X13))),s(X12,X14))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X13),s(X12,X14))),s(fun(X11,fun(X12,bool)),X15)))) ) | $spl0),
  inference(forward_demodulation,[],[f3571,f200])).
fof(f3571,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X15),s(X11,X13))),s(X12,X14))) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X13),s(X12,X14))),s(fun(X11,fun(X12,bool)),X15)))) ) | $spl0),
  inference(superposition,[],[f165,f594])).
fof(f3684,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,'F_0') != s(bool,happ(s(fun(X222,bool),happ(s(fun(X221,fun(X222,bool)),X225),s(X221,X223))),s(X222,X224))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X221,X222),'const.pair.,_2'(s(X221,X223),s(X222,X224))),s(fun(X221,fun(X222,bool)),X225)))) ) | $spl0),
  inference(forward_demodulation,[],[f3636,f141])).
fof(f3636,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,happ(s(fun(X222,bool),happ(s(fun(X221,fun(X222,bool)),X225),s(X221,X223))),s(X222,X224))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X221,X222),'const.pair.,_2'(s(X221,X223),s(X222,X224))),s(fun(X221,fun(X222,bool)),X225)))) ) | $spl0),
  inference(superposition,[],[f1629,f165])).
fof(f3756,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f3704,f177])).
fof(f3704,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1236,f612])).
fof(f3757,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X11) = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3707,f177])).
fof(f3707,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X11) = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) ) | $spl0),
  inference(superposition,[],[f1236,f755])).
fof(f3762,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(forward_demodulation,[],[f3761,f200])).
fof(f3761,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))))) ) | $spl0),
  inference(condensation,[],[f3760])).
fof(f3760,plain,(
  ( ! [X39,X38] : (s(bool,X39) = s(bool,'const.combin.I_1'(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X38)) ) | $spl0),
  inference(subsumption_resolution,[],[f3722,f128])).
fof(f3722,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,X39) = s(bool,'const.combin.I_1'(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X38)) ) | $spl0),
  inference(superposition,[],[f1236,f594])).
fof(f3991,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(condensation,[],[f3990])).
fof(f3990,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3785])).
fof(f3785,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X13)))) ) | $spl0),
  inference(superposition,[],[f3762,f755])).
fof(f4105,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f4104,f117])).
fof(f4104,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f4088,f200])).
fof(f4088,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5) | s(bool,'F_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X4),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X5),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) ) | $spl0),
  inference(superposition,[],[f157,f594])).
fof(f4818,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f4762,f177])).
fof(f4762,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f4691,f612])).
fof(f4819,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6)))) ) | $spl0),
  inference(subsumption_resolution,[],[f4765,f177])).
fof(f4765,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X6)))) ) | $spl0),
  inference(superposition,[],[f4691,f755])).
fof(f4862,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'const.relation.WF_1'(s(fun(X54,fun(X54,bool)),X55))) | s(bool,X56) = s(bool,happ(s(fun(fun(X54,fun(X54,bool)),bool),'const.relation.WF_0'),s(fun(X54,fun(X54,bool)),X55))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X56)))) ) | $spl0),
  inference(superposition,[],[f1515,f143])).
fof(f4873,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X81) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),'const.semanticPrimitives.No_match_0'),s(bool,X78),s(bool,X79),s(fun(X77,bool),X80))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X81)))) ) | $spl0),
  inference(superposition,[],[f1515,f160])).
fof(f4874,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X84) | s(bool,X86) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X82),'const.semanticPrimitives.Match_type_error_0'),s(bool,X83),s(bool,X84),s(fun(X82,bool),X85))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X86)))) ) | $spl0),
  inference(superposition,[],[f1515,f161])).
fof(f4875,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(X87,bool),X91),s(X87,X88))) | s(bool,X92) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X87),'const.semanticPrimitives.Match_1'(s(X87,X88))),s(bool,X89),s(bool,X90),s(fun(X87,bool),X91))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X92)))) ) | $spl0),
  inference(superposition,[],[f1515,f163])).
fof(f4876,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,X93) | s(bool,X94) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X94)))) ) | $spl0),
  inference(superposition,[],[f1515,f172])).
fof(f4877,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s(bool,'T_0') != s(bool,happ(s(fun(X101,bool),happ(s(fun(X100,fun(X101,bool)),X104),s(X100,X102))),s(X101,X103))) | s(bool,X105) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X100,X101),'const.pair.,_2'(s(X100,X102),s(X101,X103))),s(fun(X100,fun(X101,bool)),X104))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X105)))) ) | $spl0),
  inference(superposition,[],[f1515,f165])).
fof(f4878,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,X107) = s(bool,'const.combin.I_1'(s(bool,X106))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X107)))) ) | $spl0),
  inference(superposition,[],[f1515,f141])).
fof(f4887,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f4836,f177])).
fof(f4836,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1515,f612])).
fof(f4907,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X10,fun(X11,bool)),bool),happ(s(fun('type.pair.prod'(X10,X11),fun(fun(X10,fun(X11,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X10,X11),X12))),s(fun(X10,fun(X11,bool)),X13))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X10,X11),X12),s(fun(X10,fun(X11,bool)),X13)))) )),
  inference(superposition,[],[f164,f128])).
fof(f4909,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X19),X20),s(fun(X18,fun(X19,bool)),X21))) | s(bool,X22) = s(bool,happ(s(fun(fun(X18,fun(X19,bool)),bool),happ(s(fun('type.pair.prod'(X18,X19),fun(fun(X18,fun(X19,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X18,X19),X20))),s(fun(X18,fun(X19,bool)),X21))) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f164,f612])).
fof(f4910,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,X27) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X24),X25),s(fun(X23,fun(X24,bool)),X26))) | s(bool,'F_0') = s(bool,X27) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X23,fun(X24,bool)),bool),happ(s(fun('type.pair.prod'(X23,X24),fun(fun(X23,fun(X24,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X23,X24),X25))),s(fun(X23,fun(X24,bool)),X26)))) ) | $spl0),
  inference(superposition,[],[f164,f612])).
fof(f4914,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X47) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X43,X44),X45),s(fun(X43,fun(X44,bool)),X46))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X43,fun(X44,bool)),bool),happ(s(fun('type.pair.prod'(X43,X44),fun(fun(X43,fun(X44,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X43,X44),X45))),s(fun(X43,fun(X44,bool)),X46))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X47)))) ) | $spl0),
  inference(superposition,[],[f164,f755])).
fof(f4915,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X48,X49),X50),s(fun(X48,fun(X49,bool)),X51))) | s(bool,X52) = s(bool,happ(s(fun(fun(X48,fun(X49,bool)),bool),happ(s(fun('type.pair.prod'(X48,X49),fun(fun(X48,fun(X49,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X48,X49),X50))),s(fun(X48,fun(X49,bool)),X51))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X52)))) ) | $spl0),
  inference(superposition,[],[f164,f755])).
fof(f4956,plain,(
  ( ! [X212,X210,X213,X211,X209] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X209,X210),X211),s(fun(X209,fun(X210,bool)),X212))) | s(bool,X213) = s(bool,happ(s(fun(fun(X209,fun(X210,bool)),bool),happ(s(fun('type.pair.prod'(X209,X210),fun(fun(X209,fun(X210,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X209,X210),X211))),s(fun(X209,fun(X210,bool)),X212))) | s(bool,'F_0') = s(bool,X213)) ) | $spl0),
  inference(superposition,[],[f900,f164])).
fof(f4958,plain,(
  ( ! [X222,X220,X219,X223,X221] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X219,X220),X221),s(fun(X219,fun(X220,bool)),X222))) | s(bool,X223) = s(bool,happ(s(fun(fun(X219,fun(X220,bool)),bool),happ(s(fun('type.pair.prod'(X219,X220),fun(fun(X219,fun(X220,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X219,X220),X221))),s(fun(X219,fun(X220,bool)),X222))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X223)))) ) | $spl0),
  inference(superposition,[],[f1515,f164])).
fof(f4989,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X14,X15),X16),s(fun(X14,fun(X15,bool)),X17))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,fun(X15,bool)),bool),happ(s(fun('type.pair.prod'(X14,X15),fun(fun(X14,fun(X15,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X14,X15),X16))),s(fun(X14,fun(X15,bool)),X17)))) ) | $spl0),
  inference(forward_demodulation,[],[f4908,f200])).
fof(f4908,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X14,X15),X16),s(fun(X14,fun(X15,bool)),X17))) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,fun(X15,bool)),bool),happ(s(fun('type.pair.prod'(X14,X15),fun(fun(X14,fun(X15,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X14,X15),X16))),s(fun(X14,fun(X15,bool)),X17)))) ) | $spl0),
  inference(superposition,[],[f164,f594])).
fof(f5006,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X156,X157),X158),s(fun(X156,fun(X157,bool)),X159))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X156,fun(X157,bool)),bool),happ(s(fun('type.pair.prod'(X156,X157),fun(fun(X156,fun(X157,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X156,X157),X158))),s(fun(X156,fun(X157,bool)),X159)))) )),
  inference(forward_demodulation,[],[f4945,f141])).
fof(f4945,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X156,X157),X158),s(fun(X156,fun(X157,bool)),X159))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X156,fun(X157,bool)),bool),happ(s(fun('type.pair.prod'(X156,X157),fun(fun(X156,fun(X157,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X156,X157),X158))),s(fun(X156,fun(X157,bool)),X159)))) )),
  inference(superposition,[],[f387,f164])).
fof(f5059,plain,(
  ( ! [X52,X53,X51] : (s(bool,'const.combin.I_1'(s(bool,X53))) != s(bool,'const.relation.WF_1'(s(fun(X51,fun(X51,bool)),X52))) | s(bool,X53) = s(bool,happ(s(fun(fun(X51,fun(X51,bool)),bool),'const.relation.WF_0'),s(fun(X51,fun(X51,bool)),X52))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X51,fun(X51,bool)),bool),'const.relation.WF_0'),s(fun(X51,fun(X51,bool)),X52)))) ) | $spl0),
  inference(superposition,[],[f1617,f143])).
fof(f5070,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X75) != s(bool,'const.combin.I_1'(s(bool,X78))) | s(bool,X78) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X74),'const.semanticPrimitives.No_match_0'),s(bool,X75),s(bool,X76),s(fun(X74,bool),X77))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X74),'const.semanticPrimitives.No_match_0'),s(bool,X75),s(bool,X76),s(fun(X74,bool),X77)))) ) | $spl0),
  inference(superposition,[],[f1617,f160])).
fof(f5071,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X81) != s(bool,'const.combin.I_1'(s(bool,X83))) | s(bool,X83) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X79),'const.semanticPrimitives.Match_type_error_0'),s(bool,X80),s(bool,X81),s(fun(X79,bool),X82))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X79),'const.semanticPrimitives.Match_type_error_0'),s(bool,X80),s(bool,X81),s(fun(X79,bool),X82)))) ) | $spl0),
  inference(superposition,[],[f1617,f161])).
fof(f5072,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,'const.combin.I_1'(s(bool,X89))) != s(bool,happ(s(fun(X84,bool),X88),s(X84,X85))) | s(bool,X89) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X84),'const.semanticPrimitives.Match_1'(s(X84,X85))),s(bool,X86),s(bool,X87),s(fun(X84,bool),X88))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X84),'const.semanticPrimitives.Match_1'(s(X84,X85))),s(bool,X86),s(bool,X87),s(fun(X84,bool),X88)))) ) | $spl0),
  inference(superposition,[],[f1617,f163])).
fof(f5073,plain,(
  ( ! [X90,X91] : (s(bool,X90) != s(bool,'const.combin.I_1'(s(bool,X91))) | s(bool,X91) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X90)))) ) | $spl0),
  inference(superposition,[],[f1617,f172])).
fof(f5074,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s(bool,'const.combin.I_1'(s(bool,X101))) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X97,X98),X99),s(fun(X97,fun(X98,bool)),X100))) | s(bool,X101) = s(bool,happ(s(fun(fun(X97,fun(X98,bool)),bool),happ(s(fun('type.pair.prod'(X97,X98),fun(fun(X97,fun(X98,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X97,X98),X99))),s(fun(X97,fun(X98,bool)),X100))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X97,fun(X98,bool)),bool),happ(s(fun('type.pair.prod'(X97,X98),fun(fun(X97,fun(X98,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X97,X98),X99))),s(fun(X97,fun(X98,bool)),X100)))) ) | $spl0),
  inference(superposition,[],[f1617,f164])).
fof(f5075,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(bool,'const.combin.I_1'(s(bool,X107))) != s(bool,happ(s(fun(X103,bool),happ(s(fun(X102,fun(X103,bool)),X106),s(X102,X104))),s(X103,X105))) | s(bool,X107) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X102,X103),'const.pair.,_2'(s(X102,X104),s(X103,X105))),s(fun(X102,fun(X103,bool)),X106))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X102,X103),'const.pair.,_2'(s(X102,X104),s(X103,X105))),s(fun(X102,fun(X103,bool)),X106)))) ) | $spl0),
  inference(superposition,[],[f1617,f165])).
fof(f5076,plain,(
  ( ! [X109,X108] : (s(bool,X108) != s(bool,'const.combin.I_1'(s(bool,X109))) | s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,X108))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X108)))) ) | $spl0),
  inference(superposition,[],[f1617,f141])).
fof(f5115,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(bool,X110) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X106,bool),happ(s(fun(X105,fun(X106,bool)),X109),s(X105,X107))),s(X106,X108))))) | s(bool,X110) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X105,X106),'const.pair.,_2'(s(X105,X107),s(X106,X108))),s(fun(X105,fun(X106,bool)),X109))) | s(bool,'F_0') = s(bool,X110)) ) | $spl0),
  inference(superposition,[],[f1617,f165])).
fof(f5170,plain,(
  ( ! [X54,X56,X55] : (s(bool,X56) != s(bool,'const.relation.WF_1'(s(fun(X54,fun(X54,bool)),X55))) | s(bool,X56) = s(bool,happ(s(fun(fun(X54,fun(X54,bool)),bool),'const.relation.WF_0'),s(fun(X54,fun(X54,bool)),X55))) | s(bool,'F_0') = s(bool,X56)) ) | $spl0),
  inference(forward_demodulation,[],[f5099,f141])).
fof(f5099,plain,(
  ( ! [X54,X56,X55] : (s(bool,X56) != s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X54,fun(X54,bool)),X55))))) | s(bool,X56) = s(bool,happ(s(fun(fun(X54,fun(X54,bool)),bool),'const.relation.WF_0'),s(fun(X54,fun(X54,bool)),X55))) | s(bool,'F_0') = s(bool,X56)) ) | $spl0),
  inference(superposition,[],[f1617,f143])).
fof(f5175,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) != s(bool,X81) | s(bool,X81) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),'const.semanticPrimitives.No_match_0'),s(bool,X78),s(bool,X79),s(fun(X77,bool),X80))) | s(bool,'F_0') = s(bool,X81)) ) | $spl0),
  inference(forward_demodulation,[],[f5110,f141])).
fof(f5110,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X81) != s(bool,'const.combin.I_1'(s(bool,X78))) | s(bool,X81) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X77),'const.semanticPrimitives.No_match_0'),s(bool,X78),s(bool,X79),s(fun(X77,bool),X80))) | s(bool,'F_0') = s(bool,X81)) ) | $spl0),
  inference(superposition,[],[f1617,f160])).
fof(f5176,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,X84) != s(bool,X86) | s(bool,X86) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X82),'const.semanticPrimitives.Match_type_error_0'),s(bool,X83),s(bool,X84),s(fun(X82,bool),X85))) | s(bool,'F_0') = s(bool,X86)) ) | $spl0),
  inference(forward_demodulation,[],[f5111,f141])).
fof(f5111,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,X86) != s(bool,'const.combin.I_1'(s(bool,X84))) | s(bool,X86) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X82),'const.semanticPrimitives.Match_type_error_0'),s(bool,X83),s(bool,X84),s(fun(X82,bool),X85))) | s(bool,'F_0') = s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f1617,f161])).
fof(f5177,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,X92) != s(bool,happ(s(fun(X87,bool),X91),s(X87,X88))) | s(bool,X92) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X87),'const.semanticPrimitives.Match_1'(s(X87,X88))),s(bool,X89),s(bool,X90),s(fun(X87,bool),X91))) | s(bool,'F_0') = s(bool,X92)) ) | $spl0),
  inference(forward_demodulation,[],[f5112,f141])).
fof(f5112,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,X92) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X87,bool),X91),s(X87,X88))))) | s(bool,X92) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X87),'const.semanticPrimitives.Match_1'(s(X87,X88))),s(bool,X89),s(bool,X90),s(fun(X87,bool),X91))) | s(bool,'F_0') = s(bool,X92)) ) | $spl0),
  inference(superposition,[],[f1617,f163])).
fof(f5178,plain,(
  ( ! [X94,X93] : (s(bool,X93) != s(bool,X94) | s(bool,X94) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'F_0') = s(bool,X94)) ) | $spl0),
  inference(forward_demodulation,[],[f5113,f141])).
fof(f5113,plain,(
  ( ! [X94,X93] : (s(bool,X94) != s(bool,'const.combin.I_1'(s(bool,X93))) | s(bool,X94) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'F_0') = s(bool,X94)) ) | $spl0),
  inference(superposition,[],[f1617,f172])).
fof(f5179,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X104) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X100,X101),X102),s(fun(X100,fun(X101,bool)),X103))) | s(bool,X104) = s(bool,happ(s(fun(fun(X100,fun(X101,bool)),bool),happ(s(fun('type.pair.prod'(X100,X101),fun(fun(X100,fun(X101,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X100,X101),X102))),s(fun(X100,fun(X101,bool)),X103))) | s(bool,'F_0') = s(bool,X104)) ) | $spl0),
  inference(forward_demodulation,[],[f5114,f141])).
fof(f5114,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X104) != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X100,X101),X102),s(fun(X100,fun(X101,bool)),X103))))) | s(bool,X104) = s(bool,happ(s(fun(fun(X100,fun(X101,bool)),bool),happ(s(fun('type.pair.prod'(X100,X101),fun(fun(X100,fun(X101,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X100,X101),X102))),s(fun(X100,fun(X101,bool)),X103))) | s(bool,'F_0') = s(bool,X104)) ) | $spl0),
  inference(superposition,[],[f1617,f164])).
fof(f5820,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(X19,happ(s(fun(X17,X19),X20),s(X17,X21))) = s(X19,happ(s(fun(X17,X19),happ(s(fun(fun(X17,X19),fun(X17,X19)),X22),s(fun(X17,X19),'const.relation.RESTRICT_3'(s(fun(X17,X19),X20),s(fun(X17,fun(X17,bool)),X18),s(X17,X21))))),s(X17,X21))) | s(fun(X17,X19),X20) != s(fun(X17,X19),'const.relation.WFREC_2'(s(fun(X17,fun(X17,bool)),X18),s(fun(fun(X17,X19),fun(X17,X19)),X22))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),X18)))) ) | $spl0),
  inference(subsumption_resolution,[],[f5819,f117])).
fof(f5819,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'T_0')) | s(X19,happ(s(fun(X17,X19),X20),s(X17,X21))) = s(X19,happ(s(fun(X17,X19),happ(s(fun(fun(X17,X19),fun(X17,X19)),X22),s(fun(X17,X19),'const.relation.RESTRICT_3'(s(fun(X17,X19),X20),s(fun(X17,fun(X17,bool)),X18),s(X17,X21))))),s(X17,X21))) | s(fun(X17,X19),X20) != s(fun(X17,X19),'const.relation.WFREC_2'(s(fun(X17,fun(X17,bool)),X18),s(fun(fun(X17,X19),fun(X17,X19)),X22))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),X18)))) ) | $spl0),
  inference(forward_demodulation,[],[f5791,f200])).
fof(f5791,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | s(X19,happ(s(fun(X17,X19),X20),s(X17,X21))) = s(X19,happ(s(fun(X17,X19),happ(s(fun(fun(X17,X19),fun(X17,X19)),X22),s(fun(X17,X19),'const.relation.RESTRICT_3'(s(fun(X17,X19),X20),s(fun(X17,fun(X17,bool)),X18),s(X17,X21))))),s(X17,X21))) | s(fun(X17,X19),X20) != s(fun(X17,X19),'const.relation.WFREC_2'(s(fun(X17,fun(X17,bool)),X18),s(fun(fun(X17,X19),fun(X17,X19)),X22))) | s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),X18)))) ) | $spl0),
  inference(superposition,[],[f162,f594])).
fof(f6055,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X115),s(X111,X113))),s(X112,X114))))) | s(bool,X116) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X113),s(X112,X114))),s(fun(X111,fun(X112,bool)),X115))) | s(bool,'F_0') = s(bool,X116)) ) | $spl0),
  inference(superposition,[],[f1875,f165])).
fof(f6250,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,X81) != s(bool,'const.combin.I_1'(s(bool,X84))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X80),'const.semanticPrimitives.No_match_0'),s(bool,X81),s(bool,X82),s(fun(X80,bool),X83))) | s(bool,'T_0') = s(bool,X84)) ) | $spl0),
  inference(superposition,[],[f1916,f160])).
fof(f6251,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X87) != s(bool,'const.combin.I_1'(s(bool,X89))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X85),'const.semanticPrimitives.Match_type_error_0'),s(bool,X86),s(bool,X87),s(fun(X85,bool),X88))) | s(bool,'T_0') = s(bool,X89)) ) | $spl0),
  inference(superposition,[],[f1916,f161])).
fof(f6252,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'const.combin.I_1'(s(bool,X95))) != s(bool,happ(s(fun(X90,bool),X94),s(X90,X91))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X90),'const.semanticPrimitives.Match_1'(s(X90,X91))),s(bool,X92),s(bool,X93),s(fun(X90,bool),X94))) | s(bool,'T_0') = s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f1916,f163])).
fof(f6253,plain,(
  ( ! [X97,X96] : (s(bool,X96) != s(bool,'const.combin.I_1'(s(bool,X97))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X96))) | s(bool,'T_0') = s(bool,X97)) ) | $spl0),
  inference(superposition,[],[f1916,f172])).
fof(f6254,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'const.combin.I_1'(s(bool,X107))) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X103,X104),X105),s(fun(X103,fun(X104,bool)),X106))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X103,fun(X104,bool)),bool),happ(s(fun('type.pair.prod'(X103,X104),fun(fun(X103,fun(X104,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X103,X104),X105))),s(fun(X103,fun(X104,bool)),X106))) | s(bool,'T_0') = s(bool,X107)) ) | $spl0),
  inference(superposition,[],[f1916,f164])).
fof(f6255,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'const.combin.I_1'(s(bool,X113))) != s(bool,happ(s(fun(X109,bool),happ(s(fun(X108,fun(X109,bool)),X112),s(X108,X110))),s(X109,X111))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X108,X109),'const.pair.,_2'(s(X108,X110),s(X109,X111))),s(fun(X108,fun(X109,bool)),X112))) | s(bool,'T_0') = s(bool,X113)) ) | $spl0),
  inference(superposition,[],[f1916,f165])).
fof(f6297,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,X116) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X115),s(X111,X113))),s(X112,X114))))) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X113),s(X112,X114))),s(fun(X111,fun(X112,bool)),X115)))) ) | $spl0),
  inference(superposition,[],[f1916,f165])).
fof(f6374,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X84) != s(bool,X87) | s(bool,'F_0') = s(bool,X87) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X83),'const.semanticPrimitives.No_match_0'),s(bool,X84),s(bool,X85),s(fun(X83,bool),X86)))) ) | $spl0),
  inference(forward_demodulation,[],[f6292,f141])).
fof(f6292,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X87) != s(bool,'const.combin.I_1'(s(bool,X84))) | s(bool,'F_0') = s(bool,X87) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X83),'const.semanticPrimitives.No_match_0'),s(bool,X84),s(bool,X85),s(fun(X83,bool),X86)))) ) | $spl0),
  inference(superposition,[],[f1916,f160])).
fof(f6375,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,X90) != s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X88),'const.semanticPrimitives.Match_type_error_0'),s(bool,X89),s(bool,X90),s(fun(X88,bool),X91)))) ) | $spl0),
  inference(forward_demodulation,[],[f6293,f141])).
fof(f6293,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,X92) != s(bool,'const.combin.I_1'(s(bool,X90))) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X88),'const.semanticPrimitives.Match_type_error_0'),s(bool,X89),s(bool,X90),s(fun(X88,bool),X91)))) ) | $spl0),
  inference(superposition,[],[f1916,f161])).
fof(f6376,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,X98) != s(bool,happ(s(fun(X93,bool),X97),s(X93,X94))) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X93),'const.semanticPrimitives.Match_1'(s(X93,X94))),s(bool,X95),s(bool,X96),s(fun(X93,bool),X97)))) ) | $spl0),
  inference(forward_demodulation,[],[f6294,f141])).
fof(f6294,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,X98) != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X93,bool),X97),s(X93,X94))))) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X93),'const.semanticPrimitives.Match_1'(s(X93,X94))),s(bool,X95),s(bool,X96),s(fun(X93,bool),X97)))) ) | $spl0),
  inference(superposition,[],[f1916,f163])).
fof(f6378,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s(bool,X110) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X106,X107),X108),s(fun(X106,fun(X107,bool)),X109))) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X106,fun(X107,bool)),bool),happ(s(fun('type.pair.prod'(X106,X107),fun(fun(X106,fun(X107,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X106,X107),X108))),s(fun(X106,fun(X107,bool)),X109)))) ) | $spl0),
  inference(forward_demodulation,[],[f6296,f141])).
fof(f6296,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s(bool,X110) != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X106,X107),X108),s(fun(X106,fun(X107,bool)),X109))))) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X106,fun(X107,bool)),bool),happ(s(fun('type.pair.prod'(X106,X107),fun(fun(X106,fun(X107,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X106,X107),X108))),s(fun(X106,fun(X107,bool)),X109)))) ) | $spl0),
  inference(superposition,[],[f1916,f164])).
fof(f6467,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,X116) != s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X115),s(X111,X113))),s(X112,X114))) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X113),s(X112,X114))),s(fun(X111,fun(X112,bool)),X115)))) ) | $spl0),
  inference(superposition,[],[f6320,f165])).
fof(f6504,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,X81) != s(bool,X84) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X80),'const.semanticPrimitives.No_match_0'),s(bool,X81),s(bool,X82),s(fun(X80,bool),X83))) | s(bool,'T_0') = s(bool,X84)) ) | $spl0),
  inference(superposition,[],[f6320,f160])).
fof(f6505,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X87) != s(bool,X89) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X85),'const.semanticPrimitives.Match_type_error_0'),s(bool,X86),s(bool,X87),s(fun(X85,bool),X88))) | s(bool,'T_0') = s(bool,X89)) ) | $spl0),
  inference(superposition,[],[f6320,f161])).
fof(f6506,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X95) != s(bool,happ(s(fun(X90,bool),X94),s(X90,X91))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X90),'const.semanticPrimitives.Match_1'(s(X90,X91))),s(bool,X92),s(bool,X93),s(fun(X90,bool),X94))) | s(bool,'T_0') = s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f6320,f163])).
fof(f6507,plain,(
  ( ! [X97,X96] : (s(bool,X96) != s(bool,X97) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X96))) | s(bool,'T_0') = s(bool,X97)) ) | $spl0),
  inference(superposition,[],[f6320,f172])).
fof(f6508,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,X107) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X103,X104),X105),s(fun(X103,fun(X104,bool)),X106))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X103,fun(X104,bool)),bool),happ(s(fun('type.pair.prod'(X103,X104),fun(fun(X103,fun(X104,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X103,X104),X105))),s(fun(X103,fun(X104,bool)),X106))) | s(bool,'T_0') = s(bool,X107)) ) | $spl0),
  inference(superposition,[],[f6320,f164])).
fof(f6509,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,X113) != s(bool,happ(s(fun(X109,bool),happ(s(fun(X108,fun(X109,bool)),X112),s(X108,X110))),s(X109,X111))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X108,X109),'const.pair.,_2'(s(X108,X110),s(X109,X111))),s(fun(X108,fun(X109,bool)),X112))) | s(bool,'T_0') = s(bool,X113)) ) | $spl0),
  inference(superposition,[],[f6320,f165])).
fof(f6612,plain,(
  ( ! [X94,X95,X96] : (s(bool,'const.combin.I_1'(s(bool,X96))) != s(bool,'const.relation.WF_1'(s(fun(X94,fun(X94,bool)),X95))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X94,fun(X94,bool)),bool),'const.relation.WF_0'),s(fun(X94,fun(X94,bool)),X95))) | s(bool,'T_0') = s(bool,X96)) ) | $spl0),
  inference(superposition,[],[f1916,f143])).
fof(f6625,plain,(
  ( ! [X123,X124,X122] : (s(bool,X124) != s(bool,'const.relation.WF_1'(s(fun(X122,fun(X122,bool)),X123))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X122,fun(X122,bool)),bool),'const.relation.WF_0'),s(fun(X122,fun(X122,bool)),X123))) | s(bool,'T_0') = s(bool,X124)) ) | $spl0),
  inference(superposition,[],[f6320,f143])).
fof(f6678,plain,(
  ( ! [X92,X93,X91] : (s(bool,X93) != s(bool,'const.relation.WF_1'(s(fun(X91,fun(X91,bool)),X92))) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X91,fun(X91,bool)),bool),'const.relation.WF_0'),s(fun(X91,fun(X91,bool)),X92)))) ) | $spl0),
  inference(forward_demodulation,[],[f6611,f141])).
fof(f6611,plain,(
  ( ! [X92,X93,X91] : (s(bool,X93) != s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X91,fun(X91,bool)),X92))))) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X91,fun(X91,bool)),bool),'const.relation.WF_0'),s(fun(X91,fun(X91,bool)),X92)))) ) | $spl0),
  inference(superposition,[],[f1916,f143])).
fof(f7771,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X0)),bool))) )),
  inference(superposition,[],[f2200,f128])).
fof(f7773,plain,(
  ( ! [X2,X3] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X2))) | s(bool,X3) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X2)),bool)) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f2200,f612])).
fof(f7774,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X4))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X4)),bool))) ) | $spl0),
  inference(superposition,[],[f2200,f612])).
fof(f7777,plain,(
  ( ! [X10,X11] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X10))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X10)),bool))))) ) | $spl0),
  inference(superposition,[],[f2200,f755])).
fof(f7779,plain,(
  ( ! [X14,X15] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X14))) | s(bool,X15) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X14)),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl0),
  inference(superposition,[],[f2200,f755])).
fof(f7810,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1)))) ) | $spl0),
  inference(forward_demodulation,[],[f7772,f200])).
fof(f7772,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X1)),bool))) ) | $spl0),
  inference(superposition,[],[f2200,f594])).
fof(f7812,plain,(
  ( ! [X16] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X16))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X16)),bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f7780,f183])).
fof(f7780,plain,(
  ( ! [X17,X16] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X16))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X16)),bool))))) ) | $spl0),
  inference(superposition,[],[f2200,f1796])).
fof(f8211,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(condensation,[],[f8210])).
fof(f8210,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7865])).
fof(f7865,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2195,f128])).
fof(f8227,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,'const.combin.I_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X39)))) ) | $spl0),
  inference(subsumption_resolution,[],[f7879,f177])).
fof(f7879,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,'const.combin.I_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X37) = s(bool,X39) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X39)))) ) | $spl0),
  inference(superposition,[],[f2195,f755])).
fof(f8450,plain,(
  ( ! [X86] : (s(bool,X86) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X86)))) ) | $spl0),
  inference(superposition,[],[f8211,f141])).
fof(f9367,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,X8) | s(bool,X9) = s(bool,'const.combin.I_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,X9) | s(bool,X7) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | $spl0),
  inference(superposition,[],[f1236,f205])).
fof(f9368,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X13))) | s(bool,X11) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f1515,f205])).
fof(f9372,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) = s(bool,'const.combin.I_1'(s(bool,X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f3762,f205])).
fof(f9373,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl0),
  inference(superposition,[],[f3991,f205])).
fof(f10433,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f8450,f755])).
fof(f10434,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f8450,f612])).
fof(f11021,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f10535,f177])).
fof(f10535,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f8450,f755])).
fof(f11028,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X24) = s(bool,'const.combin.I_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X24)) ) | $spl0),
  inference(subsumption_resolution,[],[f10549,f177])).
fof(f10549,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X23) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X24) = s(bool,'const.combin.I_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X24)) ) | $spl0),
  inference(superposition,[],[f8450,f612])).
fof(f11029,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X36) = s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X36)))) ) | $spl0),
  inference(subsumption_resolution,[],[f10555,f177])).
fof(f10555,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X35) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X36) = s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X36)))) ) | $spl0),
  inference(superposition,[],[f8450,f755])).
fof(f11142,plain,(
  ( ! [X198,X196,X197] : (s(bool,'F_0') != s(bool,f(s(X196,sK3(X197,X198,bool,X196)))) | s(fun(X196,bool),X197) = s(fun(X196,bool),X198) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X196,sK3(X197,X198,bool,X196))))))) ) | $spl0),
  inference(forward_demodulation,[],[f10676,f141])).
fof(f10676,plain,(
  ( ! [X198,X196,X197] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) != s(bool,f(s(X196,sK3(X197,X198,bool,X196)))) | s(fun(X196,bool),X197) = s(fun(X196,bool),X198) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(X196,sK3(X197,X198,bool,X196))))))) ) | $spl0),
  inference(superposition,[],[f159,f8450])).
fof(f11157,plain,(
  ( ! [X228] : (s('type.semanticPrimitives.match_result'(bool),X228) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X228) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X228) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK2(X228,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f10684,f141])).
fof(f10684,plain,(
  ( ! [X228] : (s('type.semanticPrimitives.match_result'(bool),X228) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X228) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X228) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK2(X228,bool))))) ) | $spl0),
  inference(superposition,[],[f144,f8450])).
fof(f11158,plain,(
  ( ! [X229] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X229))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X229)),bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f10685,f141])).
fof(f10685,plain,(
  ( ! [X229] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X229))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X229)),bool))))) ) | $spl0),
  inference(superposition,[],[f2200,f8450])).
fof(f11159,plain,(
  ( ! [X231,X230,X232] : (s(X232,f(s(bool,'F_0'))) != s(X232,g(s(bool,'F_0'))) | s(fun(bool,X232),X230) = s(fun(bool,X232),X231) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X230,X231,X232,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f10686,f141])).
fof(f10686,plain,(
  ( ! [X231,X230,X232] : (s(X232,f(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) != s(X232,g(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(fun(bool,X232),X230) = s(fun(bool,X232),X231) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X230,X231,X232,bool))))) ) | $spl0),
  inference(superposition,[],[f159,f8450])).
fof(f11207,plain,(
  ( ! [X30] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f10799,f117])).
fof(f10799,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f134,f8450])).
fof(f11330,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,X48) | s(bool,X48) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f11207,f141])).
fof(f11348,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X23)))) ) | $spl0),
  inference(subsumption_resolution,[],[f11252,f141])).
fof(f11252,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X23)))) ) | $spl0),
  inference(superposition,[],[f11207,f755])).
fof(f11475,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun(X95,bool),happ(s(fun(X94,fun(X95,bool)),X98),s(X94,X96))),s(X95,X97))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X94,X95),'const.pair.,_2'(s(X94,X96),s(X95,X97))),s(fun(X94,fun(X95,bool)),X98)))) )),
  inference(superposition,[],[f278,f165])).
fof(f11638,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f11500,f117])).
fof(f11500,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f134,f278])).
fof(f11681,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X2)))) ) | $spl0),
  inference(superposition,[],[f11330,f755])).
fof(f11682,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f11330,f612])).
fof(f11684,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f11330,f128])).
fof(f11859,plain,(
  ( ! [X80,X81,X82] : (s(bool,X82) != s(bool,'const.relation.WF_1'(s(fun(X80,fun(X80,bool)),X81))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X82))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X80,fun(X80,bool)),bool),'const.relation.WF_0'),s(fun(X80,fun(X80,bool)),X81)))) ) | $spl0),
  inference(superposition,[],[f2234,f143])).
fof(f11879,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X125) != s(bool,X128) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X124),'const.semanticPrimitives.No_match_0'),s(bool,X125),s(bool,X126),s(fun(X124,bool),X127)))) ) | $spl0),
  inference(superposition,[],[f2234,f160])).
fof(f11880,plain,(
  ( ! [X132,X130,X133,X131,X129] : (s(bool,X131) != s(bool,X133) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X133))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X129),'const.semanticPrimitives.Match_type_error_0'),s(bool,X130),s(bool,X131),s(fun(X129,bool),X132)))) ) | $spl0),
  inference(superposition,[],[f2234,f161])).
fof(f11881,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,X139) != s(bool,happ(s(fun(X134,bool),X138),s(X134,X135))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X139))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X134),'const.semanticPrimitives.Match_1'(s(X134,X135))),s(bool,X136),s(bool,X137),s(fun(X134,bool),X138)))) ) | $spl0),
  inference(superposition,[],[f2234,f163])).
fof(f11883,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,X151) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X147,X148),X149),s(fun(X147,fun(X148,bool)),X150))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X151))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X147,fun(X148,bool)),bool),happ(s(fun('type.pair.prod'(X147,X148),fun(fun(X147,fun(X148,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X147,X148),X149))),s(fun(X147,fun(X148,bool)),X150)))) ) | $spl0),
  inference(superposition,[],[f2234,f164])).
fof(f11884,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (s(bool,X157) != s(bool,happ(s(fun(X153,bool),happ(s(fun(X152,fun(X153,bool)),X156),s(X152,X154))),s(X153,X155))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X157))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X152,X153),'const.pair.,_2'(s(X152,X154),s(X153,X155))),s(fun(X152,fun(X153,bool)),X156)))) ) | $spl0),
  inference(superposition,[],[f2234,f165])).
fof(f12070,plain,(
  ( ! [X154,X152,X151,X153,X150] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X152),'const.semanticPrimitives.No_match_0'),s(bool,'const.relation.WF_1'(s(fun(X150,fun(X150,bool)),X151))),s(bool,X153),s(fun(X152,bool),X154))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X150,fun(X150,bool)),bool),'const.relation.WF_0'),s(fun(X150,fun(X150,bool)),X151)))) )),
  inference(superposition,[],[f316,f143])).
fof(f12090,plain,(
  ( ! [X239,X237,X235,X240,X238,X236,X234] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X238),'const.semanticPrimitives.No_match_0'),s(bool,X235),s(bool,X239),s(fun(X238,bool),X240))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X234),'const.semanticPrimitives.No_match_0'),s(bool,X235),s(bool,X236),s(fun(X234,bool),X237)))) )),
  inference(superposition,[],[f316,f160])).
fof(f12091,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X245),'const.semanticPrimitives.No_match_0'),s(bool,X243),s(bool,X246),s(fun(X245,bool),X247))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X241),'const.semanticPrimitives.Match_type_error_0'),s(bool,X242),s(bool,X243),s(fun(X241,bool),X244)))) )),
  inference(superposition,[],[f316,f161])).
fof(f12092,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X253),'const.semanticPrimitives.No_match_0'),s(bool,happ(s(fun(X248,bool),X252),s(X248,X249))),s(bool,X254),s(fun(X253,bool),X255))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X248),'const.semanticPrimitives.Match_1'(s(X248,X249))),s(bool,X250),s(bool,X251),s(fun(X248,bool),X252)))) )),
  inference(superposition,[],[f316,f163])).
fof(f12094,plain,(
  ( ! [X269,X271,X265,X267,X268,X270,X266] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X269),'const.semanticPrimitives.No_match_0'),s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X265,X266),X267),s(fun(X265,fun(X266,bool)),X268))),s(bool,X270),s(fun(X269,bool),X271))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X265,fun(X266,bool)),bool),happ(s(fun('type.pair.prod'(X265,X266),fun(fun(X265,fun(X266,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X265,X266),X267))),s(fun(X265,fun(X266,bool)),X268)))) )),
  inference(superposition,[],[f316,f164])).
fof(f12095,plain,(
  ( ! [X277,X279,X273,X275,X276,X278,X272,X274] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X277),'const.semanticPrimitives.No_match_0'),s(bool,happ(s(fun(X273,bool),happ(s(fun(X272,fun(X273,bool)),X276),s(X272,X274))),s(X273,X275))),s(bool,X278),s(fun(X277,bool),X279))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X272,X273),'const.pair.,_2'(s(X272,X274),s(X273,X275))),s(fun(X272,fun(X273,bool)),X276)))) )),
  inference(superposition,[],[f316,f165])).
fof(f12117,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X94),'const.semanticPrimitives.No_match_0'),s(bool,X95),s(bool,X93),s(fun(X94,bool),X96))) | s(bool,'F_0') = s(bool,X95) | s(bool,X93) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f316,f205])).
fof(f12347,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.No_match_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f12187,f117])).
fof(f12187,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.No_match_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f134,f316])).
fof(f12656,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,X30) = s(bool,'const.relation.WF_1'(s(fun(X31,fun(X31,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12563,f177])).
fof(f12563,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,X30) = s(bool,'const.relation.WF_1'(s(fun(X31,fun(X31,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f12476,f755])).
fof(f12707,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X94),'const.semanticPrimitives.Match_type_error_0'),s(bool,X93),s(bool,X95),s(fun(X94,bool),X96))) | s(bool,'F_0') = s(bool,X95) | s(bool,X93) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f414,f205])).
fof(f12782,plain,(
  ( ! [X158,X161,X159,X157,X160] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X159),'const.semanticPrimitives.Match_type_error_0'),s(bool,X160),s(bool,'const.relation.WF_1'(s(fun(X157,fun(X157,bool)),X158))),s(fun(X159,bool),X161))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X157,fun(X157,bool)),bool),'const.relation.WF_0'),s(fun(X157,fun(X157,bool)),X158)))) )),
  inference(superposition,[],[f414,f143])).
fof(f12802,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X245),'const.semanticPrimitives.Match_type_error_0'),s(bool,X246),s(bool,X242),s(fun(X245,bool),X247))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X241),'const.semanticPrimitives.No_match_0'),s(bool,X242),s(bool,X243),s(fun(X241,bool),X244)))) )),
  inference(superposition,[],[f414,f160])).
fof(f12803,plain,(
  ( ! [X253,X251,X249,X254,X252,X250,X248] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X252),'const.semanticPrimitives.Match_type_error_0'),s(bool,X253),s(bool,X250),s(fun(X252,bool),X254))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X248),'const.semanticPrimitives.Match_type_error_0'),s(bool,X249),s(bool,X250),s(fun(X248,bool),X251)))) )),
  inference(superposition,[],[f414,f161])).
fof(f12804,plain,(
  ( ! [X261,X255,X257,X259,X260,X262,X256,X258] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X260),'const.semanticPrimitives.Match_type_error_0'),s(bool,X261),s(bool,happ(s(fun(X255,bool),X259),s(X255,X256))),s(fun(X260,bool),X262))) | s(bool,'F_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X255),'const.semanticPrimitives.Match_1'(s(X255,X256))),s(bool,X257),s(bool,X258),s(fun(X255,bool),X259)))) )),
  inference(superposition,[],[f414,f163])).
fof(f12806,plain,(
  ( ! [X277,X273,X275,X276,X278,X272,X274] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X276),'const.semanticPrimitives.Match_type_error_0'),s(bool,X277),s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X272,X273),X274),s(fun(X272,fun(X273,bool)),X275))),s(fun(X276,bool),X278))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X272,fun(X273,bool)),bool),happ(s(fun('type.pair.prod'(X272,X273),fun(fun(X272,fun(X273,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X272,X273),X274))),s(fun(X272,fun(X273,bool)),X275)))) )),
  inference(superposition,[],[f414,f164])).
fof(f12807,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X284,X286] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X284),'const.semanticPrimitives.Match_type_error_0'),s(bool,X285),s(bool,happ(s(fun(X280,bool),happ(s(fun(X279,fun(X280,bool)),X283),s(X279,X281))),s(X280,X282))),s(fun(X284,bool),X286))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X279,X280),'const.pair.,_2'(s(X279,X281),s(X280,X282))),s(fun(X279,fun(X280,bool)),X283)))) )),
  inference(superposition,[],[f414,f165])).
fof(f13006,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.Match_type_error_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f12839,f117])).
fof(f12839,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.Match_type_error_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f134,f414])).
fof(f15596,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13) | s(bool,X14) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X11) = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f900,f10324])).
fof(f15597,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,X17) = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X15) = s(bool,X18)) ) | $spl0),
  inference(superposition,[],[f1236,f10324])).
fof(f15598,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X19) != s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X21))) | s(bool,X22) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X19) = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f1515,f10324])).
fof(f15602,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,X39))) | s(bool,'F_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X39) = s(bool,X41)) ) | $spl0),
  inference(superposition,[],[f3762,f10324])).
fof(f15603,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X43))) | s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X42) = s(bool,X44)) ) | $spl0),
  inference(superposition,[],[f3991,f10324])).
fof(f16228,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) = s(bool,X72) | s(bool,X72) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X71) = s(bool,X73) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X73)))) ) | $spl0),
  inference(subsumption_resolution,[],[f14872,f177])).
fof(f14872,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) = s(bool,X72) | s(bool,X72) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X71) = s(bool,X73) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X73)))) ) | $spl0),
  inference(superposition,[],[f10324,f755])).
fof(f16229,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,X75) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X74) = s(bool,X76) | s(bool,'F_0') = s(bool,X76)) ) | $spl0),
  inference(subsumption_resolution,[],[f14873,f177])).
fof(f14873,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,X75) | s(bool,X75) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X74) = s(bool,X76) | s(bool,'F_0') = s(bool,X76)) ) | $spl0),
  inference(superposition,[],[f10324,f612])).
fof(f16272,plain,(
  ( ! [X105,X106,X104] : (s(bool,X105) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X105) = s(bool,'const.combin.I_1'(s(bool,X104))) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') = s(bool,X106)) ) | $spl0),
  inference(subsumption_resolution,[],[f14958,f177])).
fof(f14958,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X105) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X105) = s(bool,'const.combin.I_1'(s(bool,X104))) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') = s(bool,X106)) ) | $spl0),
  inference(superposition,[],[f10324,f755])).
fof(f16644,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(condensation,[],[f16643])).
fof(f16643,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,'const.combin.I_1'(s(bool,X102))) | s(bool,X101) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X102) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f15446,f900])).
fof(f15446,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,'const.combin.I_1'(s(bool,X102))) | s(bool,X101) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X102) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X101)) ) | $spl0),
  inference(superposition,[],[f1617,f10324])).
fof(f16665,plain,(
  ( ! [X279,X278] : (s(bool,X279) != s(bool,'const.combin.I_1'(s(bool,X278))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X279))) | s(bool,X278) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f15515,f177])).
fof(f15515,plain,(
  ( ! [X279,X278] : (s(bool,X279) != s(bool,'const.combin.I_1'(s(bool,X278))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X279))) | s(bool,X278) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f1943,f10324])).
fof(f16731,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f15682,f117])).
fof(f15682,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f134,f10324])).
fof(f8044,plain,(
  ( ! [X325,X323,X324,X322] : (s(bool,'const.combin.I_1'(s(bool,X325))) != s(bool,f(s(X322,sK3(X323,X324,bool,X322)))) | s(fun(X322,bool),X323) = s(fun(X322,bool),X324) | s(bool,'T_0') = s(bool,g(s(X322,sK3(X323,X324,bool,X322)))) | s(bool,'T_0') = s(bool,X325)) ) | $spl0),
  inference(superposition,[],[f159,f2195])).
fof(f9207,plain,(
  ( ! [X366,X368,X367,X369] : (s(bool,X369) != s(bool,f(s(X366,sK3(X367,X368,bool,X366)))) | s(fun(X366,bool),X367) = s(fun(X366,bool),X368) | s(bool,'T_0') = s(bool,X369) | s(bool,'T_0') = s(bool,g(s(X366,sK3(X367,X368,bool,X366))))) )),
  inference(superposition,[],[f159,f205])).
fof(f9561,plain,(
  ( ! [X352,X354,X351,X353] : (s(bool,'T_0') != s(bool,f(s(X351,sK3(X352,X353,bool,X351)))) | s(fun(X351,bool),X352) = s(fun(X351,bool),X353) | s(bool,X354) = s(bool,g(s(X351,sK3(X352,X353,bool,X351)))) | s(bool,'T_0') = s(bool,X354)) )),
  inference(superposition,[],[f159,f205])).
fof(f157,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) & (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) <=> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X7),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) <=> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X7))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|0')).
fof(f158,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f4069,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X1))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))))) )),
  inference(unit_resulting_resolution,[],[f147,f157])).
fof(f4736,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X21,X20),'const.pair.,_2'(s(X21,X23),s(X20,X24))),s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22))) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))))) )),
  inference(superposition,[],[f158,f165])).
fof(f4083,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30),s(X26,X28))),s(X27,X29))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30)))) )),
  inference(superposition,[],[f157,f165])).
fof(f4737,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30),s(X26,X28))),s(X27,X29))) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))))) )),
  inference(superposition,[],[f158,f165])).
fof(f4082,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X21,X20),'const.pair.,_2'(s(X21,X23),s(X20,X24))),s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22))),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24)))) )),
  inference(superposition,[],[f157,f165])).
fof(f143,plain,(
  ( ! [X0,X1] : (s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))) = s(bool,happ(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0'),s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))) = s(bool,happ(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0'),s(fun(X0,fun(X0,bool)),X1)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1] : s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))) = s(bool,happ(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0'),s(fun(X0,fun(X0,bool)),X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X16] : s(bool,'const.relation.WF_1'(s(fun(X9,fun(X9,bool)),X16))) = s(bool,happ(s(fun(fun(X9,fun(X9,bool)),bool),'const.relation.WF_0'),s(fun(X9,fun(X9,bool)),X16)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','ARITY_0_1_const.relation.WF')).
fof(f149,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9] : ! [X21,X20] : (p(s(bool,happ(s(fun(X9,bool),X21),s(X9,X20)))) => p(s(bool,happ(s(fun(X9,bool),X21),s(X9,'const.min.@_1'(s(fun(X9,bool),X21)))))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.bool.SELECT_AX')).
fof(f3089,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,'T_0') != s(bool,happ(s(fun(X221,bool),X225),s(X221,X222))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X221),'const.semanticPrimitives.Match_1'(s(X221,X222))),s(bool,X223),s(bool,X224),s(fun(X221,bool),X225)))) ) | $spl0),
  inference(forward_demodulation,[],[f3036,f141])).
fof(f3036,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X221,bool),X225),s(X221,X222))))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X221),'const.semanticPrimitives.Match_1'(s(X221,X222))),s(bool,X223),s(bool,X224),s(fun(X221,bool),X225)))) ) | $spl0),
  inference(superposition,[],[f2047,f163])).
fof(f7860,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(bool,'const.combin.I_1'(s(bool,X108))) = s(bool,happ(s(fun(X103,bool),X107),s(X103,X104))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.Match_1'(s(X103,X104))),s(bool,X105),s(bool,X106),s(fun(X103,bool),X107))) | s(bool,'T_0') = s(bool,X108)) ) | $spl0),
  inference(superposition,[],[f2195,f163])).
fof(f8243,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (s(bool,X111) = s(bool,happ(s(fun(X106,bool),X110),s(X106,X107))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.Match_1'(s(X106,X107))),s(bool,X108),s(bool,X109),s(fun(X106,bool),X110)))) ) | $spl0),
  inference(forward_demodulation,[],[f7907,f141])).
fof(f7907,plain,(
  ( ! [X111,X109,X107,X110,X108,X106] : (s(bool,X111) = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X106,bool),X110),s(X106,X107))))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X106),'const.semanticPrimitives.Match_1'(s(X106,X107))),s(bool,X108),s(bool,X109),s(fun(X106,bool),X110)))) ) | $spl0),
  inference(superposition,[],[f2195,f163])).
fof(f9012,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X117),s(X113,X114))) | s(bool,X118) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X113),'const.semanticPrimitives.Match_1'(s(X113,X114))),s(bool,X115),s(bool,X116),s(fun(X113,bool),X117))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f205,f163])).
fof(f10327,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,happ(s(fun(X208,bool),X209),s(X208,'const.min.@_1'(s(fun(X208,bool),X209)))))) | s(bool,X211) = s(bool,happ(s(fun(X208,bool),X209),s(X208,X210))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(subsumption_resolution,[],[f9517,f117])).
fof(f9517,plain,(
  ( ! [X210,X208,X211,X209] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X208,bool),X209),s(X208,'const.min.@_1'(s(fun(X208,bool),X209)))))) | s(bool,X211) = s(bool,happ(s(fun(X208,bool),X209),s(X208,X210))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f149,f205])).
fof(f3642,plain,(
  ( ! [X255,X253,X251,X254,X252] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X252,bool),happ(s(fun(X251,fun(X252,bool)),X255),s(X251,X253))),s(X252,X254))))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X253),s(X252,X254))),s(fun(X251,fun(X252,bool)),X255)))) ) | $spl0),
  inference(superposition,[],[f2047,f165])).
fof(f7863,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'const.combin.I_1'(s(bool,X126))) = s(bool,happ(s(fun(X122,bool),happ(s(fun(X121,fun(X122,bool)),X125),s(X121,X123))),s(X122,X124))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X121,X122),'const.pair.,_2'(s(X121,X123),s(X122,X124))),s(fun(X121,fun(X122,bool)),X125))) | s(bool,'T_0') = s(bool,X126)) ) | $spl0),
  inference(superposition,[],[f2195,f165])).
fof(f7910,plain,(
  ( ! [X127,X125,X128,X126,X124,X129] : (s(bool,X129) = s(bool,'const.combin.I_1'(s(bool,happ(s(fun(X125,bool),happ(s(fun(X124,fun(X125,bool)),X128),s(X124,X126))),s(X125,X127))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X124,X125),'const.pair.,_2'(s(X124,X126),s(X125,X127))),s(fun(X124,fun(X125,bool)),X128)))) ) | $spl0),
  inference(superposition,[],[f2195,f165])).
fof(f8906,plain,(
  ( ! [X146,X144,X149,X147,X145,X148] : (s(bool,X149) = s(bool,happ(s(fun(X145,bool),happ(s(fun(X144,fun(X145,bool)),X148),s(X144,X146))),s(X145,X147))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X144,X145),'const.pair.,_2'(s(X144,X146),s(X145,X147))),s(fun(X144,fun(X145,bool)),X148)))) )),
  inference(superposition,[],[f205,f165])).
fof(f9015,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s(bool,'T_0') = s(bool,happ(s(fun(X132,bool),happ(s(fun(X131,fun(X132,bool)),X135),s(X131,X133))),s(X132,X134))) | s(bool,X136) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X131,X132),'const.pair.,_2'(s(X131,X133),s(X132,X134))),s(fun(X131,fun(X132,bool)),X135))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f205,f165])).
fof(f3592,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272] : (~p(s(bool,happ(s(fun(X267,bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X269,X268),'const.pair.,_2'(s(X269,X271),s(X268,X272))),s(fun(X269,fun(X268,fun(X267,bool))),X270))),s(X267,X273)))) | p(s(bool,happ(s(fun(X267,bool),happ(s(fun(X268,fun(X267,bool)),happ(s(fun(X269,fun(X268,fun(X267,bool))),X270),s(X269,X271))),s(X268,X272))),s(X267,'const.min.@_1'(s(fun(X267,bool),happ(s(fun(X268,fun(X267,bool)),happ(s(fun(X269,fun(X268,fun(X267,bool))),X270),s(X269,X271))),s(X268,X272))))))))) )),
  inference(superposition,[],[f149,f165])).
fof(f3647,plain,(
  ( ! [X280,X282,X277,X279,X281,X276,X278] : (~p(s(bool,happ(s(fun(X276,bool),happ(s(fun(X278,fun(X276,bool)),happ(s(fun(X277,fun(X278,fun(X276,bool))),X281),s(X277,X279))),s(X278,X280))),s(X276,X282)))) | p(s(bool,happ(s(fun(X276,bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X277,X278),'const.pair.,_2'(s(X277,X279),s(X278,X280))),s(fun(X277,fun(X278,fun(X276,bool))),X281))),s(X276,'const.min.@_1'(s(fun(X276,bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X277,X278),'const.pair.,_2'(s(X277,X279),s(X278,X280))),s(fun(X277,fun(X278,fun(X276,bool))),X281))))))))) )),
  inference(superposition,[],[f149,f165])).
fof(f4996,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X108,X109),X110),s(fun(X108,fun(X109,bool)),X111)))) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X108,X109),X110),s(fun(X108,fun(X109,bool)),'const.min.@_1'(s(fun(fun(X108,fun(X109,bool)),bool),happ(s(fun('type.pair.prod'(X108,X109),fun(fun(X108,fun(X109,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X108,X109),X110))))))))) )),
  inference(forward_demodulation,[],[f4934,f164])).
fof(f4934,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X108,X109),X110),s(fun(X108,fun(X109,bool)),X111)))) | p(s(bool,happ(s(fun(fun(X108,fun(X109,bool)),bool),happ(s(fun('type.pair.prod'(X108,X109),fun(fun(X108,fun(X109,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X108,X109),X110))),s(fun(X108,fun(X109,bool)),'const.min.@_1'(s(fun(fun(X108,fun(X109,bool)),bool),happ(s(fun('type.pair.prod'(X108,X109),fun(fun(X108,fun(X109,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X108,X109),X110))))))))) )),
  inference(superposition,[],[f149,f164])).
fof(f5021,plain,(
  ( ! [X251,X249,X250,X248] : (s(bool,'T_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X248,X249),X250),s(fun(X248,fun(X249,bool)),X251))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X248,fun(X249,bool)),bool),happ(s(fun('type.pair.prod'(X248,X249),fun(fun(X248,fun(X249,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X248,X249),X250))),s(fun(X248,fun(X249,bool)),X251)))) ) | $spl0),
  inference(forward_demodulation,[],[f4965,f141])).
fof(f4965,plain,(
  ( ! [X251,X249,X250,X248] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X248,X249),X250),s(fun(X248,fun(X249,bool)),X251))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X248,fun(X249,bool)),bool),happ(s(fun('type.pair.prod'(X248,X249),fun(fun(X248,fun(X249,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X248,X249),X250))),s(fun(X248,fun(X249,bool)),X251)))) ) | $spl0),
  inference(superposition,[],[f2047,f164])).
fof(f7862,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(bool,'const.combin.I_1'(s(bool,X120))) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X116,X117),X118),s(fun(X116,fun(X117,bool)),X119))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X116,fun(X117,bool)),bool),happ(s(fun('type.pair.prod'(X116,X117),fun(fun(X116,fun(X117,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X116,X117),X118))),s(fun(X116,fun(X117,bool)),X119))) | s(bool,'T_0') = s(bool,X120)) ) | $spl0),
  inference(superposition,[],[f2195,f164])).
fof(f8245,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,X123) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X119,X120),X121),s(fun(X119,fun(X120,bool)),X122))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X119,fun(X120,bool)),bool),happ(s(fun('type.pair.prod'(X119,X120),fun(fun(X119,fun(X120,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X119,X120),X121))),s(fun(X119,fun(X120,bool)),X122)))) ) | $spl0),
  inference(forward_demodulation,[],[f7909,f141])).
fof(f7909,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,X123) = s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X119,X120),X121),s(fun(X119,fun(X120,bool)),X122))))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X119,fun(X120,bool)),bool),happ(s(fun('type.pair.prod'(X119,X120),fun(fun(X119,fun(X120,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X119,X120),X121))),s(fun(X119,fun(X120,bool)),X122)))) ) | $spl0),
  inference(superposition,[],[f2195,f164])).
fof(f9014,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X126,X127),X128),s(fun(X126,fun(X127,bool)),X129))) | s(bool,X130) = s(bool,happ(s(fun(fun(X126,fun(X127,bool)),bool),happ(s(fun('type.pair.prod'(X126,X127),fun(fun(X126,fun(X127,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X126,X127),X128))),s(fun(X126,fun(X127,bool)),X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f205,f164])).
fof(f162,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10,X9] : ! [X18,X19,X13] : (s(fun(X9,X10),X13) = s(fun(X9,X10),'const.relation.WFREC_2'(s(fun(X9,fun(X9,bool)),X19),s(fun(fun(X9,X10),fun(X9,X10)),X18))) => (p(s(bool,'const.relation.WF_1'(s(fun(X9,fun(X9,bool)),X19)))) => ! [X20] : s(X10,happ(s(fun(X9,X10),X13),s(X9,X20))) = s(X10,happ(s(fun(X9,X10),happ(s(fun(fun(X9,X10),fun(X9,X10)),X18),s(fun(X9,X10),'const.relation.RESTRICT_3'(s(fun(X9,X10),X13),s(fun(X9,fun(X9,bool)),X19),s(X9,X20))))),s(X9,X20)))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.relation.WFREC_COROLLARY')).
fof(f2367,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'const.relation.WF_1'(s(fun(X27,fun(X27,bool)),X28))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X27,fun(X27,bool)),bool),'const.relation.WF_0'),s(fun(X27,fun(X27,bool)),X28)))) ) | $spl0),
  inference(forward_demodulation,[],[f2340,f141])).
fof(f2340,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X27,fun(X27,bool)),X28))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X27,fun(X27,bool)),bool),'const.relation.WF_0'),s(fun(X27,fun(X27,bool)),X28)))) ) | $spl0),
  inference(superposition,[],[f2047,f143])).
fof(f7842,plain,(
  ( ! [X59,X57,X58] : (s(bool,'const.combin.I_1'(s(bool,X59))) = s(bool,'const.relation.WF_1'(s(fun(X57,fun(X57,bool)),X58))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X57,fun(X57,bool)),bool),'const.relation.WF_0'),s(fun(X57,fun(X57,bool)),X58))) | s(bool,'T_0') = s(bool,X59)) ) | $spl0),
  inference(superposition,[],[f2195,f143])).
fof(f8230,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) = s(bool,'const.relation.WF_1'(s(fun(X60,fun(X60,bool)),X61))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X60,fun(X60,bool)),bool),'const.relation.WF_0'),s(fun(X60,fun(X60,bool)),X61)))) ) | $spl0),
  inference(forward_demodulation,[],[f7889,f141])).
fof(f7889,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'const.relation.WF_1'(s(fun(X60,fun(X60,bool)),X61))))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X60,fun(X60,bool)),bool),'const.relation.WF_0'),s(fun(X60,fun(X60,bool)),X61)))) ) | $spl0),
  inference(superposition,[],[f2195,f143])).
fof(f8998,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X70,fun(X70,bool)),X71))) | s(bool,X72) = s(bool,happ(s(fun(fun(X70,fun(X70,bool)),bool),'const.relation.WF_0'),s(fun(X70,fun(X70,bool)),X71))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f205,f143])).
fof(f10330,plain,(
  ( ! [X229,X227,X225,X223,X228,X226,X224] : (s(X225,happ(s(fun(X223,X225),X226),s(X223,X227))) = s(X225,happ(s(fun(X223,X225),happ(s(fun(fun(X223,X225),fun(X223,X225)),X228),s(fun(X223,X225),'const.relation.RESTRICT_3'(s(fun(X223,X225),X226),s(fun(X223,fun(X223,bool)),X224),s(X223,X227))))),s(X223,X227))) | s(fun(X223,X225),X226) != s(fun(X223,X225),'const.relation.WFREC_2'(s(fun(X223,fun(X223,bool)),X224),s(fun(fun(X223,X225),fun(X223,X225)),X228))) | s(bool,X229) = s(bool,'const.relation.WF_1'(s(fun(X223,fun(X223,bool)),X224))) | s(bool,'T_0') = s(bool,X229)) )),
  inference(subsumption_resolution,[],[f9521,f117])).
fof(f9521,plain,(
  ( ! [X229,X227,X225,X223,X228,X226,X224] : (~p(s(bool,'T_0')) | s(X225,happ(s(fun(X223,X225),X226),s(X223,X227))) = s(X225,happ(s(fun(X223,X225),happ(s(fun(fun(X223,X225),fun(X223,X225)),X228),s(fun(X223,X225),'const.relation.RESTRICT_3'(s(fun(X223,X225),X226),s(fun(X223,fun(X223,bool)),X224),s(X223,X227))))),s(X223,X227))) | s(fun(X223,X225),X226) != s(fun(X223,X225),'const.relation.WFREC_2'(s(fun(X223,fun(X223,bool)),X224),s(fun(fun(X223,X225),fun(X223,X225)),X228))) | s(bool,X229) = s(bool,'const.relation.WF_1'(s(fun(X223,fun(X223,bool)),X224))) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f162,f205])).
fof(f5786,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.WF_1'(s(fun(X53,fun(X53,bool)),happ(s(fun(X55,fun(X53,fun(X53,bool))),happ(s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58),s(X54,X56))),s(X55,X57)))))) | s(X59,happ(s(fun(X53,X59),X60),s(X53,X61))) = s(X59,happ(s(fun(X53,X59),happ(s(fun(fun(X53,X59),fun(X53,X59)),X62),s(fun(X53,X59),'const.relation.RESTRICT_3'(s(fun(X53,X59),X60),s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(X53,X61))))),s(X53,X61))) | s(fun(X53,X59),X60) != s(fun(X53,X59),'const.relation.WFREC_2'(s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(fun(fun(X53,X59),fun(X53,X59)),X62)))) )),
  inference(superposition,[],[f162,f165])).
fof(f5784,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (~p(s(bool,'const.relation.WF_1'(s(fun(X34,fun(X34,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X36,X35),'const.pair.,_2'(s(X36,X38),s(X35,X39))),s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37)))))) | s(X40,happ(s(fun(X34,X40),X41),s(X34,X42))) = s(X40,happ(s(fun(X34,X40),happ(s(fun(fun(X34,X40),fun(X34,X40)),X43),s(fun(X34,X40),'const.relation.RESTRICT_3'(s(fun(X34,X40),X41),s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(X34,X42))))),s(X34,X42))) | s(fun(X34,X40),X41) != s(fun(X34,X40),'const.relation.WFREC_2'(s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(fun(fun(X34,X40),fun(X34,X40)),X43)))) )),
  inference(superposition,[],[f162,f165])).
fof(f8213,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,X0))) = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(condensation,[],[f8212])).
fof(f8212,plain,(
  ( ! [X30,X31,X32] : (s(bool,'const.combin.I_1'(s(bool,X31))) = s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7833])).
fof(f7833,plain,(
  ( ! [X30,X31,X32] : (s(bool,'const.combin.I_1'(s(bool,X31))) = s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f2195,f2195])).
fof(f16753,plain,(
  ( ! [X161,X162] : (s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X162) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X161) = s(bool,X162)) ) | $spl0),
  inference(subsumption_resolution,[],[f15737,f141])).
fof(f15737,plain,(
  ( ! [X161,X162] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X162) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X161) = s(bool,X162)) ) | $spl0),
  inference(superposition,[],[f4502,f10324])).
fof(f15604,plain,(
  ( ! [X47,X48,X49] : (s(bool,X47) != s(bool,X48) | s(bool,X48) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X49) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X47) = s(bool,X49)) ) | $spl0),
  inference(superposition,[],[f4691,f10324])).
fof(f16666,plain,(
  ( ! [X282,X283] : (s(bool,X283) = s(bool,'const.combin.I_1'(s(bool,X282))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X283))) | s(bool,X282) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f15517,f177])).
fof(f15517,plain,(
  ( ! [X282,X283] : (s(bool,X283) = s(bool,'const.combin.I_1'(s(bool,X282))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X283))) | s(bool,X282) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl0),
  inference(superposition,[],[f8450,f10324])).
fof(f16428,plain,(
  ( ! [X151,X149,X150] : (s(bool,'T_0') = s(bool,X150) | s(bool,X149) != s(bool,'const.combin.I_1'(s(bool,X150))) | s(bool,X149) = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | $spl0),
  inference(subsumption_resolution,[],[f16427,f177])).
fof(f16427,plain,(
  ( ! [X151,X149,X150] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X150) | s(bool,X149) != s(bool,'const.combin.I_1'(s(bool,X150))) | s(bool,X149) = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | $spl0),
  inference(forward_demodulation,[],[f16426,f141])).
fof(f16426,plain,(
  ( ! [X151,X149,X150] : (s(bool,'T_0') = s(bool,X150) | s(bool,X149) != s(bool,'const.combin.I_1'(s(bool,X150))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X149) = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | $spl0),
  inference(forward_demodulation,[],[f15181,f141])).
fof(f15181,plain,(
  ( ! [X151,X149,X150] : (s(bool,X149) != s(bool,'const.combin.I_1'(s(bool,X150))) | s(bool,X150) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X149) = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | $spl0),
  inference(superposition,[],[f1617,f10324])).
fof(f16277,plain,(
  ( ! [X118,X117] : (s(bool,X118) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X118) = s(bool,'const.combin.I_1'(s(bool,X117))) | s(bool,X117) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f14963,f177])).
fof(f14963,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X118) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X118) = s(bool,'const.combin.I_1'(s(bool,X117))) | s(bool,X117) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f10324,f3991])).
fof(f14913,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f10324,f205])).
fof(f14912,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f10324,f10324])).
fof(f16251,plain,(
  ( ! [X147,X148] : (s(bool,X147) = s(bool,X148) | s(bool,X148) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X147)))) ) | $spl0),
  inference(subsumption_resolution,[],[f14903,f1881])).
fof(f14903,plain,(
  ( ! [X147,X148] : (s(bool,X147) = s(bool,X148) | s(bool,X148) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X147)))) ) | $spl0),
  inference(superposition,[],[f10324,f8450])).
fof(f16060,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(condensation,[],[f16059])).
fof(f16059,plain,(
  ( ! [X43,X44,X42] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,X43))) | s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X43)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f14860])).
fof(f14860,plain,(
  ( ! [X43,X44,X42] : (s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,X43))) | s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f10324,f2195])).
fof(f14734,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f10324,f205])).
fof(f14733,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X0) = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f10324,f10324])).
fof(f12554,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.relation.WF_1'(s(fun(X2,fun(X2,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f12476,f205])).
fof(f12476,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,X80) | s(bool,X80) = s(bool,'const.relation.WF_1'(s(fun(X81,fun(X81,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f5758,f141])).
fof(f10449,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) ) | $spl0),
  inference(superposition,[],[f8450,f8450])).
fof(f9571,plain,(
  ( ! [X397,X395,X396,X398] : (s(X397,f(s(bool,'T_0'))) != s(X397,g(s(bool,'T_0'))) | s(fun(bool,X397),X395) = s(fun(bool,X397),X396) | s(bool,X398) = s(bool,sK3(X395,X396,X397,bool)) | s(bool,'T_0') = s(bool,X398)) )),
  inference(superposition,[],[f159,f205])).
fof(f9570,plain,(
  ( ! [X393,X394] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X393))) | s(bool,X394) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X393)),bool)) | s(bool,'T_0') = s(bool,X394)) )),
  inference(superposition,[],[f2200,f205])).
fof(f9569,plain,(
  ( ! [X391,X392] : (s('type.semanticPrimitives.match_result'(bool),X391) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X391) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X391) | s(bool,X392) = s(bool,sK2(X391,bool)) | s(bool,'T_0') = s(bool,X392)) )),
  inference(superposition,[],[f144,f205])).
fof(f10326,plain,(
  ( ! [X206,X207,X205] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X206) | s(bool,X207) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X205),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X206),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X207)) )),
  inference(subsumption_resolution,[],[f9516,f117])).
fof(f9516,plain,(
  ( ! [X206,X207,X205] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X206) | s(bool,X207) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X205),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X206),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X207)) )),
  inference(superposition,[],[f157,f205])).
fof(f10324,plain,(
  ( ! [X120,X119] : (s(bool,X119) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X119) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) ) | $spl0),
  inference(subsumption_resolution,[],[f9479,f141])).
fof(f9479,plain,(
  ( ! [X120,X119] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X119) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X119) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) ) | $spl0),
  inference(superposition,[],[f4502,f205])).
fof(f10314,plain,(
  ( ! [X78,X79,X77] : (s(bool,X77) = s(bool,'const.relation.WF_1'(s(fun(X78,fun(X78,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(subsumption_resolution,[],[f10313,f177])).
fof(f10313,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X77) = s(bool,'const.relation.WF_1'(s(fun(X78,fun(X78,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(forward_demodulation,[],[f9461,f141])).
fof(f9461,plain,(
  ( ! [X78,X79,X77] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X77) = s(bool,'const.relation.WF_1'(s(fun(X78,fun(X78,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(superposition,[],[f1796,f205])).
fof(f10300,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f9439,f117])).
fof(f9439,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f156,f205])).
fof(f10299,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f9435,f117])).
fof(f9435,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f134,f205])).
fof(f9374,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl0),
  inference(superposition,[],[f4691,f205])).
fof(f9369,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X21))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f2368,f205])).
fof(f9366,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f900,f205])).
fof(f9013,plain,(
  ( ! [X120,X119] : (s(bool,'T_0') = s(bool,X119) | s(bool,X120) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X119))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f205,f172])).
fof(f9011,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,X110) | s(bool,X112) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X108),'const.semanticPrimitives.Match_type_error_0'),s(bool,X109),s(bool,X110),s(fun(X108,bool),X111))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f205,f161])).
fof(f9010,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') = s(bool,X104) | s(bool,X107) = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X103),'const.semanticPrimitives.No_match_0'),s(bool,X104),s(bool,X105),s(fun(X103,bool),X106))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f205,f160])).
fof(f8963,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f205,f205])).
fof(f8907,plain,(
  ( ! [X151,X150] : (s(bool,X150) = s(bool,X151) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X150)))) )),
  inference(superposition,[],[f205,f141])).
fof(f8353,plain,(
  ( ! [X368,X370,X367,X369] : (s(X369,f(s(bool,X370))) != s(X369,g(s(bool,X370))) | s(fun(bool,X369),X367) = s(fun(bool,X369),X368) | s(bool,'T_0') = s(bool,sK3(X367,X368,X369,bool)) | s(bool,'T_0') = s(bool,X370)) ) | $spl0),
  inference(forward_demodulation,[],[f8054,f141])).
fof(f8054,plain,(
  ( ! [X368,X370,X367,X369] : (s(X369,f(s(bool,'const.combin.I_1'(s(bool,X370))))) != s(X369,g(s(bool,'const.combin.I_1'(s(bool,X370))))) | s(fun(bool,X369),X367) = s(fun(bool,X369),X368) | s(bool,'T_0') = s(bool,sK3(X367,X368,X369,bool)) | s(bool,'T_0') = s(bool,X370)) ) | $spl0),
  inference(superposition,[],[f159,f2195])).
fof(f8352,plain,(
  ( ! [X366,X365] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X365))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X366))) | s(bool,'T_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X365)),bool)) | s(bool,'T_0') = s(bool,X366)) ) | $spl0),
  inference(forward_demodulation,[],[f8053,f141])).
fof(f8053,plain,(
  ( ! [X366,X365] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X365))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X366))))) | s(bool,'T_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X365)),bool)) | s(bool,'T_0') = s(bool,X366)) ) | $spl0),
  inference(superposition,[],[f2200,f2195])).
fof(f8351,plain,(
  ( ! [X364,X363] : (s('type.semanticPrimitives.match_result'(bool),X363) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X364))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X363) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X363) | s(bool,'T_0') = s(bool,sK2(X363,bool)) | s(bool,'T_0') = s(bool,X364)) ) | $spl0),
  inference(forward_demodulation,[],[f8052,f141])).
fof(f8052,plain,(
  ( ! [X364,X363] : (s('type.semanticPrimitives.match_result'(bool),X363) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X364))))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X363) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X363) | s(bool,'T_0') = s(bool,sK2(X363,bool)) | s(bool,'T_0') = s(bool,X364)) ) | $spl0),
  inference(superposition,[],[f144,f2195])).
fof(f8244,plain,(
  ( ! [X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X112)))) ) | $spl0),
  inference(forward_demodulation,[],[f7908,f141])).
fof(f7908,plain,(
  ( ! [X112,X113] : (s(bool,X113) = s(bool,'const.combin.I_1'(s(bool,X112))) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X112)))) ) | $spl0),
  inference(superposition,[],[f2195,f172])).
fof(f8242,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X101),'const.semanticPrimitives.Match_type_error_0'),s(bool,X102),s(bool,X103),s(fun(X101,bool),X104)))) ) | $spl0),
  inference(forward_demodulation,[],[f7906,f141])).
fof(f7906,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s(bool,X105) = s(bool,'const.combin.I_1'(s(bool,X103))) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X101),'const.semanticPrimitives.Match_type_error_0'),s(bool,X102),s(bool,X103),s(fun(X101,bool),X104)))) ) | $spl0),
  inference(superposition,[],[f2195,f161])).
fof(f8241,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X96),'const.semanticPrimitives.No_match_0'),s(bool,X97),s(bool,X98),s(fun(X96,bool),X99)))) ) | $spl0),
  inference(forward_demodulation,[],[f7905,f141])).
fof(f7905,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X100) = s(bool,'const.combin.I_1'(s(bool,X97))) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X96),'const.semanticPrimitives.No_match_0'),s(bool,X97),s(bool,X98),s(fun(X96,bool),X99)))) ) | $spl0),
  inference(superposition,[],[f2195,f160])).
fof(f7864,plain,(
  ( ! [X127,X128] : (s(bool,X127) = s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X127))) | s(bool,'T_0') = s(bool,X128)) ) | $spl0),
  inference(superposition,[],[f2195,f141])).
fof(f7861,plain,(
  ( ! [X109,X110] : (s(bool,X109) = s(bool,'const.combin.I_1'(s(bool,X110))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X109))) | s(bool,'T_0') = s(bool,X110)) ) | $spl0),
  inference(superposition,[],[f2195,f172])).
fof(f7859,plain,(
  ( ! [X101,X99,X102,X100,X98] : (s(bool,X100) = s(bool,'const.combin.I_1'(s(bool,X102))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X98),'const.semanticPrimitives.Match_type_error_0'),s(bool,X99),s(bool,X100),s(fun(X98,bool),X101))) | s(bool,'T_0') = s(bool,X102)) ) | $spl0),
  inference(superposition,[],[f2195,f161])).
fof(f7858,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X94) = s(bool,'const.combin.I_1'(s(bool,X97))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X93),'const.semanticPrimitives.No_match_0'),s(bool,X94),s(bool,X95),s(fun(X93,bool),X96))) | s(bool,'T_0') = s(bool,X97)) ) | $spl0),
  inference(superposition,[],[f2195,f160])).
fof(f7816,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,happ(s(fun(X8,X10),X11),s(X8,X9))) = s(X10,happ(s(fun(X8,X10),X11),s(X8,sK2('const.semanticPrimitives.Match_1'(s(X8,X9)),X8))))) )),
  inference(forward_demodulation,[],[f7787,f163])).
fof(f7787,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(X10,happ(s(fun(X8,X10),X11),s(X8,sK2('const.semanticPrimitives.Match_1'(s(X8,X9)),X8)))) = s(X10,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X8),'const.semanticPrimitives.Match_1'(s(X8,X9))),s(X10,X12),s(X10,X13),s(fun(X8,X10),X11)))) )),
  inference(superposition,[],[f163,f2200])).
fof(f6320,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X52)) ) | $spl0),
  inference(superposition,[],[f1916,f141])).
fof(f6298,plain,(
  ( ! [X118,X117] : (s(bool,X118) != s(bool,'const.combin.I_1'(s(bool,X117))) | s(bool,'F_0') = s(bool,X118) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X117)))) ) | $spl0),
  inference(superposition,[],[f1916,f141])).
fof(f568,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0')))))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(forward_demodulation,[],[f547,f143])).
fof(f547,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,happ(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0'),s(fun(X0,fun(X0,bool)),'const.min.@_1'(s(fun(fun(X0,fun(X0,bool)),bool),'const.relation.WF_0'))))))) )),
  inference(superposition,[],[f149,f143])).
fof(f5818,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,X9))) = s(X7,happ(s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),X10),s(fun(X6,X7),'const.relation.RESTRICT_3'(s(fun(X6,X7),X8),s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(X6,X9))))),s(X6,X9))) | s(fun(X6,X7),X8) != s(fun(X6,X7),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X6,X7),fun(X6,X7)),X10)))) )),
  inference(subsumption_resolution,[],[f5789,f117])).
fof(f5789,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | s(X7,happ(s(fun(X6,X7),X8),s(X6,X9))) = s(X7,happ(s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),X10),s(fun(X6,X7),'const.relation.RESTRICT_3'(s(fun(X6,X7),X8),s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(X6,X9))))),s(X6,X9))) | s(fun(X6,X7),X8) != s(fun(X6,X7),'const.relation.WFREC_2'(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X6,X7),fun(X6,X7)),X10)))) )),
  inference(superposition,[],[f162,f183])).
fof(f5765,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,'const.relation.WF_1'(s(fun(X48,fun(X48,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f5461,f171])).
fof(f5461,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'F_0')) | s(bool,X47) = s(bool,'const.combin.I_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,'const.relation.WF_1'(s(fun(X48,fun(X48,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f256,f1796])).
fof(f5758,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.relation.WF_1'(s(fun(X28,fun(X28,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f5453,f171])).
fof(f5453,plain,(
  ( ! [X28,X27] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X27))) | s(bool,X27) = s(bool,'const.relation.WF_1'(s(fun(X28,fun(X28,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f130,f1796])).
fof(f4906,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X8),s(fun(X6,fun(X7,X5)),X9))) = s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X6,X7),fun(X6,fun(X7,X5))),'const.pair.,_2'(s('type.pair.prod'(X6,X7),X8),s(fun(X6,fun(X7,X5)),X9))),s(fun('type.pair.prod'(X6,X7),fun(fun(X6,fun(X7,X5)),X5)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f164,f165])).
fof(f4691,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f4502,f141])).
fof(f4509,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f4232,f171])).
fof(f4232,plain,(
  ( ! [X23,X22] : (p(s(bool,'F_0')) | s(bool,X23) = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f256,f3991])).
fof(f4502,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,X11) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f4224,f171])).
fof(f4224,plain,(
  ( ! [X11] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X11))) | s(bool,X11) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f130,f3991])).
fof(f3490,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X5,happ(s(fun(X4,X5),'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X9),s(X6,X10))),s(fun(X7,fun(X6,fun(X4,X5))),X8))),s(X4,X11))) = s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X4),'const.pair.,_2'(s(X6,X10),s(X4,X11))),s(fun(X6,fun(X4,X5)),happ(s(fun(X7,fun(X6,fun(X4,X5))),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f165,f165])).
fof(f3322,plain,(
  ( ! [X23,X21,X22] : (s(bool,X23) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f3265,f171])).
fof(f3265,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | s(bool,X23) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f256,f152])).
fof(f3317,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)))))) )),
  inference(subsumption_resolution,[],[f3257,f171])).
fof(f3257,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_1'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X3)))))) )),
  inference(superposition,[],[f130,f152])).
fof(f2917,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK2(X1,X0)))) = s(X2,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),X1),s(X2,X4),s(X2,X5),s(fun(X0,X2),X3))) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),X1)) )),
  inference(superposition,[],[f163,f144])).
fof(f2370,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X40),'const.semanticPrimitives.Match_type_error_0'),s(bool,X41),s(bool,X42),s(fun(X40,bool),X43)))) ) | $spl0),
  inference(forward_demodulation,[],[f2349,f141])).
fof(f2349,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X40),'const.semanticPrimitives.Match_type_error_0'),s(bool,X41),s(bool,X42),s(fun(X40,bool),X43)))) ) | $spl0),
  inference(superposition,[],[f2047,f161])).
fof(f2369,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X36),'const.semanticPrimitives.No_match_0'),s(bool,X37),s(bool,X38),s(fun(X36,bool),X39)))) ) | $spl0),
  inference(forward_demodulation,[],[f2348,f141])).
fof(f2348,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X36),'const.semanticPrimitives.No_match_0'),s(bool,X37),s(bool,X38),s(fun(X36,bool),X39)))) ) | $spl0),
  inference(superposition,[],[f2047,f160])).
fof(f2368,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) | $spl0),
  inference(forward_demodulation,[],[f2347,f141])).
fof(f2347,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) | $spl0),
  inference(superposition,[],[f2047,f172])).
fof(f2200,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1))) = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,sK2('const.semanticPrimitives.Match_1'(s(X0,X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f148,f147,f144])).
fof(f2195,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | $spl0),
  inference(subsumption_resolution,[],[f2129,f171])).
fof(f2129,plain,(
  ( ! [X17,X16] : (p(s(bool,'F_0')) | s(bool,X17) = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f256,f1764])).
fof(f2047,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f1881,f128])).
fof(f1303,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f1273,f171])).
fof(f1273,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f256,f126])).
fof(f1298,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0')))) )),
  inference(subsumption_resolution,[],[f1265,f171])).
fof(f1265,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0')))) )),
  inference(superposition,[],[f130,f126])).
fof(f1256,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f1218,f177])).
fof(f1218,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f900,f612])).
fof(f900,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f177,f612])).
fof(f781,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f612,f128])).
fof(f779,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f612,f612])).
fof(f735,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f716,f171])).
fof(f716,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X12),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f256,f125])).
fof(f730,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(subsumption_resolution,[],[f708,f171])).
fof(f708,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X1),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.No_match_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(superposition,[],[f130,f125])).
fof(f612,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f594,f594])).
fof(f256,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f249,f171])).
fof(f249,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f153,f128])).
fof(f205,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f128,f128])).
fof(f200,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f167,f131])).
fof(f183,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f124,f131])).
fof(f177,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f171,f130])).
fof(f165,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9,X10,X15] : ! [X22,X20,X13] : s(X9,happ(s(fun(X15,X9),happ(s(fun(X10,fun(X15,X9)),X13),s(X10,X20))),s(X15,X22))) = s(X9,'const.pair.pair_CASE_2'(s('type.pair.prod'(X10,X15),'const.pair.,_2'(s(X10,X20),s(X15,X22))),s(fun(X10,fun(X15,X9)),X13)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.pair.pair_case_thm')).
fof(f164,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10,X15,X9] : ! [X16,X17] : s(X9,'const.pair.pair_CASE_2'(s('type.pair.prod'(X10,X15),X16),s(fun(X10,fun(X15,X9)),X17))) = s(X9,happ(s(fun(fun(X10,fun(X15,X9)),X9),happ(s(fun('type.pair.prod'(X10,X15),fun(fun(X10,fun(X15,X9)),X9)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X10,X15),X16))),s(fun(X10,fun(X15,X9)),X17)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','ARITY_0_2_const.pair.pair_CASE')).
fof(f163,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X5),s(X1,X2))) = s(X0,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X1),'const.semanticPrimitives.Match_1'(s(X1,X2))),s(X0,X3),s(X0,X4),s(fun(X1,X0),X5)))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X5),s(X1,X2))) = s(X0,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X1),'const.semanticPrimitives.Match_1'(s(X1,X2))),s(X0,X3),s(X0,X4),s(fun(X1,X0),X5)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X5),s(X1,X2))) = s(X0,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X1),'const.semanticPrimitives.Match_1'(s(X1,X2))),s(X0,X3),s(X0,X4),s(fun(X1,X0),X5)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10,X9] : ! [X14,X11,X12,X13] : s(X10,happ(s(fun(X9,X10),X13),s(X9,X14))) = s(X10,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_1'(s(X9,X14))),s(X10,X11),s(X10,X12),s(fun(X9,X10),X13)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_case_def|split|2')).
fof(f161,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X3) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X3) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X10] : ! [X11,X12,X13] : s(X10,X12) = s(X10,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0'),s(X10,X11),s(X10,X12),s(fun(X9,X10),X13)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_case_def|split|1')).
fof(f160,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X2) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X2) = s(X1,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0'),s(X1,X2),s(X1,X3),s(fun(X0,X1),X4)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X10] : ! [X11,X12,X13] : s(X10,X11) = s(X10,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.No_match_0'),s(X10,X11),s(X10,X12),s(fun(X9,X10),X13)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_case_def|split|0')).
fof(f159,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1',aHL_EXT)).
fof(f153,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X25,X26] : ((p(s(bool,X25)) => p(s(bool,X26))) => ((p(s(bool,X26)) => p(s(bool,X25))) => s(bool,X25) = s(bool,X26)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f156,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f148,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : ~s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X14] : ~s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_1'(s(X9,X14)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_distinct|split|1')).
fof(f147,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : ~s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X1)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X14] : ~s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_1'(s(X9,X14)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_distinct|split|2')).
fof(f144,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),X1) = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,sK2(X1,X0)))) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),X1) = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,sK2(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),X1) | ? [X2] : s('type.semanticPrimitives.match_result'(X0),X1) = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X2))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X0),X1) | s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),X1) | ? [X2] : s('type.semanticPrimitives.match_result'(X0),X1) = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_1'(s(X0,X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X23] : (s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(X9),X23) | s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X9),X23) | ? [X14] : s('type.semanticPrimitives.match_result'(X9),X23) = s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_1'(s(X9,X14))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_nchotomy')).
fof(f172,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f142,f141])).
fof(f142,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ! [X16] : s(X9,'const.combin.I_1'(s(X9,X16))) = s(X9,happ(s(fun(X9,X9),'const.combin.I_0'),s(X9,X16)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','ARITY_0_1_const.combin.I')).
fof(f141,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : ! [X20] : s(X9,X20) = s(X9,'const.combin.I_1'(s(X9,X20)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.combin.I_THM')).
fof(f140,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0')) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0')),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ~s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X0),'const.semanticPrimitives.No_match_0')),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : ~s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(X9),'const.semanticPrimitives.No_match_0')),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.semanticPrimitives.match_result_distinct|split|0')).
fof(f130,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f131,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f128,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1',aHL_BOOL_CASES)).
fof(f127,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.match_result_rel_def|split|2')).
fof(f124,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : p(s(bool,'const.relation.WF_1'(s(fun(X9,fun(X9,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.relation.WF_EMPTY_REL')).
fof(f117,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1',aHL_TRUTH)).
fof(f45643,plain,(
  ( ! [X133] : (p(s(bool,X133))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f45415,f44738])).
fof(f45415,plain,(
  ( ! [X133] : (s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | p(s(bool,X133))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f44738,f44358])).
fof(f44358,plain,(
  ( ! [X133] : (p(s(bool,X133)) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X133)),bool))),bool))) ) | $spl68),
  inference(subsumption_resolution,[],[f44290,f130])).
fof(f44290,plain,(
  ( ! [X133] : (p(s(bool,X133)) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X133)),bool))),bool)) | s(bool,'T_0') = s(bool,X133)) ) | $spl68),
  inference(superposition,[],[f43883,f42969])).
fof(f42969,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X8)),bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl68),
  inference(superposition,[],[f33847,f128])).
fof(f43883,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X3)),bool))) ) | $spl68),
  inference(subsumption_resolution,[],[f43700,f130])).
fof(f43700,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X3)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl68),
  inference(superposition,[],[f130,f42969])).
fof(f45642,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f45641,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633,f45635,f45637])).
fof(f45641,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f45417,f130])).
fof(f45417,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(condensation,[],[f45413])).
fof(f45413,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,X8) | p(s(bool,X7)) | p(s(bool,X8))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f44738,f43885])).
fof(f43885,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X7)),bool))) ) | $spl68),
  inference(subsumption_resolution,[],[f43704,f130])).
fof(f43704,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X7)),bool)) | s(bool,'T_0') = s(bool,X7)) ) | $spl68),
  inference(superposition,[],[f153,f42969])).
fof(f45640,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(global_subsumption,[],[f45418,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633,f45635,f45637])).
fof(f45418,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(trivial_inequality_removal,[],[f45412])).
fof(f45412,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X3))) ) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f44738,f43883])).
fof(f45638,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633,f45635,f45637])).
fof(f45636,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633,f45635])).
fof(f45634,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631,f45633])).
fof(f45632,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628,f45631])).
fof(f45629,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599,f45628])).
fof(f45600,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597,f45599])).
fof(f45598,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595,f45597])).
fof(f45596,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593,f45595])).
fof(f45594,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591,f45593])).
fof(f45592,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579,f45591])).
fof(f45587,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(subsumption_resolution,[],[f45278,f44738])).
fof(f45278,plain,(
  s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | ($spl0 | $spl3 | $spl12 | $spl68)),
  inference(backward_demodulation,[],[f44738,f43184])).
fof(f43184,plain,(
  s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | $spl68),
  inference(subsumption_resolution,[],[f43018,f171])).
fof(f43018,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | $spl68),
  inference(superposition,[],[f130,f33847])).
fof(f45580,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573,f45579])).
fof(f45574,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552,f45573])).
fof(f45569,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl84 | $spl125)),
  inference(global_subsumption,[],[f45568,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552])).
fof(f45568,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X7) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X7)) ) | ($spl0 | $spl3 | $spl12 | $spl84 | $spl125)),
  inference(subsumption_resolution,[],[f45128,f40557])).
fof(f40557,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s('type.semanticPrimitives.match_result'(bool),X7))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X7) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X7)) ) | ($spl0 | $spl84 | $spl125)),
  inference(superposition,[],[f39135,f144])).
fof(f39135,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0)))))) ) | ($spl0 | $spl84 | $spl125)),
  inference(forward_demodulation,[],[f38989,f141])).
fof(f38989,plain,(
  ( ! [X0] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0)))))) ) | ($spl0 | $spl84 | $spl125)),
  inference(unit_resulting_resolution,[],[f177,f38675,f10324])).
fof(f38675,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0)))))) ) | ($spl0 | $spl84 | $spl125)),
  inference(superposition,[],[f37236,f33966])).
fof(f33966,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0)))) ) | ($spl0 | $spl84)),
  inference(subsumption_resolution,[],[f33961,f33929])).
fof(f33961,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X0)),bool))) ) | $spl84),
  inference(backward_demodulation,[],[f33883,f7771])).
fof(f45128,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),X7))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X7) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X7)) ) | ($spl0 | $spl3 | $spl12 | $spl84 | $spl125)),
  inference(backward_demodulation,[],[f44738,f39078])).
fof(f39078,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),X7))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),X7) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),X7)) ) | ($spl0 | $spl84 | $spl125)),
  inference(superposition,[],[f38675,f144])).
fof(f45565,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl107)),
  inference(subsumption_resolution,[],[f45093,f44738])).
fof(f45093,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl0 | $spl3 | $spl12 | $spl107)),
  inference(backward_demodulation,[],[f44738,f37210])).
fof(f45560,plain,(
  $false | ($spl0 | $spl3 | $spl12 | $spl43)),
  inference(subsumption_resolution,[],[f45079,f44738])).
fof(f45079,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl0 | $spl3 | $spl12 | $spl43)),
  inference(backward_demodulation,[],[f44738,f33590])).
fof(f45559,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(global_subsumption,[],[f45558,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552])).
fof(f45558,plain,(
  ( ! [X212,X213] : (s(bool,X213) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X213) = s(bool,'const.combin.I_1'(s(bool,X212)))) ) | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(subsumption_resolution,[],[f45076,f44738])).
fof(f45076,plain,(
  ( ! [X212,X213] : (s(bool,'T_0') != s(bool,X212) | s(bool,X213) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X213) = s(bool,'const.combin.I_1'(s(bool,X212)))) ) | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(backward_demodulation,[],[f44738,f33166])).
fof(f33166,plain,(
  ( ! [X212,X213] : (s(bool,'F_0') != s(bool,X212) | s(bool,X213) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X213) = s(bool,'const.combin.I_1'(s(bool,X212)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f32769,f141])).
fof(f32769,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f31634,f117])).
fof(f31634,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f134,f10316])).
fof(f10316,plain,(
  ( ! [X83,X82] : (s(bool,X82) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f10315,f177])).
fof(f10315,plain,(
  ( ! [X83,X82] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X82) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9463,f141])).
fof(f9463,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X82) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1797,f205])).
fof(f1797,plain,(
  ( ! [X32] : (s(bool,X32) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X32)))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1378,f177])).
fof(f1378,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X32)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f755,f510])).
fof(f510,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f504,f131])).
fof(f45557,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(global_subsumption,[],[f45556,f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552])).
fof(f45556,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(condensation,[],[f45555])).
fof(f45555,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(subsumption_resolution,[],[f45075,f44738])).
fof(f45075,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl3 | $spl6 | $spl12)),
  inference(backward_demodulation,[],[f44738,f32769])).
fof(f45553,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546,f45552])).
fof(f45547,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528,f45546])).
fof(f45529,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524,f45528])).
fof(f45525,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518,f45524])).
fof(f45519,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511,f45518])).
fof(f45512,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509,f45511])).
fof(f45510,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506,f45509])).
fof(f45507,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503,f45506])).
fof(f45504,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451,f45503])).
fof(f45464,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44802,f44738])).
fof(f44802,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f1881])).
fof(f45452,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446,f45451])).
fof(f45448,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f44744,f117])).
fof(f44744,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl12)),
  inference(backward_demodulation,[],[f44738,f171])).
fof(f45447,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(global_subsumption,[],[f116,f115,f117,f124,f127,f128,f131,f130,f140,f141,f172,f144,f147,f148,f156,f153,f159,f160,f161,f163,f164,f165,f170,f167,f177,f183,f200,f205,f256,f501,f612,f730,f735,f779,f781,f900,f1256,f1298,f1303,f2047,f2195,f2200,f2368,f2369,f2370,f2917,f3317,f3322,f3490,f4502,f4509,f4691,f4906,f5758,f5765,f5816,f5818,f568,f6298,f6320,f6519,f6549,f6512,f6554,f6545,f7816,f7858,f7859,f7861,f7864,f8241,f8242,f8244,f8351,f8352,f8353,f8907,f8963,f9010,f9011,f9013,f9366,f9369,f9374,f10299,f10300,f10314,f10324,f10326,f9569,f9570,f9571,f10449,f12476,f12554,f14733,f14734,f16084,f16060,f16054,f16251,f14912,f14913,f16045,f16277,f16428,f16666,f15604,f16753,f8213,f5784,f5786,f10330,f8998,f8230,f7842,f2367,f162,f9014,f8245,f7862,f5021,f4996,f3647,f3592,f9015,f8906,f7910,f7863,f3642,f10327,f9012,f8243,f7860,f3089,f149,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f8044,f16731,f16665,f16644,f16272,f16229,f16228,f15603,f15602,f15598,f15597,f15596,f13006,f12807,f12806,f12804,f12803,f12802,f12782,f12707,f12656,f12347,f12117,f12095,f12094,f12092,f12091,f12090,f12070,f11884,f11883,f11881,f11880,f11879,f11859,f11684,f11682,f11681,f11638,f11475,f11348,f11330,f11207,f11159,f11158,f11157,f11142,f11029,f11028,f11021,f10434,f10433,f9373,f9372,f9368,f9367,f8450,f8227,f8211,f7812,f7810,f7779,f7777,f7774,f7773,f7771,f6678,f6625,f6612,f6509,f6508,f6507,f6506,f6505,f6504,f6467,f6378,f6376,f6375,f6374,f6297,f6255,f6254,f6253,f6252,f6251,f6250,f6055,f5820,f5179,f5178,f5177,f5176,f5175,f5170,f5115,f5076,f5075,f5074,f5073,f5072,f5071,f5070,f5059,f5006,f4989,f4958,f4956,f4915,f4914,f4910,f4909,f4907,f4887,f4878,f4877,f4876,f4875,f4874,f4873,f4862,f4819,f4818,f4105,f3991,f3762,f3757,f3756,f3684,f3669,f3637,f3635,f3631,f3624,f3622,f3578,f3577,f3573,f3572,f3570,f3469,f3468,f3250,f3075,f3062,f3029,f2997,f2996,f2992,f2991,f2989,f2434,f2432,f2431,f2397,f2396,f2395,f2394,f2393,f2390,f2389,f2387,f2386,f2385,f2384,f2381,f2380,f2378,f2234,f2231,f2223,f2211,f2209,f2206,f2205,f2203,f2044,f2024,f2023,f1943,f1916,f1914,f1913,f1881,f1880,f1875,f1858,f1850,f1798,f1796,f1764,f1757,f1629,f1617,f1515,f1383,f1382,f1381,f1376,f1360,f1358,f1357,f1356,f1355,f1354,f1348,f1239,f1238,f1237,f1236,f1232,f1204,f799,f798,f797,f792,f758,f757,f756,f755,f751,f691,f611,f610,f609,f608,f598,f596,f594,f578,f576,f546,f464,f463,f462,f440,f414,f389,f387,f386,f385,f316,f278,f207,f171,f152,f135,f134,f126,f125,f23419,f26185,f26186,f26187,f26189,f27586,f26209,f27515,f27488,f27671,f26383,f26384,f26385,f27477,f27724,f27727,f27870,f27962,f27963,f28094,f28103,f27062,f27066,f27067,f27068,f27069,f27070,f27071,f27072,f28192,f29050,f28932,f29076,f29087,f29088,f29089,f29090,f29091,f29092,f29017,f29190,f30503,f30504,f30505,f30513,f33583,f33535,f33560,f33561,f33562,f33563,f33564,f33565,f33566,f33567,f33568,f33569,f33570,f33571,f33606,f33607,f33608,f33609,f33610,f33611,f33612,f33613,f33614,f33615,f33616,f33617,f33618,f33619,f33620,f33855,f33856,f33623,f33862,f33866,f33868,f33873,f33874,f33875,f33876,f33877,f33878,f33879,f33880,f33881,f33648,f33649,f33650,f33651,f33895,f33896,f33897,f33901,f33903,f33906,f33909,f33910,f33682,f33929,f33930,f33933,f33810,f34230,f34229,f34228,f34306,f34714,f34716,f34908,f35367,f36011,f36087,f37141,f37144,f37145,f37146,f37147,f37148,f37149,f37150,f37151,f37217,f37219,f37220,f37221,f37226,f37227,f37229,f37230,f37231,f37232,f37207,f37206,f37204,f37243,f37863,f37865,f37868,f37869,f37873,f38408,f37903,f37906,f38426,f37948,f37955,f37957,f37958,f38128,f38130,f38480,f39386,f40021,f39388,f39389,f39390,f39391,f39392,f39393,f39396,f39397,f39400,f39401,f39402,f40023,f39405,f40028,f40046,f40049,f40051,f40053,f40055,f39435,f40057,f39437,f39438,f40059,f39440,f40061,f40063,f39655,f40064,f39660,f40069,f39695,f39698,f39700,f39701,f39702,f39703,f39704,f39706,f39712,f39713,f39714,f39717,f39726,f39727,f39728,f39729,f39730,f39731,f40083,f39738,f39739,f39740,f39741,f39743,f39744,f39746,f39747,f39748,f40122,f40127,f39804,f39805,f39808,f39809,f39810,f39828,f39829,f39849,f39854,f39855,f39856,f39857,f39858,f39860,f39866,f39867,f39868,f39871,f39880,f39881,f39882,f39883,f39884,f39885,f40132,f39892,f39893,f39894,f39895,f39897,f39898,f39900,f39901,f39902,f40171,f40176,f39958,f39961,f39962,f39963,f39981,f39982,f40834,f40836,f40837,f40839,f40841,f40842,f40843,f40846,f40847,f40852,f41219,f41243,f40877,f40880,f41259,f41275,f41288,f41290,f41292,f41314,f40931,f40933,f40934,f40936,f40938,f40939,f40940,f40943,f40944,f40949,f41327,f41351,f40974,f40977,f41368,f41378,f41022,f41028,f41030,f41031,f41033,f41424,f41885,f41886,f41887,f41890,f41891,f41895,f41896,f41899,f42083,f41954,f41957,f41966,f41969,f41970,f42137,f42138,f42153,f41981,f42157,f41986,f42222,f42237,f43233,f43235,f43236,f43238,f43294,f43295,f43296,f43298,f43300,f43303,f43325,f43326,f43923,f43925,f43926,f43928,f43934,f43936,f43937,f43939,f43970,f44153,f44539,f44540,f44541,f44542,f44543,f44544,f44545,f44546,f44547,f44548,f44728,f44730,f44731,f44732,f44733,f44734,f44738,f45446])).
fof(f44737,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f44729,f1298])).
fof(f44729,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.No_match_0'))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f44728,f200])).
fof(f502,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl4),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f43297,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),'const.min.@_1'(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),X36)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),X37)))) ) | $spl134),
  inference(cnf_transformation,[],[f43297_D])).
fof(f43297_D,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),'const.min.@_1'(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),X36)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),X37)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f43300,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),X52))) | ~p(s(bool,X53)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),'const.min.@_1'(s(fun(fun(X49,fun(X50,bool)),bool),happ(s(fun('type.pair.prod'(X49,X50),fun(fun(X49,fun(X50,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X49,X50),X51))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f43299,f141])).
fof(f43299,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X53)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),'const.min.@_1'(s(fun(fun(X49,fun(X50,bool)),bool),happ(s(fun('type.pair.prod'(X49,X50),fun(fun(X49,fun(X50,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X49,X50),X51)))))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),X52)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f43262,f134])).
fof(f43262,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,X53)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),'const.min.@_1'(s(fun(fun(X49,fun(X50,bool)),bool),happ(s(fun('type.pair.prod'(X49,X50),fun(fun(X49,fun(X50,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X49,X50),X51)))))))) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),X51),s(fun(X49,fun(X50,bool)),X52)))))) ) | $spl0),
  inference(superposition,[],[f4996,f755])).
fof(f43298,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,X43)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X40),X41),s(fun(X39,fun(X40,bool)),'const.min.@_1'(s(fun(fun(X39,fun(X40,bool)),bool),happ(s(fun('type.pair.prod'(X39,X40),fun(fun(X39,fun(X40,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X39,X40),X41)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X40),X41),s(fun(X39,fun(X40,bool)),X42)))) ) | $spl0),
  inference(subsumption_resolution,[],[f43260,f134])).
fof(f43260,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,X43)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X40),X41),s(fun(X39,fun(X40,bool)),'const.min.@_1'(s(fun(fun(X39,fun(X40,bool)),bool),happ(s(fun('type.pair.prod'(X39,X40),fun(fun(X39,fun(X40,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X39,X40),X41)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X39,X40),X41),s(fun(X39,fun(X40,bool)),X42))) | s(bool,'F_0') = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f4996,f612])).
fof(f43296,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,X38)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),'const.min.@_1'(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),X36)))))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),X37)))) ) | $spl0),
  inference(subsumption_resolution,[],[f43259,f134])).
fof(f43259,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,X38)) | p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),'const.min.@_1'(s(fun(fun(X34,fun(X35,bool)),bool),happ(s(fun('type.pair.prod'(X34,X35),fun(fun(X34,fun(X35,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X34,X35),X36)))))))) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X35),X36),s(fun(X34,fun(X35,bool)),X37)))) ) | $spl0),
  inference(superposition,[],[f4996,f612])).
fof(f6544,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f6544_D])).
fof(f6544_D,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f42832,plain,(
  ( ! [X92] : (s(bool,X92) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X92)) ) | $spl130),
  inference(cnf_transformation,[],[f42832_D])).
fof(f42832_D,plain,(
  ( ! [X92] : (s(bool,X92) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X92)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f42878,plain,(
  ( ! [X271,X270] : (s(bool,X271) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X271) | s(bool,'F_0') = s(bool,X270)) ) | ($spl0 | $spl18 | $spl69)),
  inference(subsumption_resolution,[],[f42527,f387])).
fof(f42527,plain,(
  ( ! [X271,X270] : (s(bool,X271) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X271) | s(bool,'F_0') = s(bool,X270) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X270)))) ) | ($spl0 | $spl18 | $spl69)),
  inference(superposition,[],[f6809,f37436])).
fof(f37436,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl69)),
  inference(trivial_inequality_removal,[],[f37408])).
fof(f37408,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X4) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X4)) ) | ($spl0 | $spl69)),
  inference(superposition,[],[f33846,f612])).
fof(f33846,plain,(
  s(bool,'F_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | $spl69),
  inference(cnf_transformation,[],[f33846_D])).
fof(f33846_D,plain,(
  s(bool,'F_0') != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f6809,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f6741,f177])).
fof(f6741,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl0 | $spl18)),
  inference(superposition,[],[f6553,f612])).
fof(f42861,plain,(
  ( ! [X175,X174] : (s(bool,'F_0') = s(bool,X174) | s(bool,X175) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X175)) ) | ($spl0 | $spl69)),
  inference(forward_demodulation,[],[f42860,f141])).
fof(f42860,plain,(
  ( ! [X175,X174] : (s(bool,X175) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X174) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X175)) ) | ($spl0 | $spl69)),
  inference(subsumption_resolution,[],[f42475,f11330])).
fof(f42475,plain,(
  ( ! [X175,X174] : (s(bool,X175) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X174) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X175) | s(bool,'F_0') = s(bool,X174)) ) | ($spl0 | $spl69)),
  inference(superposition,[],[f11684,f37436])).
fof(f42712,plain,(
  ( ! [X128,X129] : (s(bool,X129) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X129)) ) | ($spl0 | $spl69)),
  inference(duplicate_literal_removal,[],[f42447])).
fof(f42447,plain,(
  ( ! [X128,X129] : (s(bool,X129) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X128)) ) | ($spl0 | $spl69)),
  inference(superposition,[],[f6320,f37436])).
fof(f42838,plain,(
  ( ! [X121,X120] : (s(bool,'T_0') = s(bool,X121) | s(bool,X121) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X120)) ) | ($spl0 | $spl69)),
  inference(forward_demodulation,[],[f42713,f141])).
fof(f42713,plain,(
  ( ! [X121,X120] : (s(bool,X121) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X121) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X120)) ) | ($spl0 | $spl69)),
  inference(duplicate_literal_removal,[],[f42443])).
fof(f42443,plain,(
  ( ! [X121,X120] : (s(bool,X121) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X121) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X120) | s(bool,'F_0') = s(bool,X120)) ) | ($spl0 | $spl69)),
  inference(superposition,[],[f4818,f37436])).
fof(f42834,plain,(
  ( ! [X112] : (s(bool,X112) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X112)))) ) | $spl132),
  inference(cnf_transformation,[],[f42834_D])).
fof(f42834_D,plain,(
  ( ! [X112] : (s(bool,X112) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X112)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f42715,plain,(
  ( ! [X111,X112] : (s(bool,X112) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X112))) | s(bool,'F_0') = s(bool,X111)) ) | ($spl0 | $spl6 | $spl69)),
  inference(duplicate_literal_removal,[],[f42436])).
fof(f42436,plain,(
  ( ! [X111,X112] : (s(bool,X112) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X112))) | s(bool,'F_0') = s(bool,X111) | s(bool,'F_0') = s(bool,X111)) ) | ($spl0 | $spl6 | $spl69)),
  inference(superposition,[],[f3392,f37436])).
fof(f3392,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3341,f177])).
fof(f3341,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3233,f612])).
fof(f3233,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f3207,f510])).
fof(f3207,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'const.combin.I_1'(s(bool,X37))) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3138,f141])).
fof(f3138,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | s(bool,X22) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2870,f141])).
fof(f2870,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X8))) | s(bool,X8) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2681,f171])).
fof(f2681,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X8))) | s(bool,X8) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f130,f1797])).
fof(f42833,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') = s(bool,X107) | s(bool,X107) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X106)) ) | ($spl0 | $spl6 | $spl69)),
  inference(forward_demodulation,[],[f42716,f510])).
fof(f42716,plain,(
  ( ! [X107,X106] : (s(bool,X107) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X107) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X106)) ) | ($spl0 | $spl6 | $spl69)),
  inference(duplicate_literal_removal,[],[f42433])).
fof(f42433,plain,(
  ( ! [X107,X106] : (s(bool,X107) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X107) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X106) | s(bool,'F_0') = s(bool,X106)) ) | ($spl0 | $spl6 | $spl69)),
  inference(superposition,[],[f3224,f37436])).
fof(f3224,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3176,f177])).
fof(f3176,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3138,f612])).
fof(f42831,plain,(
  ( ! [X91] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X91)))) ) | $spl128),
  inference(cnf_transformation,[],[f42831_D])).
fof(f42831_D,plain,(
  ( ! [X91] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X91)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f42830,plain,(
  ( ! [X92,X91] : (s(bool,X92) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X91))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl0 | $spl69)),
  inference(subsumption_resolution,[],[f42427,f1858])).
fof(f42427,plain,(
  ( ! [X92,X91] : (s(bool,X92) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X91))) | s(bool,'T_0') = s(bool,X92) | s(bool,'F_0') = s(bool,X91)) ) | ($spl0 | $spl69)),
  inference(superposition,[],[f2234,f37436])).
fof(f41626,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl126),
  inference(cnf_transformation,[],[f41626_D])).
fof(f41626_D,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f41605,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,'F_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl112)),
  inference(duplicate_literal_removal,[],[f41512])).
fof(f41512,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,'F_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl0 | $spl112)),
  inference(superposition,[],[f37213,f11021])).
fof(f37234,plain,(
  ( ! [X88,X87] : (s(fun('type.semanticPrimitives.match_result'(bool),bool),X87) = s(fun('type.semanticPrimitives.match_result'(bool),bool),X88) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s(bool,'T_0') = s(bool,sK2(sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)),bool)) | s(bool,'T_0') = s(bool,g(s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)))))) ) | $spl122),
  inference(cnf_transformation,[],[f37234_D])).
fof(f37234_D,plain,(
  ( ! [X88,X87] : (s(fun('type.semanticPrimitives.match_result'(bool),bool),X87) = s(fun('type.semanticPrimitives.match_result'(bool),bool),X88) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s(bool,'T_0') = s(bool,sK2(sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)),bool)) | s(bool,'T_0') = s(bool,g(s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f37236,plain,(
  s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))))) | $spl125),
  inference(cnf_transformation,[],[f37236_D])).
fof(f37236_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f37233,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))))) | s(fun('type.semanticPrimitives.match_result'(bool),bool),X87) = s(fun('type.semanticPrimitives.match_result'(bool),bool),X88) | s(bool,'T_0') = s(bool,g(s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))))) | s(bool,'T_0') = s(bool,sK2(sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)))) ) | $spl84),
  inference(forward_demodulation,[],[f37171,f33883])).
fof(f37171,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,f(s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))))) | s(fun('type.semanticPrimitives.match_result'(bool),bool),X87) = s(fun('type.semanticPrimitives.match_result'(bool),bool),X88) | s(bool,'T_0') = s(bool,g(s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))))) | s(bool,'T_0') = s(bool,sK2(sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.No_match_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'(bool),sK3(X87,X88,bool,'type.semanticPrimitives.match_result'(bool)))) )),
  inference(superposition,[],[f2387,f2203])).
fof(f37212,plain,(
  ( ! [X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'F_0') = s(bool,sK3(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X17,X18,bool,bool))))) ) | $spl110),
  inference(cnf_transformation,[],[f37212_D])).
fof(f37212_D,plain,(
  ( ! [X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'F_0') = s(bool,sK3(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X17,X18,bool,bool))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f37213,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,f(s(bool,X19))) | s(bool,'F_0') = s(bool,X19)) ) | $spl112),
  inference(cnf_transformation,[],[f37213_D])).
fof(f37213_D,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,f(s(bool,X19))) | s(bool,'F_0') = s(bool,X19)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f37232,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') != s(bool,f(s(bool,X86))) | s(fun(bool,bool),X84) = s(fun(bool,bool),X85) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X84,X85,bool,bool)))) | s(bool,'F_0') = s(bool,sK3(X84,X85,bool,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f37170,f141])).
fof(f37170,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,f(s(bool,X86))) | s(fun(bool,bool),X84) = s(fun(bool,bool),X85) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X84,X85,bool,bool)))) | s(bool,X86) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK3(X84,X85,bool,bool))) ) | $spl0),
  inference(superposition,[],[f2387,f11021])).
fof(f37231,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') = s(bool,sK3(X81,X82,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X83))) | s(fun(bool,bool),X81) = s(fun(bool,bool),X82) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X81,X82,bool,bool)))) | s(bool,'F_0') = s(bool,X83)) ) | $spl0),
  inference(forward_demodulation,[],[f37169,f141])).
fof(f37169,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,f(s(bool,X83))) | s(fun(bool,bool),X81) = s(fun(bool,bool),X82) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X81,X82,bool,bool)))) | s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,sK3(X81,X82,bool,bool)) | s(bool,'F_0') = s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f2387,f11021])).
fof(f33590,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl43),
  inference(cnf_transformation,[],[f33590_D])).
fof(f33590_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f37211,plain,(
  ( ! [X14,X15,X16] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'F_0') = s(bool,X16) | s(bool,X16) = s(bool,sK3(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X14,X15,bool,bool))))) ) | $spl108),
  inference(cnf_transformation,[],[f37211_D])).
fof(f37211_D,plain,(
  ( ! [X14,X15,X16] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'F_0') = s(bool,X16) | s(bool,X16) = s(bool,sK3(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X14,X15,bool,bool))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f37230,plain,(
  ( ! [X80,X78,X79] : (s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X78) = s(fun(bool,bool),X79) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X78,X79,bool,bool)))) | s(bool,X80) = s(bool,sK3(X78,X79,bool,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f37168,f141])).
fof(f37168,plain,(
  ( ! [X80,X78,X79] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X78) = s(fun(bool,bool),X79) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X78,X79,bool,bool)))) | s(bool,X80) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X80) = s(bool,sK3(X78,X79,bool,bool))) ) | $spl0),
  inference(superposition,[],[f2387,f11021])).
fof(f37229,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X75) = s(fun(bool,bool),X76) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X75,X76,bool,bool)))) | s(bool,X77) = s(bool,sK3(X75,X76,bool,bool)) | s(bool,'F_0') = s(bool,X77)) ) | $spl0),
  inference(forward_demodulation,[],[f37167,f141])).
fof(f37167,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(fun(bool,bool),X75) = s(fun(bool,bool),X76) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X75,X76,bool,bool)))) | s(bool,X77) = s(bool,sK3(X75,X76,bool,bool)) | s(bool,'F_0') = s(bool,X77)) ) | $spl0),
  inference(superposition,[],[f2387,f11021])).
fof(f37210,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl107),
  inference(cnf_transformation,[],[f37210_D])).
fof(f37210_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f37208,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) | $spl104),
  inference(cnf_transformation,[],[f37208_D])).
fof(f37208_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f37227,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X66) = s(fun(bool,bool),X67) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X66,X67,bool,bool)))) | s(bool,X68) = s(bool,sK3(X66,X67,bool,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f37164,f141])).
fof(f37164,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X66) = s(fun(bool,bool),X67) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X66,X67,bool,bool)))) | s(bool,X68) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X68) = s(bool,sK3(X66,X67,bool,bool))) ) | $spl0),
  inference(superposition,[],[f2387,f10324])).
fof(f37226,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X63) = s(fun(bool,bool),X64) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X63,X64,bool,bool)))) | s(bool,X65) = s(bool,sK3(X63,X64,bool,bool)) | s(bool,'T_0') = s(bool,X65)) ) | $spl0),
  inference(forward_demodulation,[],[f37163,f141])).
fof(f37163,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) | s(fun(bool,bool),X63) = s(fun(bool,bool),X64) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X63,X64,bool,bool)))) | s(bool,X65) = s(bool,sK3(X63,X64,bool,bool)) | s(bool,'T_0') = s(bool,X65)) ) | $spl0),
  inference(superposition,[],[f2387,f10324])).
fof(f37225,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X60) = s(fun(bool,bool),X61) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X60,X61,bool,bool)))) | s(bool,X62) = s(bool,sK3(X60,X61,bool,bool))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f37162,f510])).
fof(f37162,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X60) = s(fun(bool,bool),X61) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X60,X61,bool,bool)))) | s(bool,X62) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X62) = s(bool,sK3(X60,X61,bool,bool))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2387,f10316])).
fof(f37223,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X51,X52,bool,bool)))) | s(bool,X53) = s(bool,sK3(X51,X52,bool,bool)) | s(bool,'T_0') = s(bool,X53)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f37159,f510])).
fof(f37159,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X51,X52,bool,bool)))) | s(bool,X53) = s(bool,sK3(X51,X52,bool,bool)) | s(bool,'T_0') = s(bool,X53)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2387,f10316])).
fof(f37222,plain,(
  ( ! [X50,X49] : (s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X49,X50,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X49,X50,bool,bool))))) ) | $spl120),
  inference(cnf_transformation,[],[f37222_D])).
fof(f37222_D,plain,(
  ( ! [X50,X49] : (s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X49,X50,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X49,X50,bool,bool))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f37221,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X49,X50,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X49,X50,bool,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f37158,f141])).
fof(f37158,plain,(
  ( ! [X50,X49] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X49,X50,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X49,X50,bool,bool))))) ) | $spl0),
  inference(superposition,[],[f2387,f8450])).
fof(f37214,plain,(
  ( ! [X26,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X26,X27,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X26,X27,bool,bool))))) ) | $spl114),
  inference(cnf_transformation,[],[f37214_D])).
fof(f37214_D,plain,(
  ( ! [X26,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X26,X27,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X26,X27,bool,bool))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f37220,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X45) = s(fun(bool,bool),X46) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X45,X46,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X45,X46,bool,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f37156,f141])).
fof(f37156,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) | s(fun(bool,bool),X45) = s(fun(bool,bool),X46) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X45,X46,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X45,X46,bool,bool))))) ) | $spl0),
  inference(superposition,[],[f2387,f3991])).
fof(f37219,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X43,X44,bool,bool)))) | s(bool,'F_0') = s(bool,sK3(X43,X44,bool,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f37155,f141])).
fof(f37155,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X43,X44,bool,bool)))) | s(bool,'F_0') = s(bool,sK3(X43,X44,bool,bool))) ) | $spl0),
  inference(superposition,[],[f2387,f3762])).
fof(f37218,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X38,X39,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X38,X39,bool,bool))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f37153,f510])).
fof(f37153,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X38,X39,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X38,X39,bool,bool))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2387,f1797])).
fof(f37217,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X35,X36,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X35,X36,bool,bool))))) ) | $spl0),
  inference(forward_demodulation,[],[f37152,f183])).
fof(f37152,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,f(s(bool,'const.relation.WF_1'(s(fun(X37,fun(X37,bool)),'const.relation.EMPTY_REL_0'))))) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X35,X36,bool,bool)))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X35,X36,bool,bool))))) ) | $spl0),
  inference(superposition,[],[f2387,f1796])).
fof(f37216,plain,(
  ( ! [X33,X34,X32] : (s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,X34) = s(bool,sK3(X32,X33,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X32,X33,bool,bool))))) ) | $spl118),
  inference(cnf_transformation,[],[f37216_D])).
fof(f37216_D,plain,(
  ( ! [X33,X34,X32] : (s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,X34) = s(bool,sK3(X32,X33,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X32,X33,bool,bool))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f37151,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X32,X33,bool,bool)))) | s(bool,X34) = s(bool,sK3(X32,X33,bool,bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34)))) ) | $spl0),
  inference(superposition,[],[f2387,f755])).
fof(f37215,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,f(s(bool,X31))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl116),
  inference(cnf_transformation,[],[f37215_D])).
fof(f37215_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,f(s(bool,X31))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f37150,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,f(s(bool,X31))) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X29,X30,bool,bool)))) | s(bool,'F_0') = s(bool,sK3(X29,X30,bool,bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl0),
  inference(superposition,[],[f2387,f755])).
fof(f37149,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,f(s(bool,X28))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X26,X27,bool,bool)))) | s(bool,'F_0') = s(bool,X28) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,sK3(X26,X27,bool,bool))))) ) | $spl0),
  inference(superposition,[],[f2387,f755])).
fof(f37148,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X23,X24,bool,bool)))) | s(bool,X25) = s(bool,sK3(X23,X24,bool,bool)) | s(bool,'F_0') = s(bool,X25)) ) | $spl0),
  inference(superposition,[],[f2387,f612])).
fof(f37147,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,f(s(bool,X22))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X20,X21,bool,bool)))) | s(bool,'F_0') = s(bool,sK3(X20,X21,bool,bool)) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f2387,f612])).
fof(f37146,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,f(s(bool,X19))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X17,X18,bool,bool)))) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,sK3(X17,X18,bool,bool))) ) | $spl0),
  inference(superposition,[],[f2387,f612])).
fof(f37145,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X14,X15,bool,bool)))) | s(bool,X16) = s(bool,sK3(X14,X15,bool,bool)) | s(bool,'F_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f2387,f612])).
fof(f37144,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2387,f205])).
fof(f37141,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2387,f205])).
fof(f16176,plain,(
  ( ! [X57] : (s(bool,'T_0') = s(bool,X57)) ) | $spl24),
  inference(cnf_transformation,[],[f16176_D])).
fof(f16176_D,plain,(
  ( ! [X57] : (s(bool,'T_0') = s(bool,X57)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f37027,plain,(
  ( ! [X99] : (s(bool,X99) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X99)) ) | $spl100),
  inference(cnf_transformation,[],[f37027_D])).
fof(f37027_D,plain,(
  ( ! [X99] : (s(bool,X99) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X99)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f37079,plain,(
  ( ! [X237,X238] : (s(bool,X238) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X237)) ) | ($spl0 | $spl68)),
  inference(subsumption_resolution,[],[f36734,f2047])).
fof(f36734,plain,(
  ( ! [X237,X238] : (s(bool,X238) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X237) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X237)))) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f1916,f34164])).
fof(f34164,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl68),
  inference(subsumption_resolution,[],[f34000,f177])).
fof(f34000,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl68),
  inference(superposition,[],[f33847,f205])).
fof(f37030,plain,(
  ( ! [X103] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X103)))) ) | $spl102),
  inference(cnf_transformation,[],[f37030_D])).
fof(f37030_D,plain,(
  ( ! [X103] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X103)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f37078,plain,(
  ( ! [X229,X230] : (s(bool,X230) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X230) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X229)))) ) | ($spl0 | $spl6 | $spl68)),
  inference(subsumption_resolution,[],[f36729,f3392])).
fof(f36729,plain,(
  ( ! [X229,X230] : (s(bool,X230) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X229) = s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X229)))) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f1617,f34164])).
fof(f37064,plain,(
  ( ! [X163,X162] : (s(bool,'F_0') = s(bool,X163) | s(bool,X163) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X162)) ) | ($spl0 | $spl68)),
  inference(forward_demodulation,[],[f36909,f141])).
fof(f36909,plain,(
  ( ! [X163,X162] : (s(bool,X163) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X163) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X162)) ) | ($spl0 | $spl68)),
  inference(duplicate_literal_removal,[],[f36692])).
fof(f36692,plain,(
  ( ! [X163,X162] : (s(bool,X163) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X163) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f11684,f34164])).
fof(f36917,plain,(
  ( ! [X116,X117] : (s(bool,X117) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X116)) ) | ($spl0 | $spl68)),
  inference(duplicate_literal_removal,[],[f36664])).
fof(f36664,plain,(
  ( ! [X116,X117] : (s(bool,X117) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f6320,f34164])).
fof(f37036,plain,(
  ( ! [X112,X113] : (s(bool,X113) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X112)) ) | ($spl0 | $spl18 | $spl68)),
  inference(subsumption_resolution,[],[f36662,f6553])).
fof(f36662,plain,(
  ( ! [X112,X113] : (s(bool,X113) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X112) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X112)) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f4818,f34164])).
fof(f37029,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X103))) | s(bool,'F_0') = s(bool,X104)) ) | ($spl0 | $spl6 | $spl68)),
  inference(subsumption_resolution,[],[f36655,f3233])).
fof(f36655,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X103))) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X103)) ) | ($spl0 | $spl6 | $spl68)),
  inference(superposition,[],[f3392,f34164])).
fof(f37026,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) ) | ($spl0 | $spl6 | $spl16 | $spl68)),
  inference(subsumption_resolution,[],[f36652,f6552])).
fof(f36652,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X98) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) ) | ($spl0 | $spl6 | $spl68)),
  inference(superposition,[],[f3224,f34164])).
fof(f37022,plain,(
  ( ! [X90] : (s(bool,X90) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) ) | $spl98),
  inference(cnf_transformation,[],[f37022_D])).
fof(f37022_D,plain,(
  ( ! [X90] : (s(bool,X90) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f36918,plain,(
  ( ! [X90,X89] : (s(bool,X90) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X90))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl0 | $spl68)),
  inference(duplicate_literal_removal,[],[f36646])).
fof(f36646,plain,(
  ( ! [X90,X89] : (s(bool,X90) != s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X90))) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl0 | $spl68)),
  inference(superposition,[],[f2234,f34164])).
fof(f33911,plain,(
  ( ! [X112] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X112)),bool)))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X112)))) ) | $spl96),
  inference(cnf_transformation,[],[f33911_D])).
fof(f33911_D,plain,(
  ( ! [X112] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X112)),bool)))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X112)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f33847,plain,(
  s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | $spl68),
  inference(cnf_transformation,[],[f33847_D])).
fof(f33847_D,plain,(
  s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f33682,plain,(
  ( ! [X112] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X112)),bool)))) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X112)))) ) | $spl0),
  inference(superposition,[],[f7810,f7810])).
fof(f33851,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl74),
  inference(cnf_transformation,[],[f33851_D])).
fof(f33851_D,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f33883,plain,(
  s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | $spl84),
  inference(cnf_transformation,[],[f33883_D])).
fof(f33883_D,plain,(
  s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f33910,plain,(
  ( ! [X111] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X111)),bool)) | s(bool,'T_0') = s(bool,X111)) ) | $spl0),
  inference(forward_demodulation,[],[f33681,f141])).
fof(f33681,plain,(
  ( ! [X111] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X111)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X111)) ) | $spl0),
  inference(superposition,[],[f7810,f8211])).
fof(f33909,plain,(
  ( ! [X110] : (s(bool,'T_0') = s(bool,X110) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X110)),bool))) ) | $spl0),
  inference(forward_demodulation,[],[f33908,f141])).
fof(f33908,plain,(
  ( ! [X110] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X110)),bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X110)))) ) | $spl0),
  inference(forward_demodulation,[],[f33680,f141])).
fof(f33680,plain,(
  ( ! [X110] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X110)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X110)))) ) | $spl0),
  inference(superposition,[],[f7810,f8450])).
fof(f33907,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X108)),bool)) | s(bool,'F_0') = s(bool,X109) | s(bool,X108) = s(bool,X109)) ) | $spl94),
  inference(cnf_transformation,[],[f33907_D])).
fof(f33907_D,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X108)),bool)) | s(bool,'F_0') = s(bool,X109) | s(bool,X108) = s(bool,X109)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f33906,plain,(
  ( ! [X109,X108] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X108)),bool)) | s(bool,X108) = s(bool,X109) | s(bool,'F_0') = s(bool,X109)) ) | $spl0),
  inference(forward_demodulation,[],[f33679,f141])).
fof(f33679,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X108)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(bool,X108) = s(bool,X109) | s(bool,'F_0') = s(bool,X109)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33845,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl66),
  inference(cnf_transformation,[],[f33845_D])).
fof(f33845_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f33903,plain,(
  ( ! [X102] : (s(bool,'T_0') = s(bool,X102) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X102))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool))) ) | $spl0),
  inference(forward_demodulation,[],[f33902,f141])).
fof(f33902,plain,(
  ( ! [X102] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X102))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X102) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33674,f141])).
fof(f33674,plain,(
  ( ! [X102] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X102))))) | s(bool,X102) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(superposition,[],[f7810,f3991])).
fof(f33901,plain,(
  ( ! [X101,X100] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X100))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X101)),bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X100)) ) | $spl0),
  inference(forward_demodulation,[],[f33673,f141])).
fof(f33673,plain,(
  ( ! [X101,X100] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X101)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X100))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X100)) ) | $spl0),
  inference(superposition,[],[f7810,f2195])).
fof(f33900,plain,(
  ( ! [X99] : (s(bool,'T_0') = s(bool,X99) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X99))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33899,f510])).
fof(f33899,plain,(
  ( ! [X99] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X99))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X99) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33672,f141])).
fof(f33672,plain,(
  ( ! [X99] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X99))))) | s(bool,X99) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f7810,f1797])).
fof(f33898,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X97))) | s(bool,X97) = s(bool,'const.relation.WF_1'(s(fun(X98,fun(X98,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f33898_D])).
fof(f33898_D,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X97))) | s(bool,X97) = s(bool,'const.relation.WF_1'(s(fun(X98,fun(X98,bool)),'const.relation.EMPTY_REL_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f33897,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X97))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X97) = s(bool,'const.relation.WF_1'(s(fun(X98,fun(X98,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33671,f141])).
fof(f33671,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X97))))) | s(bool,X97) = s(bool,'const.relation.WF_1'(s(fun(X98,fun(X98,bool)),'const.relation.EMPTY_REL_0')))) ) | $spl0),
  inference(superposition,[],[f7810,f1796])).
fof(f33896,plain,(
  ( ! [X96] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X96))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X96)) ) | $spl0),
  inference(forward_demodulation,[],[f33670,f141])).
fof(f33670,plain,(
  ( ! [X96] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,X96)) ) | $spl0),
  inference(superposition,[],[f7810,f1764])).
fof(f33852,plain,(
  ( ! [X10,X9] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) = s(bool,X10)) ) | $spl76),
  inference(cnf_transformation,[],[f33852_D])).
fof(f33852_D,plain,(
  ( ! [X10,X9] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) = s(bool,X10)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f33895,plain,(
  ( ! [X94,X95] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X94))) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) ) | $spl0),
  inference(forward_demodulation,[],[f33669,f141])).
fof(f33669,plain,(
  ( ! [X94,X95] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X94))))) | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) ) | $spl0),
  inference(superposition,[],[f7810,f755])).
fof(f33886,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X68)),bool)) | s(bool,'T_0') = s(bool,X68)) ) | $spl90),
  inference(cnf_transformation,[],[f33886_D])).
fof(f33886_D,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X68)),bool)) | s(bool,'T_0') = s(bool,X68)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f33651,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X68)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f7810,f128])).
fof(f33884,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') = s(bool,X61) | s(bool,X60) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X60)),bool)) | s(bool,X60) = s(bool,X61)) ) | $spl86),
  inference(cnf_transformation,[],[f33884_D])).
fof(f33884_D,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') = s(bool,X61) | s(bool,X60) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X60)),bool)) | s(bool,X60) = s(bool,X61)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f33650,plain,(
  ( ! [X66,X67] : (s(bool,X66) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X66)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,X66) = s(bool,X67) | s(bool,'F_0') = s(bool,X67)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33649,plain,(
  ( ! [X64,X65] : (s(bool,X64) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X64)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33885,plain,(
  ( ! [X62,X63] : (s(bool,X62) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X62)),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X63))) | s(bool,X62) = s(bool,X63)) ) | $spl88),
  inference(cnf_transformation,[],[f33885_D])).
fof(f33885_D,plain,(
  ( ! [X62,X63] : (s(bool,X62) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X62)),bool)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X63))) | s(bool,X62) = s(bool,X63)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f33648,plain,(
  ( ! [X62,X63] : (s(bool,X62) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X62)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,X62) = s(bool,X63) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X63)))) ) | $spl0),
  inference(superposition,[],[f7810,f755])).
fof(f33881,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') = s(bool,X61) | s(bool,X60) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X60)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,X60) = s(bool,X61)) ) | $spl0),
  inference(forward_demodulation,[],[f33647,f141])).
fof(f33647,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X60)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'F_0'))) | s(bool,X61) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X60) = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33853,plain,(
  ( ! [X11] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) ) | $spl78),
  inference(cnf_transformation,[],[f33853_D])).
fof(f33853_D,plain,(
  ( ! [X11] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f33854,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X12)),bool)) | s(bool,'F_0') = s(bool,X12)) ) | $spl80),
  inference(cnf_transformation,[],[f33854_D])).
fof(f33854_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X12)),bool)) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f33880,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X55)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X54))) | s(bool,'F_0') = s(bool,X54)) ) | $spl0),
  inference(forward_demodulation,[],[f33642,f141])).
fof(f33642,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X55)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X54))) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X54)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33879,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X53)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X53)) ) | $spl0),
  inference(forward_demodulation,[],[f33641,f141])).
fof(f33641,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X53)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X52))) | s(bool,X52) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X53)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33878,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X50))) | s(bool,X50) = s(bool,X51)) ) | $spl0),
  inference(forward_demodulation,[],[f33640,f141])).
fof(f33640,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X50))) | s(bool,X51) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X50) = s(bool,X51)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33877,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X48))) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl0),
  inference(forward_demodulation,[],[f33639,f141])).
fof(f33639,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0')))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X48))) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl0),
  inference(superposition,[],[f7810,f11021])).
fof(f33876,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X47)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X46)) ) | $spl0),
  inference(forward_demodulation,[],[f33638,f141])).
fof(f33638,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X47)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X46))) | s(bool,X47) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X46)) ) | $spl0),
  inference(superposition,[],[f7810,f10324])).
fof(f33875,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X45)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,X45)) ) | $spl0),
  inference(forward_demodulation,[],[f33637,f141])).
fof(f33637,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X45)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X44))) | s(bool,X44) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X45)) ) | $spl0),
  inference(superposition,[],[f7810,f10324])).
fof(f33850,plain,(
  s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | $spl72),
  inference(cnf_transformation,[],[f33850_D])).
fof(f33850_D,plain,(
  s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f33848,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl70),
  inference(cnf_transformation,[],[f33848_D])).
fof(f33848_D,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f33874,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X42))) | s(bool,X42) = s(bool,X43)) ) | $spl0),
  inference(forward_demodulation,[],[f33636,f141])).
fof(f33636,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X42))) | s(bool,X43) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X42) = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f7810,f10324])).
fof(f33873,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X40))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl0),
  inference(forward_demodulation,[],[f33635,f141])).
fof(f33635,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'T_0')))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X40))) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl0),
  inference(superposition,[],[f7810,f10324])).
fof(f33872,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,X39) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X38))) | s(bool,X38) = s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33634,f510])).
fof(f33634,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X38))) | s(bool,X39) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X38) = s(bool,X39)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f7810,f10316])).
fof(f33871,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X37)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33633,f510])).
fof(f33633,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X37)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X36))) | s(bool,X37) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f7810,f10316])).
fof(f33870,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X35)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X34))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33632,f510])).
fof(f33632,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X35)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X34))) | s(bool,X34) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f7810,f10316])).
fof(f33869,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X32))) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33631,f510])).
fof(f33631,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X32))) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f7810,f10316])).
fof(f33868,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X31)))) ) | $spl0),
  inference(forward_demodulation,[],[f33867,f141])).
fof(f33867,plain,(
  ( ! [X31] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl0),
  inference(forward_demodulation,[],[f33630,f141])).
fof(f33630,plain,(
  ( ! [X31] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0')))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X31)))) ) | $spl0),
  inference(superposition,[],[f7810,f8450])).
fof(f33866,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | $spl0),
  inference(forward_demodulation,[],[f33629,f141])).
fof(f33629,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,'F_0')))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | $spl0),
  inference(superposition,[],[f7810,f8211])).
fof(f33862,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X27)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl0),
  inference(forward_demodulation,[],[f33626,f141])).
fof(f33626,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'const.combin.I_1'(s(bool,X27)))),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X26))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl0),
  inference(superposition,[],[f7810,f2195])).
fof(f33857,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X21))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X21) = s(bool,X22)) ) | $spl82),
  inference(cnf_transformation,[],[f33857_D])).
fof(f33857_D,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X21))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X21) = s(bool,X22)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f33623,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X21))) | s(bool,X21) = s(bool,X22) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) | $spl0),
  inference(superposition,[],[f7810,f755])).
fof(f33856,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') = s(bool,X20) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X20)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X19))) | s(bool,'F_0') = s(bool,X19)) ) | $spl0),
  inference(forward_demodulation,[],[f33622,f141])).
fof(f33622,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X20)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X19))) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X20)))) ) | $spl0),
  inference(superposition,[],[f7810,f755])).
fof(f33855,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X18)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X18)) ) | $spl0),
  inference(forward_demodulation,[],[f33621,f141])).
fof(f33621,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X18)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl0),
  inference(superposition,[],[f7810,f755])).
fof(f33620,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X15))) | s(bool,X15) = s(bool,X16) | s(bool,'F_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33619,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X14)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X13))) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33618,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X12)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33617,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X9))) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) ) | $spl0),
  inference(superposition,[],[f7810,f612])).
fof(f33616,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X7))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f7810,f205])).
fof(f33615,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X6)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f7810,f205])).
fof(f33614,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,X4)),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f7810,f205])).
fof(f33613,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f7810,f205])).
fof(f33612,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK2('const.semanticPrimitives.Match_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.match_result'(bool),'const.semanticPrimitives.Match_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f7810,f128])).
fof(f33597,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f33597_D])).
fof(f33597_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f33588,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl40),
  inference(cnf_transformation,[],[f33588_D])).
fof(f33588_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f33598,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) | $spl54),
  inference(cnf_transformation,[],[f33598_D])).
fof(f33598_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f33611,plain,(
  ( ! [X61,X62,X63] : (s(bool,'F_0') = s(bool,X61) | s(bool,X61) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X62) = s(fun(bool,bool),X63) | s(bool,'T_0') = s(bool,sK3(X62,X63,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33580,f141])).
fof(f33580,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X62) = s(fun(bool,bool),X63) | s(bool,'T_0') = s(bool,sK3(X62,X63,bool,bool)) | s(bool,X61) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f2378,f11021])).
fof(f33610,plain,(
  ( ! [X59,X60,X58] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X58) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,'T_0') = s(bool,sK3(X59,X60,bool,bool)) | s(bool,'F_0') = s(bool,X58)) ) | $spl0),
  inference(forward_demodulation,[],[f33579,f141])).
fof(f33579,plain,(
  ( ! [X59,X60,X58] : (s(bool,X58) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,'T_0') = s(bool,sK3(X59,X60,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl0),
  inference(superposition,[],[f2378,f11021])).
fof(f33595,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl50),
  inference(cnf_transformation,[],[f33595_D])).
fof(f33595_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f33609,plain,(
  ( ! [X57,X56,X55] : (s(bool,'F_0') = s(bool,X57) | s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X55) = s(fun(bool,bool),X56) | s(bool,'T_0') = s(bool,sK3(X55,X56,bool,bool)) | s(bool,X57) = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33578,f141])).
fof(f33578,plain,(
  ( ! [X57,X56,X55] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X55) = s(fun(bool,bool),X56) | s(bool,'T_0') = s(bool,sK3(X55,X56,bool,bool)) | s(bool,X57) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X57) = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f2378,f11021])).
fof(f33587,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl38),
  inference(cnf_transformation,[],[f33587_D])).
fof(f33587_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f33594,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl48),
  inference(cnf_transformation,[],[f33594_D])).
fof(f33594_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f33608,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X53) = s(fun(bool,bool),X54) | s(bool,'T_0') = s(bool,sK3(X53,X54,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33577,f141])).
fof(f33577,plain,(
  ( ! [X54,X52,X53] : (s(bool,X52) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X53) = s(fun(bool,bool),X54) | s(bool,'T_0') = s(bool,sK3(X53,X54,bool,bool)) | s(bool,X52) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f2378,f10324])).
fof(f33607,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X49) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(bool,'T_0') = s(bool,sK3(X50,X51,bool,bool)) | s(bool,'T_0') = s(bool,X49)) ) | $spl0),
  inference(forward_demodulation,[],[f33576,f141])).
fof(f33576,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(bool,'T_0') = s(bool,sK3(X50,X51,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X49)) ) | $spl0),
  inference(superposition,[],[f2378,f10324])).
fof(f33593,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl47),
  inference(cnf_transformation,[],[f33593_D])).
fof(f33593_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f33591,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl44),
  inference(cnf_transformation,[],[f33591_D])).
fof(f33591_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f33606,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,sK3(X46,X47,bool,bool)) | s(bool,X48) = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f33575,f141])).
fof(f33575,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,sK3(X46,X47,bool,bool)) | s(bool,X48) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X48) = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f2378,f10324])).
fof(f33605,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f33605_D])).
fof(f33605_D,plain,(
  ( ! [X45] : (s(bool,X45) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f33574,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,'T_0') = s(bool,sK3(X43,X44,bool,bool)) | s(bool,X45) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2378,f10316])).
fof(f33604,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(bool,'F_0'))) | s(bool,X40) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | $spl62),
  inference(cnf_transformation,[],[f33604_D])).
fof(f33604_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(bool,'F_0'))) | s(bool,X40) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f33573,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'T_0') = s(bool,sK3(X41,X42,bool,bool)) | s(bool,X40) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2378,f10316])).
fof(f33603,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X37) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,sK3(X38,X39,bool,bool)) | s(bool,'T_0') = s(bool,X37)) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f33572,f510])).
fof(f33572,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,sK3(X38,X39,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f2378,f10316])).
fof(f33602,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34)))) ) | $spl60),
  inference(cnf_transformation,[],[f33602_D])).
fof(f33602_D,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f33571,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,sK3(X32,X33,bool,bool)) | s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X34)))) ) | $spl0),
  inference(superposition,[],[f2378,f755])).
fof(f33601,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl58),
  inference(cnf_transformation,[],[f33601_D])).
fof(f33601_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f33570,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,'T_0') = s(bool,sK3(X30,X31,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X29)))) ) | $spl0),
  inference(superposition,[],[f2378,f755])).
fof(f33600,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) | $spl56),
  inference(cnf_transformation,[],[f33600_D])).
fof(f33600_D,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f33569,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK3(X27,X28,bool,bool)) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0')))))) ) | $spl0),
  inference(superposition,[],[f2378,f755])).
fof(f33568,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,sK3(X23,X24,bool,bool)) | s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) | $spl0),
  inference(superposition,[],[f2378,f612])).
fof(f33567,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK3(X21,X22,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X20)) ) | $spl0),
  inference(superposition,[],[f2378,f612])).
fof(f33566,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK3(X18,X19,bool,bool)) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f2378,f612])).
fof(f33565,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,sK3(X14,X15,bool,bool)) | s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f2378,f612])).
fof(f33564,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2378,f205])).
fof(f33563,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2378,f205])).
fof(f33562,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2378,f205])).
fof(f33561,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2378,f205])).
fof(f33560,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2378,f128])).
fof(f32666,plain,(
  ( ! [X194,X195] : (s(bool,X194) != s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,X194) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X195)) ) | $spl36),
  inference(cnf_transformation,[],[f32666_D])).
fof(f32666_D,plain,(
  ( ! [X194,X195] : (s(bool,X194) != s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,X194) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X195)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f31925,plain,(
  ( ! [X194,X195,X193] : (s(bool,X194) != s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,X195) | s(bool,X194) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6 | $spl18)),
  inference(duplicate_literal_removal,[],[f31396])).
fof(f31396,plain,(
  ( ! [X194,X195,X193] : (s(bool,X194) != s(bool,'const.combin.I_1'(s(bool,X195))) | s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,X195) | s(bool,X194) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X193)) ) | ($spl0 | $spl6 | $spl18)),
  inference(superposition,[],[f6809,f10316])).
fof(f32163,plain,(
  ( ! [X62,X60] : (s(bool,X60) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X60) = s(bool,X62) | s(bool,X62) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | $spl34),
  inference(cnf_transformation,[],[f32163_D])).
fof(f32163_D,plain,(
  ( ! [X62,X60] : (s(bool,X60) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X60) = s(bool,X62) | s(bool,X62) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f32032,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X50) = s(bool,X52) | s(bool,X50) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f30817])).
fof(f30817,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X50) = s(bool,X52) | s(bool,X50) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X51)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f10316,f10316])).
fof(f32042,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f30723])).
fof(f30723,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f10316,f10316])).
fof(f32073,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') = s(bool,X61) | s(bool,X60) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X62) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X60) = s(bool,X62)) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f30647])).
fof(f30647,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') = s(bool,X61) | s(bool,X60) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,X60) = s(bool,X62)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f10316,f10316])).
fof(f28081,plain,(
  ( ! [X161,X162] : (s(bool,X162) != s(bool,'const.combin.I_1'(s(bool,X161))) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X162)) ) | $spl32),
  inference(cnf_transformation,[],[f28081_D])).
fof(f28081_D,plain,(
  ( ! [X161,X162] : (s(bool,X162) != s(bool,'const.combin.I_1'(s(bool,X161))) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X162)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f27373,plain,(
  ( ! [X161,X162,X160] : (s(bool,X162) != s(bool,'const.combin.I_1'(s(bool,X161))) | s(bool,'T_0') = s(bool,X162) | s(bool,'F_0') = s(bool,X160) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl18)),
  inference(duplicate_literal_removal,[],[f26943])).
fof(f26943,plain,(
  ( ! [X161,X162,X160] : (s(bool,X162) != s(bool,'const.combin.I_1'(s(bool,X161))) | s(bool,'T_0') = s(bool,X162) | s(bool,'F_0') = s(bool,X160) | s(bool,X161) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X160)) ) | ($spl0 | $spl18)),
  inference(superposition,[],[f6809,f11021])).
fof(f28076,plain,(
  ( ! [X127,X128] : (s(bool,X127) != s(bool,X128) | s(bool,X127) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X128)))) ) | $spl30),
  inference(cnf_transformation,[],[f28076_D])).
fof(f28076_D,plain,(
  ( ! [X127,X128] : (s(bool,X127) != s(bool,X128) | s(bool,X127) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X128)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f27376,plain,(
  ( ! [X127,X128,X126] : (s(bool,X127) != s(bool,X128) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,'F_0') = s(bool,X126) | s(bool,X127) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f26929])).
fof(f26929,plain,(
  ( ! [X127,X128,X126] : (s(bool,X127) != s(bool,X128) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X128))) | s(bool,'F_0') = s(bool,X126) | s(bool,X127) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X126)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3392,f11021])).
fof(f27616,plain,(
  ( ! [X68,X66] : (s(bool,X66) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X66) = s(bool,X68) | s(bool,X68) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl28),
  inference(cnf_transformation,[],[f27616_D])).
fof(f27616_D,plain,(
  ( ! [X68,X66] : (s(bool,X66) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X66) = s(bool,X68) | s(bool,X68) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f27477,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X60) = s(bool,X62) | s(bool,X60) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f26426])).
fof(f26426,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X60) = s(bool,X62) | s(bool,X60) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f11021,f11021])).
fof(f27488,plain,(
  ( ! [X74,X72,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f26340])).
fof(f26340,plain,(
  ( ! [X74,X72,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X72)) ) | $spl0),
  inference(superposition,[],[f11021,f11021])).
fof(f27515,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X66) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X68) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X66) = s(bool,X68)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f26269])).
fof(f26269,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X66) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X67) | s(bool,X68) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X66) = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f11021,f11021])).
fof(f17247,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225,f17228,f17230,f17241,f17244,f17246])).
fof(f17246,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X40,happ(s(fun(X34,X40),X41),s(X34,X42))) = s(X40,happ(s(fun(X34,X40),happ(s(fun(fun(X34,X40),fun(X34,X40)),X43),s(fun(X34,X40),'const.relation.RESTRICT_3'(s(fun(X34,X40),X41),s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(X34,X42))))),s(X34,X42))) | s(fun(X34,X40),X41) != s(fun(X34,X40),'const.relation.WFREC_2'(s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(fun(fun(X34,X40),fun(X34,X40)),X43)))) ) | $spl24),
  inference(subsumption_resolution,[],[f17123,f117])).
fof(f17123,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (~p(s(bool,'T_0')) | s(X40,happ(s(fun(X34,X40),X41),s(X34,X42))) = s(X40,happ(s(fun(X34,X40),happ(s(fun(fun(X34,X40),fun(X34,X40)),X43),s(fun(X34,X40),'const.relation.RESTRICT_3'(s(fun(X34,X40),X41),s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(X34,X42))))),s(X34,X42))) | s(fun(X34,X40),X41) != s(fun(X34,X40),'const.relation.WFREC_2'(s(fun(X34,fun(X34,bool)),happ(s(fun(X35,fun(X34,fun(X34,bool))),happ(s(fun(X36,fun(X35,fun(X34,fun(X34,bool)))),X37),s(X36,X38))),s(X35,X39))),s(fun(fun(X34,X40),fun(X34,X40)),X43)))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f5784])).
fof(f17244,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (s(X59,happ(s(fun(X53,X59),X60),s(X53,X61))) = s(X59,happ(s(fun(X53,X59),happ(s(fun(fun(X53,X59),fun(X53,X59)),X62),s(fun(X53,X59),'const.relation.RESTRICT_3'(s(fun(X53,X59),X60),s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(X53,X61))))),s(X53,X61))) | s(fun(X53,X59),X60) != s(fun(X53,X59),'const.relation.WFREC_2'(s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(fun(fun(X53,X59),fun(X53,X59)),X62)))) ) | $spl24),
  inference(subsumption_resolution,[],[f17122,f117])).
fof(f17122,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (~p(s(bool,'T_0')) | s(X59,happ(s(fun(X53,X59),X60),s(X53,X61))) = s(X59,happ(s(fun(X53,X59),happ(s(fun(fun(X53,X59),fun(X53,X59)),X62),s(fun(X53,X59),'const.relation.RESTRICT_3'(s(fun(X53,X59),X60),s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(X53,X61))))),s(X53,X61))) | s(fun(X53,X59),X60) != s(fun(X53,X59),'const.relation.WFREC_2'(s(fun(X53,fun(X53,bool)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X56),s(X55,X57))),s(fun(X54,fun(X55,fun(X53,fun(X53,bool)))),X58))),s(fun(fun(X53,X59),fun(X53,X59)),X62)))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f5786])).
fof(f17241,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) ) | $spl24),
  inference(subsumption_resolution,[],[f17117,f117])).
fof(f17117,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f162])).
fof(f17230,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24)))) ) | $spl24),
  inference(subsumption_resolution,[],[f17099,f117])).
fof(f17099,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),happ(s(fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')))),happ(s(fun(X21,fun(X20,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X22),s(X21,X23))),s(X20,X24)))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f4082])).
fof(f17228,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30)))) ) | $spl24),
  inference(subsumption_resolution,[],[f17097,f117])).
fof(f17097,plain,(
  ( ! [X30,X28,X26,X29,X27] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X28),s(X27,X29))),s(fun(X26,fun(X27,'type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))))),X30)))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f4083])).
fof(f17225,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) ) | $spl24),
  inference(subsumption_resolution,[],[f17093,f117])).
fof(f17093,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),X1)) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f157])).
fof(f17216,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) | $spl24),
  inference(subsumption_resolution,[],[f17062,f16176])).
fof(f17062,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.Match_type_error_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X14)) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f13006])).
fof(f17214,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13)) ) | $spl24),
  inference(subsumption_resolution,[],[f17053,f16176])).
fof(f17053,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') != s(bool,'const.semanticPrimitives.match_result_CASE_4'(s('type.semanticPrimitives.match_result'(X12),'const.semanticPrimitives.No_match_0'),s(bool,X13),s(bool,X14),s(fun(X12,bool),X15))) | s(bool,'F_0') = s(bool,X13)) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f12347])).
fof(f17211,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(subsumption_resolution,[],[f17036,f16176])).
fof(f17036,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f11638])).
fof(f17150,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | $spl24),
  inference(subsumption_resolution,[],[f16819,f16176])).
fof(f16819,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f387])).
fof(f17145,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f16810,f16176])).
fof(f16810,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f16176,f134])).
fof(f17245,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225,f17228,f17230,f17241,f17244])).
fof(f17242,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225,f17228,f17230,f17241])).
fof(f17231,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225,f17228,f17230])).
fof(f17229,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225,f17228])).
fof(f17227,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f17095,f117])).
fof(f17095,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f16176,f4069])).
fof(f17226,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216,f17225])).
fof(f17224,plain,(
  $false | ($spl0 | $spl24)),
  inference(global_subsumption,[],[f17223,f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216])).
fof(f17223,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl24)),
  inference(condensation,[],[f17222])).
fof(f17222,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f17089,f16176])).
fof(f17089,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X7) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X6) = s(bool,X7)) ) | ($spl0 | $spl24)),
  inference(backward_demodulation,[],[f16176,f16731])).
fof(f17217,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214,f17216])).
fof(f17215,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211,f17214])).
fof(f17212,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150,f17211])).
fof(f17209,plain,(
  $false | ($spl0 | $spl24)),
  inference(global_subsumption,[],[f17208,f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150])).
fof(f17208,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl24)),
  inference(forward_demodulation,[],[f17207,f16176])).
fof(f17207,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f17032,f16176])).
fof(f17032,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,X30) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl0 | $spl24)),
  inference(backward_demodulation,[],[f16176,f11207])).
fof(f17164,plain,(
  $false | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f16876,f16176])).
fof(f16876,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl0 | $spl24)),
  inference(backward_demodulation,[],[f16176,f1881])).
fof(f17151,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145,f17150])).
fof(f17147,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f16813,f117])).
fof(f16813,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f16176,f171])).
fof(f17146,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f116,f115,f117,f171,f124,f125,f126,f127,f128,f131,f130,f135,f134,f140,f141,f172,f144,f147,f148,f152,f156,f153,f159,f160,f161,f163,f164,f165,f177,f183,f205,f207,f256,f278,f316,f385,f386,f387,f389,f414,f440,f462,f463,f464,f501,f596,f598,f730,f735,f1298,f1303,f2200,f2203,f2378,f2917,f3250,f3317,f3322,f3490,f4906,f5818,f568,f7771,f7816,f8907,f8963,f9010,f9011,f9013,f10299,f10300,f10326,f9569,f9570,f9571,f11638,f12090,f12091,f12117,f12347,f12707,f12802,f12803,f13006,f5784,f5786,f12782,f12070,f10330,f8998,f578,f576,f546,f162,f12806,f12094,f9014,f5006,f4996,f4907,f12804,f12092,f10327,f9012,f3075,f2989,f149,f12807,f12095,f11475,f9015,f8906,f3624,f3622,f3570,f3592,f3647,f143,f4082,f4737,f4083,f4736,f4069,f158,f157,f9561,f9207,f2387,f16176,f17145])).
fof(f16652,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,'const.combin.I_1'(s(bool,X192))) | s(bool,X191) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X192)) ) | $spl26),
  inference(cnf_transformation,[],[f16652_D])).
fof(f16652_D,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,'const.combin.I_1'(s(bool,X192))) | s(bool,X191) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X192)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f15946,plain,(
  ( ! [X191,X192,X190] : (s(bool,X191) != s(bool,'const.combin.I_1'(s(bool,X192))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,X192) | s(bool,X191) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl0 | $spl18)),
  inference(duplicate_literal_removal,[],[f15482])).
fof(f15482,plain,(
  ( ! [X191,X192,X190] : (s(bool,X191) != s(bool,'const.combin.I_1'(s(bool,X192))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,X192) | s(bool,X191) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X190)) ) | ($spl0 | $spl18)),
  inference(superposition,[],[f6809,f10324])).
fof(f16175,plain,(
  ( ! [X58,X56] : (s(bool,X56) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl22),
  inference(cnf_transformation,[],[f16175_D])).
fof(f16175_D,plain,(
  ( ! [X58,X56] : (s(bool,X56) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f16045,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X53) = s(bool,X55) | s(bool,X53) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f14945])).
fof(f14945,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X53) = s(bool,X55) | s(bool,X53) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X54)) ) | $spl0),
  inference(superposition,[],[f10324,f10324])).
fof(f16054,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) = s(bool,X64) | s(bool,X64) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,X63) = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f14868])).
fof(f14868,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) = s(bool,X64) | s(bool,X64) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,X63) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62)) ) | $spl0),
  inference(superposition,[],[f10324,f10324])).
fof(f16084,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X56) = s(bool,X58)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f14799])).
fof(f14799,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s(bool,X56) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X56) = s(bool,X58)) ) | $spl0),
  inference(superposition,[],[f10324,f10324])).
fof(f13899,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13898,f13492])).
fof(f13492,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f6544,f6544])).
fof(f13898,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1)) ) | $spl10),
  inference(subsumption_resolution,[],[f13638,f171])).
fof(f13638,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f130,f6544])).
fof(f13897,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13560,f13492])).
fof(f13560,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X19)))) ) | ($spl0 | $spl10)),
  inference(superposition,[],[f1881,f6544])).
fof(f13896,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13554,f13492])).
fof(f13554,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f177,f6544])).
fof(f13895,plain,(
  $false | ($spl0 | $spl10 | $spl18)),
  inference(subsumption_resolution,[],[f13465,f1881])).
fof(f13465,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f6544,f6740])).
fof(f6740,plain,(
  ( ! [X75] : (s(bool,X75) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X75)))) ) | $spl18),
  inference(superposition,[],[f6553,f141])).
fof(f13894,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f13469,f171])).
fof(f13469,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6544,f130])).
fof(f13893,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13474,f1881])).
fof(f13474,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f6544,f1236])).
fof(f13892,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f13480,f1881])).
fof(f13480,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f6544,f3233])).
fof(f13891,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f13487,f1881])).
fof(f13487,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f6544,f3392])).
fof(f13489,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f1881,f6544,f3392])).
fof(f13481,plain,(
  $false | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f1881,f6544,f3233])).
fof(f13476,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f1881,f6544,f1236])).
fof(f13473,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f177,f6544,f900])).
fof(f13470,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f171,f6544,f130])).
fof(f13468,plain,(
  $false | ($spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f177,f6544,f6553])).
fof(f13466,plain,(
  $false | ($spl0 | $spl10 | $spl18)),
  inference(unit_resulting_resolution,[],[f1881,f6544,f6740])).
fof(f13460,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f177,f6544,f387])).
fof(f13458,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f1881,f6544,f11207])).
fof(f13456,plain,(
  $false | ($spl10 | $spl16)),
  inference(unit_resulting_resolution,[],[f177,f6544,f6552])).
fof(f13454,plain,(
  $false | ($spl0 | $spl10 | $spl16)),
  inference(unit_resulting_resolution,[],[f1881,f6544,f6864])).
fof(f6864,plain,(
  ( ! [X75] : (s(bool,X75) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X75)))) ) | $spl16),
  inference(superposition,[],[f6552,f141])).
fof(f13452,plain,(
  $false | ($spl10 | $spl14)),
  inference(unit_resulting_resolution,[],[f177,f6544,f6550])).
fof(f13450,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f177,f6544,f6320])).
fof(f13445,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f117,f6544,f134])).
fof(f13442,plain,(
  $false | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f1881,f6544,f11330])).
fof(f13440,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f177,f6544])).
fof(f13208,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl20),
  inference(cnf_transformation,[],[f13208_D])).
fof(f13208_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f13192,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f13066])).
fof(f13066,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f3392,f594])).
fof(f6553,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) ) | $spl18),
  inference(cnf_transformation,[],[f6553_D])).
fof(f6553_D,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6554,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f6481,f387])).
fof(f6481,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X32)))) ) | $spl0),
  inference(superposition,[],[f6320,f3991])).
fof(f6512,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6480])).
fof(f6480,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl0),
  inference(superposition,[],[f6320,f3762])).
fof(f6552,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl16),
  inference(cnf_transformation,[],[f6552_D])).
fof(f6552_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6551,plain,(
  ( ! [X28,X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X29)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f6479,f387])).
fof(f6479,plain,(
  ( ! [X28,X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X28)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f6320,f1797])).
fof(f6550,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X27)) ) | $spl14),
  inference(cnf_transformation,[],[f6550_D])).
fof(f6550_D,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X27)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f6549,plain,(
  ( ! [X26,X27,X25] : (s(bool,X27) != s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X27)) ) | $spl0),
  inference(subsumption_resolution,[],[f6478,f387])).
fof(f6478,plain,(
  ( ! [X26,X27,X25] : (s(bool,X27) != s(bool,'const.relation.WF_1'(s(fun(X26,fun(X26,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X25)))) ) | $spl0),
  inference(superposition,[],[f6320,f1796])).
fof(f6545,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl12),
  inference(cnf_transformation,[],[f6545_D])).
fof(f6545_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f6519,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f6470])).
fof(f6470,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f6320,f594])).
fof(f6009,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5849,f171])).
fof(f5849,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f177,f502,f153])).
fof(f6008,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f5851,f502])).
fof(f5851,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f1881,f502,f153])).
fof(f6007,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f5859,f171])).
fof(f5859,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f177,f502,f153])).
fof(f6006,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f5861,f502])).
fof(f5861,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f1881,f502,f153])).
fof(f5869,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f141,f502,f158])).
fof(f5866,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f183,f502,f130])).
fof(f5864,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f177,f502,f135])).
fof(f5862,plain,(
  $false | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f502,f1881,f502,f153])).
fof(f5860,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f171,f177,f502,f153])).
fof(f5858,plain,(
  $false | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f502,f1881,f502,f153])).
fof(f5856,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f502,f177,f502,f153])).
fof(f5852,plain,(
  $false | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f502,f1881,f502,f153])).
fof(f5850,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f171,f177,f502,f153])).
fof(f5848,plain,(
  $false | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f502,f1881,f502,f153])).
fof(f5846,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f502,f177,f502,f153])).
fof(f5842,plain,(
  $false | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f1881,f1881,f502,f256])).
fof(f5840,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f177,f177,f502,f256])).
fof(f5834,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f117,f502])).
fof(f5817,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) | $spl8),
  inference(cnf_transformation,[],[f5817_D])).
fof(f5817_D,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5816,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) | $spl0),
  inference(subsumption_resolution,[],[f5815,f134])).
fof(f5815,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X1) | ~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) | $spl0),
  inference(forward_demodulation,[],[f5788,f141])).
fof(f5788,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X1)))) ) | $spl0),
  inference(superposition,[],[f162,f1796])).
fof(f504,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl6),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f501,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f496,f172])).
fof(f496,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f149,f172])).
fof(f175,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f174,f117])).
fof(f174,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f173,f127])).
fof(f173,plain,(
  ~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f168,f169])).
fof(f167,plain,(
  p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | $spl0),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f169,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | $spl2),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f115,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)) & (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f101])).
fof(f101,plain,(
  ? [X0,X1] : ((p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)) & (~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)))) | s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ? [X0,X1] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)))) <~> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1)))) <=> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7)))) <=> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7)))) <=> s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') = s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),X7))),
  file('thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1','thm.dec_to_exhProof.dec_to_exhProof_3|48|_prove_ltb|split|1')).
fof(f168,plain,(
  ~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) | $spl1),
  inference(cnf_transformation,[],[f168_D])).
fof(f168_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f170,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | $spl3),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f116,plain,(
  s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),'const.semanticPrimitives.Match_type_error_0') != s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1) | ~p(s(bool,'const.dec_to_exhProof.match_result_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK0),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'))),'const.semanticPrimitives.Match_type_error_0'),s('type.semanticPrimitives.match_result'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'))),sK1))))),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 188
% Passive clauses: 1550
% Generated clauses: 45457
% Final active clauses: 59
% Final passive clauses: 133
% Input formulas: 37
% Initial clauses: 51
% 
% Duplicate literals: 1829
% Trivial inequalities: 685
% Fw subsumption resolutions: 4065
% Fw demodulations: 7295
% Bw demodulations: 1215
% Condensations: 3509
% Global subsumptions: 44
% 
% Simple tautologies: 677
% Equational tautologies: 213
% Forward subsumptions: 25372
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 40
% 
% Unit resulting resolution: 1575
% Factoring: 1
% Forward superposition: 12364
% Backward superposition: 11748
% Self superposition: 312
% Equality factoring: 162
% Equality resolution: 45
% 
% Split clauses: 128
% Split components: 277
% SAT solver clauses: 1946
% SAT solver unit clauses: 73
% SAT solver binary clauses: 262
% 
% Sat splits: 128
% Sat splitting refutations: 94
% 
% Memory used [KB]: 24434
% Time elapsed: 1.861 s
% ------------------------------
% ------------------------------
% Success in time 3.18 s
