Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% (145407)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 144
% Passive clauses: 1831
% Generated clauses: 18443
% Final active clauses: 128
% Final passive clauses: 1415
% Input formulas: 39
% Initial clauses: 92
% 
% Duplicate literals: 808
% Trivial inequalities: 317
% Fw subsumption resolutions: 2916
% Fw demodulations: 2851
% Bw demodulations: 13
% 
% Simple tautologies: 376
% Equational tautologies: 50
% Forward subsumptions: 8767
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 429
% Factoring: 3
% Forward superposition: 5439
% Backward superposition: 5235
% Self superposition: 149
% Equality factoring: 95
% Equality resolution: 32
% 
% Split clauses: 203
% Split components: 423
% SAT solver clauses: 530
% SAT solver unit clauses: 16
% SAT solver binary clauses: 109
% 
% Sat splits: 203
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 525
% TWLsolver calls for satisfiability: 122
% 
% Memory used [KB]: 8315
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f5333,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5328,f145,f258,f948,f299,f303,f294,f296,f144,f256,f254,f1214,f1014,f1019,f217,f290,f259,f261,f160,f270,f2665,f793,f781,f790,f153,f266,f218,f292,f3020,f338,f341,f342,f345,f780,f789,f201,f282,f1253,f1261,f143,f255,f253,f469,f483,f205,f283,f1257,f1265,f156,f268,f2663,f2681,f216,f291])).
fof(f291,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))))) | s('type.ast.id'('type.list.list'('type.string.char')),X0) = s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) | $spl42),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))))) | s('type.ast.id'('type.list.list'('type.string.char')),X0) = s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f216,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) = s('type.ast.id'('type.list.list'('type.string.char')),X1) | ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))))) | s('type.ast.id'('type.list.list'('type.string.char')),X0) = s('type.ast.id'('type.list.list'('type.string.char')),X1)) & (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1) | p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))))) <=> s('type.ast.id'('type.list.list'('type.string.char')),X0) = s('type.ast.id'('type.list.list'('type.string.char')),X1))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X20,X21] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X21))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X20)))))) <=> s('type.ast.id'('type.list.list'('type.string.char')),X20) = s('type.ast.id'('type.list.list'('type.string.char')),X21))),
  file('thm.evalProps.same_tid_tid|split|0','thm.semanticPrimitives.same_tid_def|split|0')).
fof(f2681,plain,(
  ( ! [X5] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X5))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | s('type.ast.id'('type.list.list'('type.string.char')),X5) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0))) ) | $spl432),
  inference(cnf_transformation,[],[f2681_D])).
fof(f2681_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X5))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | s('type.ast.id'('type.list.list'('type.string.char')),X5) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f2663,plain,(
  ( ! [X5] : (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X5))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | s('type.ast.id'('type.list.list'('type.string.char')),X5) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0))) ) | ($spl42 | $spl412)),
  inference(superposition,[],[f291,f2635])).
fof(f2635,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)))) | $spl412),
  inference(cnf_transformation,[],[f2635_D])).
fof(f2635_D,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)))) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f268_D])).
fof(f268_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f156,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.evalProps.same_tid_tid|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f1265,plain,(
  s(bool,'T_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) | $spl300),
  inference(cnf_transformation,[],[f1265_D])).
fof(f1265_D,plain,(
  s(bool,'T_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f1257,plain,(
  s(bool,'T_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) | ($spl0 | $spl12)),
  inference(resolution,[],[f253,f268])).
fof(f283,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f205,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.evalProps.same_tid_tid|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f483,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl90),
  inference(cnf_transformation,[],[f483_D])).
fof(f483_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f469,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | ($spl16 | $spl26)),
  inference(resolution,[],[f283,f270])).
fof(f253,plain,(
  p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | $spl0),
  inference(cnf_transformation,[],[f253_D])).
fof(f253_D,plain,(
  p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f255,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | $spl2),
  inference(cnf_transformation,[],[f255_D])).
fof(f255_D,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f143,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1)))) & (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f101])).
fof(f101,plain,(
  ? [X0,X1] : ((p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn',X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))) & (~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn',X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ? [X0,X1] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn',X0)))) <~> s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn',X0)))) <=> s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s('type.semanticPrimitives.tid_or_exn',X6)))) <=> s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s('type.semanticPrimitives.tid_or_exn',X6)))) <=> s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))))),
  file('thm.evalProps.same_tid_tid|split|0','thm.evalProps.same_tid_tid|split|0')).
fof(f1261,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) | s(bool,'F_0') = s(bool,X0)) ) | $spl294),
  inference(cnf_transformation,[],[f1261_D])).
fof(f1261_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f1253,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl0 | $spl90)),
  inference(resolution,[],[f253,f483])).
fof(f282,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))))) ) | $spl24),
  inference(cnf_transformation,[],[f282_D])).
fof(f282_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f201,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X24,X25] : s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X24))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X25)))))),
  file('thm.evalProps.same_tid_tid|split|0','thm.semanticPrimitives.same_tid_def|split|2')).
fof(f789,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s('type.semanticPrimitives.tid_or_exn',X6))) | s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X6))))) ) | $spl140),
  inference(cnf_transformation,[],[f789_D])).
fof(f789_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s('type.semanticPrimitives.tid_or_exn',X6))) | s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X6))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f780,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s('type.semanticPrimitives.tid_or_exn',X6))) | s('type.semanticPrimitives.tid_or_exn',X6) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X6))))) ) | ($spl10 | $spl24)),
  inference(superposition,[],[f282,f266])).
fof(f345,plain,(
  $false | $spl62),
  inference(equality_resolution,[],[f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) | $spl62),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f341,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) | ($spl7 | $spl40 | $spl60)),
  inference(subsumption_resolution,[],[f339,f261])).
fof(f339,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1)) ) | ($spl40 | $spl60)),
  inference(backward_demodulation,[],[f338,f290])).
fof(f338,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl60),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3020,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9)) ) | ($spl140 | $spl294 | $spl413)),
  inference(subsumption_resolution,[],[f3019,f2634])).
fof(f2634,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)))) | $spl413),
  inference(cnf_transformation,[],[f2634_D])).
fof(f2634_D,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)))) <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f3019,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0))))) ) | ($spl140 | $spl294)),
  inference(duplicate_literal_removal,[],[f2978])).
fof(f2978,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,X9) | s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl140 | $spl294)),
  inference(superposition,[],[f789,f1261])).
fof(f292,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))) ) | $spl44),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f218,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ~s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X17,X16] : ~s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X16))) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X17)))),
  file('thm.evalProps.same_tid_tid|split|0','thm.semanticPrimitives.tid_or_exn_distinct')).
fof(f266,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK3(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0))))) ) | $spl10),
  inference(cnf_transformation,[],[f266_D])).
fof(f266_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK3(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f153,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK3(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK3(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f45])).
fof(f45,plain,(
  ! [X0] : (? [X2] : s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X2))) | ? [X1] : s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X18] : (? [X19] : s('type.semanticPrimitives.tid_or_exn',X18) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X19))) | ? [X19] : s('type.semanticPrimitives.tid_or_exn',X18) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X19))))),
  file('thm.evalProps.same_tid_tid|split|0','thm.semanticPrimitives.tid_or_exn_nchotomy')).
fof(f790,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.tid_or_exn',X8) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))) | s('type.semanticPrimitives.tid_or_exn',X8) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X8))))) ) | $spl142),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.tid_or_exn',X8) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))) | s('type.semanticPrimitives.tid_or_exn',X8) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X8))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f781,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.tid_or_exn',X8) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X9))) | s('type.semanticPrimitives.tid_or_exn',X8) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X8))))) ) | ($spl10 | $spl44)),
  inference(superposition,[],[f292,f266])).
fof(f793,plain,(
  ( ! [X12] : (s('type.semanticPrimitives.tid_or_exn',X12) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X12)))) | s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn',X12)) ) | $spl148),
  inference(cnf_transformation,[],[f793_D])).
fof(f793_D,plain,(
  ( ! [X12] : (s('type.semanticPrimitives.tid_or_exn',X12) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X12)))) | s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn',X12)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f2665,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn',X7) | s('type.semanticPrimitives.tid_or_exn',X7) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X7))))) ) | ($spl142 | $spl412)),
  inference(superposition,[],[f790,f2635])).
fof(f270,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f160,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.evalProps.same_tid_tid|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f261,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f261_D])).
fof(f261_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f259,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f148,f145])).
fof(f148,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.evalProps.same_tid_tid|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f290,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1) | p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) ) | $spl40),
  inference(cnf_transformation,[],[f290_D])).
fof(f290_D,plain,(
  ( ! [X0,X1] : (s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1) | p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f217,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))))) | s('type.ast.id'('type.list.list'('type.string.char')),X0) != s('type.ast.id'('type.list.list'('type.string.char')),X1)) )),
  inference(cnf_transformation,[],[f137])).
fof(f1019,plain,(
  ( ! [X0] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) ) | $spl182),
  inference(cnf_transformation,[],[f1019_D])).
fof(f1019_D,plain,(
  ( ! [X0] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1014,plain,(
  ( ! [X0] : (p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))),s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))))))) ) | $spl40),
  inference(equality_resolution,[],[f290])).
fof(f1214,plain,(
  $false | ($spl2 | $spl7 | $spl50 | $spl182)),
  inference(subsumption_resolution,[],[f1213,f261])).
fof(f1213,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl50 | $spl182)),
  inference(forward_demodulation,[],[f1201,f303])).
fof(f1201,plain,(
  p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0)))) | ($spl2 | $spl182)),
  inference(superposition,[],[f1019,f255])).
fof(f254,plain,(
  ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f256,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | $spl3),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f144,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))),s('type.semanticPrimitives.tid_or_exn',sK0))))),
  inference(cnf_transformation,[],[f102])).
fof(f296,plain,(
  ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0)))) | $spl49),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f294,plain,(
  ~p(s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f254,f255])).
fof(f303,plain,(
  s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0))) | $spl50),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f299,plain,(
  s(bool,'F_0') = s(bool,'const.semanticPrimitives.same_tid_2'(s('type.semanticPrimitives.tid_or_exn',sK0),s('type.semanticPrimitives.tid_or_exn',sK0))) | ($spl16 | $spl49)),
  inference(resolution,[],[f270,f296])).
fof(f948,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | ($spl148 | $spl172)),
  inference(trivial_inequality_removal,[],[f927])).
fof(f927,plain,(
  s('type.semanticPrimitives.tid_or_exn',sK0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1))) | s('type.semanticPrimitives.tid_or_exn',sK0) != s('type.semanticPrimitives.tid_or_exn',sK0) | ($spl148 | $spl172)),
  inference(superposition,[],[f793,f925])).
fof(f925,plain,(
  s('type.ast.id'('type.list.list'('type.string.char')),sK1) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) | $spl172),
  inference(cnf_transformation,[],[f925_D])).
fof(f925_D,plain,(
  s('type.ast.id'('type.list.list'('type.string.char')),sK1) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f258,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f145,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.same_tid_tid|split|0',aHL_TRUTH)).
fof(f5328,plain,(
  $false | ($spl4 | $spl173 | $spl300 | $spl432)),
  inference(subsumption_resolution,[],[f5327,f924])).
fof(f924,plain,(
  s('type.ast.id'('type.list.list'('type.string.char')),sK1) != s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) | $spl173),
  inference(cnf_transformation,[],[f924_D])).
fof(f924_D,plain,(
  s('type.ast.id'('type.list.list'('type.string.char')),sK1) != s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f5327,plain,(
  s('type.ast.id'('type.list.list'('type.string.char')),sK1) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) | ($spl4 | $spl300 | $spl432)),
  inference(subsumption_resolution,[],[f5299,f258])).
fof(f5299,plain,(
  ~p(s(bool,'T_0')) | s('type.ast.id'('type.list.list'('type.string.char')),sK1) = s('type.ast.id'('type.list.list'('type.string.char')),sK2(sK0)) | ($spl300 | $spl432)),
  inference(superposition,[],[f2681,f1265])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 136
% Passive clauses: 376
% Generated clauses: 5982
% Final active clauses: 89
% Final passive clauses: 202
% Input formulas: 39
% Initial clauses: 92
% 
% Duplicate literals: 262
% Trivial inequalities: 94
% Fw subsumption resolutions: 476
% Fw demodulations: 541
% Bw demodulations: 41
% Condensations: 231
% Global subsumptions: 13
% 
% Simple tautologies: 89
% Equational tautologies: 38
% Forward subsumptions: 2986
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 155
% Factoring: 1
% Forward superposition: 1233
% Backward superposition: 1497
% Self superposition: 53
% Equality factoring: 42
% Equality resolution: 33
% 
% Split clauses: 71
% Split components: 143
% SAT solver clauses: 759
% SAT solver unit clauses: 107
% SAT solver binary clauses: 53
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 1
% 
% Sat splits: 71
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 705
% TWLsolver calls for satisfiability: 438
% 
% Memory used [KB]: 3965
% Time elapsed: 0.119 s
% ------------------------------
% ------------------------------
% Success in time 0.77 s
