Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSound
% (45911)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 520
% Passive clauses: 3499
% Generated clauses: 10934
% Final active clauses: 518
% Final passive clauses: 2744
% Input formulas: 133
% Initial clauses: 354
% 
% Duplicate literals: 162
% Trivial inequalities: 42
% Fw subsumption resolutions: 433
% Fw demodulations: 895
% Bw demodulations: 2
% 
% Simple tautologies: 235
% Equational tautologies: 36
% Forward subsumptions: 5640
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 3381
% Factoring: 800
% Forward superposition: 2748
% Backward superposition: 2010
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 135
% Split components: 272
% SAT solver clauses: 235
% SAT solver unit clauses: 9
% SAT solver binary clauses: 59
% 
% Sat splits: 135
% 
% TWLsolver clauses: 235
% TWLsolver calls for satisfiability: 190
% 
% Memory used [KB]: 7931
% Time elapsed: 0.612 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.typeSound
% (45925)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 405
% Passive clauses: 2211
% Generated clauses: 11797
% Final active clauses: 401
% Final passive clauses: 1792
% Input formulas: 133
% Initial clauses: 354
% 
% Duplicate literals: 108
% Trivial inequalities: 11
% Fw subsumption resolutions: 498
% Fw demodulations: 1460
% Bw demodulations: 16
% Condensations: 121
% Global subsumptions: 23
% 
% Simple tautologies: 230
% Equational tautologies: 10
% Forward subsumptions: 4515
% 
% Binary resolution: 2513
% Factoring: 336
% Forward superposition: 2612
% Backward superposition: 1467
% Self superposition: 50
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 2947
% SAT solver unit clauses: 380
% SAT solver binary clauses: 126
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 12
% 
% Sat splits: 5
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2855
% TWLsolver calls for satisfiability: 2712
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.typeSound
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSound
% SZS output start Proof for thm.typeSound
fof(f37732,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f597,f981,f979,f596,f980,f982,f14231,f14232,f14235,f14233,f14236,f14237,f14239,f14241,f14245,f14247,f14251,f14253,f14255,f14257,f14267,f14269,f14279,f14281,f14293,f14305,f14321,f14323,f14340,f14357,f14373,f14375,f14391,f14393,f14409,f14411,f14421,f14431,f14446,f14462,f14493,f14495,f14573,f14575,f14578,f14580,f14583,f14585,f14598,f14600,f14607,f14609,f14627,f14629,f14632,f14634,f14637,f14639,f14644,f14646,f14649,f14651,f14717,f14733,f14756,f14772,f14774,f14789,f14800,f14824,f14826,f14828,f14830,f14832,f14834,f16021,f16022,f14848,f14854,f14856,f14860,f14866,f14868,f14886,f14888,f14901,f14914,f14964,f14966,f14968,f14970,f14984,f14986,f14988,f14990,f14992,f14994,f14996,f14998,f15051,f15053,f15055,f15057,f15071,f15073,f15075,f15077,f15079,f15081,f15083,f15085,f15103,f15114,f15138,f15140,f15142,f15170,f15172,f15185,f15187,f15215,f15217,f15230,f15232,f15259,f15261,f15274,f15276,f15292,f15305,f15321,f15334,f15345,f15356,f15381,f15383,f15385,f15387,f15412,f15414,f15416,f15418,f15420,f15432,f15434,f15436,f15449,f15451,f15464,f15466,f15479,f15481,f15494,f15496,f15510,f15519,f15551,f15553,f15555,f15566,f15568,f15570,f15572,f15574,f15576,f15578,f15580,f15625,f15627,f15629,f15631,f15633,f15635,f15637,f15668,f15670,f15672,f15674,f15676,f15678,f15709,f15711,f15713,f15715,f15717,f15719,f15738,f15740,f15759,f15761,f15763,f15775,f15787,f15799,f15801,f15813,f15815,f15826,f15835,f15846,f15855,f15867,f15870,f15872,f15874,f15876,f15878,f15892,f15894,f15896,f15898,f15900,f15902,f15904,f15917,f15919,f15921,f15923,f15925,f15927,f15929,f15931,f16009,f16023,f16024,f16025,f16026,f16027,f16028,f16029,f16030,f16031,f16032,f16033,f16034,f16035,f16036,f16037,f16038,f16039,f16040,f16041,f16042,f16043,f16044,f16045,f16046,f16047,f16048,f16049,f16050,f16051,f16052,f16053,f16054,f16055,f16056,f16057,f16058,f16059,f16060,f16061,f16062,f16063,f16064,f16065,f16066,f16067,f16068,f16069,f16070,f16071,f16072,f16073,f16074,f16075,f16076,f16077,f16078,f16079,f16080,f16081,f16082,f16083,f16084,f16085,f16086,f16087,f16088,f16089,f16090,f16091,f16092,f16093,f16094,f16095,f16096,f16097,f16098,f16099,f16100,f16101,f16102,f16103,f16104,f16105,f16106,f16107,f16108,f16109,f16110,f16111,f16112,f16113,f16114,f16115,f16116,f16117,f16118,f16119,f16120,f16121,f16122,f16123,f16124,f16125,f16126,f16127,f16128,f16129,f16130,f16131,f16132,f16133,f16134,f16135,f16136,f16137,f16138,f16139,f16140,f16141,f16142,f16143,f16144,f16145,f16146,f16147,f16148,f16149,f16150,f16151,f16152,f16153,f16154,f16155,f16156,f16157,f16158,f16159,f16160,f16161,f16162,f16163,f16164,f16165,f16166,f16167,f16168,f16169,f16170,f16171,f16172,f16173,f16174,f16175,f16176,f16177,f16178,f16179,f16180,f16181,f16182,f16183,f16184,f16185,f16186,f16187,f16188,f16189,f16190,f16191,f16192,f16193,f16194,f16195,f16196,f16197,f16198,f16199,f16200,f16201,f16202,f16203,f16204,f16205,f16206,f16207,f16208,f16209,f16210,f16211,f16212,f16213,f16214,f16215,f16216,f16217,f16218,f16219,f16220,f16221,f16222,f16223,f16224,f16225,f16226,f16227,f16228,f16229,f16230,f16231,f16232,f16233,f16234,f16235,f16236,f16237,f16238,f16239,f16240,f16241,f16242,f16243,f16244,f16245,f16246,f16247,f16248,f16249,f16250,f16251,f16252,f16253,f16254,f16255,f16256,f16257,f16258,f16259,f16260,f16261,f16262,f16263,f16264,f16265,f16266,f16267,f16268,f16269,f16270,f16271,f16272,f16273,f16274,f16275,f16276,f16277,f16278,f16279,f16280,f16281,f16282,f16283,f16284,f16285,f16286,f16287,f16288,f16289,f16290,f16291,f16292,f16293,f16294,f16295,f16296,f16297,f16298,f16299,f16300,f16301,f16302,f16303,f16304,f16305,f16306,f16307,f16308,f16309,f16310,f16311,f16312,f16313,f16314,f16315,f16316,f16317,f16318,f16319,f16320,f16321,f16322,f16323,f16324,f16325,f16326,f16327,f16328,f16329,f16330,f16331,f16332,f16333,f16334,f16335,f16336,f16337,f16338,f16339,f16340,f16341,f16342,f16343,f16344,f16345,f16346,f16347,f16348,f16349,f16350,f16351,f16352,f16353,f16354,f16355,f16356,f16357,f16358,f16359,f16360,f16361,f16362,f16363,f16364,f16365,f16366,f16367,f16368,f16369,f16370,f16371,f16372,f16373,f16374,f16375,f16376,f16377,f16378,f16379,f16380,f16381,f16382,f16383,f16384,f16385,f16386,f16387,f16388,f16389,f16390,f16391,f16392,f16393,f16394,f16395,f16396,f16397,f16398,f16399,f16400,f16401,f16402,f16403,f16404,f16405,f16406,f16407,f16408,f16409,f16410,f16411,f16412,f16413,f16414,f16415,f16416,f16417,f16418,f16419,f16420,f16421,f16422,f16423,f16424,f16425,f16426,f16427,f16428,f16429,f16430,f16431,f16432,f16433,f16434,f16435,f16436,f16437,f16438,f16439,f16440,f16441,f16442,f16443,f16444,f16445,f16446,f16447,f16448,f16449,f16450,f16451,f16452,f16453,f16454,f16455,f16456,f16457,f16458,f16459,f16460,f16461,f16462,f16463,f16464,f16465,f16466,f16467,f16468,f16469,f16470,f16471,f16472,f16473,f16474,f16475,f16476,f16477,f16478,f16479,f16480,f16481,f16482,f16483,f16484,f16485,f16486,f16487,f16488,f16489,f16490,f16491,f16492,f16493,f16494,f16495,f16496,f16497,f16498,f16499,f16500,f16501,f16502,f16503,f16504,f16505,f16506,f16507,f16508,f16509,f16510,f16511,f16512,f16513,f16514,f16515,f16516,f16517,f16518,f16519,f16520,f16521,f16522,f16523,f16524,f16525,f16526,f16527,f16528,f16529,f16530,f16531,f16532,f16533,f16534,f16535,f16536,f16537,f16538,f16539,f16540,f16541,f16542,f16543,f16544,f16545,f16546,f16547,f16548,f16549,f16550,f16551,f16552,f16553,f16554,f16555,f16556,f16557,f16558,f16559,f16560,f16561,f16562,f16563,f16564,f16565,f16566,f16567,f16568,f16569,f16570,f16571,f16572,f16573,f16574,f16575,f16576,f16577,f16578,f16579,f16580,f16581,f16582,f16583,f16584,f16585,f16586,f16587,f16588,f16589,f16590,f16591,f16592,f16593,f16594,f16595,f16596,f16597,f16598,f16599,f16600,f16601,f16602,f16603,f16604,f16605,f16606,f16607,f16608,f16609,f16610,f37132,f37211,f37731])).
fof(f37731,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f37730,f980])).
fof(f37730,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | $spl0),
  inference(forward_demodulation,[],[f37242,f5378])).
fof(f5378,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,sK32('const.pair.,_2'(s(X0,X1),s(X2,X3)),X2,X0))) )),
  inference(unit_resulting_resolution,[],[f787,f820])).
fof(f820,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f536])).
fof(f536,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X7,X6] : ! [X8,X38,X46,X47] : (s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X8),s(X6,X38))) = s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X46),s(X6,X47))) <=> (s(X7,X8) = s(X7,X46) & s(X6,X38) = s(X6,X47)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK32(X2,X1,X0)),s(X1,sK33(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK32(X2,X1,X0)),s(X1,sK33(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK32,sK33])],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X7,X6] : ! [X8] : ? [X42,X41] : s('type.pair.prod'(X7,X6),X8) = s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X42),s(X6,X41)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f37242,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK32('const.pair.,_2'(s(sK14,sK16),s(sK15,sK18)),sK15,sK14)),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f3882,f979,f806])).
fof(f806,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) | (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X3),s(X1,sK35(X4,X3,X2,X1,X0)))) & p(s(bool,'const.bool.IN_2'(s(X1,sK35(X4,X3,X2,X1,X0)),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2)))))))) & (! [X5] : (s(X0,X4) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) | ? [X6] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2)))))))) & (! [X5] : (s(X0,X4) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))))))))))),
  inference(rectify,[],[f531])).
fof(f531,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) | ? [X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2)))))))) & (! [X5] : (s(X0,X4) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2))))))))))),
  inference(nnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) <=> ? [X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2))))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X3),s('type.list.list'(X1),X2)))))))) <=> ? [X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) & p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2))))))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6,X7] : ! [X9,X25,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X25),s('type.list.list'(X7),X9)))))))) <=> ? [X38] : (s(X6,X8) = s(X6,happ(s(fun(X7,X6),X25),s(X7,X38))) & p(s(bool,'const.bool.IN_2'(s(X7,X38),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),X9))))))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.list.MEM_MAP')).
fof(f3882,plain,(
  ( ! [X2,X0,X1] : (s(X0,sK32(X2,X1,X0)) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(backward_demodulation,[],[f3857,f784])).
fof(f784,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7,X6] : ! [X40] : s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X40))) = s(X7,happ(s(fun('type.pair.prod'(X7,X6),X7),'const.pair.FST_0'),s('type.pair.prod'(X7,X6),X40)))),
  file('thm.typeSound.typeSound_46_prove_ltb','ARITY_0_1_const.pair.FST')).
fof(f3857,plain,(
  ( ! [X2,X0,X1] : (s(X0,sK32(X2,X1,X0)) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2)))) )),
  inference(superposition,[],[f789,f787])).
fof(f789,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X7] : ! [X8,X38] : s(X7,X8) = s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X8),s(X6,X38)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.pair.FST')).
fof(f37211,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f37210,f983])).
fof(f983,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f602,f598])).
fof(f598,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSound.typeSound_46_prove_ltb',aHL_TRUTH)).
fof(f602,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f37210,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f979,f1132])).
fof(f1132,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f633])).
fof(f633,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f446])).
fof(f446,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X33] : (s(bool,'F_0') = s(bool,X33) <=> ~p(s(bool,X33)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f37132,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f37131,f983])).
fof(f37131,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f37130,f1132])).
fof(f37130,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK14,sK15),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(sK14,sK15),X4),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))),s(sK14,sK16))))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f36600,f598])).
fof(f36600,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(sK14,sK15),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(sK14,sK15),X4),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))),s(sK14,sK16))))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | $spl2),
  inference(superposition,[],[f802,f1257])).
fof(f1257,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f981,f629])).
fof(f629,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X33] : (s(bool,'T_0') = s(bool,X33) <=> p(s(bool,X33)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f802,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.list.list'('type.pair.prod'(X0,X1)),X2))),s(X0,X3))))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X2))))))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.list.list'('type.pair.prod'(X0,X1)),X2))),s(X0,X3))))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))),
  inference(ennf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.list.list'('type.pair.prod'(X0,X1)),X2))),s(X0,X3))))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(X0,X1),X4),s('type.list.list'('type.pair.prod'(X0,X1)),X2))),s(X0,X3))))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X2)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : ! [X11,X12,X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.MAP_2'(s(fun('type.pair.prod'(X7,X6),X7),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X7,X6)),X11)))))))) => p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X12),s(X6,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X7,X6),'const.finite_map.FUPDATE_LIST_2'(s('type.finite_map.fmap'(X7,X6),X13),s('type.list.list'('type.pair.prod'(X7,X6)),X11))),s(X7,X12))))),s(fun('type.pair.prod'(X7,X6),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X7,X6)),X11)))))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.finite_map.finite_map_|49|92_prove_ltb')).
fof(f16610,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14240,f16008])).
fof(f16008,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f823])).
fof(f823,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK37(X5,X4,X3,X2,X1,X0))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK37(X5,X4,X3,X2,X1,X0))) | s(X1,X3) = s(X1,X5)) & (s(X1,X3) != s(X1,X5) | p(s(bool,sK37(X5,X4,X3,X2,X1,X0)))) & s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),'const.list.CONS_2'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X2))),s('type.list.list'('type.pair.prod'(X1,X0)),X4))),s(X1,X5))) = s('type.option.option'(X0),'const.bool.COND_3'(s(bool,sK37(X5,X4,X3,X2,X1,X0)),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))),s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),X4),s(X1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37])],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(X1,X3) = s(X1,X5)) & (s(X1,X3) != s(X1,X5) | p(s(bool,X6))) & s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X6),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))),s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),X4),s(X1,X5))))) = s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),'const.list.CONS_2'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X2))),s('type.list.list'('type.pair.prod'(X1,X0)),X4))),s(X1,X5))))),
  inference(flattening,[],[f537])).
fof(f537,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(X1,X3) = s(X1,X5)) & (s(X1,X3) != s(X1,X5) | p(s(bool,X6)))) & s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X6),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))),s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),X4),s(X1,X5))))) = s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),'const.list.CONS_2'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X2))),s('type.list.list'('type.pair.prod'(X1,X0)),X4))),s(X1,X5))))),
  inference(nnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X3) = s(X1,X5)) & s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X6),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))),s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),X4),s(X1,X5))))) = s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),'const.list.CONS_2'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X2))),s('type.list.list'('type.pair.prod'(X1,X0)),X4))),s(X1,X5))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X3) = s(X1,X5)) & s('type.option.option'(X0),'const.bool.COND_3'(s(bool,X6),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))),s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),X4),s(X1,X5))))) = s('type.option.option'(X0),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X1,X0)),'const.list.CONS_2'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X2))),s('type.list.list'('type.pair.prod'(X1,X0)),X4))),s(X1,X5))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X7,X6] : ! [X38,X8,X33,X42] : ? [X66] : ((p(s(bool,X66)) <=> s(X6,X8) = s(X6,X42)) & s('type.option.option'(X7),'const.bool.COND_3'(s(bool,X66),s('type.option.option'(X7),'const.option.SOME_1'(s(X7,X38))),s('type.option.option'(X7),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X6,X7)),X33),s(X6,X42))))) = s('type.option.option'(X7),'const.alist.ALOOKUP_2'(s('type.list.list'('type.pair.prod'(X6,X7)),'const.list.CONS_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X8),s(X7,X38))),s('type.list.list'('type.pair.prod'(X6,X7)),X33))),s(X6,X42))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.alist.ALOOKUP_def|split|1')).
fof(f14240,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f632])).
fof(f632,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f446])).
fof(f16609,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14244,f598])).
fof(f14244,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1012,f14233,f734])).
fof(f734,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X34,X35] : ((p(s(bool,X34)) => p(s(bool,X35))) => ((p(s(bool,X35)) => p(s(bool,X34))) => s(bool,X34) = s(bool,X35)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f1012,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f983,f628])).
fof(f628,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f444])).
fof(f16608,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14250,f598])).
fof(f14250,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1012,f14233,f734])).
fof(f16607,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14259,f14233])).
fof(f14259,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1137,f14233,f763])).
fof(f763,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f519])).
fof(f519,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1137,plain,(
  ( ! [X0] : (sP1('T_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f764])).
fof(f764,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f16606,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14260,f14233])).
fof(f14260,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2216,f14233,f763])).
fof(f2216,plain,(
  ( ! [X2,X0,X1] : (sP1('T_0',sK37('F_0',X0,'T_0',X1,bool,X2))) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f764])).
fof(f2203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1012,f823])).
fof(f16605,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14261,f14233])).
fof(f14261,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2556,f14233,f763])).
fof(f2556,plain,(
  ( ! [X2,X0,X1] : (sP1('T_0',sK37('T_0',X0,'F_0',X1,bool,X2))) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f764])).
fof(f2204,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f1012,f823])).
fof(f16604,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14262,f598])).
fof(f14262,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1028,f14233,f763])).
fof(f1028,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f983,f598,f767])).
fof(f767,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f520])).
fof(f16603,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14263,f598])).
fof(f14263,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1138,f14233,f763])).
fof(f1138,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f767])).
fof(f16602,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14264,f598])).
fof(f14264,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2217,f14233,f763])).
fof(f2217,plain,(
  ( ! [X2,X0,X1] : (sP1(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f767])).
fof(f16601,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14265,f598])).
fof(f14265,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2557,f14233,f763])).
fof(f2557,plain,(
  ( ! [X2,X0,X1] : (sP1(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f767])).
fof(f16600,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14270,f598])).
fof(f14270,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1023,f14233,f763])).
fof(f1023,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f598,f764])).
fof(f16599,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14271,f598])).
fof(f14271,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1137,f14233,f763])).
fof(f16598,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14272,f598])).
fof(f14272,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2216,f14233,f763])).
fof(f16597,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14273,f598])).
fof(f14273,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2556,f14233,f763])).
fof(f16596,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14275,f14233])).
fof(f14275,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1138,f14233,f763])).
fof(f16595,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14276,f14233])).
fof(f14276,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2217,f14233,f763])).
fof(f16594,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14277,f14233])).
fof(f14277,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2557,f14233,f763])).
fof(f16593,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14307,f14233])).
fof(f14307,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1142,f14233,f839])).
fof(f839,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f544])).
fof(f544,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f543])).
fof(f543,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1142,plain,(
  ( ! [X0,X1] : (~sP2(X0,'F_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f836])).
fof(f836,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f545])).
fof(f16592,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14308,f14233])).
fof(f14308,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2221,f14233,f839])).
fof(f2221,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,'F_0',sK37('F_0',X1,'T_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f836])).
fof(f16591,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14309,f14233])).
fof(f14309,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2561,f14233,f839])).
fof(f2561,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,'F_0',sK37('T_0',X1,'F_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f836])).
fof(f16590,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14310,f14233])).
fof(f14310,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1040,f14233,f839])).
fof(f1040,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f983,f983,f837])).
fof(f837,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f545])).
fof(f16589,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14311,f14233])).
fof(f14311,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1145,f14233,f839])).
fof(f1145,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0',X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f837])).
fof(f16588,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14312,f14233])).
fof(f14312,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2224,f14233,f839])).
fof(f2224,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(sK37('F_0',X0,'T_0',X1,bool,X2),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f837])).
fof(f16587,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14313,f14233])).
fof(f14313,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2564,f14233,f839])).
fof(f2564,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(sK37('T_0',X0,'F_0',X1,bool,X2),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f837])).
fof(f16586,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14315,f14233])).
fof(f14315,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f14233,f839])).
fof(f1148,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f837])).
fof(f16585,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14317,f14233])).
fof(f14317,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2227,f14233,f839])).
fof(f2227,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('F_0',sK37('F_0',X0,'T_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f837])).
fof(f16584,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14319,f14233])).
fof(f14319,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2567,f14233,f839])).
fof(f2567,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('F_0',sK37('T_0',X0,'F_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f837])).
fof(f16583,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14332,f14233])).
fof(f14332,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1139,f14233,f839])).
fof(f1139,plain,(
  ( ! [X0,X1] : (~sP2(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f836])).
fof(f16582,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14333,f14233])).
fof(f14333,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f14233,f839])).
fof(f16581,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14334,f14233])).
fof(f14334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2218,f14233,f839])).
fof(f2218,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,sK37('F_0',X1,'T_0',X2,bool,X3),'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f836])).
fof(f16580,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14335,f14233])).
fof(f14335,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2227,f14233,f839])).
fof(f16579,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14336,f14233])).
fof(f14336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2558,f14233,f839])).
fof(f2558,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(X0,sK37('T_0',X1,'F_0',X2,bool,X3),'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f836])).
fof(f16578,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14337,f14233])).
fof(f14337,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2567,f14233,f839])).
fof(f16577,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14349,f14233])).
fof(f14349,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1139,f14233,f840])).
fof(f840,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f545])).
fof(f16576,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14350,f14233])).
fof(f14350,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f14233,f840])).
fof(f16575,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14351,f14233])).
fof(f14351,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2218,f14233,f840])).
fof(f16574,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14352,f14233])).
fof(f14352,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2227,f14233,f840])).
fof(f16573,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14353,f14233])).
fof(f14353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2558,f14233,f840])).
fof(f16572,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14354,f14233])).
fof(f14354,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2567,f14233,f840])).
fof(f16571,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14358,f14233])).
fof(f14358,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f840])).
fof(f1037,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f983,f836])).
fof(f16570,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14359,f14233])).
fof(f14359,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1142,f14233,f840])).
fof(f16569,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14360,f14233])).
fof(f14360,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2221,f14233,f840])).
fof(f16568,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14361,f14233])).
fof(f14361,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2561,f14233,f840])).
fof(f16567,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14363,f14233])).
fof(f14363,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1145,f14233,f840])).
fof(f16566,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14364,f14233])).
fof(f14364,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2224,f14233,f840])).
fof(f16565,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14365,f14233])).
fof(f14365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2564,f14233,f840])).
fof(f16564,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14366,f14233])).
fof(f14366,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1139,f14233,f840])).
fof(f16563,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14368,f14233])).
fof(f14368,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2218,f14233,f840])).
fof(f16562,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14370,f14233])).
fof(f14370,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2558,f14233,f840])).
fof(f16561,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14377,f14233])).
fof(f14377,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1142,f14233,f841])).
fof(f841,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f545])).
fof(f16560,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14378,f14233])).
fof(f14378,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2221,f14233,f841])).
fof(f16559,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14379,f14233])).
fof(f14379,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2561,f14233,f841])).
fof(f16558,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14380,f14233])).
fof(f14380,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1040,f14233,f841])).
fof(f16557,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14381,f14233])).
fof(f14381,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1145,f14233,f841])).
fof(f16556,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14382,f14233])).
fof(f14382,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2224,f14233,f841])).
fof(f16555,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14383,f14233])).
fof(f14383,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2564,f14233,f841])).
fof(f16554,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14385,f14233])).
fof(f14385,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f14233,f841])).
fof(f16553,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14387,f14233])).
fof(f14387,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2227,f14233,f841])).
fof(f16552,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14389,f14233])).
fof(f14389,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2567,f14233,f841])).
fof(f16551,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14394,f14233])).
fof(f14394,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f841])).
fof(f16550,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14395,f14233])).
fof(f14395,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1142,f14233,f841])).
fof(f16549,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14396,f14233])).
fof(f14396,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2221,f14233,f841])).
fof(f16548,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14397,f14233])).
fof(f14397,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2561,f14233,f841])).
fof(f16547,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14399,f14233])).
fof(f14399,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1145,f14233,f841])).
fof(f16546,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14400,f14233])).
fof(f14400,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2224,f14233,f841])).
fof(f16545,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14401,f14233])).
fof(f14401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2564,f14233,f841])).
fof(f16544,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14402,f14233])).
fof(f14402,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1139,f14233,f841])).
fof(f16543,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14404,f14233])).
fof(f14404,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2218,f14233,f841])).
fof(f16542,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14406,f14233])).
fof(f14406,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2558,f14233,f841])).
fof(f16541,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14420,f15928])).
fof(f15928,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f965])).
fof(f965,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f857])).
fof(f857,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f552])).
fof(f552,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((~p(s(bool,X2)) | p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X2)) & ~p(s(bool,X0))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f551])).
fof(f551,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & (((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f14420,plain,(
  ( ! [X0,X1] : (~sP3('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f852])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f552])).
fof(f16540,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14430,f15928])).
fof(f14430,plain,(
  ( ! [X0,X1] : (~sP3('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f853])).
fof(f853,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f552])).
fof(f16539,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14432,f598])).
fof(f14432,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1374,f983,f14233,f865])).
fof(f865,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f556])).
fof(f556,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1374,plain,(
  sP7('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f1012,f881])).
fof(f881,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f340,f428,f427,f426,f425])).
fof(f425,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f426,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f427,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f340,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X41,X42,X43] : ((p(s(bool,X43)) <=> s(bool,X41) = s(bool,X42)) <=> ((p(s(bool,X43)) | p(s(bool,X42)) | p(s(bool,X41))) & (p(s(bool,X43)) | ~p(s(bool,X41)) | ~p(s(bool,X42))) & (p(s(bool,X42)) | ~p(s(bool,X41)) | ~p(s(bool,X43))) & (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f16538,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14434,f14233])).
fof(f14434,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f983,f14233,f865])).
fof(f1409,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f598,f985,f882])).
fof(f882,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f566])).
fof(f985,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(forward_demodulation,[],[f685,f609])).
fof(f609,plain,(
  ( ! [X0] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.list.LIST_TO_SET|split|0')).
fof(f685,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X7] : ! [X8] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.NIL_0')))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.list.MEM|split|0')).
fof(f16537,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14435,f14233])).
fof(f14435,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1503,f983,f14233,f865])).
fof(f1503,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f598,f700,f882])).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X7] : ! [X34,X35] : s(X7,X35) = s(X7,'const.bool.COND_3'(s(bool,'F_0'),s(X7,X34),s(X7,X35)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f16536,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14436,f14233])).
fof(f14436,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1613,f983,f14233,f865])).
fof(f1613,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f598,f701,f882])).
fof(f701,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X7] : ! [X34,X35] : s(X7,X34) = s(X7,'const.bool.COND_3'(s(bool,'T_0'),s(X7,X34),s(X7,X35)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f16535,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14437,f598])).
fof(f14437,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2284,f983,f14233,f865])).
fof(f2284,plain,(
  ( ! [X2,X0,X1] : (sP7(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1012,f2203,f881])).
fof(f16534,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14438,f598])).
fof(f14438,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2624,f983,f14233,f865])).
fof(f2624,plain,(
  ( ! [X2,X0,X1] : (sP7(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1012,f2204,f881])).
fof(f16533,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14486,f598])).
fof(f14486,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f3352,f983,f14233,f865])).
fof(f3352,plain,(
  sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f997,f993,f998,f983,f867])).
fof(f867,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | p(s(bool,X1)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f557])).
fof(f998,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f878])).
fof(f878,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f425])).
fof(f993,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f983,f871])).
fof(f871,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f560])).
fof(f560,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f559])).
fof(f559,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f427])).
fof(f997,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f983,f876])).
fof(f876,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f562])).
fof(f562,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f561])).
fof(f561,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f426])).
fof(f16532,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14487,f598])).
fof(f14487,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1409,f983,f14233,f865])).
fof(f16531,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14488,f598])).
fof(f14488,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1503,f983,f14233,f865])).
fof(f16530,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14489,f598])).
fof(f14489,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1613,f983,f14233,f865])).
fof(f16529,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14490,f14233])).
fof(f14490,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2284,f983,f14233,f865])).
fof(f16528,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14491,f14233])).
fof(f14491,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2624,f983,f14233,f865])).
fof(f16527,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14496,f14833])).
fof(f14833,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f880])).
fof(f880,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f565])).
fof(f14496,plain,(
  ~sP4('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f992,f14233,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f557])).
fof(f992,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f870])).
fof(f870,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f560])).
fof(f995,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f874])).
fof(f874,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f563])).
fof(f1390,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f598,f1012,f883])).
fof(f883,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f566])).
fof(f16526,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14497,f14833])).
fof(f14497,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f992,f14233,f866])).
fof(f2246,plain,(
  ( ! [X2,X0,X1] : (~sP7('T_0','T_0',sK37('F_0',X0,'T_0',X1,bool,X2))) )),
  inference(unit_resulting_resolution,[],[f598,f598,f2203,f865])).
fof(f16525,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14498,f14833])).
fof(f14498,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f992,f14233,f866])).
fof(f2586,plain,(
  ( ! [X2,X0,X1] : (~sP7('T_0','T_0',sK37('T_0',X0,'F_0',X1,bool,X2))) )),
  inference(unit_resulting_resolution,[],[f598,f598,f2204,f865])).
fof(f16524,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14499,f14833])).
fof(f14499,plain,(
  ~sP4('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f992,f14233,f866])).
fof(f996,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f983,f875])).
fof(f875,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f563])).
fof(f16523,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14500,f14833])).
fof(f14500,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f992,f14233,f866])).
fof(f2263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(X0,sK37('F_0',X1,'T_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f875])).
fof(f16522,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14501,f14833])).
fof(f14501,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f992,f14233,f866])).
fof(f2603,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(X0,sK37('T_0',X1,'F_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f875])).
fof(f16521,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14502,f14833])).
fof(f14502,plain,(
  ~sP4('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f993,f14233,f866])).
fof(f16520,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14507,f14833])).
fof(f14507,plain,(
  ~sP4('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f993,f14233,f866])).
fof(f16519,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14509,f14833])).
fof(f14509,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1502,f1163,f1161,f14233,f866])).
fof(f1161,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f871])).
fof(f1163,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f875])).
fof(f1502,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f983,f700,f884])).
fof(f884,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f566])).
fof(f16518,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14510,f14833])).
fof(f14510,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))),s(bool,X1)),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1612,f1163,f1161,f14233,f866])).
fof(f1612,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f983,f701,f884])).
fof(f16517,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14511,f14833])).
fof(f14511,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f2256,f14233,f866])).
fof(f2256,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,sK37('F_0',X1,'T_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f871])).
fof(f16516,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14512,f14833])).
fof(f14512,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,sK37('F_0',X0,'T_0',X1,bool,X2))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f2263,f2256,f14233,f866])).
fof(f16515,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14513,f14833])).
fof(f14513,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)),s(bool,X3)),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f2263,f2256,f14233,f866])).
fof(f16514,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14514,f14833])).
fof(f14514,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f2256,f14233,f866])).
fof(f16513,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14515,f14833])).
fof(f14515,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f2596,f14233,f866])).
fof(f2596,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,sK37('T_0',X1,'F_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f871])).
fof(f16512,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14516,f14833])).
fof(f14516,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,sK37('T_0',X0,'F_0',X1,bool,X2))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f2603,f2596,f14233,f866])).
fof(f16511,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14517,f14833])).
fof(f14517,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)),s(bool,X3)),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f2603,f2596,f14233,f866])).
fof(f16510,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14518,f14833])).
fof(f14518,plain,(
  ( ! [X2,X0,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f2596,f14233,f866])).
fof(f16509,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14523,f14833])).
fof(f14523,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1502,f1163,f994,f14233,f866])).
fof(f994,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f983,f872])).
fof(f872,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f560])).
fof(f16508,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14524,f14833])).
fof(f14524,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))),s(bool,X1)),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1612,f1163,f994,f14233,f866])).
fof(f16507,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14525,f14833])).
fof(f14525,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,sK37('F_0',X0,'T_0',X1,bool,X2))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f2263,f994,f14233,f866])).
fof(f16506,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14526,f14833])).
fof(f14526,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('F_0',X0,'T_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)),s(bool,X3)),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f2263,f994,f14233,f866])).
fof(f16505,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14527,f14833])).
fof(f14527,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,sK37('T_0',X0,'F_0',X1,bool,X2))),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f2603,f994,f14233,f866])).
fof(f16504,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14528,f14833])).
fof(f14528,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4(sK37('T_0',X0,'F_0',X1,bool,X2),'const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)),s(bool,X3)),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f2603,f994,f14233,f866])).
fof(f16503,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14533,f14833])).
fof(f14533,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))),'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1505,f1164,f994,f14233,f866])).
fof(f1164,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f876])).
fof(f1505,plain,(
  ( ! [X0,X1] : (~sP7('F_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f983,f700,f884])).
fof(f16502,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14534,f14833])).
fof(f14534,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))),s(bool,X1)),'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1615,f1164,f994,f14233,f866])).
fof(f1615,plain,(
  ( ! [X0,X1] : (~sP7('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f983,f701,f884])).
fof(f16501,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14535,f14833])).
fof(f14535,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,sK37('F_0',X1,'T_0',X2,bool,X3))),sK37('F_0',X1,'T_0',X2,bool,X3),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1505,f2264,f994,f14233,f866])).
fof(f2264,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK37('F_0',X0,'T_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f876])).
fof(f16500,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14536,f14833])).
fof(f14536,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)),s(bool,X3)),sK37('F_0',X0,'T_0',X1,bool,X2),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1615,f2264,f994,f14233,f866])).
fof(f16499,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14537,f14833])).
fof(f14537,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,sK37('T_0',X1,'F_0',X2,bool,X3))),sK37('T_0',X1,'F_0',X2,bool,X3),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1505,f2604,f994,f14233,f866])).
fof(f2604,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5(sK37('T_0',X0,'F_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f876])).
fof(f16498,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14538,f14833])).
fof(f14538,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.bool.COND_3'(s(bool,'T_0'),s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)),s(bool,X3)),sK37('T_0',X0,'F_0',X1,bool,X2),'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1615,f2604,f994,f14233,f866])).
fof(f16497,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14540,f14773])).
fof(f14773,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f870])).
fof(f14540,plain,(
  ~sP6('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f998,f14233,f866])).
fof(f16496,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14541,f14773])).
fof(f14541,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f998,f14233,f866])).
fof(f16495,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14542,f14773])).
fof(f14542,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f998,f14233,f866])).
fof(f16494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14543,f14773])).
fof(f14543,plain,(
  ~sP6('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f998,f14233,f866])).
fof(f16493,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14544,f14773])).
fof(f14544,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f998,f14233,f866])).
fof(f16492,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14545,f14773])).
fof(f14545,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f998,f14233,f866])).
fof(f16491,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14569,f14827])).
fof(f14827,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f874])).
fof(f14569,plain,(
  ~sP5('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f992,f998,f14233,f866])).
fof(f16490,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14570,f14827])).
fof(f14570,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f992,f998,f14233,f866])).
fof(f16489,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14571,f14827])).
fof(f14571,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f992,f998,f14233,f866])).
fof(f16488,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14581,f14827])).
fof(f14581,plain,(
  ~sP5('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f993,f998,f14233,f866])).
fof(f16487,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14596,f14827])).
fof(f14596,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2256,f998,f14233,f866])).
fof(f16486,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14605,f14827])).
fof(f14605,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2596,f998,f14233,f866])).
fof(f16485,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14614,f14773])).
fof(f14614,plain,(
  ~sP6('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f999,f14233,f866])).
fof(f999,plain,(
  ( ! [X0,X1] : (sP4(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f598,f879])).
fof(f879,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f565])).
fof(f16484,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14615,f14773])).
fof(f14615,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f999,f14233,f866])).
fof(f16483,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14616,f14773])).
fof(f14616,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f999,f14233,f866])).
fof(f16482,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14617,f14773])).
fof(f14617,plain,(
  ~sP6('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f999,f14233,f866])).
fof(f16481,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14618,f14773])).
fof(f14618,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f999,f14233,f866])).
fof(f16480,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14619,f14773])).
fof(f14619,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f999,f14233,f866])).
fof(f16479,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14620,f14827])).
fof(f14620,plain,(
  ~sP5('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f992,f999,f14233,f866])).
fof(f16478,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14621,f14827])).
fof(f14621,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f992,f999,f14233,f866])).
fof(f16477,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14622,f14827])).
fof(f14622,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f992,f999,f14233,f866])).
fof(f16476,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14628,f14885])).
fof(f14885,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1012,f14233,f883])).
fof(f14628,plain,(
  ( ! [X0] : (sP7(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f996,f992,f999,f14233,f866])).
fof(f16475,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14635,f14827])).
fof(f14635,plain,(
  ~sP5('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f993,f999,f14233,f866])).
fof(f16474,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14636,f14885])).
fof(f14636,plain,(
  ( ! [X0] : (sP7(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f996,f993,f999,f14233,f866])).
fof(f16473,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14638,f1390])).
fof(f14638,plain,(
  sP7('T_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f995,f993,f999,f14233,f866])).
fof(f16472,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14642,f14827])).
fof(f14642,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2256,f999,f14233,f866])).
fof(f16471,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14645,f2246])).
fof(f14645,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0',sK37('F_0',X0,'T_0',X1,bool,X2))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f995,f2256,f999,f14233,f866])).
fof(f16470,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14647,f14827])).
fof(f14647,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2596,f999,f14233,f866])).
fof(f16469,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14650,f2586])).
fof(f14650,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0',sK37('T_0',X0,'F_0',X1,bool,X2))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f995,f2596,f999,f14233,f866])).
fof(f16468,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14661,f14885])).
fof(f14661,plain,(
  ( ! [X0] : (sP7('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0','F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f996,f1162,f999,f14233,f866])).
fof(f1162,plain,(
  ( ! [X2,X0,X1] : (sP6('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f872])).
fof(f16467,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14699,f14233])).
fof(f14699,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1161,f983,f14233,f869])).
fof(f869,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f560])).
fof(f16466,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14700,f14233])).
fof(f14700,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2256,f983,f14233,f869])).
fof(f16465,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14701,f14233])).
fof(f14701,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2596,f983,f14233,f869])).
fof(f16464,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14702,f14233])).
fof(f14702,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1488,f983,f14233,f869])).
fof(f1488,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1412,f864])).
fof(f864,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f557])).
fof(f1412,plain,(
  ( ! [X0,X1] : (sP7('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f598,f985,f882])).
fof(f16463,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14703,f14233])).
fof(f14703,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1743,f983,f14233,f869])).
fof(f1743,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f1506,f864])).
fof(f1506,plain,(
  ( ! [X0,X1] : (sP7('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f598,f700,f882])).
fof(f16462,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14704,f14233])).
fof(f14704,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1862,f983,f14233,f869])).
fof(f1862,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f1616,f864])).
fof(f1616,plain,(
  ( ! [X0,X1] : (sP7('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f598,f701,f882])).
fof(f16461,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14705,f14233])).
fof(f14705,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f994,f983,f14233,f869])).
fof(f16460,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14706,f14233])).
fof(f14706,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1162,f983,f14233,f869])).
fof(f16459,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14707,f14233])).
fof(f14707,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2257,f983,f14233,f869])).
fof(f2257,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(sK37('F_0',X0,'T_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f872])).
fof(f16458,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14708,f14233])).
fof(f14708,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2597,f983,f14233,f869])).
fof(f2597,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(sK37('T_0',X0,'F_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f872])).
fof(f16457,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14716,f14773])).
fof(f14716,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f598,f14233,f869])).
fof(f16456,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14732,f14773])).
fof(f14732,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f14233,f869])).
fof(f16455,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14734,f14233])).
fof(f14734,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f993,f983,f14233,f869])).
fof(f16454,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14735,f14233])).
fof(f14735,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1161,f983,f14233,f869])).
fof(f16453,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14736,f14233])).
fof(f14736,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2256,f983,f14233,f869])).
fof(f16452,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14737,f14233])).
fof(f14737,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2596,f983,f14233,f869])).
fof(f16451,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14738,f598])).
fof(f14738,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1488,f983,f14233,f869])).
fof(f16450,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14739,f598])).
fof(f14739,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1743,f983,f14233,f869])).
fof(f16449,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14740,f598])).
fof(f14740,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1862,f983,f14233,f869])).
fof(f16448,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14742,f14233])).
fof(f14742,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1162,f983,f14233,f869])).
fof(f16447,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14743,f14233])).
fof(f14743,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2257,f983,f14233,f869])).
fof(f16446,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14744,f14233])).
fof(f14744,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2597,f983,f14233,f869])).
fof(f16445,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14755,f14773])).
fof(f14755,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f598,f14233,f869])).
fof(f16444,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14771,f14773])).
fof(f14771,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f14233,f869])).
fof(f16443,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14776,f14233])).
fof(f14776,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1163,f983,f14233,f873])).
fof(f873,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f563])).
fof(f16442,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14777,f14233])).
fof(f14777,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2263,f983,f14233,f873])).
fof(f16441,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14778,f14233])).
fof(f14778,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2603,f983,f14233,f873])).
fof(f16440,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14779,f14233])).
fof(f14779,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f997,f983,f14233,f873])).
fof(f16439,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14780,f14233])).
fof(f14780,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1164,f983,f14233,f873])).
fof(f16438,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14781,f14233])).
fof(f14781,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2264,f983,f14233,f873])).
fof(f16437,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14782,f14233])).
fof(f14782,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2604,f983,f14233,f873])).
fof(f16436,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14788,f14827])).
fof(f14788,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f598,f14233,f873])).
fof(f16435,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14799,f14827])).
fof(f14799,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f14233,f873])).
fof(f16434,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14815,f14233])).
fof(f14815,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f996,f983,f14233,f873])).
fof(f16433,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14816,f14233])).
fof(f14816,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1163,f983,f14233,f873])).
fof(f16432,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14817,f14233])).
fof(f14817,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2263,f983,f14233,f873])).
fof(f16431,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14818,f14233])).
fof(f14818,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2603,f983,f14233,f873])).
fof(f16430,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14820,f14233])).
fof(f14820,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1164,f983,f14233,f873])).
fof(f16429,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14821,f14233])).
fof(f14821,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2264,f983,f14233,f873])).
fof(f16428,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14822,f14233])).
fof(f14822,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2604,f983,f14233,f873])).
fof(f16427,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14823,f14827])).
fof(f14823,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f598,f983,f14233,f873])).
fof(f16426,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14825,f14827])).
fof(f14825,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f983,f14233,f873])).
fof(f16425,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14839,f700])).
fof(f14839,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1505,f14233,f882])).
fof(f16424,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14840,f701])).
fof(f14840,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1615,f14233,f882])).
fof(f16423,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14842,f985])).
fof(f14842,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1411,f14233,f882])).
fof(f1411,plain,(
  ( ! [X0,X1] : (~sP7('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f983,f985,f884])).
fof(f16422,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14844,f985])).
fof(f14844,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1408,f14233,f882])).
fof(f1408,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f985,f884])).
fof(f16421,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14845,f700])).
fof(f14845,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f14233,f882])).
fof(f16420,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14846,f701])).
fof(f14846,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f14233,f882])).
fof(f16419,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14900,f15275])).
fof(f15275,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f14233,f14233,f898])).
fof(f898,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f572])).
fof(f572,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f342,f431,f430])).
fof(f430,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f431,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f342,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X41,X42,X43] : ((p(s(bool,X43)) <=> (p(s(bool,X42)) & p(s(bool,X41)))) <=> ((p(s(bool,X43)) | ~p(s(bool,X42)) | ~p(s(bool,X41))) & (p(s(bool,X42)) | ~p(s(bool,X43))) & (p(s(bool,X41)) | ~p(s(bool,X43)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.sat.sat_23_prove_ltb')).
fof(f14900,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f886])).
fof(f886,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f569])).
fof(f569,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f568])).
fof(f568,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f567])).
fof(f567,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f431])).
fof(f16418,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14913,f15275])).
fof(f14913,plain,(
  ( ! [X0,X1] : (~sP9(X0,'F_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f887])).
fof(f887,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f569])).
fof(f16417,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14916,f14233])).
fof(f14916,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1165,f14233,f888])).
fof(f888,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f569])).
fof(f1165,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f886])).
fof(f1169,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f894])).
fof(f894,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f570])).
fof(f570,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f430])).
fof(f16416,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14917,f14233])).
fof(f14917,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1001,f1165,f14233,f888])).
fof(f1001,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f893])).
fof(f893,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f571])).
fof(f16415,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14918,f14233])).
fof(f14918,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1003,f1165,f14233,f888])).
fof(f1003,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f983,f895])).
fof(f895,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f571])).
fof(f16414,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14919,f14233])).
fof(f14919,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1165,f14233,f888])).
fof(f1170,plain,(
  ( ! [X2,X0,X1] : (sP8('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f895])).
fof(f16413,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14920,f14233])).
fof(f14920,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2303,f1165,f14233,f888])).
fof(f2303,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK37('F_0',X0,'T_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f895])).
fof(f16412,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14921,f14233])).
fof(f14921,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2643,f1165,f14233,f888])).
fof(f2643,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK37('T_0',X0,'F_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f895])).
fof(f16411,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14922,f14233])).
fof(f14922,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2291,f14233,f888])).
fof(f2291,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('T_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f886])).
fof(f2302,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,sK37('F_0',X1,'T_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f894])).
fof(f16410,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14923,f14233])).
fof(f14923,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2291,f14233,f888])).
fof(f16409,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14924,f14233])).
fof(f14924,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f2291,f14233,f888])).
fof(f16408,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14925,f14233])).
fof(f14925,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f2291,f14233,f888])).
fof(f16407,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14926,f14233])).
fof(f14926,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2291,f14233,f888])).
fof(f16406,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14927,f14233])).
fof(f14927,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2291,f14233,f888])).
fof(f16405,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14928,f14233])).
fof(f14928,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2631,f14233,f888])).
fof(f2631,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('T_0',X0,sK37('T_0',X1,'F_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f886])).
fof(f2642,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,sK37('T_0',X1,'F_0',X2,bool,X3),X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f894])).
fof(f16404,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14929,f14233])).
fof(f14929,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2631,f14233,f888])).
fof(f16403,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14930,f14233])).
fof(f14930,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f2631,f14233,f888])).
fof(f16402,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14931,f14233])).
fof(f14931,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f2631,f14233,f888])).
fof(f16401,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14932,f14233])).
fof(f14932,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2631,f14233,f888])).
fof(f16400,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14933,f14233])).
fof(f14933,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2631,f14233,f888])).
fof(f16399,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14934,f14233])).
fof(f14934,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f1063,f14233,f888])).
fof(f1063,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f598,f983,f887])).
fof(f16398,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14935,f14233])).
fof(f14935,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f1063,f14233,f888])).
fof(f16397,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14937,f14233])).
fof(f14937,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1063,f14233,f888])).
fof(f16396,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14938,f14233])).
fof(f14938,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f1063,f14233,f888])).
fof(f16395,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14939,f14233])).
fof(f14939,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f1063,f14233,f888])).
fof(f16394,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14940,f14233])).
fof(f14940,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1167,f14233,f888])).
fof(f1167,plain,(
  ( ! [X0,X1] : (~sP9('T_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f887])).
fof(f16393,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14941,f14233])).
fof(f14941,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1001,f1167,f14233,f888])).
fof(f16392,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14943,f14233])).
fof(f14943,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1167,f14233,f888])).
fof(f16391,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14944,f14233])).
fof(f14944,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2302,f1167,f14233,f888])).
fof(f16390,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14945,f14233])).
fof(f14945,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2642,f1167,f14233,f888])).
fof(f16389,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14946,f14233])).
fof(f14946,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2294,f14233,f888])).
fof(f2294,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f887])).
fof(f16388,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14947,f14233])).
fof(f14947,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2294,f14233,f888])).
fof(f16387,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14949,f14233])).
fof(f14949,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f2294,f14233,f888])).
fof(f16386,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14950,f14233])).
fof(f14950,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2294,f14233,f888])).
fof(f16385,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14951,f14233])).
fof(f14951,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2294,f14233,f888])).
fof(f16384,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14952,f14233])).
fof(f14952,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2634,f14233,f888])).
fof(f2634,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f887])).
fof(f16383,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14953,f14233])).
fof(f14953,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2634,f14233,f888])).
fof(f16382,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14955,f14233])).
fof(f14955,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f2634,f14233,f888])).
fof(f16381,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14956,f14233])).
fof(f14956,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2634,f14233,f888])).
fof(f16380,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14957,f14233])).
fof(f14957,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2634,f14233,f888])).
fof(f16379,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14963,f15141])).
fof(f15141,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f893])).
fof(f14963,plain,(
  ~sP8('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f598,f14233,f888])).
fof(f16378,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14965,f15141])).
fof(f14965,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0','T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f598,f14233,f888])).
fof(f16377,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14967,f15141])).
fof(f14967,plain,(
  ( ! [X2,X0,X1] : (~sP8(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f598,f14233,f888])).
fof(f16376,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14969,f15141])).
fof(f14969,plain,(
  ( ! [X2,X0,X1] : (~sP8(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f598,f14233,f888])).
fof(f16375,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14983,f15141])).
fof(f14983,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f888])).
fof(f1062,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f598,f983,f886])).
fof(f16374,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14985,f15141])).
fof(f14985,plain,(
  ( ! [X0,X1] : (~sP8(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f888])).
fof(f16373,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14987,f15141])).
fof(f14987,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X0,sK37('F_0',X1,'T_0',X2,bool,X3),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f888])).
fof(f16372,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14989,f15141])).
fof(f14989,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X0,sK37('T_0',X1,'F_0',X2,bool,X3),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f888])).
fof(f16371,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14991,f15141])).
fof(f14991,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f888])).
fof(f16370,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f14993,f15141])).
fof(f14993,plain,(
  ( ! [X0,X1] : (~sP8('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f888])).
fof(f16369,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14995,f15141])).
fof(f14995,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f888])).
fof(f16368,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f14997,f15141])).
fof(f14997,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(sK37('T_0',X0,'F_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f888])).
fof(f16367,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15002,f14233])).
fof(f15002,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f1062,f14233,f888])).
fof(f1002,plain,(
  ( ! [X0,X1] : (sP8(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f983,f894])).
fof(f16366,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15003,f14233])).
fof(f15003,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f1062,f14233,f888])).
fof(f16365,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15005,f14233])).
fof(f15005,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1062,f14233,f888])).
fof(f16364,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15006,f14233])).
fof(f15006,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f1062,f14233,f888])).
fof(f16363,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15007,f14233])).
fof(f15007,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f1062,f14233,f888])).
fof(f16362,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15008,f14233])).
fof(f15008,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1165,f14233,f888])).
fof(f16361,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15009,f14233])).
fof(f15009,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1001,f1165,f14233,f888])).
fof(f16360,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15011,f14233])).
fof(f15011,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1165,f14233,f888])).
fof(f16359,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15012,f14233])).
fof(f15012,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2303,f1165,f14233,f888])).
fof(f16358,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15013,f14233])).
fof(f15013,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2643,f1165,f14233,f888])).
fof(f16357,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15014,f14233])).
fof(f15014,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2291,f14233,f888])).
fof(f16356,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15015,f14233])).
fof(f15015,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2291,f14233,f888])).
fof(f16355,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15017,f14233])).
fof(f15017,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f2291,f14233,f888])).
fof(f16354,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15018,f14233])).
fof(f15018,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2291,f14233,f888])).
fof(f16353,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15019,f14233])).
fof(f15019,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2291,f14233,f888])).
fof(f16352,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15020,f14233])).
fof(f15020,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2631,f14233,f888])).
fof(f16351,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15021,f14233])).
fof(f15021,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2631,f14233,f888])).
fof(f16350,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15023,f14233])).
fof(f15023,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f2631,f14233,f888])).
fof(f16349,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15024,f14233])).
fof(f15024,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2631,f14233,f888])).
fof(f16348,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15025,f14233])).
fof(f15025,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2631,f14233,f888])).
fof(f16347,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15027,f14233])).
fof(f15027,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1167,f14233,f888])).
fof(f16346,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15028,f14233])).
fof(f15028,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1001,f1167,f14233,f888])).
fof(f16345,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15029,f14233])).
fof(f15029,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1002,f1167,f14233,f888])).
fof(f16344,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15030,f14233])).
fof(f15030,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1167,f14233,f888])).
fof(f16343,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15031,f14233])).
fof(f15031,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2302,f1167,f14233,f888])).
fof(f16342,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15032,f14233])).
fof(f15032,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2642,f1167,f14233,f888])).
fof(f16341,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15033,f14233])).
fof(f15033,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2294,f14233,f888])).
fof(f16340,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15034,f14233])).
fof(f15034,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2294,f14233,f888])).
fof(f16339,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15035,f14233])).
fof(f15035,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f2294,f14233,f888])).
fof(f16338,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15036,f14233])).
fof(f15036,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f2294,f14233,f888])).
fof(f16337,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15037,f14233])).
fof(f15037,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2294,f14233,f888])).
fof(f16336,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15038,f14233])).
fof(f15038,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2294,f14233,f888])).
fof(f16335,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15039,f14233])).
fof(f15039,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2634,f14233,f888])).
fof(f16334,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15040,f14233])).
fof(f15040,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1001,f2634,f14233,f888])).
fof(f16333,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15041,f14233])).
fof(f15041,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f2634,f14233,f888])).
fof(f16332,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15042,f14233])).
fof(f15042,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f2634,f14233,f888])).
fof(f16331,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15043,f14233])).
fof(f15043,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2634,f14233,f888])).
fof(f16330,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15044,f14233])).
fof(f15044,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2634,f14233,f888])).
fof(f16329,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15050,f15141])).
fof(f15050,plain,(
  ~sP8('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f14233,f888])).
fof(f16328,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15052,f15141])).
fof(f15052,plain,(
  ( ! [X0] : (~sP8('T_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f598,f14233,f888])).
fof(f16327,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15054,f15141])).
fof(f15054,plain,(
  ( ! [X2,X0,X1] : (~sP8('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f598,f14233,f888])).
fof(f16326,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15056,f15141])).
fof(f15056,plain,(
  ( ! [X2,X0,X1] : (~sP8('T_0',sK37('T_0',X0,'F_0',X1,bool,X2),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f598,f14233,f888])).
fof(f16325,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15070,f15141])).
fof(f15070,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f888])).
fof(f16324,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15072,f15141])).
fof(f15072,plain,(
  ( ! [X0,X1] : (~sP8(X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f888])).
fof(f16323,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15074,f15141])).
fof(f15074,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X0,sK37('F_0',X1,'T_0',X2,bool,X3),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f888])).
fof(f16322,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15076,f15141])).
fof(f15076,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X0,sK37('T_0',X1,'F_0',X2,bool,X3),'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f888])).
fof(f16321,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15078,f15141])).
fof(f15078,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f888])).
fof(f16320,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15080,f15141])).
fof(f15080,plain,(
  ( ! [X0,X1] : (~sP8('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f888])).
fof(f16319,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15082,f15141])).
fof(f15082,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f888])).
fof(f16318,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15084,f15141])).
fof(f15084,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(sK37('T_0',X0,'F_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f888])).
fof(f16317,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15090,f14233])).
fof(f15090,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f983,f14233,f892])).
fof(f892,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f571])).
fof(f16316,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15091,f14233])).
fof(f15091,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f983,f14233,f892])).
fof(f16315,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15092,f14233])).
fof(f15092,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f983,f14233,f892])).
fof(f16314,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15093,f14233])).
fof(f15093,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f983,f14233,f892])).
fof(f16313,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15094,f14233])).
fof(f15094,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f983,f14233,f892])).
fof(f16312,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15095,f14233])).
fof(f15095,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f983,f14233,f892])).
fof(f16311,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15096,f14233])).
fof(f15096,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f983,f14233,f892])).
fof(f16310,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15102,f15141])).
fof(f15102,plain,(
  ( ! [X0] : (~sP8(X0,'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f598,f14233,f892])).
fof(f16309,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15113,f15141])).
fof(f15113,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f14233,f892])).
fof(f16308,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15129,f14233])).
fof(f15129,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f983,f14233,f892])).
fof(f16307,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15130,f14233])).
fof(f15130,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f983,f14233,f892])).
fof(f16306,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15131,f14233])).
fof(f15131,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f983,f14233,f892])).
fof(f16305,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15132,f14233])).
fof(f15132,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f983,f14233,f892])).
fof(f16304,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15134,f14233])).
fof(f15134,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f983,f14233,f892])).
fof(f16303,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15135,f14233])).
fof(f15135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f983,f14233,f892])).
fof(f16302,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15136,f14233])).
fof(f15136,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f983,f14233,f892])).
fof(f16301,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15137,f15141])).
fof(f15137,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f598,f983,f14233,f892])).
fof(f16300,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15139,f15141])).
fof(f15139,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f983,f14233,f892])).
fof(f16299,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15144,f14233])).
fof(f15144,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f598,f14233,f898])).
fof(f16298,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15145,f14233])).
fof(f15145,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f598,f14233,f898])).
fof(f16297,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15146,f14233])).
fof(f15146,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f598,f14233,f898])).
fof(f16296,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15147,f14233])).
fof(f15147,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f598,f14233,f898])).
fof(f16295,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15148,f14233])).
fof(f15148,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f598,f14233,f898])).
fof(f16294,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15149,f14233])).
fof(f15149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f598,f14233,f898])).
fof(f16293,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15150,f14233])).
fof(f15150,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f598,f14233,f898])).
fof(f16292,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15152,f14233])).
fof(f15152,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16291,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15153,f14233])).
fof(f15153,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16290,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15154,f14233])).
fof(f15154,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16289,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15155,f14233])).
fof(f15155,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f898])).
fof(f16288,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15156,f14233])).
fof(f15156,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16287,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15157,f14233])).
fof(f15157,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16286,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15158,f14233])).
fof(f15158,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16285,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15159,f598])).
fof(f15159,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f14233,f14233,f898])).
fof(f1850,plain,(
  ~sP9('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1845,f885])).
fof(f885,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f569])).
fof(f1845,plain,(
  ~sP8('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f598,f983,f598,f892])).
fof(f16284,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15160,f598])).
fof(f15160,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f2316,plain,(
  ( ! [X2,X0,X1] : (~sP9(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f598,f2203,f901])).
fof(f901,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f16283,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15161,f598])).
fof(f15161,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f2656,plain,(
  ( ! [X2,X0,X1] : (~sP9(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f598,f2204,f901])).
fof(f16282,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15162,f598])).
fof(f15162,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f598,f14233,f898])).
fof(f16281,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15163,f598])).
fof(f15163,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f598,f14233,f898])).
fof(f16280,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15164,f598])).
fof(f15164,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f598,f14233,f898])).
fof(f16279,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15165,f598])).
fof(f15165,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f598,f14233,f898])).
fof(f16278,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15167,f14233])).
fof(f15167,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f598,f14233,f898])).
fof(f16277,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15168,f14233])).
fof(f15168,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f598,f14233,f898])).
fof(f16276,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15173,f598])).
fof(f15173,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f898])).
fof(f16275,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15174,f598])).
fof(f15174,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16274,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15175,f598])).
fof(f15175,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16273,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15176,f598])).
fof(f15176,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16272,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15177,f598])).
fof(f15177,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f898])).
fof(f16271,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15178,f598])).
fof(f15178,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16270,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15179,f598])).
fof(f15179,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16269,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15180,f598])).
fof(f15180,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16268,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15182,f14233])).
fof(f15182,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f16267,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15183,f14233])).
fof(f15183,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f16266,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15188,f14233])).
fof(f15188,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f14233,f898])).
fof(f16265,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15189,f14233])).
fof(f15189,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f598,f14233,f898])).
fof(f16264,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15190,f14233])).
fof(f15190,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f598,f14233,f898])).
fof(f16263,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15191,f14233])).
fof(f15191,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f598,f14233,f898])).
fof(f16262,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15193,f14233])).
fof(f15193,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f598,f14233,f898])).
fof(f16261,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15194,f14233])).
fof(f15194,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f598,f14233,f898])).
fof(f16260,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15195,f14233])).
fof(f15195,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f598,f14233,f898])).
fof(f16259,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15196,f14233])).
fof(f15196,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f898])).
fof(f16258,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15197,f14233])).
fof(f15197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16257,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15198,f14233])).
fof(f15198,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16256,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15199,f14233])).
fof(f15199,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16255,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15201,f14233])).
fof(f15201,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16254,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15202,f14233])).
fof(f15202,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16253,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15203,f14233])).
fof(f15203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16252,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15204,f598])).
fof(f15204,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f14233,f14233,f898])).
fof(f16251,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15205,f598])).
fof(f15205,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f16250,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15206,f598])).
fof(f15206,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f16249,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15207,f598])).
fof(f15207,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f14233,f898])).
fof(f16248,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15208,f598])).
fof(f15208,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f598,f14233,f898])).
fof(f16247,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15209,f598])).
fof(f15209,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f598,f14233,f898])).
fof(f16246,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15210,f598])).
fof(f15210,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f598,f14233,f898])).
fof(f16245,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15212,f14233])).
fof(f15212,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f598,f14233,f898])).
fof(f16244,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15213,f14233])).
fof(f15213,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f598,f14233,f898])).
fof(f16243,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15218,f598])).
fof(f15218,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f898])).
fof(f16242,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15219,f598])).
fof(f15219,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16241,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15220,f598])).
fof(f15220,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16240,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15221,f598])).
fof(f15221,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16239,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15222,f598])).
fof(f15222,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f898])).
fof(f16238,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15223,f598])).
fof(f15223,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16237,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15224,f598])).
fof(f15224,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16236,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15225,f598])).
fof(f15225,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16235,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15227,f14233])).
fof(f15227,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f16234,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15228,f14233])).
fof(f15228,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f16233,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15234,f14233])).
fof(f15234,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f598,f14233,f898])).
fof(f16232,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15235,f14233])).
fof(f15235,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f598,f14233,f898])).
fof(f16231,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15236,f14233])).
fof(f15236,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f598,f14233,f898])).
fof(f16230,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15237,f598])).
fof(f15237,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f598,f14233,f898])).
fof(f16229,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15238,f598])).
fof(f15238,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f598,f14233,f898])).
fof(f16228,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15239,f598])).
fof(f15239,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f598,f14233,f898])).
fof(f16227,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15240,f14233])).
fof(f15240,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f898])).
fof(f16226,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15241,f14233])).
fof(f15241,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16225,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15242,f14233])).
fof(f15242,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16224,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15243,f14233])).
fof(f15243,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16223,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15245,f14233])).
fof(f15245,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16222,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15246,f14233])).
fof(f15246,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16221,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15247,f14233])).
fof(f15247,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16220,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15248,f598])).
fof(f15248,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f14233,f14233,f898])).
fof(f16219,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15249,f598])).
fof(f15249,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f16218,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15250,f598])).
fof(f15250,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f16217,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15252,f14233])).
fof(f15252,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f598,f14233,f898])).
fof(f16216,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15253,f14233])).
fof(f15253,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f598,f14233,f898])).
fof(f16215,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15254,f14233])).
fof(f15254,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f598,f14233,f898])).
fof(f16214,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15255,f598])).
fof(f15255,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f598,f14233,f898])).
fof(f16213,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15256,f598])).
fof(f15256,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f598,f14233,f898])).
fof(f16212,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15257,f598])).
fof(f15257,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f598,f14233,f898])).
fof(f16211,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15263,f14233])).
fof(f15263,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1165,f14233,f14233,f898])).
fof(f16210,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15264,f14233])).
fof(f15264,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2291,f14233,f14233,f898])).
fof(f16209,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15265,f14233])).
fof(f15265,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2631,f14233,f14233,f898])).
fof(f16208,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15266,f14233])).
fof(f15266,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f14233,f898])).
fof(f16207,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15267,f14233])).
fof(f15267,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1167,f14233,f14233,f898])).
fof(f16206,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15268,f14233])).
fof(f15268,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2294,f14233,f14233,f898])).
fof(f16205,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15269,f14233])).
fof(f15269,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2634,f14233,f14233,f898])).
fof(f16204,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15270,f598])).
fof(f15270,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f14233,f14233,f898])).
fof(f16203,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15271,f598])).
fof(f15271,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2316,f14233,f14233,f898])).
fof(f16202,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15272,f598])).
fof(f15272,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2656,f14233,f14233,f898])).
fof(f16201,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15278,f14233])).
fof(f15278,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1171,f983,f14233,f901])).
fof(f1171,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f896])).
fof(f896,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f16200,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15279,f14233])).
fof(f15279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2304,f983,f14233,f901])).
fof(f2304,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('F_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f896])).
fof(f16199,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15280,f14233])).
fof(f15280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2644,f983,f14233,f901])).
fof(f2644,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('F_0',X0,sK37('T_0',X1,'F_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f896])).
fof(f16198,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15281,f14233])).
fof(f15281,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1072,f983,f14233,f901])).
fof(f1072,plain,(
  ( ! [X0] : (sP9('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f983,f983,f897])).
fof(f897,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f16197,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15282,f14233])).
fof(f15282,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1176,f983,f14233,f901])).
fof(f1176,plain,(
  ( ! [X0,X1] : (sP9('F_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f897])).
fof(f16196,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15283,f14233])).
fof(f15283,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2310,f983,f14233,f901])).
fof(f2310,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('F_0',sK37('F_0',X0,'T_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f897])).
fof(f16195,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15284,f14233])).
fof(f15284,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2650,f983,f14233,f901])).
fof(f2650,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('F_0',sK37('T_0',X0,'F_0',X1,bool,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f983,f2204,f897])).
fof(f16194,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15306,f14233])).
fof(f15306,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1066,f983,f14233,f901])).
fof(f1066,plain,(
  ( ! [X0] : (sP9('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f983,f896])).
fof(f16193,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15307,f14233])).
fof(f15307,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1171,f983,f14233,f901])).
fof(f16192,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15308,f14233])).
fof(f15308,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2304,f983,f14233,f901])).
fof(f16191,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15309,f14233])).
fof(f15309,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2644,f983,f14233,f901])).
fof(f16190,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15311,f14233])).
fof(f15311,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1176,f983,f14233,f901])).
fof(f16189,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15312,f14233])).
fof(f15312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2310,f983,f14233,f901])).
fof(f16188,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15313,f14233])).
fof(f15313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2650,f983,f14233,f901])).
fof(f16187,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15344,f15920])).
fof(f15920,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1006,f14233,f963])).
fof(f963,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f905])).
fof(f905,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f575])).
fof(f575,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1006,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f983,f912])).
fof(f912,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f578])).
fof(f578,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f577])).
fof(f577,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f15344,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f902])).
fof(f902,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f576])).
fof(f16186,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15355,f15920])).
fof(f15355,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f903])).
fof(f903,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f576])).
fof(f16185,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15380,f15435])).
fof(f15435,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f911])).
fof(f911,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f579])).
fof(f15380,plain,(
  ~sP10('F_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f983,f14233,f906])).
fof(f906,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f576])).
fof(f1077,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f598,f983,f903])).
fof(f16184,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15382,f15435])).
fof(f15382,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0','F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f983,f14233,f906])).
fof(f1182,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f903])).
fof(f16183,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15384,f15435])).
fof(f15384,plain,(
  ( ! [X2,X0,X1] : (~sP10(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f983,f14233,f906])).
fof(f2318,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f903])).
fof(f16182,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15386,f15435])).
fof(f15386,plain,(
  ( ! [X2,X0,X1] : (~sP10(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f983,f14233,f906])).
fof(f2658,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f903])).
fof(f16181,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15407,f15435])).
fof(f15407,plain,(
  ~sP10('F_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f983,f14233,f907])).
fof(f907,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f576])).
fof(f1073,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f983,f902])).
fof(f16180,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15408,f15435])).
fof(f15408,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0','T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f983,f14233,f907])).
fof(f1181,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f902])).
fof(f16179,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15409,f15435])).
fof(f15409,plain,(
  ( ! [X2,X0,X1] : (~sP10(sK37('F_0',X0,'T_0',X1,bool,X2),'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f983,f14233,f907])).
fof(f2317,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f902])).
fof(f16178,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15410,f15435])).
fof(f15410,plain,(
  ( ! [X2,X0,X1] : (~sP10(sK37('T_0',X0,'F_0',X1,bool,X2),'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f983,f14233,f907])).
fof(f2657,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11(sK37('T_0',X0,'F_0',X1,bool,X2),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f2204,f902])).
fof(f16177,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15431,f15435])).
fof(f15431,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f983,f14233,f909])).
fof(f909,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f579])).
fof(f16176,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15437,f598])).
fof(f15437,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f914])).
fof(f914,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f344,f434,f433])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X41,X42,X43] : ((p(s(bool,X43)) <=> (p(s(bool,X42)) | p(s(bool,X41)))) <=> ((p(s(bool,X43)) | ~p(s(bool,X42))) & (p(s(bool,X43)) | ~p(s(bool,X41))) & (p(s(bool,X42)) | p(s(bool,X41)) | ~p(s(bool,X43)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f1973,plain,(
  ~sP11('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1968,f904])).
fof(f904,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f576])).
fof(f1968,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f598,f983,f983,f909])).
fof(f16175,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15438,f598])).
fof(f15438,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2370,f14233,f914])).
fof(f2370,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f598,f2203,f916])).
fof(f916,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f581])).
fof(f16174,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15439,f598])).
fof(f15439,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2373,f14233,f914])).
fof(f2373,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0','F_0',sK37('F_0',X0,'T_0',X1,bool,X2))) )),
  inference(unit_resulting_resolution,[],[f983,f598,f2203,f916])).
fof(f16173,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15442,f14233])).
fof(f15442,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f14233,f914])).
fof(f16172,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15443,f14233])).
fof(f15443,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f14233,f914])).
fof(f16171,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15444,f14233])).
fof(f15444,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f14233,f914])).
fof(f16170,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15445,f14233])).
fof(f15445,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f14233,f914])).
fof(f16169,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15446,f14233])).
fof(f15446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f14233,f914])).
fof(f16168,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15447,f14233])).
fof(f15447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f14233,f914])).
fof(f16167,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15454,f14233])).
fof(f15454,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2373,f14233,f914])).
fof(f16166,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15455,f598])).
fof(f15455,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f14233,f914])).
fof(f16165,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15456,f14233])).
fof(f15456,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f14233,f914])).
fof(f16164,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15457,f598])).
fof(f15457,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f14233,f914])).
fof(f16163,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15458,f14233])).
fof(f15458,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f14233,f914])).
fof(f16162,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15459,f598])).
fof(f15459,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f14233,f914])).
fof(f16161,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15460,f14233])).
fof(f15460,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f14233,f914])).
fof(f16160,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15461,f598])).
fof(f15461,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f14233,f914])).
fof(f16159,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15462,f14233])).
fof(f15462,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f14233,f914])).
fof(f16158,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15467,f598])).
fof(f15467,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f915])).
fof(f915,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f581])).
fof(f16157,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15468,f598])).
fof(f15468,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2370,f14233,f915])).
fof(f16156,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15469,f598])).
fof(f15469,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2373,f14233,f915])).
fof(f16155,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15472,f14233])).
fof(f15472,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f14233,f915])).
fof(f16154,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15473,f14233])).
fof(f15473,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f14233,f915])).
fof(f16153,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15474,f14233])).
fof(f15474,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f14233,f915])).
fof(f16152,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15475,f14233])).
fof(f15475,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f14233,f915])).
fof(f16151,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15476,f14233])).
fof(f15476,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f14233,f915])).
fof(f16150,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15477,f14233])).
fof(f15477,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('T_0',X0,'F_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f14233,f915])).
fof(f16149,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15483,f14233])).
fof(f15483,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2370,f14233,f915])).
fof(f16148,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15485,f14233])).
fof(f15485,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f14233,f915])).
fof(f16147,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15486,f598])).
fof(f15486,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f14233,f915])).
fof(f16146,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15487,f14233])).
fof(f15487,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f14233,f915])).
fof(f16145,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15488,f598])).
fof(f15488,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f14233,f915])).
fof(f16144,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15489,f14233])).
fof(f15489,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f14233,f915])).
fof(f16143,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15490,f598])).
fof(f15490,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f14233,f915])).
fof(f16142,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15491,f14233])).
fof(f15491,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f14233,f915])).
fof(f16141,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15492,f598])).
fof(f15492,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f14233,f915])).
fof(f16140,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15518,f15899])).
fof(f15899,plain,(
  ( ! [X0,X1] : (sP13('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f14233,f962])).
fof(f962,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP12(X0,X1,X2) | sP13(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f922])).
fof(f922,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP12(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP12(X0,X1,X2) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f583])).
fof(f583,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP12(X2,X0,X1) | sP13(X2,X0,X1)))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X2,X0,X1] : ((~sP13(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP12(X2,X0,X1)) | sP13(X2,X0,X1)))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X2,X0,X1] : (sP13(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP12(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f1009,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f983,f929])).
fof(f929,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f586])).
fof(f586,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f585])).
fof(f585,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f15518,plain,(
  ( ! [X0,X1] : (~sP13('F_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f920])).
fof(f920,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f584])).
fof(f16139,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15521,f14233])).
fof(f15521,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1009,f1186,f14233,f923])).
fof(f923,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f584])).
fof(f1186,plain,(
  ( ! [X0,X1] : (~sP13('F_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f919])).
fof(f919,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f584])).
fof(f16138,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15522,f14233])).
fof(f15522,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f1186,f14233,f923])).
fof(f1190,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X2))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f927])).
fof(f927,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f587])).
fof(f16137,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15523,f14233])).
fof(f15523,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1008,f1186,f14233,f923])).
fof(f1008,plain,(
  ( ! [X0,X1] : (sP12(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f598,f928])).
fof(f928,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f587])).
fof(f16136,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15524,f14233])).
fof(f15524,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f2378,f14233,f923])).
fof(f2378,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13('F_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f919])).
fof(f16135,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15525,f14233])).
fof(f15525,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2401,f2378,f14233,f923])).
fof(f2401,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(X0,X1,sK37('F_0',X2,'T_0',X3,bool,X4))) )),
  inference(unit_resulting_resolution,[],[f2203,f927])).
fof(f16134,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15526,f14233])).
fof(f15526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1008,f2378,f14233,f923])).
fof(f16133,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15527,f14233])).
fof(f15527,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f1091,f14233,f923])).
fof(f1091,plain,(
  ( ! [X0] : (~sP13('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f598,f983,f920])).
fof(f16132,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15528,f14233])).
fof(f15528,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1008,f1091,f14233,f923])).
fof(f16131,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15530,f14233])).
fof(f15530,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f1091,f14233,f923])).
fof(f16130,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15531,f14233])).
fof(f15531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2401,f1091,f14233,f923])).
fof(f16129,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15533,f14233])).
fof(f15533,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f1189,f14233,f923])).
fof(f1189,plain,(
  ( ! [X0,X1] : (~sP13('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f920])).
fof(f1191,plain,(
  ( ! [X2,X0,X1] : (sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f929])).
fof(f16128,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15534,f14233])).
fof(f15534,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1008,f1189,f14233,f923])).
fof(f16127,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15536,f14233])).
fof(f15536,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f1189,f14233,f923])).
fof(f16126,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15537,f14233])).
fof(f15537,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2401,f1189,f14233,f923])).
fof(f16125,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15539,f14233])).
fof(f15539,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f2381,f14233,f923])).
fof(f2381,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f920])).
fof(f2402,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(sK37('F_0',X0,'T_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f929])).
fof(f16124,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15540,f14233])).
fof(f15540,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1008,f2381,f14233,f923])).
fof(f16123,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15542,f14233])).
fof(f15542,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f2381,f14233,f923])).
fof(f16122,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15543,f14233])).
fof(f15543,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2401,f2381,f14233,f923])).
fof(f16121,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15550,f15762])).
fof(f15762,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f928])).
fof(f15550,plain,(
  ~sP12('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f598,f14233,f923])).
fof(f16120,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15552,f15762])).
fof(f15552,plain,(
  ( ! [X0] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0','T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f598,f14233,f923])).
fof(f16119,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15554,f15762])).
fof(f15554,plain,(
  ( ! [X2,X0,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f598,f14233,f923])).
fof(f16118,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15565,f15762])).
fof(f15565,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f14233,f14233,f923])).
fof(f1090,plain,(
  ( ! [X0] : (~sP13('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f983,f919])).
fof(f16117,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15567,f15762])).
fof(f15567,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f14233,f14233,f923])).
fof(f16116,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15569,f15762])).
fof(f15569,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('F_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f14233,f14233,f923])).
fof(f16115,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15571,f15762])).
fof(f15571,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f14233,f14233,f923])).
fof(f16114,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15573,f15762])).
fof(f15573,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f14233,f14233,f923])).
fof(f1184,plain,(
  ( ! [X0,X1] : (~sP13('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f983,f982,f919])).
fof(f16113,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15575,f15762])).
fof(f15575,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f14233,f14233,f923])).
fof(f16112,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15577,f15762])).
fof(f15577,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f14233,f14233,f923])).
fof(f2376,plain,(
  ( ! [X2,X0,X3,X1] : (~sP13(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f2203,f919])).
fof(f16111,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15579,f15762])).
fof(f15579,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f14233,f14233,f923])).
fof(f16110,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15583,f14233])).
fof(f15583,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1007,f14233,f923])).
fof(f1007,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f927])).
fof(f16109,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15584,f14233])).
fof(f15584,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f1007,f14233,f923])).
fof(f16108,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15585,f14233])).
fof(f15585,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f1007,f14233,f923])).
fof(f16107,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15586,f14233])).
fof(f15586,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f1007,f14233,f923])).
fof(f16106,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15589,f14233])).
fof(f15589,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f1190,f14233,f923])).
fof(f16105,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15590,f14233])).
fof(f15590,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f2381,f1190,f14233,f923])).
fof(f16104,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15593,f14233])).
fof(f15593,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f2401,f14233,f923])).
fof(f16103,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15594,f14233])).
fof(f15594,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f2401,f14233,f923])).
fof(f16102,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15599,f14233])).
fof(f15599,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f1008,f14233,f923])).
fof(f16101,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15600,f14233])).
fof(f15600,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1008,f14233,f923])).
fof(f16100,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15601,f14233])).
fof(f15601,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f1008,f14233,f923])).
fof(f16099,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15602,f14233])).
fof(f15602,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f1008,f14233,f923])).
fof(f16098,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15604,f14233])).
fof(f15604,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1191,f14233,f923])).
fof(f16097,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15605,f14233])).
fof(f15605,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f1191,f14233,f923])).
fof(f16096,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15606,f14233])).
fof(f15606,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f2402,f14233,f923])).
fof(f16095,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15607,f14233])).
fof(f15607,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f2402,f14233,f923])).
fof(f16094,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15616,f15762])).
fof(f15616,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f14233,f14233,f923])).
fof(f16093,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15617,f15762])).
fof(f15617,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f14233,f14233,f923])).
fof(f16092,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15618,f15762])).
fof(f15618,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('F_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f14233,f14233,f923])).
fof(f16091,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15619,f15762])).
fof(f15619,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f14233,f14233,f923])).
fof(f16090,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15620,f15762])).
fof(f15620,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f14233,f14233,f923])).
fof(f16089,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15621,f15762])).
fof(f15621,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f14233,f14233,f923])).
fof(f16088,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15622,f15762])).
fof(f15622,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f14233,f14233,f923])).
fof(f16087,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15623,f15762])).
fof(f15623,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f14233,f14233,f923])).
fof(f16086,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15662,f15762])).
fof(f15662,plain,(
  ~sP12('F_0','F_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f983,f14233,f924])).
fof(f924,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f584])).
fof(f16085,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15663,f15762])).
fof(f15663,plain,(
  ( ! [X0] : (~sP12('F_0','F_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f983,f14233,f924])).
fof(f16084,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15664,f15762])).
fof(f15664,plain,(
  ( ! [X2,X0,X1] : (~sP12('F_0','F_0',sK37('F_0',X0,'T_0',X1,bool,X2))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f983,f14233,f924])).
fof(f16083,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15665,f15762])).
fof(f15665,plain,(
  ( ! [X0] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0','F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f983,f14233,f924])).
fof(f16082,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15666,f15762])).
fof(f15666,plain,(
  ( ! [X2,X0,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'F_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f983,f14233,f924])).
fof(f16081,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15703,f15762])).
fof(f15703,plain,(
  ~sP12('F_0','F_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f983,f14233,f925])).
fof(f925,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~sP12(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f584])).
fof(f16080,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15704,f15762])).
fof(f15704,plain,(
  ( ! [X0] : (~sP12('F_0','F_0','const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f983,f14233,f925])).
fof(f16079,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15705,f15762])).
fof(f15705,plain,(
  ( ! [X2,X0,X1] : (~sP12('F_0','F_0',sK37('F_0',X0,'T_0',X1,bool,X2))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f983,f14233,f925])).
fof(f16078,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15706,f15762])).
fof(f15706,plain,(
  ( ! [X0] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'F_0','F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f983,f14233,f925])).
fof(f16077,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15707,f15762])).
fof(f15707,plain,(
  ( ! [X2,X0,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'F_0','F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f983,f14233,f925])).
fof(f16076,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15732,f14233])).
fof(f15732,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f983,f14233,f926])).
fof(f926,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f587])).
fof(f16075,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15733,f14233])).
fof(f15733,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2401,f983,f14233,f926])).
fof(f16074,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15734,f14233])).
fof(f15734,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f983,f14233,f926])).
fof(f16073,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15735,f14233])).
fof(f15735,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f983,f14233,f926])).
fof(f16072,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15736,f14233])).
fof(f15736,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f983,f14233,f926])).
fof(f16071,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15737,f15762])).
fof(f15737,plain,(
  ( ! [X0] : (~sP12(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f598,f983,f14233,f926])).
fof(f16070,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15739,f15762])).
fof(f15739,plain,(
  ( ! [X0,X1] : (~sP12(X0,'F_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f983,f14233,f926])).
fof(f16069,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15752,f14233])).
fof(f15752,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1007,f983,f14233,f926])).
fof(f16068,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15753,f14233])).
fof(f15753,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1190,f983,f14233,f926])).
fof(f16067,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15754,f14233])).
fof(f15754,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2401,f983,f14233,f926])).
fof(f16066,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15756,f14233])).
fof(f15756,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f983,f14233,f926])).
fof(f16065,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15757,f14233])).
fof(f15757,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f983,f14233,f926])).
fof(f16064,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15758,f15762])).
fof(f15758,plain,(
  ( ! [X0] : (~sP12('T_0','F_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f598,f983,f14233,f926])).
fof(f16063,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15760,f15762])).
fof(f15760,plain,(
  ( ! [X0,X1] : (~sP12(X0,'F_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f983,f14233,f926])).
fof(f16062,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15788,f598])).
fof(f15788,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2098,f14233,f932])).
fof(f932,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP13(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP13(X2,X0,X1)) & (~sP13(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP13(X2,X0,X1))),
  inference(definition_folding,[],[f411,f437,f436])).
fof(f411,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X41,X42,X43] : ((p(s(bool,X43)) <=> (p(s(bool,X42)) => p(s(bool,X41)))) <=> ((p(s(bool,X43)) | p(s(bool,X42))) & (p(s(bool,X43)) | ~p(s(bool,X41))) & (~p(s(bool,X42)) | p(s(bool,X41)) | ~p(s(bool,X43)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f2098,plain,(
  ~sP13('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2093,f921])).
fof(f921,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f584])).
fof(f2093,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f598,f598,f983,f926])).
fof(f16061,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15789,f598])).
fof(f15789,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2412,f14233,f932])).
fof(f2412,plain,(
  ( ! [X2,X0,X1] : (~sP13('T_0',sK37('F_0',X0,'T_0',X1,bool,X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f598,f2203,f933])).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (~sP13(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f589])).
fof(f16060,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15794,f14233])).
fof(f15794,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f14233,f932])).
fof(f16059,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15795,f14233])).
fof(f15795,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f14233,f932])).
fof(f16058,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15796,f14233])).
fof(f15796,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f14233,f932])).
fof(f16057,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15797,f14233])).
fof(f15797,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f14233,f932])).
fof(f16056,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15803,f14233])).
fof(f15803,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2412,f14233,f932])).
fof(f16055,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15804,f14233])).
fof(f15804,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f14233,f932])).
fof(f16054,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15805,f14233])).
fof(f15805,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f14233,f932])).
fof(f16053,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15806,f14233])).
fof(f15806,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f14233,f932])).
fof(f16052,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15807,f598])).
fof(f15807,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f14233,f932])).
fof(f16051,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15808,f14233])).
fof(f15808,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f14233,f932])).
fof(f16050,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15809,f598])).
fof(f15809,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f14233,f932])).
fof(f16049,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15810,f14233])).
fof(f15810,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f14233,f932])).
fof(f16048,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15811,f598])).
fof(f15811,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f14233,f932])).
fof(f16047,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15816,f598])).
fof(f15816,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1093,f983,f14233,f933])).
fof(f1093,plain,(
  ( ! [X0] : (sP13('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f598,f931])).
fof(f931,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP13(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f589])).
fof(f16046,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15817,f598])).
fof(f15817,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1192,f983,f14233,f933])).
fof(f1192,plain,(
  ( ! [X0,X1] : (sP13('T_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f598,f982,f931])).
fof(f16045,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15818,f598])).
fof(f15818,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2409,f983,f14233,f933])).
fof(f2409,plain,(
  ( ! [X2,X0,X3,X1] : (sP13('T_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) )),
  inference(unit_resulting_resolution,[],[f598,f2203,f931])).
fof(f16044,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15820,f14233])).
fof(f15820,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2407,f983,f14233,f933])).
fof(f2407,plain,(
  ( ! [X2,X0,X1] : (sP13(sK37('F_0',X0,'T_0',X1,bool,X2),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f983,f598,f2203,f930])).
fof(f930,plain,(
  ( ! [X2,X0,X1] : (sP13(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f589])).
fof(f16043,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15837,f14233])).
fof(f15837,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1192,f983,f14233,f933])).
fof(f16042,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15838,f14233])).
fof(f15838,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK37('F_0',X0,'T_0',X1,bool,X2)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2409,f983,f14233,f933])).
fof(f16041,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15839,f598])).
fof(f15839,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2080,f983,f14233,f933])).
fof(f2080,plain,(
  sP13('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1009,f598,f983,f925])).
fof(f16040,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15840,f598])).
fof(f15840,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2407,f983,f14233,f933])).
fof(f16039,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15856,f16008])).
fof(f15856,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233,f961])).
fof(f961,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f960])).
fof(f960,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f943])).
fof(f943,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8,X37,X38,X39] : ((s(bool,X8) = s(bool,X37) & (p(s(bool,X37)) => s(bool,X38) = s(bool,X39))) => ((p(s(bool,X8)) => p(s(bool,X38))) <=> (p(s(bool,X37)) => p(s(bool,X39)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.bool.IMP_CONG')).
fof(f16038,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15883,f1007])).
fof(f15883,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f14233,f962])).
fof(f16037,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15884,f1190])).
fof(f15884,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X1))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f14233,f962])).
fof(f16036,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15885,f2401])).
fof(f15885,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('F_0',X0,sK37('F_0',X1,'T_0',X2,bool,X3))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f14233,f962])).
fof(f16035,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15886,f1008])).
fof(f15886,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f14233,f962])).
fof(f16034,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15887,f1007])).
fof(f15887,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'F_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f14233,f962])).
fof(f16033,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15888,f1008])).
fof(f15888,plain,(
  ( ! [X0,X1] : (~sP12('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f14233,f962])).
fof(f16032,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15889,f1007])).
fof(f15889,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f14233,f962])).
fof(f16031,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15890,f1008])).
fof(f15890,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2381,f14233,f962])).
fof(f16030,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15908,f1004])).
fof(f1004,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f598,f910])).
fof(f910,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f579])).
fof(f15908,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f14233,f963])).
fof(f16029,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15909,f1005])).
fof(f1005,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f598,f911])).
fof(f15909,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f14233,f963])).
fof(f16028,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15910,f1004])).
fof(f15910,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,'T_0')) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f14233,f963])).
fof(f16027,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f15911,f1005])).
fof(f15911,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),'T_0',X1)) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f14233,f963])).
fof(f16026,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15912,f1004])).
fof(f15912,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10(sK37('F_0',X0,'T_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f14233,f963])).
fof(f16025,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15913,f1005])).
fof(f15913,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10(sK37('F_0',X0,'T_0',X1,bool,X2),'T_0',X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2318,f14233,f963])).
fof(f16024,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15914,f1004])).
fof(f15914,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10(sK37('T_0',X0,'F_0',X1,bool,X2),X3,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f14233,f963])).
fof(f16023,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f15915,f1005])).
fof(f15915,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10(sK37('T_0',X0,'F_0',X1,bool,X2),'T_0',X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2658,f14233,f963])).
fof(f16009,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1012,f14233,f823])).
fof(f15931,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f966])).
fof(f966,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f838])).
fof(f838,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f545])).
fof(f15929,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1052,f14233,f965])).
fof(f1052,plain,(
  ( ! [X0] : (~sP3('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f983,f598,f852])).
fof(f15927,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f2695,f14233,f963])).
fof(f2695,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK37('T_0',X0,'F_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2204,f912])).
fof(f15925,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f2355,f14233,f963])).
fof(f2355,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK37('F_0',X0,'T_0',X1,bool,X2),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f2203,f912])).
fof(f15923,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f1183,f14233,f963])).
fof(f1183,plain,(
  ( ! [X2,X0,X1] : (sP10('const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X0))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))),X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f982,f912])).
fof(f15921,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f1006,f14233,f963])).
fof(f15919,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f1005,f14233,f963])).
fof(f15917,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f1004,f14233,f963])).
fof(f15904,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f2402,f14233,f962])).
fof(f15902,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1191,f14233,f962])).
fof(f15900,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1009,f14233,f962])).
fof(f15898,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f1008,f14233,f962])).
fof(f15896,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f2401,f14233,f962])).
fof(f15894,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f1190,f14233,f962])).
fof(f15892,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1007,f14233,f962])).
fof(f15878,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f701,f14233,f961])).
fof(f15876,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f700,f14233,f961])).
fof(f15874,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f788,f14233,f961])).
fof(f788,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X7,X6] : ! [X8,X38] : s(X6,X38) = s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X8),s(X6,X38)))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.pair.SND')).
fof(f15872,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f789,f14233,f961])).
fof(f15870,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f794,f14233,f961])).
fof(f794,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X4) = s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4))))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X4) = s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4))))),s(X0,X3)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X4) = s(X1,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X0,X1),'const.finite_map.FUPDATE_2'(s('type.finite_map.fmap'(X0,X1),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4))))),s(X0,X3)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X7,X6] : ! [X25,X8,X38] : s(X6,X38) = s(X6,'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'(X7,X6),'const.finite_map.FUPDATE_2'(s('type.finite_map.fmap'(X7,X6),X25),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X8),s(X6,X38))))),s(X7,X8)))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.finite_map.FAPPLY_FUPDATE')).
fof(f15867,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f985,f14233,f961])).
fof(f15855,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1093,f983,f14233,f14233,f933])).
fof(f15846,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2080,f983,f598,f14233,f933])).
fof(f15835,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1093,f983,f14233,f14233,f933])).
fof(f15826,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1093,f983,f598,f14233,f933])).
fof(f15815,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2098,f14233,f14233,f932])).
fof(f15813,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f598,f14233,f932])).
fof(f15801,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2098,f14233,f14233,f932])).
fof(f15799,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2098,f598,f14233,f932])).
fof(f15787,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f983,f14233,f931])).
fof(f15775,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f983,f983,f14233,f930])).
fof(f15763,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2093,f14233,f928])).
fof(f15761,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1007,f14233,f983,f14233,f926])).
fof(f15759,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1007,f598,f983,f14233,f926])).
fof(f15740,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1007,f14233,f983,f14233,f926])).
fof(f15738,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f598,f983,f14233,f926])).
fof(f15719,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f2402,f983,f14233,f925])).
fof(f15717,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1191,f983,f14233,f925])).
fof(f15715,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1009,f983,f14233,f925])).
fof(f15713,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f2401,f983,f14233,f925])).
fof(f15711,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f1190,f983,f14233,f925])).
fof(f15709,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1007,f983,f14233,f925])).
fof(f15678,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f2402,f983,f14233,f924])).
fof(f15676,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1191,f983,f14233,f924])).
fof(f15674,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1009,f983,f14233,f924])).
fof(f15672,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f2401,f983,f14233,f924])).
fof(f15670,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f1190,f983,f14233,f924])).
fof(f15668,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1007,f983,f14233,f924])).
fof(f15637,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2376,f2402,f14233,f14233,f923])).
fof(f15635,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f1191,f14233,f14233,f923])).
fof(f15633,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1009,f14233,f14233,f923])).
fof(f15631,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1091,f1008,f14233,f14233,f923])).
fof(f15629,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f2401,f14233,f14233,f923])).
fof(f15627,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1186,f1190,f14233,f14233,f923])).
fof(f15625,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1090,f1007,f14233,f14233,f923])).
fof(f15580,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f2381,f14233,f14233,f923])).
fof(f15578,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f2376,f14233,f14233,f923])).
fof(f15576,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f1189,f14233,f14233,f923])).
fof(f15574,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f1184,f14233,f14233,f923])).
fof(f15572,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f1091,f14233,f14233,f923])).
fof(f15570,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f2378,f14233,f14233,f923])).
fof(f15568,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1009,f1186,f14233,f14233,f923])).
fof(f15566,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f1090,f14233,f14233,f923])).
fof(f15555,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2402,f2381,f598,f14233,f923])).
fof(f15553,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1191,f1189,f598,f14233,f923])).
fof(f15551,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1009,f1091,f598,f14233,f923])).
fof(f15519,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2080,f983,f14233,f920])).
fof(f15510,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1954,f983,f983,f14233,f916])).
fof(f1954,plain,(
  sP11('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1006,f983,f983,f908])).
fof(f908,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f576])).
fof(f15496,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f14233,f915])).
fof(f15494,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1077,f598,f14233,f915])).
fof(f15481,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f14233,f915])).
fof(f15479,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f598,f14233,f915])).
fof(f15466,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f14233,f914])).
fof(f15464,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f598,f14233,f914])).
fof(f15451,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f14233,f14233,f914])).
fof(f15449,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1973,f598,f14233,f914])).
fof(f15436,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1968,f14233,f911])).
fof(f15434,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1968,f14233,f910])).
fof(f15432,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1006,f983,f983,f14233,f909])).
fof(f15420,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2657,f2695,f983,f14233,f907])).
fof(f15418,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2317,f2355,f983,f14233,f907])).
fof(f15416,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1181,f1183,f983,f14233,f907])).
fof(f15414,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f1006,f983,f14233,f907])).
fof(f15412,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1073,f1004,f983,f14233,f907])).
fof(f15387,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2695,f2658,f983,f14233,f906])).
fof(f15385,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2355,f2318,f983,f14233,f906])).
fof(f15383,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1183,f1182,f983,f14233,f906])).
fof(f15381,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1006,f1077,f983,f14233,f906])).
fof(f15356,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1954,f983,f14233,f903])).
fof(f15345,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1954,f983,f14233,f902])).
fof(f15334,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1066,f983,f14233,f14233,f901])).
fof(f15321,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1066,f983,f598,f14233,f901])).
fof(f15305,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1066,f983,f14233,f14233,f901])).
fof(f15292,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1072,f983,f598,f14233,f901])).
fof(f15276,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f14233,f898])).
fof(f15274,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f598,f14233,f14233,f898])).
fof(f15261,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f598,f14233,f898])).
fof(f15259,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1850,f598,f598,f14233,f898])).
fof(f15232,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f14233,f898])).
fof(f15230,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f14233,f14233,f898])).
fof(f15217,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f598,f14233,f898])).
fof(f15215,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f598,f14233,f898])).
fof(f15187,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f14233,f14233,f14233,f898])).
fof(f15185,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1062,f598,f14233,f14233,f898])).
fof(f15172,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f14233,f598,f14233,f898])).
fof(f15170,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1063,f598,f598,f14233,f898])).
fof(f15142,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1845,f14233,f893])).
fof(f15140,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f14233,f983,f14233,f892])).
fof(f15138,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f598,f983,f14233,f892])).
fof(f15114,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f983,f14233,f14233,f892])).
fof(f15103,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f983,f598,f14233,f892])).
fof(f15085,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2634,f14233,f14233,f888])).
fof(f15083,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2294,f14233,f14233,f888])).
fof(f15081,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1167,f14233,f14233,f888])).
fof(f15079,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f1063,f14233,f14233,f888])).
fof(f15077,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2631,f14233,f14233,f888])).
fof(f15075,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2291,f14233,f14233,f888])).
fof(f15073,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1165,f14233,f14233,f888])).
fof(f15071,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f1062,f14233,f14233,f888])).
fof(f15057,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2631,f598,f14233,f888])).
fof(f15055,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2291,f598,f14233,f888])).
fof(f15053,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1165,f598,f14233,f888])).
fof(f15051,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f1062,f598,f14233,f888])).
fof(f14998,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2634,f14233,f14233,f888])).
fof(f14996,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2294,f14233,f14233,f888])).
fof(f14994,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1167,f14233,f14233,f888])).
fof(f14992,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f1063,f14233,f14233,f888])).
fof(f14990,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2642,f2631,f14233,f14233,f888])).
fof(f14988,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2302,f2291,f14233,f14233,f888])).
fof(f14986,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1169,f1165,f14233,f14233,f888])).
fof(f14984,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1002,f1062,f14233,f14233,f888])).
fof(f14970,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2643,f2634,f598,f14233,f888])).
fof(f14968,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2303,f2294,f598,f14233,f888])).
fof(f14966,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1170,f1167,f598,f14233,f888])).
fof(f14964,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1003,f1063,f598,f14233,f888])).
fof(f14914,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1072,f983,f14233,f887])).
fof(f14901,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1066,f983,f14233,f886])).
fof(f14888,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1373,f1012,f14233,f883])).
fof(f1373,plain,(
  sP7('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f983,f1012,f881])).
fof(f14886,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1374,f1012,f14233,f883])).
fof(f14868,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1615,f701,f14233,f882])).
fof(f14866,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1505,f700,f14233,f882])).
fof(f14860,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1411,f985,f14233,f882])).
fof(f14856,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1612,f701,f14233,f882])).
fof(f14854,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1502,f700,f14233,f882])).
fof(f14848,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1408,f985,f14233,f882])).
fof(f16022,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f14843])).
fof(f14843,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1823,f14233,f882])).
fof(f1823,plain,(
  ~sP7('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1818,f862])).
fof(f862,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f557])).
fof(f1818,plain,(
  ~sP4('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f983,f983,f877])).
fof(f877,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f565])).
fof(f16021,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f14841])).
fof(f14841,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1814,f14233,f882])).
fof(f1814,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1809,f863])).
fof(f863,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP5(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f557])).
fof(f1809,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f598,f983,f598,f873])).
fof(f14834,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1818,f14233,f880])).
fof(f14832,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1818,f14233,f879])).
fof(f14830,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1818,f14233,f878])).
fof(f14828,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1809,f14233,f874])).
fof(f14826,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f996,f14233,f983,f14233,f873])).
fof(f14824,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f996,f598,f983,f14233,f873])).
fof(f14800,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f996,f983,f14233,f14233,f873])).
fof(f14789,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f997,f983,f598,f14233,f873])).
fof(f14774,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1777,f14233,f870])).
fof(f1777,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f598,f598,f869])).
fof(f14772,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f993,f983,f14233,f14233,f869])).
fof(f14756,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1488,f983,f598,f14233,f869])).
fof(f14733,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f993,f983,f14233,f14233,f869])).
fof(f14717,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f994,f983,f598,f14233,f869])).
fof(f14651,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f2596,f999,f14233,f866])).
fof(f14649,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f2596,f999,f14233,f866])).
fof(f14646,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f2256,f999,f14233,f866])).
fof(f14644,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f2256,f999,f14233,f866])).
fof(f14639,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f993,f999,f14233,f866])).
fof(f14637,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f993,f999,f14233,f866])).
fof(f14634,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f992,f999,f14233,f866])).
fof(f14632,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f992,f999,f14233,f866])).
fof(f14629,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f992,f999,f14233,f866])).
fof(f14627,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f992,f999,f14233,f866])).
fof(f14609,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f995,f2596,f998,f14233,f866])).
fof(f14607,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f2596,f998,f14233,f866])).
fof(f14600,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f995,f2256,f998,f14233,f866])).
fof(f14598,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f2256,f998,f14233,f866])).
fof(f14585,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f993,f998,f14233,f866])).
fof(f14583,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f993,f998,f14233,f866])).
fof(f14580,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f2603,f992,f998,f14233,f866])).
fof(f14578,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2246,f2263,f992,f998,f14233,f866])).
fof(f14575,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f996,f992,f998,f14233,f866])).
fof(f14573,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1390,f995,f992,f998,f14233,f866])).
fof(f14495,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1374,f14233,f983,f14233,f865])).
fof(f14493,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f3352,f598,f983,f14233,f865])).
fof(f14462,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1374,f983,f14233,f14233,f865])).
fof(f14446,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1374,f983,f598,f14233,f865])).
fof(f14431,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1059,f983,f14233,f853])).
fof(f1059,plain,(
  ( ! [X0] : (sP3(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f983,f983,f854])).
fof(f854,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f552])).
fof(f14421,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1059,f983,f14233,f852])).
fof(f14411,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f841])).
fof(f14409,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f598,f14233,f841])).
fof(f14393,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f841])).
fof(f14391,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1040,f598,f14233,f841])).
fof(f14375,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f840])).
fof(f14373,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f598,f14233,f840])).
fof(f14357,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f840])).
fof(f14340,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f839])).
fof(f14323,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1037,f14233,f14233,f839])).
fof(f14321,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1040,f598,f14233,f839])).
fof(f14305,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1019,f983,f14233,f767])).
fof(f1019,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f983,f762])).
fof(f762,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f520])).
fof(f14293,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1019,f983,f14233,f764])).
fof(f14281,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1023,f14233,f14233,f763])).
fof(f14279,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1023,f598,f14233,f763])).
fof(f14269,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1023,f14233,f14233,f763])).
fof(f14267,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1028,f598,f14233,f763])).
fof(f14257,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1031,f14233,f749])).
fof(f749,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f513])).
fof(f513,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f512])).
fof(f512,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1031,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f983,f983,f988])).
fof(f988,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f747,f983])).
fof(f747,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f513])).
fof(f14255,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1031,f14233,f748])).
fof(f748,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f513])).
fof(f14253,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f1012,f14233,f734])).
fof(f14251,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f598,f1012,f14233,f734])).
fof(f14247,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f14233,f1012,f14233,f734])).
fof(f14245,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f598,f1012,f14233,f734])).
fof(f14241,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f985,f14233,f632])).
fof(f14239,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1012,f14233,f629])).
fof(f14237,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f983,f14233])).
fof(f14236,plain,(
  ( ! [X8] : (~p(s(bool,sK33(sK29('const.pair.SND_0',bool,X8),bool,X8)))) ) | $spl10),
  inference(cnf_transformation,[],[f14236_D])).
fof(f14236_D,plain,(
  ( ! [X8] : (~p(s(bool,sK33(sK29('const.pair.SND_0',bool,X8),bool,X8)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f14233,plain,(
  ( ! [X6] : (p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f14233_D])).
fof(f14233_D,plain,(
  ( ! [X6] : (p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f14235,plain,(
  ( ! [X10,X8] : (p(s(bool,X10)) | ~p(s(bool,sK33(sK29('const.pair.SND_0',bool,X8),bool,X8)))) )),
  inference(forward_demodulation,[],[f14234,f5512])).
fof(f5512,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,sK33('const.pair.,_2'(s(X2,X3),s(X0,X1)),X0,X2))) )),
  inference(unit_resulting_resolution,[],[f787,f821])).
fof(f821,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f536])).
fof(f14234,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,sK33('const.pair.,_2'(s(X8,X9),s(bool,X10)),bool,X8))) | ~p(s(bool,sK33(sK29('const.pair.SND_0',bool,X8),bool,X8)))) )),
  inference(forward_demodulation,[],[f14219,f3889])).
fof(f3889,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK33(X2,X1,X0)) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(backward_demodulation,[],[f3858,f783])).
fof(f783,plain,(
  ( ! [X2,X0,X1] : (s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ! [X2] : s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,happ(s(fun('type.pair.prod'(X0,X1),X1),'const.pair.SND_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X7,X6] : ! [X40] : s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X40))) = s(X6,happ(s(fun('type.pair.prod'(X7,X6),X6),'const.pair.SND_0'),s('type.pair.prod'(X7,X6),X40)))),
  file('thm.typeSound.typeSound_46_prove_ltb','ARITY_0_1_const.pair.SND')).
fof(f3858,plain,(
  ( ! [X4,X5,X3] : (s(X4,sK33(X5,X4,X3)) = s(X4,'const.pair.SND_1'(s('type.pair.prod'(X3,X4),X5)))) )),
  inference(superposition,[],[f788,f787])).
fof(f14219,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,sK33(sK29('const.pair.SND_0',bool,X8),bool,X8))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),'const.pair.SND_0'),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X9),s(bool,X10))))))) )),
  inference(superposition,[],[f785,f3889])).
fof(f785,plain,(
  ( ! [X6,X2,X0,X7,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK29(X2,X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) )),
  inference(cnf_transformation,[],[f528])).
fof(f528,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK29(X2,X1,X0))))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK30(X2,X1,X0)),s(X1,sK31(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30,sK31])],[f527])).
fof(f527,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f526])).
fof(f526,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X7,X6] : ! [X26] : (! [X43] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X26),s('type.pair.prod'(X7,X6),X43)))) <=> ! [X50,X51] : p(s(bool,happ(s(fun('type.pair.prod'(X7,X6),bool),X26),s('type.pair.prod'(X7,X6),'const.pair.,_2'(s(X7,X50),s(X6,X51)))))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.pair.FORALL_PROD')).
fof(f14232,plain,(
  ( ! [X5] : (~p(s(bool,sK32(sK29('const.pair.FST_0',X5,bool),X5,bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f14232_D])).
fof(f14232_D,plain,(
  ( ! [X5] : (~p(s(bool,sK32(sK29('const.pair.FST_0',X5,bool),X5,bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f14231,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | ~p(s(bool,sK32(sK29('const.pair.FST_0',X5,bool),X5,bool)))) )),
  inference(forward_demodulation,[],[f14230,f5378])).
fof(f14230,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,sK32('const.pair.,_2'(s(bool,X6),s(X5,X7)),X5,bool))) | ~p(s(bool,sK32(sK29('const.pair.FST_0',X5,bool),X5,bool)))) )),
  inference(forward_demodulation,[],[f14218,f3882])).
fof(f14218,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,sK32(sK29('const.pair.FST_0',X5,bool),X5,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X5),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X5),'const.pair.,_2'(s(bool,X6),s(X5,X7))))))) )),
  inference(superposition,[],[f785,f3882])).
fof(f982,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X5))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) | $spl4),
  inference(cnf_transformation,[],[f982_D])).
fof(f982_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X5))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f980,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | $spl3),
  inference(cnf_transformation,[],[f980_D])).
fof(f980_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f596,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X5))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  (! [X5] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,X5))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))))) & (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,sK18))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17,sK18])],[f440])).
fof(f440,plain,(
  ? [X0,X1,X2,X3] : ((! [X5] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X5))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3))))))))) & (? [X4] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3))))))))))),
  inference(rectify,[],[f439])).
fof(f439,plain,(
  ? [X0,X1,X2,X3] : ((! [X4] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3))))))))) & (? [X4] : p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3))))))))))),
  inference(nnf_transformation,[],[f352])).
fof(f352,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) <~> ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : ! [X2,X3] : (! [X4] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X4))),s(fun('type.pair.prod'(X0,X1),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X0,X1)),X3)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (! [X10] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X8),s(X7,X10))),s(fun('type.pair.prod'(X6,X7),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X6,X7)),X9)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun('type.pair.prod'(X6,X7),X6),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X6,X7)),X9)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (! [X10] : ~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X8),s(X7,X10))),s(fun('type.pair.prod'(X6,X7),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(X6,X7)),X9)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun('type.pair.prod'(X6,X7),X6),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(X6,X7)),X9)))))))))),
  file('thm.typeSound.typeSound_46_prove_ltb','thm.typeSound.typeSound_46_prove_ltb')).
fof(f979,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,sK18))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) | $spl0),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,sK18))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f981,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | $spl2),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f597,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK16),s(fun(sK14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK14),'const.list.MAP_2'(s(fun('type.pair.prod'(sK14,sK15),sK14),'const.pair.FST_0'),s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17)))))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK14,sK15),'const.pair.,_2'(s(sK14,sK16),s(sK15,sK18))),s(fun('type.pair.prod'(sK14,sK15),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.pair.prod'(sK14,sK15)),sK17))))))),
  inference(cnf_transformation,[],[f441])).
% SZS output end Proof for thm.typeSound
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 635
% Passive clauses: 3419
% Generated clauses: 37131
% Final active clauses: 562
% Final passive clauses: 2487
% Input formulas: 133
% Initial clauses: 354
% 
% Duplicate literals: 10
% Trivial inequalities: 5
% Fw subsumption resolutions: 645
% Fw demodulations: 13311
% Bw demodulations: 25
% Condensations: 19
% 
% Simple tautologies: 95
% Equational tautologies: 24
% Forward subsumptions: 10892
% Fw demodulations to eq. taut.: 7904
% 
% Unit resulting resolution: 14165
% Factoring: 1
% Forward superposition: 4824
% Backward superposition: 3675
% Self superposition: 79
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 4232
% SAT solver unit clauses: 1768
% SAT solver binary clauses: 1661
% 
% Sat splits: 4
% Sat splitting refutations: 803
% 
% Memory used [KB]: 41577
% Time elapsed: 2.431 s
% ------------------------------
% ------------------------------
% Success in time 3.78 s
