Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f16066,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f16065,f15918,f225,f355,f353,f15884,f15896,f274,f356,f364,f237,f357,f224,f354,f352])).
fof(f352,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) | $spl0),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f354,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0) | $spl2),
  inference(cnf_transformation,[],[f354_D])).
fof(f354_D,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f224,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0)))),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') != s('type.list.list'('type.closLang.exp'),sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f187])).
fof(f187,plain,(
  ? [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X0))) | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X0))) | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') != s('type.list.list'('type.closLang.exp'),X0)))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ? [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X0))) <~> s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X0))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ~! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X0))) <=> s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X6))) <=> s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X6))) <=> s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X6))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.clos_relationProps.exp3_size_EQ|48|')).
fof(f357,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f357_D])).
fof(f357_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f237,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relationProps.exp3_size_EQ|48|',aHL_FALSITY)).
fof(f364,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) ) | $spl16),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f356,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(cnf_transformation,[],[f356_D])).
fof(f356_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f274,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))) <=> p(s(bool,'F_0')))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.numeral.numeral_eq|split|0')).
fof(f15896,plain,(
  ( ! [X149,X150] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X149),s('type.list.list'('type.closLang.exp'),X150)))))) ) | $spl94),
  inference(cnf_transformation,[],[f15896_D])).
fof(f15896_D,plain,(
  ( ! [X149,X150] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X149),s('type.list.list'('type.closLang.exp'),X150)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f15884,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X11),s('type.list.list'('type.closLang.exp'),X12)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f15773,f801])).
fof(f801,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(subsumption_resolution,[],[f792,f364])).
fof(f792,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f280,f232])).
fof(f232,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f58])).
fof(f58,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.numeral.numeral_distrib|split|13')).
fof(f280,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X8))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.numeral.numeral_distrib|split|17')).
fof(f15773,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X11),s('type.list.list'('type.closLang.exp'),X12))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f341,f15544])).
fof(f15544,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X0),s('type.list.list'('type.closLang.exp'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X0))),s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X1)))))))) )),
  inference(forward_demodulation,[],[f15543,f10546])).
fof(f10546,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f10503,f232])).
fof(f10503,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(trivial_inequality_removal,[],[f10407])).
fof(f10407,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(superposition,[],[f280,f10269])).
fof(f10269,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f10222,f247])).
fof(f247,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.numeral.numeral_distrib|split|4')).
fof(f10222,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f10111,f246])).
fof(f246,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.arithmetic.MULT|split|0')).
fof(f10111,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f317,f814])).
fof(f814,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(equality_resolution,[],[f281])).
fof(f281,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f199])).
fof(f317,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X8,X7] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.numeral.numeral_distrib|split|5')).
fof(f15543,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X0),s('type.list.list'('type.closLang.exp'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X0))),s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X1)))))))) )),
  inference(forward_demodulation,[],[f325,f10503])).
fof(f325,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X0),s('type.list.list'('type.closLang.exp'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X0))),s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X1)))))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X0),s('type.list.list'('type.closLang.exp'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X0))),s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X1)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17,X18] : s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.CONS_2'(s('type.closLang.exp',X17),s('type.list.list'('type.closLang.exp'),X18))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.closLang.exp_size_1'(s('type.closLang.exp',X17))),s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X18)))))))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.closLang.exp_size_def|split|15')).
fof(f341,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.arithmetic.ADD_EQ_|48|')).
fof(f353,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) | $spl1),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f355,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') != s('type.list.list'('type.closLang.exp'),sK0) | $spl3),
  inference(cnf_transformation,[],[f355_D])).
fof(f355_D,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') != s('type.list.list'('type.closLang.exp'),sK0) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f225,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') != s('type.list.list'('type.closLang.exp'),sK0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),sK0)))),
  inference(cnf_transformation,[],[f188])).
fof(f15918,plain,(
  $false | ($spl0 | $spl3 | $spl94)),
  inference(subsumption_resolution,[],[f15917,f355])).
fof(f15917,plain,(
  s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0) | ($spl0 | $spl94)),
  inference(trivial_inequality_removal,[],[f15915])).
fof(f15915,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),sK0) | ($spl0 | $spl94)),
  inference(superposition,[],[f15912,f352])).
fof(f15912,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),X0))) | s('type.list.list'('type.closLang.exp'),'const.list.NIL_0') = s('type.list.list'('type.closLang.exp'),X0)) ) | $spl94),
  inference(superposition,[],[f15896,f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK1(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK1(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f143])).
fof(f143,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X13] : ! [X6] : (s('type.list.list'(X13),'const.list.NIL_0') = s('type.list.list'(X13),X6) | ? [X21,X12] : s('type.list.list'(X13),X6) = s('type.list.list'(X13),'const.list.CONS_2'(s(X13,X21),s('type.list.list'(X13),X12))))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.list.list_CASES')).
fof(f16065,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f16064,f230])).
fof(f230,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.NIL_0')))),
  file('thm.clos_relationProps.exp3_size_EQ|48|','thm.closLang.exp_size_def|split|14')).
fof(f16064,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.closLang.exp3_size_1'(s('type.list.list'('type.closLang.exp'),'const.list.NIL_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f353,f354])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 302
% Passive clauses: 2254
% Generated clauses: 15914
% Final active clauses: 148
% Final passive clauses: 1203
% Input formulas: 89
% Initial clauses: 126
% 
% Duplicate literals: 81
% Trivial inequalities: 226
% Fw subsumption resolutions: 720
% Fw demodulations: 2618
% Bw demodulations: 74
% 
% Simple tautologies: 74
% Equational tautologies: 22
% Forward subsumptions: 10024
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 135
% Factoring: 1
% Forward superposition: 4804
% Backward superposition: 6842
% Self superposition: 77
% Equality factoring: 24
% Equality resolution: 17
% 
% Split clauses: 71
% Split components: 142
% SAT solver clauses: 143
% SAT solver unit clauses: 20
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 71
% Sat splitting refutations: 25
% 
% TWLsolver clauses: 140
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 10234
% Time elapsed: 0.276 s
% ------------------------------
% ------------------------------
% Success in time 0.28 s
