Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% (58221)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 2035
% Generated clauses: 18428
% Final active clauses: 101
% Final passive clauses: 1637
% Input formulas: 21
% Initial clauses: 93
% 
% Duplicate literals: 1062
% Trivial inequalities: 209
% Fw subsumption resolutions: 2874
% Fw demodulations: 1369
% Bw demodulations: 9
% 
% Simple tautologies: 341
% Equational tautologies: 130
% Forward subsumptions: 7282
% 
% Binary resolution: 153
% Forward superposition: 5654
% Backward superposition: 5665
% Self superposition: 424
% Equality factoring: 380
% Equality resolution: 15
% 
% Split clauses: 1459
% Split components: 2972
% SAT solver clauses: 3133
% SAT solver unit clauses: 218
% SAT solver binary clauses: 735
% 
% Sat splits: 1459
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 3126
% TWLsolver calls for satisfiability: 108
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relationProps
% (58253)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 230
% Generated clauses: 5371
% Final active clauses: 93
% Final passive clauses: 136
% Input formulas: 21
% Initial clauses: 93
% 
% Duplicate literals: 658
% Trivial inequalities: 74
% Fw subsumption resolutions: 368
% Fw demodulations: 84
% Bw demodulations: 1
% Condensations: 975
% Global subsumptions: 12
% 
% Simple tautologies: 201
% Equational tautologies: 31
% Forward subsumptions: 1554
% 
% Binary resolution: 101
% Forward superposition: 1423
% Backward superposition: 1087
% Self superposition: 152
% Equality factoring: 120
% Equality resolution: 14
% 
% Split clauses: 64
% Split components: 128
% SAT solver clauses: 550
% SAT solver unit clauses: 106
% SAT solver binary clauses: 88
% SAT solver learnt clauses: 24
% 
% Sat splits: 64
% 
% TWLsolver clauses: 502
% TWLsolver calls for satisfiability: 1229
% 
% Memory used [KB]: 4733
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relationProps
% (58294)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 604
% Passive clauses: 5590
% Generated clauses: 91571
% Final active clauses: 115
% Final passive clauses: 274
% Input formulas: 21
% Initial clauses: 93
% 
% Duplicate literals: 4202
% Trivial inequalities: 303
% Fw subsumption resolutions: 5842
% Fw demodulations: 10122
% Bw demodulations: 3098
% Condensations: 8008
% Global subsumptions: 916
% 
% Simple tautologies: 444
% Equational tautologies: 1085
% Forward subsumptions: 53258
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 20809
% Forward superposition: 11146
% Backward superposition: 11063
% Self superposition: 425
% Equality factoring: 400
% Equality resolution: 26
% 
% Split clauses: 758
% Split components: 1533
% SAT solver clauses: 9254
% SAT solver unit clauses: 40
% SAT solver binary clauses: 921
% 
% Sat splits: 758
% Sat splitting refutations: 1230
% 
% Memory used [KB]: 38250
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f13147,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f155,f204,f154,f205,f203,f207,f229,f234,f239,f212,f272,f273,f274,f295,f275,f299,f278,f279,f280,f303,f281,f307,f284,f285,f286,f311,f287,f315,f472,f473,f474,f475,f478,f479,f480,f481,f605,f606,f607,f608,f611,f612,f613,f614,f733,f734,f735,f736,f745,f746,f747,f748,f861,f862,f863,f489,f864,f875,f876,f877,f878,f990,f991,f992,f491,f993,f996,f997,f998,f999,f884,f1121,f1122,f1123,f493,f1124,f1127,f1128,f1129,f1130,f886,f1370,f1371,f1372,f1373,f1727,f1731,f1753,f1822,f1839,f1840,f1913,f1914,f1915,f1916,f1919,f1920,f1921,f1922,f1925,f1926,f1927,f1928,f1940,f1992,f2130,f2251,f2252,f2253,f2254,f2257,f2258,f2259,f2260,f492,f2551,f2552,f2553,f2554,f2739,f2745,f2749,f2753,f2757,f2761,f2764,f2767,f2770,f2774,f2886,f2887,f2888,f2889,f2892,f2893,f2894,f2895,f2898,f2899,f2900,f2901,f3087,f3089,f488,f3101,f3102,f3103,f3104,f885,f3240,f3323,f3325,f3329,f3330,f3331,f3332,f306,f3335,f3336,f3337,f3338,f3423,f3425,f3485,f3487,f3777,f3778,f3779,f3780,f3783,f3784,f3785,f3786,f3834,f3835,f3836,f3837,f4080,f4084,f4087,f4089,f4093,f4096,f4098,f4101,f4103,f4105,f4107,f4109,f4373,f4374,f4375,f4376,f4377,f4395,f4398,f4400,f4401,f167,f163,f159,f157,f4405,f4467,f4468,f4469,f490,f4470,f4473,f4475,f4569,f4571,f4615,f4617,f4572,f4570,f4476,f4474,f883,f4676,f4677,f4678,f4679,f1394,f5149,f5166,f5183,f5652,f5653,f5693,f5654,f5655,f5695,f5661,f5662,f5699,f5663,f5664,f5700,f5670,f5671,f5703,f5672,f5673,f5704,f5873,f5786,f5787,f5788,f5789,f5877,f5830,f5831,f5880,f5832,f5833,f5881,f5839,f5840,f5884,f5841,f5842,f5885,f5848,f5849,f5876,f5889,f5891,f5915,f6115,f5999,f6000,f6001,f6002,f6118,f6071,f6072,f6121,f6073,f6074,f6122,f6080,f6081,f6117,f6125,f6126,f6089,f6090,f6130,f6091,f6092,f6131,f6327,f6328,f6196,f6197,f6198,f6199,f6331,f6236,f6237,f6238,f6239,f6333,f6284,f6285,f6336,f6286,f6287,f6337,f6293,f6294,f6330,f6340,f6341,f6302,f6303,f6332,f6345,f6346,f6566,f6410,f6411,f6412,f6413,f6568,f6521,f6522,f6567,f6571,f6572,f6530,f6531,f6576,f6532,f6533,f6577,f6539,f6540,f6580,f6541,f6542,f6581,f6718,f6742,f6745,f6747,f6749,f6753,f6756,f6758,f6760,f6902,f6910,f6738,f6736,f6734,f6732,f6730,f6728,f6726,f6985,f6977,f6976,f6975,f7036,f7037,f7038,f1392,f7039,f1993,f7050,f7064,f7076,f2775,f7142,f7161,f7177,f7193,f7210,f7588,f7591,f7422,f7423,f7424,f7425,f7599,f7495,f7496,f7497,f7498,f7601,f7542,f7543,f7598,f7604,f7605,f7552,f7553,f7609,f7554,f7555,f7610,f7562,f7563,f7600,f7614,f7615,f7867,f7869,f7695,f7696,f7697,f7698,f7872,f7735,f7736,f7737,f7738,f7874,f7820,f7821,f7871,f7877,f7878,f7830,f7831,f5888,f7873,f7882,f7883,f5890,f7840,f7841,f5694,f7887,f7842,f7843,f5696,f7888,f7967,f7995,f7997,f7999,f8003,f8006,f8008,f8009,f8010,f8012,f6387,f3324,f3326,f2043,f1979,f1935,f7593,f2037,f2031,f1986,f5563,f5719,f235,f5952,f6145,f240,f8163,f8164,f8165,f294,f8166,f298,f8169,f8171,f302,f8175,f8176,f8177,f310,f8178,f314,f8289,f8291,f8295,f8296,f8297,f623,f8298,f625,f8904,f8906,f8292,f8290,f8907,f8908,f621,f8918,f8937,f8958,f8978,f9001,f9074,f9279,f9280,f9281,f8058,f9282,f8056,f9328,f9339,f2740,f9352,f9364,f2746,f9430,f2750,f9446,f9459,f2741,f2758,f9473,f2762,f9489,f2751,f2768,f9565,f9592,f241,f230,f9728,f9729,f9730,f486,f9731,f619,f9740,f9741,f9742,f753,f9743,f755,f9867,f9868,f9869,f1845,f9870,f1941,f9881,f9885,f9892,f9898,f9903,f9909,f9911,f9912,f9918,f231,f206,f9958,f10006,f10007,f10008,f487,f10009,f10018,f10019,f10020,f754,f10021,f756,f10180,f10181,f10182,f1846,f10183,f10332,f10334,f10483,f10532,f10537,f10540,f10545,f10558,f10561,f214,f10572,f10628,f10629,f10630,f1942,f10631,f10927,f10973,f10979,f10984,f10987,f10992,f10998,f11003,f11006,f11011,f11017,f11021,f218,f11076,f11077,f11078,f11079,f1994,f11414,f11467,f11486,f11491,f11496,f11501,f11506,f11511,f11516,f11531,f11536,f216,f11547,f11599,f11600,f11601,f1393,f293,f11602,f297,f11904,f11917,f11922,f11928,f11932,f11936,f11942,f11946,f11950,f11954,f11958,f11962,f11966,f11970,f11974,f224,f11978,f11982,f12095,f12096,f12097,f12098,f1395,f12234,f12285,f12290,f12295,f12300,f12305,f12310,f12315,f12320,f12325,f12330,f12335,f12340,f12346,f12351,f220,f12392,f12393,f12394,f8059,f12395,f12412,f12519,f12521,f622,f12524,f12525,f12526,f624,f309,f12527,f626,f313,f12645,f12646,f12647,f8057,f12648,f225,f305,f12954,f12969,f12983,f12988,f13003,f13019,f13023,f13029,f13031,f13043,f13055,f13061,f13066,f13071,f13076,f236,f12520,f12518,f620,f301,f13082,f13088,f13094,f13100,f13106,f226,f222,f13111,f13112,f13146])).
fof(f13146,plain,(
  $false | ($spl14 | $spl23)),
  inference(subsumption_resolution,[],[f12430,f236])).
fof(f12430,plain,(
  p(s(bool,sK9)) | $spl14),
  inference(resolution,[],[f222,f126])).
fof(f126,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | ~sP5(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((~sP5(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X2)) & p(s(bool,X1)) & ~p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | sP5(X0,X1,X2,X3)))),
  inference(rectify,[],[f63])).
fof(f63,plain,(
  ! [X2,X1,X0,X3] : ((~sP5(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | sP5(X2,X1,X0,X3)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X2,X1,X0,X3] : ((~sP5(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) | sP5(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X2,X1,X0,X3] : (sP5(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f13112,plain,(
  $false | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f218,f3258])).
fof(f3258,plain,(
  ( ! [X80,X78,X79] : (~sP3(sK10,X78,X79,X80)) ) | $spl26),
  inference(subsumption_resolution,[],[f3208,f170])).
fof(f170,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb',aHL_FALSITY)).
fof(f3208,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'F_0')) | ~sP3(sK10,X78,X79,X80)) ) | $spl26),
  inference(superposition,[],[f137,f293])).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~sP3(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((~sP3(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | sP3(X0,X1,X2,X3)))),
  inference(rectify,[],[f69])).
fof(f69,plain,(
  ! [X2,X1,X0,X3] : ((~sP3(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | sP3(X2,X1,X0,X3)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X2,X1,X0,X3] : ((~sP3(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2))) | sP3(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X2,X1,X0,X3] : (sP3(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f13111,plain,(
  $false | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f13110,f170])).
fof(f13110,plain,(
  p(s(bool,'F_0')) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f10592,f293])).
fof(f10592,plain,(
  p(s(bool,sK10)) | $spl10),
  inference(resolution,[],[f218,f137])).
fof(f222,plain,(
  sP5(sK10,sK9,sK8,sK11) | $spl14),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  sP5(sK10,sK9,sK8,sK11) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f226,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl18),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f13106,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13105,f205])).
fof(f13105,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13104,f3268])).
fof(f3268,plain,(
  ( ! [X116,X114,X115] : (~sP1(sK10,X114,X115,X116)) ) | $spl26),
  inference(subsumption_resolution,[],[f3220,f170])).
fof(f3220,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'F_0')) | ~sP1(sK10,X114,X115,X116)) ) | $spl26),
  inference(superposition,[],[f147,f293])).
fof(f147,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((~sP1(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | sP1(X0,X1,X2,X3)))),
  inference(rectify,[],[f75])).
fof(f75,plain,(
  ! [X2,X1,X0,X3] : ((~sP1(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP1(X2,X1,X0,X3)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X2,X1,X0,X3] : ((~sP1(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) | sP1(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X2,X1,X0,X3] : (sP1(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f13104,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13103,f4116])).
fof(f4116,plain,(
  ( ! [X92,X90,X91] : (~sP2(X90,X91,sK8,X92)) ) | $spl42),
  inference(subsumption_resolution,[],[f3961,f170])).
fof(f3961,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,'F_0')) | ~sP2(X90,X91,sK8,X92)) ) | $spl42),
  inference(superposition,[],[f140,f309])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~sP2(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : ((~sP2(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2,X3)))),
  inference(rectify,[],[f72])).
fof(f72,plain,(
  ! [X2,X1,X0,X3] : ((~sP2(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | sP2(X2,X1,X0,X3)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X2,X1,X0,X3] : ((~sP2(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2))) | sP2(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X2,X1,X0,X3] : (sP2(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f13103,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13102,f3258])).
fof(f13102,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f13101,f4119])).
fof(f4119,plain,(
  ( ! [X54,X56,X55] : (~sP4(X54,X55,sK8,X56)) ) | $spl42),
  inference(subsumption_resolution,[],[f3949,f170])).
fof(f3949,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'F_0')) | ~sP4(X54,X55,sK8,X56)) ) | $spl42),
  inference(superposition,[],[f130,f309])).
fof(f130,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~sP4(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : ((~sP4(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X2)) & p(s(bool,X1)) & ~p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | sP4(X0,X1,X2,X3)))),
  inference(rectify,[],[f66])).
fof(f66,plain,(
  ! [X2,X1,X0,X3] : ((~sP4(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | sP4(X2,X1,X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X2,X1,X0,X3] : ((~sP4(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) | sP4(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X2,X1,X0,X3] : (sP4(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f13101,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f8460,f4121])).
fof(f4121,plain,(
  ( ! [X19,X20,X18] : (~sP6(X18,X19,sK8,X20)) ) | $spl42),
  inference(subsumption_resolution,[],[f3937,f170])).
fof(f3937,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'F_0')) | ~sP6(X18,X19,sK8,X20)) ) | $spl42),
  inference(superposition,[],[f120,f309])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~sP6(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ((~sP6(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | sP6(X0,X1,X2,X3)))),
  inference(rectify,[],[f60])).
fof(f60,plain,(
  ! [X2,X1,X0,X3] : ((~sP6(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X2,X1,X0,X3)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X2,X1,X0,X3] : ((~sP6(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) | sP6(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X2,X1,X0,X3] : (sP6(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f8460,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(resolution,[],[f107,f211])).
fof(f211,plain,(
  ~sP0(sK10,sK9,sK8,sK11) | $spl5),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ~sP0(sK10,sK9,sK8,sK11) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(X0,X1,X2,X3) | sP5(X0,X1,X2,X3) | sP4(X0,X1,X2,X3) | sP3(X0,X1,X2,X3) | sP2(X0,X1,X2,X3) | sP1(X0,X1,X2,X3) | sP0(X0,X1,X2,X3) | ~sP7(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ((~sP7(X0,X1,X2,X3) | sP0(X0,X1,X2,X3) | sP1(X0,X1,X2,X3) | sP2(X0,X1,X2,X3) | sP3(X0,X1,X2,X3) | sP4(X0,X1,X2,X3) | sP5(X0,X1,X2,X3) | sP6(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0)))) & ((~sP0(X0,X1,X2,X3) & ~sP1(X0,X1,X2,X3) & ~sP2(X0,X1,X2,X3) & ~sP3(X0,X1,X2,X3) & ~sP4(X0,X1,X2,X3) & ~sP5(X0,X1,X2,X3) & ~sP6(X0,X1,X2,X3) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) | sP7(X0,X1,X2,X3)))),
  inference(rectify,[],[f57])).
fof(f57,plain,(
  ! [X2,X1,X0,X3] : ((~sP7(X2,X1,X0,X3) | sP0(X2,X1,X0,X3) | sP1(X2,X1,X0,X3) | sP2(X2,X1,X0,X3) | sP3(X2,X1,X0,X3) | sP4(X2,X1,X0,X3) | sP5(X2,X1,X0,X3) | sP6(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))) & ((~sP0(X2,X1,X0,X3) & ~sP1(X2,X1,X0,X3) & ~sP2(X2,X1,X0,X3) & ~sP3(X2,X1,X0,X3) & ~sP4(X2,X1,X0,X3) & ~sP5(X2,X1,X0,X3) & ~sP6(X2,X1,X0,X3) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)))) | sP7(X2,X1,X0,X3)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X2,X1,X0,X3] : ((~sP7(X2,X1,X0,X3) | (sP0(X2,X1,X0,X3) | sP1(X2,X1,X0,X3) | sP2(X2,X1,X0,X3) | sP3(X2,X1,X0,X3) | sP4(X2,X1,X0,X3) | sP5(X2,X1,X0,X3) | sP6(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2))))) & ((~sP0(X2,X1,X0,X3) & ~sP1(X2,X1,X0,X3) & ~sP2(X2,X1,X0,X3) & ~sP3(X2,X1,X0,X3) & ~sP4(X2,X1,X0,X3) & ~sP5(X2,X1,X0,X3) & ~sP6(X2,X1,X0,X3) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)))) | sP7(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X2,X1,X0,X3] : (sP7(X2,X1,X0,X3) <=> (sP0(X2,X1,X0,X3) | sP1(X2,X1,X0,X3) | sP2(X2,X1,X0,X3) | sP3(X2,X1,X0,X3) | sP4(X2,X1,X0,X3) | sP5(X2,X1,X0,X3) | sP6(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f13100,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13099,f205])).
fof(f13099,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13098,f3268])).
fof(f13098,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13097,f4116])).
fof(f13097,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13096,f3258])).
fof(f13096,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f13095,f4119])).
fof(f13095,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f1841,f4121])).
fof(f1841,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(resolution,[],[f211,f107])).
fof(f13094,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13093,f4112])).
fof(f4112,plain,(
  ( ! [X127,X128,X126] : (~sP0(X126,X127,sK8,X128)) ) | $spl42),
  inference(subsumption_resolution,[],[f3973,f170])).
fof(f3973,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,'F_0')) | ~sP0(X126,X127,sK8,X128)) ) | $spl42),
  inference(superposition,[],[f150,f309])).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((~sP0(X0,X1,X2,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | sP0(X0,X1,X2,X3)))),
  inference(rectify,[],[f78])).
fof(f78,plain,(
  ! [X2,X1,X0,X3] : ((~sP0(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP0(X2,X1,X0,X3)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X2,X1,X0,X3] : ((~sP0(X2,X1,X0,X3) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) | sP0(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X2,X1,X0,X3] : (sP0(X2,X1,X0,X3) <=> (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f13093,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13092,f3268])).
fof(f13092,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13091,f4116])).
fof(f13091,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13090,f3258])).
fof(f13090,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl42)),
  inference(subsumption_resolution,[],[f13089,f4119])).
fof(f13089,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl42)),
  inference(subsumption_resolution,[],[f9949,f4121])).
fof(f9949,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f107])).
fof(f13088,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13087,f205])).
fof(f13087,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13086,f4112])).
fof(f13086,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13085,f3268])).
fof(f13085,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13084,f3258])).
fof(f13084,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl42)),
  inference(subsumption_resolution,[],[f13083,f4119])).
fof(f13083,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl42)),
  inference(subsumption_resolution,[],[f11537,f4121])).
fof(f11537,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(resolution,[],[f215,f107])).
fof(f215,plain,(
  ~sP2(sK10,sK9,sK8,sK11) | $spl9),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  ~sP2(sK10,sK9,sK8,sK11) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f13082,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13081,f205])).
fof(f13081,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13080,f4112])).
fof(f13080,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13079,f3268])).
fof(f13079,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13078,f4116])).
fof(f13078,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13077,f3258])).
fof(f13077,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl42)),
  inference(subsumption_resolution,[],[f12402,f4121])).
fof(f12402,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl13),
  inference(resolution,[],[f219,f107])).
fof(f219,plain,(
  ~sP4(sK10,sK9,sK8,sK11) | $spl13),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  ~sP4(sK10,sK9,sK8,sK11) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f301,plain,(
  s(bool,'F_0') = s(bool,sK9) | $spl34),
  inference(cnf_transformation,[],[f301_D])).
fof(f301_D,plain,(
  s(bool,'F_0') = s(bool,sK9) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f620,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | $spl59),
  inference(cnf_transformation,[],[f620_D])).
fof(f620_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f12518,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f169])).
fof(f169,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb',aHL_BOOL_CASES)).
fof(f12251,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | ($spl1 | $spl26)),
  inference(forward_demodulation,[],[f204,f293])).
fof(f12520,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f168])).
fof(f168,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) | s(bool,'F_0') = s(bool,X10))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.BOOL_CASES_AX')).
fof(f236,plain,(
  ~p(s(bool,sK9)) | $spl23),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  ~p(s(bool,sK9)) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f13076,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13075,f205])).
fof(f13075,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13074,f3268])).
fof(f13074,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f13073,f141])).
fof(f141,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | ~sP2(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f13073,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f13072,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | ~sP3(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f70])).
fof(f13072,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f9886,f4119])).
fof(f9886,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(subsumption_resolution,[],[f1843,f121])).
fof(f121,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | ~sP6(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f61])).
fof(f1843,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(resolution,[],[f211,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | sP6(X0,X1,X2,X3) | sP5(X0,X1,X2,X3) | sP4(X0,X1,X2,X3) | sP3(X0,X1,X2,X3) | sP2(X0,X1,X2,X3) | sP1(X0,X1,X2,X3) | sP0(X0,X1,X2,X3) | ~sP7(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f13071,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13070,f4112])).
fof(f13070,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13069,f3268])).
fof(f13069,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl42)),
  inference(subsumption_resolution,[],[f13068,f141])).
fof(f13068,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl42)),
  inference(subsumption_resolution,[],[f13067,f136])).
fof(f13067,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl42)),
  inference(subsumption_resolution,[],[f9959,f4119])).
fof(f9959,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f9951,f121])).
fof(f9951,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f109])).
fof(f13066,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13065,f205])).
fof(f13065,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13064,f4112])).
fof(f13064,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13063,f3268])).
fof(f13063,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl42)),
  inference(subsumption_resolution,[],[f13062,f136])).
fof(f13062,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl42)),
  inference(subsumption_resolution,[],[f11556,f4119])).
fof(f11556,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(subsumption_resolution,[],[f11540,f121])).
fof(f11540,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(resolution,[],[f215,f109])).
fof(f13061,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13060,f205])).
fof(f13060,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13059,f4112])).
fof(f13059,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13058,f3268])).
fof(f13058,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl42)),
  inference(subsumption_resolution,[],[f13057,f141])).
fof(f13057,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl42)),
  inference(subsumption_resolution,[],[f13056,f136])).
fof(f13056,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl42)),
  inference(subsumption_resolution,[],[f12040,f4119])).
fof(f12040,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl17),
  inference(resolution,[],[f223,f109])).
fof(f223,plain,(
  ~sP6(sK10,sK9,sK8,sK11) | $spl17),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ~sP6(sK10,sK9,sK8,sK11) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f13055,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13054,f205])).
fof(f13054,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13053,f4112])).
fof(f13053,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl26)),
  inference(subsumption_resolution,[],[f13052,f3268])).
fof(f13052,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl13),
  inference(subsumption_resolution,[],[f13051,f141])).
fof(f13051,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl13),
  inference(subsumption_resolution,[],[f13050,f136])).
fof(f13050,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl13),
  inference(subsumption_resolution,[],[f12405,f121])).
fof(f12405,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl13),
  inference(resolution,[],[f219,f109])).
fof(f13043,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13042,f205])).
fof(f13042,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13041,f4112])).
fof(f13041,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13040,f3268])).
fof(f13040,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13039,f4116])).
fof(f13039,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl26 | $spl42)),
  inference(subsumption_resolution,[],[f13038,f3258])).
fof(f13038,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl42)),
  inference(subsumption_resolution,[],[f12037,f4119])).
fof(f12037,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl17),
  inference(resolution,[],[f223,f107])).
fof(f13031,plain,(
  $false | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f214,f3268])).
fof(f13029,plain,(
  $false | ($spl6 | $spl26)),
  inference(subsumption_resolution,[],[f13028,f170])).
fof(f13028,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl26)),
  inference(forward_demodulation,[],[f9977,f293])).
fof(f9977,plain,(
  p(s(bool,sK10)) | $spl6),
  inference(resolution,[],[f214,f147])).
fof(f13023,plain,(
  $false | ($spl1 | $spl2 | $spl5 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13022,f205])).
fof(f13022,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl5 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13021,f3268])).
fof(f13021,plain,(
  sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl5 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13020,f4119])).
fof(f13020,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl5 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12288,f12902])).
fof(f12902,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2,sK11)) ) | ($spl1 | $spl26 | $spl38 | $spl42)),
  inference(resolution,[],[f12819,f124])).
fof(f124,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~sP5(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f64])).
fof(f12819,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl1 | $spl26 | $spl38 | $spl42)),
  inference(backward_demodulation,[],[f305,f12254])).
fof(f12254,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | ($spl1 | $spl26 | $spl42)),
  inference(forward_demodulation,[],[f10492,f293])).
fof(f10492,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f204,f309])).
fof(f12288,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12287,f215])).
fof(f12287,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12286,f217])).
fof(f217,plain,(
  ~sP3(sK10,sK9,sK8,sK11) | $spl11),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  ~sP3(sK10,sK9,sK8,sK11) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f12286,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl22)),
  inference(subsumption_resolution,[],[f9886,f235])).
fof(f13019,plain,(
  $false | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13018,f4112])).
fof(f13018,plain,(
  sP0(sK10,sK9,sK8,sK11) | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13017,f3268])).
fof(f13017,plain,(
  sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13016,f4119])).
fof(f13016,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12293,f12902])).
fof(f12293,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12292,f215])).
fof(f12292,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12291,f217])).
fof(f12291,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f9959,f235])).
fof(f13003,plain,(
  $false | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13002,f205])).
fof(f13002,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13001,f4112])).
fof(f13001,plain,(
  sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f13000,f3268])).
fof(f13000,plain,(
  sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12999,f4119])).
fof(f12999,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12317,f12902])).
fof(f12317,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12316,f215])).
fof(f12316,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl22)),
  inference(subsumption_resolution,[],[f11517,f235])).
fof(f11517,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl11),
  inference(subsumption_resolution,[],[f11032,f121])).
fof(f11032,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl11),
  inference(resolution,[],[f217,f109])).
fof(f12988,plain,(
  $false | ($spl1 | $spl2 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12987,f205])).
fof(f12987,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12986,f4112])).
fof(f12986,plain,(
  sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12985,f3268])).
fof(f12985,plain,(
  sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12984,f4119])).
fof(f12984,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12332,f12902])).
fof(f12332,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12331,f217])).
fof(f12331,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl22)),
  inference(subsumption_resolution,[],[f11556,f235])).
fof(f12983,plain,(
  $false | ($spl1 | $spl2 | $spl9 | $spl11 | $spl17 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12982,f205])).
fof(f12982,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl17 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12981,f4112])).
fof(f12981,plain,(
  sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl17 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12980,f3268])).
fof(f12980,plain,(
  sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl17 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12979,f4119])).
fof(f12979,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl17 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12343,f12902])).
fof(f12343,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12342,f215])).
fof(f12342,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12341,f217])).
fof(f12341,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12040,f235])).
fof(f12969,plain,(
  $false | ($spl1 | $spl2 | $spl9 | $spl11 | $spl13 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12968,f205])).
fof(f12968,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl13 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12967,f4112])).
fof(f12967,plain,(
  sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl13 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12966,f3268])).
fof(f12966,plain,(
  sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl1 | $spl9 | $spl11 | $spl13 | $spl22 | $spl26 | $spl38 | $spl42)),
  inference(subsumption_resolution,[],[f12416,f12902])).
fof(f12416,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl13 | $spl22 | $spl38)),
  inference(subsumption_resolution,[],[f12415,f215])).
fof(f12415,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl13 | $spl22 | $spl38)),
  inference(subsumption_resolution,[],[f12414,f217])).
fof(f12414,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl22 | $spl38)),
  inference(subsumption_resolution,[],[f12413,f2736])).
fof(f2736,plain,(
  ( ! [X23,X21,X22] : (~sP6(X21,sK9,X22,X23)) ) | $spl38),
  inference(subsumption_resolution,[],[f2672,f200])).
fof(f200,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f175])).
fof(f175,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : ! [X12] : (s(X11,X12) = s(X11,X12) <=> p(s(bool,'T_0')))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f2672,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | ~sP6(X21,sK9,X22,X23)) ) | $spl38),
  inference(superposition,[],[f121,f305])).
fof(f12413,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl22)),
  inference(subsumption_resolution,[],[f12405,f235])).
fof(f12954,plain,(
  $false | ($spl1 | $spl14 | $spl26 | $spl38 | $spl42)),
  inference(resolution,[],[f12902,f222])).
fof(f305,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl38),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  s(bool,'T_0') = s(bool,sK9) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f225,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl19),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f12648,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl42)),
  inference(superposition,[],[f12254,f168])).
fof(f8057,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl187),
  inference(cnf_transformation,[],[f8057_D])).
fof(f8057_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f12647,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl42)),
  inference(superposition,[],[f12254,f168])).
fof(f12646,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl42)),
  inference(superposition,[],[f12254,f169])).
fof(f12645,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl42)),
  inference(superposition,[],[f12254,f169])).
fof(f313,plain,(
  s(bool,'T_0') = s(bool,sK8) | $spl46),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  s(bool,'T_0') = s(bool,sK8) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f626,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | $spl65),
  inference(cnf_transformation,[],[f626_D])).
fof(f626_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f12527,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f168])).
fof(f309,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl42),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  s(bool,'F_0') = s(bool,sK8) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f624,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | $spl63),
  inference(cnf_transformation,[],[f624_D])).
fof(f624_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f12526,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f168])).
fof(f12525,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f169])).
fof(f12524,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f169])).
fof(f622,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | $spl61),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f12521,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f168])).
fof(f12519,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26)),
  inference(superposition,[],[f12251,f169])).
fof(f12412,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12411,f205])).
fof(f12411,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12410,f211])).
fof(f12410,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12409,f213])).
fof(f213,plain,(
  ~sP1(sK10,sK9,sK8,sK11) | $spl7),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ~sP1(sK10,sK9,sK8,sK11) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f12409,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12408,f215])).
fof(f12408,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12407,f217])).
fof(f12407,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl13 | $spl38)),
  inference(subsumption_resolution,[],[f12402,f2736])).
fof(f12395,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl26 | $spl38)),
  inference(superposition,[],[f12250,f168])).
fof(f12250,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl1 | $spl26 | $spl38)),
  inference(forward_demodulation,[],[f9910,f293])).
fof(f9910,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl1 | $spl38)),
  inference(forward_demodulation,[],[f204,f305])).
fof(f8059,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl189),
  inference(cnf_transformation,[],[f8059_D])).
fof(f8059_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f12394,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl26 | $spl38)),
  inference(superposition,[],[f12250,f168])).
fof(f12393,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl26 | $spl38)),
  inference(superposition,[],[f12250,f169])).
fof(f12392,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl26 | $spl38)),
  inference(superposition,[],[f12250,f169])).
fof(f220,plain,(
  sP4(sK10,sK9,sK8,sK11) | $spl12),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  sP4(sK10,sK9,sK8,sK11) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12351,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl17)),
  inference(subsumption_resolution,[],[f12350,f205])).
fof(f12350,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl17)),
  inference(subsumption_resolution,[],[f12349,f211])).
fof(f12349,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl17)),
  inference(subsumption_resolution,[],[f12348,f213])).
fof(f12348,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl17)),
  inference(subsumption_resolution,[],[f12347,f215])).
fof(f12347,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f12037,f217])).
fof(f12346,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12345,f205])).
fof(f12345,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12344,f211])).
fof(f12344,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl17 | $spl22)),
  inference(subsumption_resolution,[],[f12343,f213])).
fof(f12340,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12339,f205])).
fof(f12339,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12338,f211])).
fof(f12338,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12337,f215])).
fof(f12337,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12336,f217])).
fof(f12336,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl22)),
  inference(subsumption_resolution,[],[f11912,f235])).
fof(f11912,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(subsumption_resolution,[],[f10565,f121])).
fof(f10565,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(resolution,[],[f213,f109])).
fof(f12335,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12334,f205])).
fof(f12334,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12333,f211])).
fof(f12333,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12332,f213])).
fof(f12330,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12329,f205])).
fof(f12329,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12328,f211])).
fof(f12328,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12327,f213])).
fof(f12327,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12326,f217])).
fof(f12326,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl38)),
  inference(subsumption_resolution,[],[f11537,f2736])).
fof(f12325,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12324,f205])).
fof(f12324,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12323,f211])).
fof(f12323,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12322,f213])).
fof(f12322,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12321,f215])).
fof(f12321,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl38)),
  inference(subsumption_resolution,[],[f11029,f2736])).
fof(f11029,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl11),
  inference(resolution,[],[f217,f107])).
fof(f12320,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12319,f205])).
fof(f12319,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12318,f211])).
fof(f12318,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12317,f213])).
fof(f12315,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12314,f205])).
fof(f12314,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12313,f211])).
fof(f12313,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12312,f215])).
fof(f12312,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12311,f217])).
fof(f12311,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl38)),
  inference(subsumption_resolution,[],[f10562,f2736])).
fof(f10562,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(resolution,[],[f213,f107])).
fof(f12310,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12309,f211])).
fof(f12309,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12308,f213])).
fof(f12308,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12307,f215])).
fof(f12307,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12306,f217])).
fof(f12306,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl38)),
  inference(subsumption_resolution,[],[f9949,f2736])).
fof(f12305,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12304,f205])).
fof(f12304,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12303,f213])).
fof(f12303,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12302,f215])).
fof(f12302,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12301,f217])).
fof(f12301,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f1841,f2736])).
fof(f12300,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12299,f205])).
fof(f12299,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12298,f213])).
fof(f12298,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12297,f215])).
fof(f12297,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl11 | $spl38)),
  inference(subsumption_resolution,[],[f12296,f217])).
fof(f12296,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f8460,f2736])).
fof(f12295,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12294,f211])).
fof(f12294,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12293,f213])).
fof(f12290,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12289,f205])).
fof(f12289,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl11 | $spl22)),
  inference(subsumption_resolution,[],[f12288,f213])).
fof(f12285,plain,(
  $false | ($spl12 | $spl21)),
  inference(subsumption_resolution,[],[f12031,f231])).
fof(f12031,plain,(
  p(s(bool,sK8)) | $spl12),
  inference(resolution,[],[f220,f130])).
fof(f12234,plain,(
  $false | ($spl1 | $spl12 | $spl26 | $spl38 | $spl46)),
  inference(resolution,[],[f12182,f220])).
fof(f12182,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2,sK11)) ) | ($spl1 | $spl26 | $spl38 | $spl46)),
  inference(resolution,[],[f11990,f129])).
fof(f129,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~sP4(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f67])).
fof(f11990,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl1 | $spl26 | $spl38 | $spl46)),
  inference(forward_demodulation,[],[f11989,f293])).
fof(f11989,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl1 | $spl38 | $spl46)),
  inference(forward_demodulation,[],[f9910,f313])).
fof(f1395,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl77),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f12098,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl46)),
  inference(superposition,[],[f11986,f168])).
fof(f11986,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | ($spl1 | $spl26 | $spl46)),
  inference(forward_demodulation,[],[f11985,f293])).
fof(f11985,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | ($spl1 | $spl46)),
  inference(forward_demodulation,[],[f204,f313])).
fof(f12097,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl46)),
  inference(superposition,[],[f11986,f168])).
fof(f12096,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl46)),
  inference(superposition,[],[f11986,f169])).
fof(f12095,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl26 | $spl46)),
  inference(superposition,[],[f11986,f169])).
fof(f11982,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11981,f205])).
fof(f11981,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11980,f4099])).
fof(f4099,plain,(
  ( ! [X109,X110,X108] : (~sP1(X108,X109,sK8,X110)) ) | $spl46),
  inference(subsumption_resolution,[],[f4040,f200])).
fof(f4040,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'T_0')) | ~sP1(X108,X109,sK8,X110)) ) | $spl46),
  inference(superposition,[],[f145,f313])).
fof(f145,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f76])).
fof(f11980,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11979,f141])).
fof(f11979,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11468,f136])).
fof(f11468,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f9886,f4081])).
fof(f4081,plain,(
  ( ! [X37,X38,X36] : (~sP5(X36,X37,sK8,X38)) ) | $spl46),
  inference(subsumption_resolution,[],[f4016,f200])).
fof(f4016,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~sP5(X36,X37,sK8,X38)) ) | $spl46),
  inference(superposition,[],[f125,f313])).
fof(f125,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | ~sP5(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f64])).
fof(f11978,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11977,f211])).
fof(f11977,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f11976,f4099])).
fof(f11976,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f11975,f141])).
fof(f11975,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f11475,f136])).
fof(f11475,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f9959,f4081])).
fof(f224,plain,(
  sP6(sK10,sK9,sK8,sK11) | $spl16),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  sP6(sK10,sK9,sK8,sK11) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f11974,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11973,f211])).
fof(f11973,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11972,f4099])).
fof(f11972,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11971,f215])).
fof(f11971,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11482,f4090])).
fof(f4090,plain,(
  ( ! [X74,X72,X73] : (~sP3(X72,X73,sK8,X74)) ) | $spl46),
  inference(subsumption_resolution,[],[f4028,f200])).
fof(f4028,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | ~sP3(X72,X73,sK8,X74)) ) | $spl46),
  inference(superposition,[],[f135,f313])).
fof(f135,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | ~sP3(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f70])).
fof(f11482,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f10974,f4081])).
fof(f10974,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f9950,f230])).
fof(f9950,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2,X3) | sP5(X0,X1,X2,X3) | sP4(X0,X1,X2,X3) | sP3(X0,X1,X2,X3) | sP2(X0,X1,X2,X3) | sP1(X0,X1,X2,X3) | sP0(X0,X1,X2,X3) | ~sP7(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f11970,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11969,f205])).
fof(f11969,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11968,f4099])).
fof(f11968,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11967,f215])).
fof(f11967,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11487,f4090])).
fof(f11487,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f8460,f4081])).
fof(f11966,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11965,f205])).
fof(f11965,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11964,f4099])).
fof(f11964,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11963,f215])).
fof(f11963,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f11492,f4090])).
fof(f11492,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl46)),
  inference(subsumption_resolution,[],[f1841,f4081])).
fof(f11962,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11961,f205])).
fof(f11961,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11960,f4099])).
fof(f11960,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11959,f215])).
fof(f11959,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11497,f4090])).
fof(f11497,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f10993,f4081])).
fof(f10993,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20)),
  inference(subsumption_resolution,[],[f1842,f230])).
fof(f1842,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(resolution,[],[f211,f108])).
fof(f11958,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11957,f211])).
fof(f11957,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11956,f4099])).
fof(f11956,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11955,f215])).
fof(f11955,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f11502,f4090])).
fof(f11502,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl46)),
  inference(subsumption_resolution,[],[f9949,f4081])).
fof(f11954,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11953,f205])).
fof(f11953,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11952,f211])).
fof(f11952,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11951,f215])).
fof(f11951,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11507,f4090])).
fof(f11507,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl46)),
  inference(subsumption_resolution,[],[f10562,f4081])).
fof(f11950,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11949,f205])).
fof(f11949,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11948,f211])).
fof(f11948,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl9 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11947,f215])).
fof(f11947,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11512,f4090])).
fof(f11512,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl46)),
  inference(subsumption_resolution,[],[f11012,f4081])).
fof(f11012,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20)),
  inference(subsumption_resolution,[],[f10564,f230])).
fof(f10564,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(resolution,[],[f213,f108])).
fof(f11946,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11945,f205])).
fof(f11945,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11944,f211])).
fof(f11944,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11943,f4099])).
fof(f11943,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11518,f141])).
fof(f11518,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11517,f4081])).
fof(f11942,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11941,f200])).
fof(f11941,plain,(
  ~p(s(bool,'T_0')) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl11 | $spl46)),
  inference(forward_demodulation,[],[f11940,f313])).
fof(f11940,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f11939,f205])).
fof(f11939,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl11)),
  inference(subsumption_resolution,[],[f11938,f211])).
fof(f11938,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f11937,f145])).
fof(f11937,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11)),
  inference(subsumption_resolution,[],[f11525,f215])).
fof(f11525,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl11),
  inference(subsumption_resolution,[],[f11031,f125])).
fof(f11031,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl11),
  inference(resolution,[],[f217,f108])).
fof(f11936,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11935,f205])).
fof(f11935,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11934,f211])).
fof(f11934,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11933,f4099])).
fof(f11933,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11532,f215])).
fof(f11532,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl46)),
  inference(subsumption_resolution,[],[f11029,f4081])).
fof(f11932,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11931,f205])).
fof(f11931,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11930,f211])).
fof(f11930,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11929,f4099])).
fof(f11929,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11542,f4090])).
fof(f11542,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11537,f4081])).
fof(f11928,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11927,f200])).
fof(f11927,plain,(
  ~p(s(bool,'T_0')) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(forward_demodulation,[],[f11926,f313])).
fof(f11926,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9)),
  inference(subsumption_resolution,[],[f11925,f205])).
fof(f11925,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9)),
  inference(subsumption_resolution,[],[f11924,f211])).
fof(f11924,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(subsumption_resolution,[],[f11923,f145])).
fof(f11923,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(subsumption_resolution,[],[f11548,f135])).
fof(f11548,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(subsumption_resolution,[],[f11539,f125])).
fof(f11539,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl9),
  inference(resolution,[],[f215,f108])).
fof(f11922,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11921,f205])).
fof(f11921,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11920,f211])).
fof(f11920,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11919,f4099])).
fof(f11919,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11557,f136])).
fof(f11557,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl46)),
  inference(subsumption_resolution,[],[f11556,f4081])).
fof(f11917,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11916,f205])).
fof(f11916,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11915,f211])).
fof(f11915,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11914,f141])).
fof(f11914,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11913,f136])).
fof(f11913,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl46)),
  inference(subsumption_resolution,[],[f11912,f4081])).
fof(f11904,plain,(
  $false | ($spl1 | $spl16 | $spl26 | $spl34 | $spl46)),
  inference(resolution,[],[f11852,f224])).
fof(f11852,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2,sK11)) ) | ($spl1 | $spl26 | $spl34 | $spl46)),
  inference(resolution,[],[f11674,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~sP6(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f61])).
fof(f11674,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl1 | $spl26 | $spl34 | $spl46)),
  inference(backward_demodulation,[],[f293,f11427])).
fof(f11427,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl1 | $spl34 | $spl46)),
  inference(forward_demodulation,[],[f10937,f313])).
fof(f10937,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | ($spl1 | $spl34)),
  inference(forward_demodulation,[],[f204,f301])).
fof(f297,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl30),
  inference(cnf_transformation,[],[f297_D])).
fof(f297_D,plain,(
  s(bool,'T_0') = s(bool,sK10) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f11602,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl1 | $spl34 | $spl46)),
  inference(superposition,[],[f11427,f168])).
fof(f293,plain,(
  s(bool,'F_0') = s(bool,sK10) | $spl26),
  inference(cnf_transformation,[],[f293_D])).
fof(f293_D,plain,(
  s(bool,'F_0') = s(bool,sK10) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1393,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl75),
  inference(cnf_transformation,[],[f1393_D])).
fof(f1393_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f11601,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl1 | $spl34 | $spl46)),
  inference(superposition,[],[f11427,f168])).
fof(f11600,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl1 | $spl34 | $spl46)),
  inference(superposition,[],[f11427,f169])).
fof(f11599,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl1 | $spl34 | $spl46)),
  inference(superposition,[],[f11427,f169])).
fof(f11547,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | ($spl2 | $spl9 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11546,f205])).
fof(f11546,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11545,f2429])).
fof(f2429,plain,(
  ( ! [X130,X131,X129] : (~sP0(X129,sK9,X130,X131)) ) | $spl34),
  inference(subsumption_resolution,[],[f2392,f170])).
fof(f2392,plain,(
  ( ! [X130,X131,X129] : (p(s(bool,'F_0')) | ~sP0(X129,sK9,X130,X131)) ) | $spl34),
  inference(superposition,[],[f151,f301])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f79])).
fof(f11545,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11544,f4099])).
fof(f11544,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11543,f4090])).
fof(f11543,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl9 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11542,f2417])).
fof(f2417,plain,(
  ( ! [X59,X57,X58] : (~sP4(X57,sK9,X58,X59)) ) | $spl34),
  inference(subsumption_resolution,[],[f2368,f170])).
fof(f2368,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,'F_0')) | ~sP4(X57,sK9,X58,X59)) ) | $spl34),
  inference(superposition,[],[f131,f301])).
fof(f131,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | ~sP4(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f67])).
fof(f216,plain,(
  sP2(sK10,sK9,sK8,sK11) | $spl8),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  sP2(sK10,sK9,sK8,sK11) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f11536,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11535,f205])).
fof(f11535,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11534,f2429])).
fof(f11534,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11533,f4099])).
fof(f11533,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11532,f2417])).
fof(f11531,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11530,f200])).
fof(f11530,plain,(
  ~p(s(bool,'T_0')) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl34 | $spl46)),
  inference(forward_demodulation,[],[f11529,f313])).
fof(f11529,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl11 | $spl34)),
  inference(subsumption_resolution,[],[f11528,f205])).
fof(f11528,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34)),
  inference(subsumption_resolution,[],[f11527,f2429])).
fof(f11527,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34)),
  inference(subsumption_resolution,[],[f11526,f145])).
fof(f11526,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl11 | $spl34)),
  inference(subsumption_resolution,[],[f11525,f2417])).
fof(f11516,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl7 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11515,f205])).
fof(f11515,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11514,f2429])).
fof(f11514,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11513,f4090])).
fof(f11513,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11512,f2417])).
fof(f11511,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl7 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11510,f205])).
fof(f11510,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11509,f2429])).
fof(f11509,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11508,f4090])).
fof(f11508,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11507,f2417])).
fof(f11506,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11505,f2429])).
fof(f11505,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11504,f4099])).
fof(f11504,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11503,f4090])).
fof(f11503,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11502,f2417])).
fof(f11501,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11500,f205])).
fof(f11500,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11499,f4099])).
fof(f11499,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11498,f4090])).
fof(f11498,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11497,f2417])).
fof(f11496,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11495,f205])).
fof(f11495,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11494,f4099])).
fof(f11494,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11493,f4090])).
fof(f11493,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11492,f2417])).
fof(f11491,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11490,f205])).
fof(f11490,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11489,f4099])).
fof(f11489,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11488,f4090])).
fof(f11488,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11487,f2417])).
fof(f11486,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11485,f2429])).
fof(f11485,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11484,f4099])).
fof(f11484,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11483,f4090])).
fof(f11483,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl34 | $spl46)),
  inference(subsumption_resolution,[],[f11482,f2417])).
fof(f11467,plain,(
  $false | ($spl8 | $spl25)),
  inference(subsumption_resolution,[],[f11025,f241])).
fof(f11025,plain,(
  p(s(bool,sK10)) | $spl8),
  inference(resolution,[],[f216,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~sP2(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f11414,plain,(
  $false | ($spl1 | $spl8 | $spl30 | $spl34 | $spl46)),
  inference(resolution,[],[f11362,f216])).
fof(f11362,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2,sK11)) ) | ($spl1 | $spl30 | $spl34 | $spl46)),
  inference(resolution,[],[f11280,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~sP2(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f11280,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl1 | $spl30 | $spl34 | $spl46)),
  inference(backward_demodulation,[],[f313,f10938])).
fof(f10938,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | ($spl1 | $spl30 | $spl34)),
  inference(forward_demodulation,[],[f10937,f297])).
fof(f1994,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl85),
  inference(cnf_transformation,[],[f1994_D])).
fof(f1994_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f11079,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl34)),
  inference(superposition,[],[f10938,f168])).
fof(f11078,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl34)),
  inference(superposition,[],[f10938,f168])).
fof(f11077,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl34)),
  inference(superposition,[],[f10938,f169])).
fof(f11076,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl34)),
  inference(superposition,[],[f10938,f169])).
fof(f218,plain,(
  sP3(sK10,sK9,sK8,sK11) | $spl10),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  sP3(sK10,sK9,sK8,sK11) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f11021,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl30)),
  inference(subsumption_resolution,[],[f11020,f200])).
fof(f11020,plain,(
  ~p(s(bool,'T_0')) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl30)),
  inference(forward_demodulation,[],[f11019,f297])).
fof(f11019,plain,(
  ~p(s(bool,sK10)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7)),
  inference(subsumption_resolution,[],[f11018,f205])).
fof(f11018,plain,(
  ~p(s(bool,sK10)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7)),
  inference(subsumption_resolution,[],[f10583,f211])).
fof(f10583,plain,(
  ~p(s(bool,sK10)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(subsumption_resolution,[],[f10582,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | ~sP6(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f61])).
fof(f10582,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(subsumption_resolution,[],[f10581,f132])).
fof(f132,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | ~sP4(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f67])).
fof(f10581,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(subsumption_resolution,[],[f10566,f127])).
fof(f127,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | ~sP5(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f64])).
fof(f10566,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl7),
  inference(resolution,[],[f213,f110])).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | sP6(X0,X1,X2,X3) | sP5(X0,X1,X2,X3) | sP4(X0,X1,X2,X3) | sP3(X0,X1,X2,X3) | sP2(X0,X1,X2,X3) | sP1(X0,X1,X2,X3) | sP0(X0,X1,X2,X3) | ~sP7(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f11017,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11016,f205])).
fof(f11016,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11015,f211])).
fof(f11015,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11014,f2417])).
fof(f11014,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11013,f2413])).
fof(f2413,plain,(
  ( ! [X39,X41,X40] : (~sP5(X39,sK9,X40,X41)) ) | $spl34),
  inference(subsumption_resolution,[],[f2362,f170])).
fof(f2362,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'F_0')) | ~sP5(X39,sK9,X40,X41)) ) | $spl34),
  inference(superposition,[],[f126,f301])).
fof(f11013,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl20 | $spl30)),
  inference(subsumption_resolution,[],[f11012,f2135])).
fof(f2135,plain,(
  ( ! [X26,X24,X25] : (~sP6(sK10,X24,X25,X26)) ) | $spl30),
  inference(subsumption_resolution,[],[f2080,f200])).
fof(f2080,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | ~sP6(sK10,X24,X25,X26)) ) | $spl30),
  inference(superposition,[],[f122,f297])).
fof(f11011,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl7 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11010,f205])).
fof(f11010,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl7 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11009,f211])).
fof(f11009,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11008,f2417])).
fof(f11008,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11007,f2413])).
fof(f11007,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl30)),
  inference(subsumption_resolution,[],[f10562,f2135])).
fof(f11006,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl24 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11005,f211])).
fof(f11005,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl24 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11004,f2425])).
fof(f2425,plain,(
  ( ! [X111,X112,X113] : (~sP1(X111,sK9,X112,X113)) ) | $spl34),
  inference(subsumption_resolution,[],[f2386,f170])).
fof(f2386,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'F_0')) | ~sP1(X111,sK9,X112,X113)) ) | $spl34),
  inference(superposition,[],[f146,f301])).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f76])).
fof(f11004,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f9969,f2135])).
fof(f9969,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f9968,f240])).
fof(f9968,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f9967,f132])).
fof(f9967,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f9952,f127])).
fof(f9952,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f110])).
fof(f11003,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11002,f211])).
fof(f11002,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11001,f2425])).
fof(f11001,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f11000,f2417])).
fof(f11000,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10999,f2413])).
fof(f10999,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f9949,f2135])).
fof(f10998,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10997,f205])).
fof(f10997,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10996,f2425])).
fof(f10996,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10995,f2417])).
fof(f10995,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10994,f2413])).
fof(f10994,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30)),
  inference(subsumption_resolution,[],[f10993,f2135])).
fof(f10992,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10991,f205])).
fof(f10991,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10990,f2425])).
fof(f10990,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10989,f2417])).
fof(f10989,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10988,f2413])).
fof(f10988,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f1841,f2135])).
fof(f10987,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl24 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10986,f205])).
fof(f10986,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10985,f2425])).
fof(f10985,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f9882,f2135])).
fof(f9882,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f1848,f240])).
fof(f1848,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(subsumption_resolution,[],[f1847,f132])).
fof(f1847,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(subsumption_resolution,[],[f1844,f127])).
fof(f1844,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | $spl5),
  inference(resolution,[],[f211,f110])).
fof(f10984,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10983,f205])).
fof(f10983,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10982,f2425])).
fof(f10982,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10981,f2417])).
fof(f10981,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10980,f2413])).
fof(f10980,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f8460,f2135])).
fof(f10979,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10978,f211])).
fof(f10978,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10977,f2425])).
fof(f10977,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10976,f2417])).
fof(f10976,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f10975,f2413])).
fof(f10975,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl20 | $spl30)),
  inference(subsumption_resolution,[],[f10974,f2135])).
fof(f10973,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f10590,f230])).
fof(f10590,plain,(
  ~p(s(bool,sK8)) | $spl10),
  inference(resolution,[],[f218,f135])).
fof(f10927,plain,(
  $false | ($spl1 | $spl10 | $spl30 | $spl34 | $spl42)),
  inference(resolution,[],[f10875,f218])).
fof(f10875,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2,sK11)) ) | ($spl1 | $spl30 | $spl34 | $spl42)),
  inference(resolution,[],[f10697,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ~sP3(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f70])).
fof(f10697,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl1 | $spl30 | $spl34 | $spl42)),
  inference(backward_demodulation,[],[f301,f10493])).
fof(f10493,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | ($spl1 | $spl30 | $spl42)),
  inference(forward_demodulation,[],[f10492,f297])).
fof(f10631,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl30 | $spl42)),
  inference(superposition,[],[f10493,f168])).
fof(f1942,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl83),
  inference(cnf_transformation,[],[f1942_D])).
fof(f1942_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f10630,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl30 | $spl42)),
  inference(superposition,[],[f10493,f168])).
fof(f10629,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl1 | $spl30 | $spl42)),
  inference(superposition,[],[f10493,f169])).
fof(f10628,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl1 | $spl30 | $spl42)),
  inference(superposition,[],[f10493,f169])).
fof(f10572,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | ($spl2 | $spl7 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10571,f205])).
fof(f10571,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10570,f9940])).
fof(f9940,plain,(
  ( ! [X64,X65,X63] : (~sP0(X63,X64,sK8,X65)) ) | $spl21),
  inference(resolution,[],[f231,f150])).
fof(f10570,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10569,f9934])).
fof(f9934,plain,(
  ( ! [X47,X45,X46] : (~sP2(X45,X46,sK8,X47)) ) | $spl21),
  inference(resolution,[],[f231,f140])).
fof(f10569,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10568,f9928])).
fof(f9928,plain,(
  ( ! [X28,X29,X27] : (~sP4(X27,X28,sK8,X29)) ) | $spl21),
  inference(resolution,[],[f231,f130])).
fof(f10568,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10567,f2138])).
fof(f2138,plain,(
  ( ! [X43,X44,X42] : (~sP5(sK10,X42,X43,X44)) ) | $spl30),
  inference(subsumption_resolution,[],[f2086,f200])).
fof(f2086,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | ~sP5(sK10,X42,X43,X44)) ) | $spl30),
  inference(superposition,[],[f127,f297])).
fof(f10567,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl7 | $spl21)),
  inference(subsumption_resolution,[],[f10562,f9922])).
fof(f9922,plain,(
  ( ! [X10,X11,X9] : (~sP6(X9,X10,sK8,X11)) ) | $spl21),
  inference(resolution,[],[f231,f120])).
fof(f214,plain,(
  sP1(sK10,sK9,sK8,sK11) | $spl6),
  inference(cnf_transformation,[],[f214_D])).
fof(f214_D,plain,(
  sP1(sK10,sK9,sK8,sK11) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f10561,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f10560,f9940])).
fof(f10560,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f10559,f9934])).
fof(f10559,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f9969,f9922])).
fof(f10558,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10557,f9940])).
fof(f10557,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10556,f9934])).
fof(f10556,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10555,f9928])).
fof(f10555,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10554,f2138])).
fof(f10554,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl21)),
  inference(subsumption_resolution,[],[f9949,f9922])).
fof(f10545,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10544,f205])).
fof(f10544,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10543,f9934])).
fof(f10543,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10542,f9928])).
fof(f10542,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10541,f2138])).
fof(f10541,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21)),
  inference(subsumption_resolution,[],[f1841,f9922])).
fof(f10540,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f10539,f205])).
fof(f10539,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f10538,f9934])).
fof(f10538,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f9882,f9922])).
fof(f10537,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10536,f205])).
fof(f10536,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10535,f9934])).
fof(f10535,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10534,f9928])).
fof(f10534,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21 | $spl30)),
  inference(subsumption_resolution,[],[f10533,f2138])).
fof(f10533,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl21)),
  inference(subsumption_resolution,[],[f8460,f9922])).
fof(f10532,plain,(
  $false | ($spl6 | $spl23)),
  inference(subsumption_resolution,[],[f9976,f236])).
fof(f9976,plain,(
  p(s(bool,sK9)) | $spl6),
  inference(resolution,[],[f214,f146])).
fof(f10483,plain,(
  $false | ($spl1 | $spl6 | $spl30 | $spl38 | $spl42)),
  inference(resolution,[],[f10431,f214])).
fof(f10431,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,sK11)) ) | ($spl1 | $spl30 | $spl38 | $spl42)),
  inference(resolution,[],[f10354,f144])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f76])).
fof(f10354,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl1 | $spl30 | $spl38 | $spl42)),
  inference(backward_demodulation,[],[f309,f10055])).
fof(f10055,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl1 | $spl30 | $spl38)),
  inference(backward_demodulation,[],[f297,f9910])).
fof(f10334,plain,(
  $false | ($spl43 | $spl47)),
  inference(subsumption_resolution,[],[f10330,f308])).
fof(f308,plain,(
  s(bool,'F_0') != s(bool,sK8) | $spl43),
  inference(cnf_transformation,[],[f308_D])).
fof(f308_D,plain,(
  s(bool,'F_0') != s(bool,sK8) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f10330,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl47),
  inference(trivial_inequality_removal,[],[f10328])).
fof(f10328,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,sK8) | $spl47),
  inference(superposition,[],[f312,f168])).
fof(f312,plain,(
  s(bool,'T_0') != s(bool,sK8) | $spl47),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s(bool,'T_0') != s(bool,sK8) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f10332,plain,(
  $false | ($spl43 | $spl47)),
  inference(subsumption_resolution,[],[f10331,f308])).
fof(f10331,plain,(
  s(bool,'F_0') = s(bool,sK8) | $spl47),
  inference(trivial_inequality_removal,[],[f10326])).
fof(f10326,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,sK8) | $spl47),
  inference(superposition,[],[f312,f169])).
fof(f10183,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl38)),
  inference(superposition,[],[f10055,f168])).
fof(f1846,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl81),
  inference(cnf_transformation,[],[f1846_D])).
fof(f1846_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f10182,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl38)),
  inference(superposition,[],[f10055,f168])).
fof(f10181,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl38)),
  inference(superposition,[],[f10055,f169])).
fof(f10180,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl30 | $spl38)),
  inference(superposition,[],[f10055,f169])).
fof(f756,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl69),
  inference(cnf_transformation,[],[f756_D])).
fof(f756_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f10021,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f168])).
fof(f754,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl67),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f10020,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f168])).
fof(f10019,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f169])).
fof(f10018,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f169])).
fof(f10009,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f168])).
fof(f487,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | $spl51),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f10008,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f168])).
fof(f10007,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f169])).
fof(f10006,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl1 | $spl38)),
  inference(superposition,[],[f9910,f169])).
fof(f9958,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP1(sK10,sK9,sK8,sK11) | ($spl2 | $spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9957,f211])).
fof(f9957,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9956,f2759])).
fof(f2759,plain,(
  ( ! [X94,X95,X93] : (~sP2(X93,sK9,X94,X95)) ) | $spl38),
  inference(subsumption_resolution,[],[f2696,f200])).
fof(f2696,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | ~sP2(X93,sK9,X94,X95)) ) | $spl38),
  inference(superposition,[],[f141,f305])).
fof(f9956,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9955,f2755])).
fof(f2755,plain,(
  ( ! [X76,X77,X75] : (~sP3(X75,sK9,X76,X77)) ) | $spl38),
  inference(subsumption_resolution,[],[f2690,f200])).
fof(f2690,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'T_0')) | ~sP3(X75,sK9,X76,X77)) ) | $spl38),
  inference(superposition,[],[f136,f305])).
fof(f9955,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9954,f2140])).
fof(f2140,plain,(
  ( ! [X61,X62,X60] : (~sP4(sK10,X60,X61,X62)) ) | $spl30),
  inference(subsumption_resolution,[],[f2092,f200])).
fof(f2092,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~sP4(sK10,X60,X61,X62)) ) | $spl30),
  inference(superposition,[],[f132,f297])).
fof(f9954,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9953,f2138])).
fof(f9953,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | ($spl2 | $spl38)),
  inference(subsumption_resolution,[],[f9949,f2736])).
fof(f206,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | $spl3),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  ~sP7(sK10,sK9,sK8,sK11) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f231,plain,(
  ~p(s(bool,sK8)) | $spl21),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  ~p(s(bool,sK8)) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f9918,plain,(
  ~p(s(bool,sK8)) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9917,f145])).
fof(f9917,plain,(
  ~p(s(bool,sK8)) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9916,f2759])).
fof(f9916,plain,(
  ~p(s(bool,sK8)) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9915,f135])).
fof(f9915,plain,(
  ~p(s(bool,sK8)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9914,f2140])).
fof(f9914,plain,(
  ~p(s(bool,sK8)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f9913,f125])).
fof(f9913,plain,(
  ~p(s(bool,sK8)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f1842,f2736])).
fof(f9912,plain,(
  $false | ($spl8 | $spl38)),
  inference(subsumption_resolution,[],[f216,f2759])).
fof(f9911,plain,(
  $false | ($spl12 | $spl30)),
  inference(subsumption_resolution,[],[f220,f2140])).
fof(f9909,plain,(
  $false | ($spl16 | $spl38)),
  inference(subsumption_resolution,[],[f224,f2736])).
fof(f9903,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9902,f2759])).
fof(f9902,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9901,f2755])).
fof(f9901,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9900,f2140])).
fof(f9900,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9899,f2138])).
fof(f9899,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f1841,f2736])).
fof(f9898,plain,(
  sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9897,f2759])).
fof(f9897,plain,(
  sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9896,f2755])).
fof(f9896,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9895,f2140])).
fof(f9895,plain,(
  sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9894,f2138])).
fof(f9894,plain,(
  sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20 | $spl38)),
  inference(subsumption_resolution,[],[f9893,f2736])).
fof(f9893,plain,(
  sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl20)),
  inference(subsumption_resolution,[],[f1842,f230])).
fof(f9892,plain,(
  sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9891,f200])).
fof(f9891,plain,(
  ~p(s(bool,'T_0')) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(forward_demodulation,[],[f9890,f305])).
fof(f9890,plain,(
  ~p(s(bool,sK9)) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f9889,f141])).
fof(f9889,plain,(
  ~p(s(bool,sK9)) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f9888,f136])).
fof(f9888,plain,(
  ~p(s(bool,sK9)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f9887,f2140])).
fof(f9887,plain,(
  ~p(s(bool,sK9)) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30)),
  inference(subsumption_resolution,[],[f9886,f2138])).
fof(f9885,plain,(
  sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f9884,f2759])).
fof(f9884,plain,(
  sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f9883,f2755])).
fof(f9883,plain,(
  sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f9882,f2736])).
fof(f9881,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9880,f2759])).
fof(f9880,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9879,f2755])).
fof(f9879,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9878,f2140])).
fof(f9878,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl30 | $spl38)),
  inference(subsumption_resolution,[],[f9877,f2138])).
fof(f9877,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | ~sP7(sK10,sK9,sK8,sK11) | ($spl5 | $spl38)),
  inference(subsumption_resolution,[],[f8460,f2736])).
fof(f1941,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl82),
  inference(cnf_transformation,[],[f1941_D])).
fof(f1941_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9870,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl38)),
  inference(superposition,[],[f9594,f168])).
fof(f9594,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl0 | $spl30 | $spl38)),
  inference(forward_demodulation,[],[f2069,f305])).
fof(f2069,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f297,f203])).
fof(f1845,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl80),
  inference(cnf_transformation,[],[f1845_D])).
fof(f1845_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9869,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl38)),
  inference(superposition,[],[f9594,f168])).
fof(f9868,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl38)),
  inference(superposition,[],[f9594,f169])).
fof(f9867,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl38)),
  inference(superposition,[],[f9594,f169])).
fof(f755,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl68),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9743,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f168])).
fof(f9593,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl0 | $spl38)),
  inference(forward_demodulation,[],[f203,f305])).
fof(f753,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl66),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9742,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f168])).
fof(f9741,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f169])).
fof(f9740,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f169])).
fof(f619,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | $spl58),
  inference(cnf_transformation,[],[f619_D])).
fof(f619_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9731,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f168])).
fof(f486,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | $spl50),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9730,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f168])).
fof(f9729,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f169])).
fof(f9728,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl38)),
  inference(superposition,[],[f9593,f169])).
fof(f230,plain,(
  p(s(bool,sK8)) | $spl20),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  p(s(bool,sK8)) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f241,plain,(
  ~p(s(bool,sK10)) | $spl25),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ~p(s(bool,sK10)) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f9592,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl3 | $spl38)),
  inference(subsumption_resolution,[],[f9591,f200])).
fof(f9591,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl3 | $spl38)),
  inference(forward_demodulation,[],[f7670,f305])).
fof(f7670,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f148,f1746])).
fof(f1746,plain,(
  ~sP1(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f112])).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f148,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,X2,X3) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f9565,plain,(
  $false | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9507,f190])).
fof(f190,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb',aHL_TRUTH)).
fof(f9507,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl46)),
  inference(backward_demodulation,[],[f313,f231])).
fof(f2768,plain,(
  ( ! [X122] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl100),
  inference(cnf_transformation,[],[f2768_D])).
fof(f2768_D,plain,(
  ( ! [X122] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2751,plain,(
  ( ! [X67] : (~p(s(bool,X67))) ) | $spl94),
  inference(cnf_transformation,[],[f2751_D])).
fof(f2751_D,plain,(
  ( ! [X67] : (~p(s(bool,X67))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f9489,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9483,f231])).
fof(f9483,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(resolution,[],[f4099,f148])).
fof(f2762,plain,(
  ( ! [X104] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl98),
  inference(cnf_transformation,[],[f2762_D])).
fof(f2762_D,plain,(
  ( ! [X104] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f9473,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9467,f231])).
fof(f9467,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,sK8)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(resolution,[],[f4094,f143])).
fof(f143,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,X2,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f4094,plain,(
  ( ! [X94,X95,X93] : (~sP2(X93,sK8,X94,X95)) ) | $spl46),
  inference(subsumption_resolution,[],[f4035,f200])).
fof(f4035,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | ~sP2(X93,sK8,X94,X95)) ) | $spl46),
  inference(superposition,[],[f141,f313])).
fof(f2758,plain,(
  ( ! [X89] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl96),
  inference(cnf_transformation,[],[f2758_D])).
fof(f2758_D,plain,(
  ( ! [X89] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f2741,plain,(
  ( ! [X34] : (p(s(bool,X34))) ) | $spl88),
  inference(cnf_transformation,[],[f2741_D])).
fof(f2741_D,plain,(
  ( ! [X34] : (p(s(bool,X34))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9459,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9454,f231])).
fof(f9454,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,sK8)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(resolution,[],[f4091,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,X2,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f4091,plain,(
  ( ! [X76,X77,X75] : (~sP3(X75,sK8,X76,X77)) ) | $spl46),
  inference(subsumption_resolution,[],[f4029,f200])).
fof(f4029,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'T_0')) | ~sP3(X75,sK8,X76,X77)) ) | $spl46),
  inference(superposition,[],[f136,f313])).
fof(f9446,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9441,f231])).
fof(f9441,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(resolution,[],[f4090,f138])).
fof(f2750,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl92),
  inference(cnf_transformation,[],[f2750_D])).
fof(f2750_D,plain,(
  ( ! [X68] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f9430,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9424,f231])).
fof(f9424,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK8)) | ~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(resolution,[],[f4085,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1,X2,X3) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f4085,plain,(
  ( ! [X61,X62,X60] : (~sP4(sK8,X60,X61,X62)) ) | $spl46),
  inference(subsumption_resolution,[],[f4024,f200])).
fof(f4024,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~sP4(sK8,X60,X61,X62)) ) | $spl46),
  inference(superposition,[],[f132,f313])).
fof(f2746,plain,(
  ( ! [X53] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl90),
  inference(cnf_transformation,[],[f2746_D])).
fof(f2746_D,plain,(
  ( ! [X53] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f9364,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9359,f231])).
fof(f9359,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK8)) | ~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(resolution,[],[f4082,f128])).
fof(f128,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,X1,X2,X3) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f4082,plain,(
  ( ! [X43,X44,X42] : (~sP5(sK8,X42,X43,X44)) ) | $spl46),
  inference(subsumption_resolution,[],[f4018,f200])).
fof(f4018,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | ~sP5(sK8,X42,X43,X44)) ) | $spl46),
  inference(superposition,[],[f127,f313])).
fof(f9352,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9348,f231])).
fof(f9348,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(resolution,[],[f4081,f128])).
fof(f2740,plain,(
  ( ! [X35] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl86),
  inference(cnf_transformation,[],[f2740_D])).
fof(f2740_D,plain,(
  ( ! [X35] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f9339,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9335,f231])).
fof(f9335,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK8)) | p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(resolution,[],[f4078,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f4078,plain,(
  ( ! [X26,X24,X25] : (~sP6(sK8,X24,X25,X26)) ) | $spl46),
  inference(subsumption_resolution,[],[f4012,f200])).
fof(f4012,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | ~sP6(sK8,X24,X25,X26)) ) | $spl46),
  inference(superposition,[],[f122,f313])).
fof(f9328,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f9324,f231])).
fof(f9324,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,sK8)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(resolution,[],[f4077,f123])).
fof(f4077,plain,(
  ( ! [X23,X21,X22] : (~sP6(X21,sK8,X22,X23)) ) | $spl46),
  inference(subsumption_resolution,[],[f4011,f200])).
fof(f4011,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | ~sP6(X21,sK8,X22,X23)) ) | $spl46),
  inference(superposition,[],[f121,f313])).
fof(f8056,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl186),
  inference(cnf_transformation,[],[f8056_D])).
fof(f8056_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f9282,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl38)),
  inference(superposition,[],[f9016,f168])).
fof(f9016,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | ($spl0 | $spl26 | $spl38)),
  inference(backward_demodulation,[],[f305,f3548])).
fof(f3548,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | ($spl0 | $spl26)),
  inference(backward_demodulation,[],[f293,f203])).
fof(f8058,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl188),
  inference(cnf_transformation,[],[f8058_D])).
fof(f8058_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f9281,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl38)),
  inference(superposition,[],[f9016,f168])).
fof(f9280,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl38)),
  inference(superposition,[],[f9016,f169])).
fof(f9279,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl38)),
  inference(superposition,[],[f9016,f169])).
fof(f9074,plain,(
  $false | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f9017,f190])).
fof(f9017,plain,(
  ~p(s(bool,'T_0')) | ($spl23 | $spl38)),
  inference(backward_demodulation,[],[f305,f236])).
fof(f9001,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f8991,f236])).
fof(f8991,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK9)) | ~p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f2747,f133])).
fof(f2747,plain,(
  ( ! [X61,X62,X60] : (~sP4(sK9,X60,X61,X62)) ) | $spl38),
  inference(subsumption_resolution,[],[f2685,f200])).
fof(f2685,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~sP4(sK9,X60,X61,X62)) ) | $spl38),
  inference(superposition,[],[f132,f305])).
fof(f8978,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f8969,f236])).
fof(f8969,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK9)) | ~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f2743,f128])).
fof(f2743,plain,(
  ( ! [X43,X44,X42] : (~sP5(sK9,X42,X43,X44)) ) | $spl38),
  inference(subsumption_resolution,[],[f2679,f200])).
fof(f2679,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | ~sP5(sK9,X42,X43,X44)) ) | $spl38),
  inference(superposition,[],[f127,f305])).
fof(f8958,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f8950,f236])).
fof(f8950,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f2742,f128])).
fof(f2742,plain,(
  ( ! [X37,X38,X36] : (~sP5(X36,X37,sK9,X38)) ) | $spl38),
  inference(subsumption_resolution,[],[f2677,f200])).
fof(f2677,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~sP5(X36,X37,sK9,X38)) ) | $spl38),
  inference(superposition,[],[f125,f305])).
fof(f8937,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f8929,f236])).
fof(f8929,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,sK9)) | p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f2737,f123])).
fof(f2737,plain,(
  ( ! [X26,X24,X25] : (~sP6(sK9,X24,X25,X26)) ) | $spl38),
  inference(subsumption_resolution,[],[f2673,f200])).
fof(f2673,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | ~sP6(sK9,X24,X25,X26)) ) | $spl38),
  inference(superposition,[],[f122,f305])).
fof(f8918,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f8910,f236])).
fof(f8910,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,sK9)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X5),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f2736,f123])).
fof(f621,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | $spl60),
  inference(cnf_transformation,[],[f621_D])).
fof(f621_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8908,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(forward_demodulation,[],[f8170,f293])).
fof(f8170,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f8907,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(forward_demodulation,[],[f8172,f293])).
fof(f8172,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f8290,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f8292,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f8906,plain,(
  $false | ($spl35 | $spl39)),
  inference(subsumption_resolution,[],[f8901,f304])).
fof(f304,plain,(
  s(bool,'T_0') != s(bool,sK9) | $spl39),
  inference(cnf_transformation,[],[f304_D])).
fof(f304_D,plain,(
  s(bool,'T_0') != s(bool,sK9) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f8901,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl35),
  inference(trivial_inequality_removal,[],[f8900])).
fof(f8900,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK9) | $spl35),
  inference(superposition,[],[f300,f168])).
fof(f300,plain,(
  s(bool,'F_0') != s(bool,sK9) | $spl35),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  s(bool,'F_0') != s(bool,sK9) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f8904,plain,(
  $false | ($spl35 | $spl39)),
  inference(subsumption_resolution,[],[f8902,f304])).
fof(f8902,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl35),
  inference(trivial_inequality_removal,[],[f8898])).
fof(f8898,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK9) | $spl35),
  inference(superposition,[],[f300,f169])).
fof(f625,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | $spl64),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8298,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f623,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | $spl62),
  inference(cnf_transformation,[],[f623_D])).
fof(f623_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8297,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f8296,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f8295,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f8291,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f8289,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f314,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | $spl48),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f8178,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f310,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | $spl44),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8177,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f8176,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f8175,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f302,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | $spl36),
  inference(cnf_transformation,[],[f302_D])).
fof(f302_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8171,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f8169,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f298,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | $spl32),
  inference(cnf_transformation,[],[f298_D])).
fof(f298_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8166,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f294,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | $spl28),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8165,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f8164,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f8163,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f240,plain,(
  p(s(bool,sK10)) | $spl24),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  p(s(bool,sK10)) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6145,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f133,f1749])).
fof(f1749,plain,(
  ~sP4(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP4(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f5952,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f128,f1750])).
fof(f1750,plain,(
  ~sP5(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f116])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP5(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f235,plain,(
  p(s(bool,sK9)) | $spl22),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  p(s(bool,sK9)) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5719,plain,(
  p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f123,f1751])).
fof(f1751,plain,(
  ~sP6(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f117])).
fof(f117,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP6(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f5563,plain,(
  p(s(bool,sK10)) | p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f118,f206])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f1986,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f1748,f138])).
fof(f1748,plain,(
  ~sP3(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP3(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f2031,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f1749,f133])).
fof(f2037,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f1750,f128])).
fof(f7593,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl20)),
  inference(subsumption_resolution,[],[f7398,f230])).
fof(f7398,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f143,f1747])).
fof(f1747,plain,(
  ~sP2(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f113])).
fof(f113,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP2(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f1935,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f1746,f148])).
fof(f1979,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f1747,f143])).
fof(f2043,plain,(
  p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl3),
  inference(resolution,[],[f1751,f123])).
fof(f3326,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3324,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f6387,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f138,f1748])).
fof(f8012,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f8011,f170])).
fof(f8011,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f6387,f301])).
fof(f8010,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl34)),
  inference(forward_demodulation,[],[f3324,f301])).
fof(f8009,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl34)),
  inference(forward_demodulation,[],[f3326,f301])).
fof(f8008,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f8007,f170])).
fof(f8007,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f2043,f301])).
fof(f8006,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f8005,f170])).
fof(f8005,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f1979,f301])).
fof(f8003,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl20 | $spl34)),
  inference(subsumption_resolution,[],[f8002,f170])).
fof(f8002,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK10)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl20 | $spl34)),
  inference(forward_demodulation,[],[f7593,f301])).
fof(f7999,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f7998,f170])).
fof(f7998,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f1986,f301])).
fof(f7997,plain,(
  p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f7996,f170])).
fof(f7996,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f5563,f301])).
fof(f7995,plain,(
  p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f7994,f170])).
fof(f7994,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK10)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f5719,f301])).
fof(f7967,plain,(
  $false | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7910,f170])).
fof(f7910,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl42)),
  inference(backward_demodulation,[],[f309,f230])).
fof(f7888,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl184),
  inference(cnf_transformation,[],[f7888_D])).
fof(f7888_D,plain,(
  ( ! [X26,X27,X25] : (sP1(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f5696,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X24)) ) | $spl118),
  inference(cnf_transformation,[],[f5696_D])).
fof(f5696_D,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f7843,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP1(X29,X30,'F_0',X31) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f148,f168])).
fof(f7842,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP1(X25,X26,'F_0',X27) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f148,f169])).
fof(f7887,plain,(
  ( ! [X19,X17,X18] : (sP1(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl182),
  inference(cnf_transformation,[],[f7887_D])).
fof(f7887_D,plain,(
  ( ! [X19,X17,X18] : (sP1(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f5694,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | $spl114),
  inference(cnf_transformation,[],[f5694_D])).
fof(f5694_D,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f7841,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP1(X21,X22,'T_0',X23) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f148,f168])).
fof(f7840,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP1(X17,X18,'T_0',X19) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f148,f169])).
fof(f5890,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'T_0') = s(bool,X24)) ) | $spl140),
  inference(cnf_transformation,[],[f5890_D])).
fof(f5890_D,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7883,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7833,f4417])).
fof(f4417,plain,(
  ( ! [X23,X21,X22] : (~sP1(X21,'F_0',X22,X23)) )),
  inference(resolution,[],[f170,f146])).
fof(f7833,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP1(X29,'F_0',X30,X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f148,f168])).
fof(f7882,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f7832,f4417])).
fof(f7832,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP1(X25,'F_0',X26,X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f148,f169])).
fof(f7873,plain,(
  ( ! [X169,X170,X168] : (sP1(X168,'T_0',X169,X170) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X169)) | ~p(s(bool,X168))) ) | $spl180),
  inference(cnf_transformation,[],[f7873_D])).
fof(f7873_D,plain,(
  ( ! [X169,X170,X168] : (sP1(X168,'T_0',X169,X170) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X169)) | ~p(s(bool,X168))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f5888,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | $spl138),
  inference(cnf_transformation,[],[f5888_D])).
fof(f5888_D,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7831,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP1(X21,'T_0',X22,X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f148,f168])).
fof(f7830,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP1(X17,'T_0',X18,X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f148,f169])).
fof(f7878,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7823,f4418])).
fof(f4418,plain,(
  ( ! [X26,X24,X25] : (~sP1('F_0',X24,X25,X26)) )),
  inference(resolution,[],[f170,f147])).
fof(f7823,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP1('F_0',X29,X30,X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f148,f168])).
fof(f7877,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f7822,f4418])).
fof(f7822,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP1('F_0',X25,X26,X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f148,f169])).
fof(f7871,plain,(
  ( ! [X169,X170,X168] : (sP1('T_0',X168,X169,X170) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X169)) | ~p(s(bool,X168))) ) | $spl178),
  inference(cnf_transformation,[],[f7871_D])).
fof(f7871_D,plain,(
  ( ! [X169,X170,X168] : (sP1('T_0',X168,X169,X170) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X169)) | ~p(s(bool,X168))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f7821,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP1('T_0',X21,X22,X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f148,f168])).
fof(f7820,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP1('T_0',X17,X18,X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f148,f169])).
fof(f7874,plain,(
  ( ! [X246,X247,X245] : (~p(s(bool,X245)) | p(s(bool,X246)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X247),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7750,f4114])).
fof(f4114,plain,(
  ( ! [X111,X112,X113] : (~sP1(X111,sK8,X112,X113)) ) | $spl42),
  inference(subsumption_resolution,[],[f3968,f170])).
fof(f3968,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'F_0')) | ~sP1(X111,sK8,X112,X113)) ) | $spl42),
  inference(superposition,[],[f146,f309])).
fof(f7750,plain,(
  ( ! [X246,X247,X245] : (sP1(X245,sK8,X246,X247) | ~p(s(bool,X245)) | p(s(bool,X246)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X247),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl20),
  inference(resolution,[],[f148,f230])).
fof(f7738,plain,(
  ( ! [X182,X180,X183,X181] : (sP1(X180,'T_0',X181,X182) | ~p(s(bool,X180)) | p(s(bool,X181)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X182),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X183))) )),
  inference(resolution,[],[f148,f188])).
fof(f188,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : ((p(s(bool,X10)) & p(s(bool,'T_0'))) <=> p(s(bool,X10)))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.AND_CLAUSES|split|1')).
fof(f7737,plain,(
  ( ! [X177,X178,X176,X179] : (sP1(X176,'T_0',X177,X178) | ~p(s(bool,X176)) | p(s(bool,X177)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X178),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X179))) )),
  inference(resolution,[],[f148,f184])).
fof(f184,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) & p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.AND_CLAUSES|split|0')).
fof(f7736,plain,(
  ( ! [X175,X173,X174,X172] : (sP1(X172,'T_0',X173,X174) | ~p(s(bool,X172)) | p(s(bool,X173)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X174),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X175))) )),
  inference(resolution,[],[f148,f181])).
fof(f181,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) | p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.OR_CLAUSES|split|0')).
fof(f7735,plain,(
  ( ! [X171,X169,X170,X168] : (sP1(X168,'T_0',X169,X170) | ~p(s(bool,X168)) | p(s(bool,X169)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X171))) )),
  inference(resolution,[],[f148,f177])).
fof(f177,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : ((p(s(bool,X10)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.OR_CLAUSES|split|1')).
fof(f7872,plain,(
  ( ! [X246,X247,X245] : (~p(s(bool,X245)) | p(s(bool,X246)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X247),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7710,f4113])).
fof(f4113,plain,(
  ( ! [X116,X114,X115] : (~sP1(sK8,X114,X115,X116)) ) | $spl42),
  inference(subsumption_resolution,[],[f3969,f170])).
fof(f3969,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'F_0')) | ~sP1(sK8,X114,X115,X116)) ) | $spl42),
  inference(superposition,[],[f147,f309])).
fof(f7710,plain,(
  ( ! [X246,X247,X245] : (sP1(sK8,X245,X246,X247) | ~p(s(bool,X245)) | p(s(bool,X246)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X247),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl20),
  inference(resolution,[],[f148,f230])).
fof(f7698,plain,(
  ( ! [X182,X180,X183,X181] : (sP1('T_0',X180,X181,X182) | ~p(s(bool,X180)) | p(s(bool,X181)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X182),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X183))) )),
  inference(resolution,[],[f148,f188])).
fof(f7697,plain,(
  ( ! [X177,X178,X176,X179] : (sP1('T_0',X176,X177,X178) | ~p(s(bool,X176)) | p(s(bool,X177)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X178),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X179))) )),
  inference(resolution,[],[f148,f184])).
fof(f7696,plain,(
  ( ! [X175,X173,X174,X172] : (sP1('T_0',X172,X173,X174) | ~p(s(bool,X172)) | p(s(bool,X173)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X174),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X175))) )),
  inference(resolution,[],[f148,f181])).
fof(f7695,plain,(
  ( ! [X171,X169,X170,X168] : (sP1('T_0',X168,X169,X170) | ~p(s(bool,X168)) | p(s(bool,X169)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X170),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X171))) )),
  inference(resolution,[],[f148,f177])).
fof(f7869,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7668,f230])).
fof(f7668,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,sK8)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl42),
  inference(resolution,[],[f148,f4113])).
fof(f7867,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7665,f230])).
fof(f7665,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | ~p(s(bool,sK8)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl42),
  inference(resolution,[],[f148,f4114])).
fof(f7615,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7565,f4415])).
fof(f4415,plain,(
  ( ! [X17,X15,X16] : (~sP2(X15,X16,'F_0',X17)) )),
  inference(resolution,[],[f170,f140])).
fof(f7565,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP2(X29,X30,'F_0',X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f143,f168])).
fof(f7614,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f7564,f4415])).
fof(f7564,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP2(X25,X26,'F_0',X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f143,f169])).
fof(f7600,plain,(
  ( ! [X163,X161,X162] : (sP2(X161,X162,'T_0',X163) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | p(s(bool,X162)) | ~p(s(bool,X161))) ) | $spl172),
  inference(cnf_transformation,[],[f7600_D])).
fof(f7600_D,plain,(
  ( ! [X163,X161,X162] : (sP2(X161,X162,'T_0',X163) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | p(s(bool,X162)) | ~p(s(bool,X161))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7563,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP2(X21,X22,'T_0',X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f143,f168])).
fof(f7562,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP2(X17,X18,'T_0',X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f143,f169])).
fof(f7610,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl176),
  inference(cnf_transformation,[],[f7610_D])).
fof(f7610_D,plain,(
  ( ! [X26,X27,X25] : (sP2(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7555,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP2(X29,'F_0',X30,X31) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f143,f168])).
fof(f7554,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP2(X25,'F_0',X26,X27) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f143,f169])).
fof(f7609,plain,(
  ( ! [X19,X17,X18] : (sP2(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl174),
  inference(cnf_transformation,[],[f7609_D])).
fof(f7609_D,plain,(
  ( ! [X19,X17,X18] : (sP2(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7553,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP2(X21,'T_0',X22,X23) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f143,f168])).
fof(f7552,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP2(X17,'T_0',X18,X19) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f143,f169])).
fof(f7605,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7545,f4416])).
fof(f4416,plain,(
  ( ! [X19,X20,X18] : (~sP2('F_0',X18,X19,X20)) )),
  inference(resolution,[],[f170,f142])).
fof(f7545,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP2('F_0',X29,X30,X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f143,f168])).
fof(f7604,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f7544,f4416])).
fof(f7544,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP2('F_0',X25,X26,X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f143,f169])).
fof(f7598,plain,(
  ( ! [X163,X161,X162] : (sP2('T_0',X161,X162,X163) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X162)) | p(s(bool,X161))) ) | $spl170),
  inference(cnf_transformation,[],[f7598_D])).
fof(f7598_D,plain,(
  ( ! [X163,X161,X162] : (sP2('T_0',X161,X162,X163) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X162)) | p(s(bool,X161))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7543,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP2('T_0',X21,X22,X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f143,f168])).
fof(f7542,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP2('T_0',X17,X18,X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f143,f169])).
fof(f7601,plain,(
  ( ! [X239,X240,X238] : (~p(s(bool,X238)) | p(s(bool,X239)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X240),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7510,f4116])).
fof(f7510,plain,(
  ( ! [X239,X240,X238] : (sP2(X238,X239,sK8,X240) | ~p(s(bool,X238)) | p(s(bool,X239)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X240),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl20),
  inference(resolution,[],[f143,f230])).
fof(f7498,plain,(
  ( ! [X175,X173,X176,X174] : (sP2(X173,X174,'T_0',X175) | ~p(s(bool,X173)) | p(s(bool,X174)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X175),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X176))) )),
  inference(resolution,[],[f143,f188])).
fof(f7497,plain,(
  ( ! [X171,X169,X172,X170] : (sP2(X169,X170,'T_0',X171) | ~p(s(bool,X169)) | p(s(bool,X170)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X171),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X172))) )),
  inference(resolution,[],[f143,f184])).
fof(f7496,plain,(
  ( ! [X167,X165,X166,X168] : (sP2(X165,X166,'T_0',X167) | ~p(s(bool,X165)) | p(s(bool,X166)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X167),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X168))) )),
  inference(resolution,[],[f143,f181])).
fof(f7495,plain,(
  ( ! [X163,X161,X164,X162] : (sP2(X161,X162,'T_0',X163) | ~p(s(bool,X161)) | p(s(bool,X162)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X164))) )),
  inference(resolution,[],[f143,f177])).
fof(f7599,plain,(
  ( ! [X239,X240,X238] : (p(s(bool,X238)) | ~p(s(bool,X239)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X240),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7437,f4115])).
fof(f4115,plain,(
  ( ! [X97,X98,X96] : (~sP2(sK8,X96,X97,X98)) ) | $spl42),
  inference(subsumption_resolution,[],[f3963,f170])).
fof(f3963,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,'F_0')) | ~sP2(sK8,X96,X97,X98)) ) | $spl42),
  inference(superposition,[],[f142,f309])).
fof(f7437,plain,(
  ( ! [X239,X240,X238] : (sP2(sK8,X238,X239,X240) | p(s(bool,X238)) | ~p(s(bool,X239)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X240),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl20),
  inference(resolution,[],[f143,f230])).
fof(f7425,plain,(
  ( ! [X175,X173,X176,X174] : (sP2('T_0',X173,X174,X175) | p(s(bool,X173)) | ~p(s(bool,X174)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X175),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X176))) )),
  inference(resolution,[],[f143,f188])).
fof(f7424,plain,(
  ( ! [X171,X169,X172,X170] : (sP2('T_0',X169,X170,X171) | p(s(bool,X169)) | ~p(s(bool,X170)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X171),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X172))) )),
  inference(resolution,[],[f143,f184])).
fof(f7423,plain,(
  ( ! [X167,X165,X166,X168] : (sP2('T_0',X165,X166,X167) | p(s(bool,X165)) | ~p(s(bool,X166)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X167),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X168))) )),
  inference(resolution,[],[f143,f181])).
fof(f7422,plain,(
  ( ! [X163,X161,X164,X162] : (sP2('T_0',X161,X162,X163) | p(s(bool,X161)) | ~p(s(bool,X162)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X163),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X164))) )),
  inference(resolution,[],[f143,f177])).
fof(f7591,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7396,f230])).
fof(f7396,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,sK8)) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f143,f4115])).
fof(f7588,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7391,f230])).
fof(f7391,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f143,f4116])).
fof(f7210,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7206,f230])).
fof(f7206,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4121,f123])).
fof(f7193,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7189,f230])).
fof(f7189,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,sK8)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl42),
  inference(resolution,[],[f4120,f128])).
fof(f4120,plain,(
  ( ! [X39,X41,X40] : (~sP5(X39,sK8,X40,X41)) ) | $spl42),
  inference(subsumption_resolution,[],[f3944,f170])).
fof(f3944,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'F_0')) | ~sP5(X39,sK8,X40,X41)) ) | $spl42),
  inference(superposition,[],[f126,f309])).
fof(f7177,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7173,f230])).
fof(f7173,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4119,f133])).
fof(f7161,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7157,f230])).
fof(f7157,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,sK8)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4118,f133])).
fof(f4118,plain,(
  ( ! [X59,X57,X58] : (~sP4(X57,sK8,X58,X59)) ) | $spl42),
  inference(subsumption_resolution,[],[f3950,f170])).
fof(f3950,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,'F_0')) | ~sP4(X57,sK8,X58,X59)) ) | $spl42),
  inference(superposition,[],[f131,f309])).
fof(f7142,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7138,f230])).
fof(f7138,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK8)) | p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl42),
  inference(resolution,[],[f4117,f138])).
fof(f4117,plain,(
  ( ! [X80,X78,X79] : (~sP3(sK8,X78,X79,X80)) ) | $spl42),
  inference(subsumption_resolution,[],[f3957,f170])).
fof(f3957,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'F_0')) | ~sP3(sK8,X78,X79,X80)) ) | $spl42),
  inference(superposition,[],[f137,f309])).
fof(f2775,plain,(
  ( ! [X143] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl102),
  inference(cnf_transformation,[],[f2775_D])).
fof(f2775_D,plain,(
  ( ! [X143] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7076,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7072,f230])).
fof(f7072,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4112,f153])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,X2,X3) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f7064,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7060,f230])).
fof(f7060,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,sK8)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4111,f153])).
fof(f4111,plain,(
  ( ! [X130,X131,X129] : (~sP0(X129,sK8,X130,X131)) ) | $spl42),
  inference(subsumption_resolution,[],[f3974,f170])).
fof(f3974,plain,(
  ( ! [X130,X131,X129] : (p(s(bool,'F_0')) | ~sP0(X129,sK8,X130,X131)) ) | $spl42),
  inference(superposition,[],[f151,f309])).
fof(f7050,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl20 | $spl42)),
  inference(subsumption_resolution,[],[f7046,f230])).
fof(f7046,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK8)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl42),
  inference(resolution,[],[f4110,f153])).
fof(f4110,plain,(
  ( ! [X134,X132,X133] : (~sP0(sK8,X132,X133,X134)) ) | $spl42),
  inference(subsumption_resolution,[],[f3975,f170])).
fof(f3975,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,'F_0')) | ~sP0(sK8,X132,X133,X134)) ) | $spl42),
  inference(superposition,[],[f152,f309])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f79])).
fof(f1993,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl84),
  inference(cnf_transformation,[],[f1993_D])).
fof(f1993_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7039,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f168])).
fof(f2345,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | ($spl0 | $spl30 | $spl34)),
  inference(backward_demodulation,[],[f301,f2069])).
fof(f1392,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl74),
  inference(cnf_transformation,[],[f1392_D])).
fof(f1392_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7038,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f168])).
fof(f7037,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f169])).
fof(f7036,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f169])).
fof(f6975,plain,(
  $false | $spl94),
  inference(resolution,[],[f2751,f189])).
fof(f189,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.TRUTH')).
fof(f6976,plain,(
  $false | $spl94),
  inference(resolution,[],[f2751,f190])).
fof(f6977,plain,(
  $false | $spl94),
  inference(resolution,[],[f2751,f200])).
fof(f6985,plain,(
  $false | ($spl20 | $spl94)),
  inference(resolution,[],[f2751,f230])).
fof(f6726,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X29),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6725,f2736])).
fof(f6725,plain,(
  ( ! [X28,X29,X27] : (sP6(X27,sK9,X28,X29) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X29),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6679,f170])).
fof(f6679,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'F_0')) | sP6(X27,sK9,X28,X29) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X29),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(superposition,[],[f123,f301])).
fof(f6728,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X32),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6727,f2737])).
fof(f6727,plain,(
  ( ! [X30,X31,X32] : (sP6(sK9,X30,X31,X32) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X32),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6680,f170])).
fof(f6680,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'F_0')) | sP6(sK9,X30,X31,X32) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X32),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(superposition,[],[f123,f301])).
fof(f6730,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X47),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6729,f2742])).
fof(f6729,plain,(
  ( ! [X47,X45,X46] : (sP5(X45,X46,sK9,X47) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X47),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6685,f170])).
fof(f6685,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'F_0')) | sP5(X45,X46,sK9,X47) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X47),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(superposition,[],[f128,f301])).
fof(f6732,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X50),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6731,f2743])).
fof(f6731,plain,(
  ( ! [X50,X48,X49] : (sP5(sK9,X48,X49,X50) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X50),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6686,f170])).
fof(f6686,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'F_0')) | sP5(sK9,X48,X49,X50) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X50),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(superposition,[],[f128,f301])).
fof(f6734,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X65),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6733,f2747])).
fof(f6733,plain,(
  ( ! [X64,X65,X63] : (sP4(sK9,X63,X64,X65) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X65),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6691,f170])).
fof(f6691,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'F_0')) | sP4(sK9,X63,X64,X65) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X65),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(superposition,[],[f133,f301])).
fof(f6736,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X83),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6735,f2754])).
fof(f2754,plain,(
  ( ! [X74,X72,X73] : (~sP3(X72,X73,sK9,X74)) ) | $spl38),
  inference(subsumption_resolution,[],[f2689,f200])).
fof(f2689,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | ~sP3(X72,X73,sK9,X74)) ) | $spl38),
  inference(superposition,[],[f135,f305])).
fof(f6735,plain,(
  ( ! [X83,X81,X82] : (sP3(X81,X82,sK9,X83) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X83),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6697,f170])).
fof(f6697,plain,(
  ( ! [X83,X81,X82] : (p(s(bool,'F_0')) | sP3(X81,X82,sK9,X83) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X83),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(superposition,[],[f138,f301])).
fof(f6738,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X86),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6737,f2755])).
fof(f6737,plain,(
  ( ! [X85,X86,X84] : (sP3(X84,sK9,X85,X86) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X86),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(subsumption_resolution,[],[f6698,f170])).
fof(f6698,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,'F_0')) | sP3(X84,sK9,X85,X86) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X86),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(superposition,[],[f138,f301])).
fof(f6910,plain,(
  ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f6909,f170])).
fof(f6909,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl30 | $spl34)),
  inference(forward_demodulation,[],[f6908,f301])).
fof(f6908,plain,(
  p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f6907,f200])).
fof(f6907,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl30)),
  inference(forward_demodulation,[],[f1979,f297])).
fof(f6902,plain,(
  p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f6901,f170])).
fof(f6901,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl24 | $spl34)),
  inference(forward_demodulation,[],[f4587,f301])).
fof(f4587,plain,(
  p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f1986,f240])).
fof(f6760,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6759,f170])).
fof(f6759,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'F_0')) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f5715,f301])).
fof(f5715,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | p(s(bool,sK9)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f123,f2736])).
fof(f6758,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6757,f170])).
fof(f6757,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'F_0')) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f5718,f301])).
fof(f5718,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,sK9)) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f123,f2737])).
fof(f6756,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6755,f170])).
fof(f6755,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'F_0')) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f5946,f301])).
fof(f5946,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f128,f2742])).
fof(f6753,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6752,f170])).
fof(f6752,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'F_0')) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f5951,f301])).
fof(f5951,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,sK9)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X40),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f128,f2743])).
fof(f6749,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X37),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6748,f170])).
fof(f6748,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X37),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f6144,f301])).
fof(f6144,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,sK9)) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X37),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(resolution,[],[f133,f2747])).
fof(f6747,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6746,f170])).
fof(f6746,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'F_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f6380,f301])).
fof(f6380,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f138,f2754])).
fof(f6745,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X34),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(subsumption_resolution,[],[f6744,f170])).
fof(f6744,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'F_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X34),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl34 | $spl38)),
  inference(forward_demodulation,[],[f6383,f301])).
fof(f6383,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,sK9)) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X34),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(resolution,[],[f138,f2755])).
fof(f6742,plain,(
  p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl30 | $spl34)),
  inference(subsumption_resolution,[],[f6741,f170])).
fof(f6741,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl30 | $spl34)),
  inference(forward_demodulation,[],[f6740,f301])).
fof(f6740,plain,(
  p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f6739,f200])).
fof(f6739,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl30)),
  inference(forward_demodulation,[],[f6387,f297])).
fof(f6718,plain,(
  $false | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6666,f170])).
fof(f6666,plain,(
  p(s(bool,'F_0')) | ($spl22 | $spl34)),
  inference(backward_demodulation,[],[f301,f235])).
fof(f6581,plain,(
  ( ! [X26,X27,X25] : (sP3(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl168),
  inference(cnf_transformation,[],[f6581_D])).
fof(f6581_D,plain,(
  ( ! [X26,X27,X25] : (sP3(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6542,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP3(X29,X30,'F_0',X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f138,f168])).
fof(f6541,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP3(X25,X26,'F_0',X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f138,f169])).
fof(f6580,plain,(
  ( ! [X19,X17,X18] : (sP3(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl166),
  inference(cnf_transformation,[],[f6580_D])).
fof(f6580_D,plain,(
  ( ! [X19,X17,X18] : (sP3(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6540,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP3(X21,X22,'T_0',X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f138,f168])).
fof(f6539,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP3(X17,X18,'T_0',X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f138,f169])).
fof(f6577,plain,(
  ( ! [X26,X27,X25] : (sP3(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl164),
  inference(cnf_transformation,[],[f6577_D])).
fof(f6577_D,plain,(
  ( ! [X26,X27,X25] : (sP3(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6533,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP3(X29,'F_0',X30,X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f138,f168])).
fof(f6532,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP3(X25,'F_0',X26,X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f138,f169])).
fof(f6576,plain,(
  ( ! [X19,X17,X18] : (sP3(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl162),
  inference(cnf_transformation,[],[f6576_D])).
fof(f6576_D,plain,(
  ( ! [X19,X17,X18] : (sP3(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6531,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP3(X21,'T_0',X22,X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f138,f168])).
fof(f6530,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP3(X17,'T_0',X18,X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f138,f169])).
fof(f6572,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6524,f4414])).
fof(f4414,plain,(
  ( ! [X14,X12,X13] : (~sP3('F_0',X12,X13,X14)) )),
  inference(resolution,[],[f170,f137])).
fof(f6524,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP3('F_0',X29,X30,X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f138,f168])).
fof(f6571,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f6523,f4414])).
fof(f6523,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP3('F_0',X25,X26,X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f138,f169])).
fof(f6567,plain,(
  ( ! [X156,X154,X155] : (sP3('T_0',X154,X155,X156) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X156),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X155)) | p(s(bool,X154))) ) | $spl160),
  inference(cnf_transformation,[],[f6567_D])).
fof(f6567_D,plain,(
  ( ! [X156,X154,X155] : (sP3('T_0',X154,X155,X156) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X156),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X155)) | p(s(bool,X154))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6522,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP3('T_0',X21,X22,X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f138,f168])).
fof(f6521,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP3('T_0',X17,X18,X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f138,f169])).
fof(f6568,plain,(
  ( ! [X241,X242,X240] : (p(s(bool,X240)) | p(s(bool,X241)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X242),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6428,f2419])).
fof(f2419,plain,(
  ( ! [X80,X78,X79] : (~sP3(sK9,X78,X79,X80)) ) | $spl34),
  inference(subsumption_resolution,[],[f2375,f170])).
fof(f2375,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'F_0')) | ~sP3(sK9,X78,X79,X80)) ) | $spl34),
  inference(superposition,[],[f137,f301])).
fof(f6428,plain,(
  ( ! [X241,X242,X240] : (sP3(sK9,X240,X241,X242) | p(s(bool,X240)) | p(s(bool,X241)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X242),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl22),
  inference(resolution,[],[f138,f235])).
fof(f6413,plain,(
  ( ! [X167,X169,X166,X168] : (sP3('T_0',X166,X167,X168) | p(s(bool,X166)) | p(s(bool,X167)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X168),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ~p(s(bool,X169))) )),
  inference(resolution,[],[f138,f188])).
fof(f6412,plain,(
  ( ! [X165,X163,X164,X162] : (sP3('T_0',X162,X163,X164) | p(s(bool,X162)) | p(s(bool,X163)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X164),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ~p(s(bool,X165))) )),
  inference(resolution,[],[f138,f184])).
fof(f6411,plain,(
  ( ! [X158,X161,X159,X160] : (sP3('T_0',X158,X159,X160) | p(s(bool,X158)) | p(s(bool,X159)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X160),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ~p(s(bool,X161))) )),
  inference(resolution,[],[f138,f181])).
fof(f6410,plain,(
  ( ! [X156,X154,X157,X155] : (sP3('T_0',X154,X155,X156) | p(s(bool,X154)) | p(s(bool,X155)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X156),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ~p(s(bool,X157))) )),
  inference(resolution,[],[f138,f177])).
fof(f6566,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X43),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6386,f235])).
fof(f6386,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,sK9)) | p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X43),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(resolution,[],[f138,f2419])).
fof(f6346,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6305,f4412])).
fof(f4412,plain,(
  ( ! [X6,X8,X7] : (~sP4(X6,X7,'F_0',X8)) )),
  inference(resolution,[],[f170,f130])).
fof(f6305,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP4(X29,X30,'F_0',X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f133,f168])).
fof(f6345,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f6304,f4412])).
fof(f6304,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP4(X25,X26,'F_0',X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f133,f169])).
fof(f6332,plain,(
  ( ! [X142,X140,X141] : (sP4(X140,X141,'T_0',X142) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X141)) | p(s(bool,X140))) ) | $spl154),
  inference(cnf_transformation,[],[f6332_D])).
fof(f6332_D,plain,(
  ( ! [X142,X140,X141] : (sP4(X140,X141,'T_0',X142) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X141)) | p(s(bool,X140))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6303,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP4(X21,X22,'T_0',X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f133,f168])).
fof(f6302,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP4(X17,X18,'T_0',X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f133,f169])).
fof(f6341,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6296,f4413])).
fof(f4413,plain,(
  ( ! [X10,X11,X9] : (~sP4(X9,'F_0',X10,X11)) )),
  inference(resolution,[],[f170,f131])).
fof(f6296,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP4(X29,'F_0',X30,X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f133,f168])).
fof(f6340,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f6295,f4413])).
fof(f6295,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP4(X25,'F_0',X26,X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f133,f169])).
fof(f6330,plain,(
  ( ! [X142,X140,X141] : (sP4(X140,'T_0',X141,X142) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X141)) | p(s(bool,X140))) ) | $spl152),
  inference(cnf_transformation,[],[f6330_D])).
fof(f6330_D,plain,(
  ( ! [X142,X140,X141] : (sP4(X140,'T_0',X141,X142) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X141)) | p(s(bool,X140))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6294,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP4(X21,'T_0',X22,X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f133,f168])).
fof(f6293,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP4(X17,'T_0',X18,X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f133,f169])).
fof(f6337,plain,(
  ( ! [X26,X27,X25] : (sP4('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl158),
  inference(cnf_transformation,[],[f6337_D])).
fof(f6337_D,plain,(
  ( ! [X26,X27,X25] : (sP4('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6287,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP4('F_0',X29,X30,X31) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f133,f168])).
fof(f6286,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP4('F_0',X25,X26,X27) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f133,f169])).
fof(f6336,plain,(
  ( ! [X19,X17,X18] : (sP4('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl156),
  inference(cnf_transformation,[],[f6336_D])).
fof(f6336_D,plain,(
  ( ! [X19,X17,X18] : (sP4('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6285,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP4('T_0',X21,X22,X23) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f133,f168])).
fof(f6284,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP4('T_0',X17,X18,X19) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f133,f169])).
fof(f6333,plain,(
  ( ! [X227,X228,X226] : (p(s(bool,X226)) | ~p(s(bool,X227)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X228),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6254,f2416])).
fof(f2416,plain,(
  ( ! [X54,X56,X55] : (~sP4(X54,X55,sK9,X56)) ) | $spl34),
  inference(subsumption_resolution,[],[f2367,f170])).
fof(f2367,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'F_0')) | ~sP4(X54,X55,sK9,X56)) ) | $spl34),
  inference(superposition,[],[f130,f301])).
fof(f6254,plain,(
  ( ! [X227,X228,X226] : (sP4(X226,X227,sK9,X228) | p(s(bool,X226)) | ~p(s(bool,X227)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X228),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl22),
  inference(resolution,[],[f133,f235])).
fof(f6239,plain,(
  ( ! [X154,X152,X155,X153] : (sP4(X152,X153,'T_0',X154) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X154),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X155))) )),
  inference(resolution,[],[f133,f188])).
fof(f6238,plain,(
  ( ! [X151,X149,X150,X148] : (sP4(X148,X149,'T_0',X150) | p(s(bool,X148)) | ~p(s(bool,X149)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X150),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X151))) )),
  inference(resolution,[],[f133,f184])).
fof(f6237,plain,(
  ( ! [X146,X144,X147,X145] : (sP4(X144,X145,'T_0',X146) | p(s(bool,X144)) | ~p(s(bool,X145)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X146),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X147))) )),
  inference(resolution,[],[f133,f181])).
fof(f6236,plain,(
  ( ! [X142,X140,X143,X141] : (sP4(X140,X141,'T_0',X142) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X143))) )),
  inference(resolution,[],[f133,f177])).
fof(f6331,plain,(
  ( ! [X227,X228,X226] : (p(s(bool,X226)) | ~p(s(bool,X227)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X228),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6214,f2417])).
fof(f6214,plain,(
  ( ! [X227,X228,X226] : (sP4(X226,sK9,X227,X228) | p(s(bool,X226)) | ~p(s(bool,X227)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X228),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl22),
  inference(resolution,[],[f133,f235])).
fof(f6199,plain,(
  ( ! [X154,X152,X155,X153] : (sP4(X152,'T_0',X153,X154) | p(s(bool,X152)) | ~p(s(bool,X153)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X154),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X155))) )),
  inference(resolution,[],[f133,f188])).
fof(f6198,plain,(
  ( ! [X151,X149,X150,X148] : (sP4(X148,'T_0',X149,X150) | p(s(bool,X148)) | ~p(s(bool,X149)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X150),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X151))) )),
  inference(resolution,[],[f133,f184])).
fof(f6197,plain,(
  ( ! [X146,X144,X147,X145] : (sP4(X144,'T_0',X145,X146) | p(s(bool,X144)) | ~p(s(bool,X145)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X146),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X147))) )),
  inference(resolution,[],[f133,f181])).
fof(f6196,plain,(
  ( ! [X142,X140,X143,X141] : (sP4(X140,'T_0',X141,X142) | p(s(bool,X140)) | ~p(s(bool,X141)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X142),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~p(s(bool,X143))) )),
  inference(resolution,[],[f133,f177])).
fof(f6328,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6142,f235])).
fof(f6142,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,sK9)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f133,f2417])).
fof(f6327,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6140,f235])).
fof(f6140,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f133,f2416])).
fof(f6131,plain,(
  ( ! [X26,X27,X25] : (sP5(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) | $spl150),
  inference(cnf_transformation,[],[f6131_D])).
fof(f6131_D,plain,(
  ( ! [X26,X27,X25] : (sP5(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6092,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP5(X29,X30,'F_0',X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f128,f168])).
fof(f6091,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP5(X25,X26,'F_0',X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f128,f169])).
fof(f6130,plain,(
  ( ! [X19,X17,X18] : (sP5(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) | $spl148),
  inference(cnf_transformation,[],[f6130_D])).
fof(f6130_D,plain,(
  ( ! [X19,X17,X18] : (sP5(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f6090,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP5(X21,X22,'T_0',X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f128,f168])).
fof(f6089,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP5(X17,X18,'T_0',X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f128,f169])).
fof(f6126,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6083,f4411])).
fof(f4411,plain,(
  ( ! [X4,X5,X3] : (~sP5(X3,'F_0',X4,X5)) )),
  inference(resolution,[],[f170,f126])).
fof(f6083,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP5(X29,'F_0',X30,X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f128,f168])).
fof(f6125,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f6082,f4411])).
fof(f6082,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP5(X25,'F_0',X26,X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f128,f169])).
fof(f6117,plain,(
  ( ! [X134,X135,X133] : (sP5(X133,'T_0',X134,X135) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X135),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X134)) | p(s(bool,X133))) ) | $spl142),
  inference(cnf_transformation,[],[f6117_D])).
fof(f6117_D,plain,(
  ( ! [X134,X135,X133] : (sP5(X133,'T_0',X134,X135) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X135),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X134)) | p(s(bool,X133))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f6081,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP5(X21,'T_0',X22,X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f128,f168])).
fof(f6080,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP5(X17,'T_0',X18,X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f128,f169])).
fof(f6122,plain,(
  ( ! [X26,X27,X25] : (sP5('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl146),
  inference(cnf_transformation,[],[f6122_D])).
fof(f6122_D,plain,(
  ( ! [X26,X27,X25] : (sP5('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f6074,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP5('F_0',X29,X30,X31) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f128,f168])).
fof(f6073,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP5('F_0',X25,X26,X27) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f128,f169])).
fof(f6121,plain,(
  ( ! [X19,X17,X18] : (sP5('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) | $spl144),
  inference(cnf_transformation,[],[f6121_D])).
fof(f6121_D,plain,(
  ( ! [X19,X17,X18] : (sP5('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | ~p(s(bool,X17))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f6072,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP5('T_0',X21,X22,X23) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f128,f168])).
fof(f6071,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP5('T_0',X17,X18,X19) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f128,f169])).
fof(f6118,plain,(
  ( ! [X220,X219,X221] : (p(s(bool,X219)) | p(s(bool,X220)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X221),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f6017,f2413])).
fof(f6017,plain,(
  ( ! [X220,X219,X221] : (sP5(X219,sK9,X220,X221) | p(s(bool,X219)) | p(s(bool,X220)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X221),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl22),
  inference(resolution,[],[f128,f235])).
fof(f6002,plain,(
  ( ! [X146,X147,X145,X148] : (sP5(X145,'T_0',X146,X147) | p(s(bool,X145)) | p(s(bool,X146)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X147),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X148))) )),
  inference(resolution,[],[f128,f188])).
fof(f6001,plain,(
  ( ! [X144,X142,X143,X141] : (sP5(X141,'T_0',X142,X143) | p(s(bool,X141)) | p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X144))) )),
  inference(resolution,[],[f128,f184])).
fof(f6000,plain,(
  ( ! [X140,X138,X139,X137] : (sP5(X137,'T_0',X138,X139) | p(s(bool,X137)) | p(s(bool,X138)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X139),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X140))) )),
  inference(resolution,[],[f128,f181])).
fof(f5999,plain,(
  ( ! [X134,X136,X135,X133] : (sP5(X133,'T_0',X134,X135) | p(s(bool,X133)) | p(s(bool,X134)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X135),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,X136))) )),
  inference(resolution,[],[f128,f177])).
fof(f6115,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X34),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5949,f235])).
fof(f5949,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,sK9)) | p(s(bool,X33)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X34),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl34),
  inference(resolution,[],[f128,f2413])).
fof(f5915,plain,(
  $false | $spl88),
  inference(resolution,[],[f2741,f170])).
fof(f5891,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5851,f4410])).
fof(f4410,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,'F_0',X2)) )),
  inference(resolution,[],[f170,f120])).
fof(f5851,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,X28)) | sP6(X29,X30,'F_0',X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f123,f168])).
fof(f5889,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f5850,f4410])).
fof(f5850,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X24)) | sP6(X25,X26,'F_0',X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f123,f169])).
fof(f5876,plain,(
  ( ! [X121,X120,X119] : (sP6(X119,X120,'T_0',X121) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X121),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | p(s(bool,X120)) | p(s(bool,X119))) ) | $spl128),
  inference(cnf_transformation,[],[f5876_D])).
fof(f5876_D,plain,(
  ( ! [X121,X120,X119] : (sP6(X119,X120,'T_0',X121) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X121),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | p(s(bool,X120)) | p(s(bool,X119))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f5849,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X20)) | sP6(X21,X22,'T_0',X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f123,f168])).
fof(f5848,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP6(X17,X18,'T_0',X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f123,f169])).
fof(f5885,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) | $spl136),
  inference(cnf_transformation,[],[f5885_D])).
fof(f5885_D,plain,(
  ( ! [X26,X27,X25] : (sP6(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5842,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP6(X29,'F_0',X30,X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f123,f168])).
fof(f5841,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP6(X25,'F_0',X26,X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f123,f169])).
fof(f5884,plain,(
  ( ! [X19,X17,X18] : (sP6(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) | $spl134),
  inference(cnf_transformation,[],[f5884_D])).
fof(f5884_D,plain,(
  ( ! [X19,X17,X18] : (sP6(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f5840,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP6(X21,'T_0',X22,X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f123,f168])).
fof(f5839,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP6(X17,'T_0',X18,X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f123,f169])).
fof(f5881,plain,(
  ( ! [X26,X27,X25] : (sP6('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) | $spl132),
  inference(cnf_transformation,[],[f5881_D])).
fof(f5881_D,plain,(
  ( ! [X26,X27,X25] : (sP6('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5833,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP6('F_0',X29,X30,X31) | p(s(bool,X29)) | ~p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f123,f168])).
fof(f5832,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP6('F_0',X25,X26,X27) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f123,f169])).
fof(f5880,plain,(
  ( ! [X19,X17,X18] : (sP6('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) | $spl130),
  inference(cnf_transformation,[],[f5880_D])).
fof(f5880_D,plain,(
  ( ! [X19,X17,X18] : (sP6('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f5831,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP6('T_0',X21,X22,X23) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f123,f168])).
fof(f5830,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP6('T_0',X17,X18,X19) | p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f123,f169])).
fof(f5877,plain,(
  ( ! [X206,X207,X205] : (p(s(bool,X205)) | p(s(bool,X206)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X207),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5804,f2410])).
fof(f2410,plain,(
  ( ! [X19,X20,X18] : (~sP6(X18,X19,sK9,X20)) ) | $spl34),
  inference(subsumption_resolution,[],[f2355,f170])).
fof(f2355,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'F_0')) | ~sP6(X18,X19,sK9,X20)) ) | $spl34),
  inference(superposition,[],[f120,f301])).
fof(f5804,plain,(
  ( ! [X206,X207,X205] : (sP6(X205,X206,sK9,X207) | p(s(bool,X205)) | p(s(bool,X206)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X207),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl22),
  inference(resolution,[],[f123,f235])).
fof(f5789,plain,(
  ( ! [X134,X132,X133,X131] : (sP6(X131,X132,'T_0',X133) | p(s(bool,X131)) | p(s(bool,X132)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X133),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X134))) )),
  inference(resolution,[],[f123,f188])).
fof(f5788,plain,(
  ( ! [X127,X130,X128,X129] : (sP6(X127,X128,'T_0',X129) | p(s(bool,X127)) | p(s(bool,X128)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X129),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X130))) )),
  inference(resolution,[],[f123,f184])).
fof(f5787,plain,(
  ( ! [X125,X123,X126,X124] : (sP6(X123,X124,'T_0',X125) | p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X126))) )),
  inference(resolution,[],[f123,f181])).
fof(f5786,plain,(
  ( ! [X121,X122,X120,X119] : (sP6(X119,X120,'T_0',X121) | p(s(bool,X119)) | p(s(bool,X120)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X121),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ~p(s(bool,X122))) )),
  inference(resolution,[],[f123,f177])).
fof(f5873,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5713,f235])).
fof(f5713,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | ~p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X25),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f123,f2410])).
fof(f5704,plain,(
  ( ! [X26,X27,X25] : (sP7(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) | $spl126),
  inference(cnf_transformation,[],[f5704_D])).
fof(f5704_D,plain,(
  ( ! [X26,X27,X25] : (sP7(X25,X26,'F_0',X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f5673,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP7(X29,X30,'F_0',X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f118,f168])).
fof(f5672,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP7(X25,X26,'F_0',X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f118,f169])).
fof(f5703,plain,(
  ( ! [X19,X17,X18] : (sP7(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) | $spl124),
  inference(cnf_transformation,[],[f5703_D])).
fof(f5703_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X17,X18,'T_0',X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f5671,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP7(X21,X22,'T_0',X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f118,f168])).
fof(f5670,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP7(X17,X18,'T_0',X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f118,f169])).
fof(f5700,plain,(
  ( ! [X26,X27,X25] : (sP7(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) | $spl122),
  inference(cnf_transformation,[],[f5700_D])).
fof(f5700_D,plain,(
  ( ! [X26,X27,X25] : (sP7(X25,'F_0',X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f5664,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP7(X29,'F_0',X30,X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f118,f168])).
fof(f5663,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP7(X25,'F_0',X26,X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f118,f169])).
fof(f5699,plain,(
  ( ! [X19,X17,X18] : (sP7(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) | $spl120),
  inference(cnf_transformation,[],[f5699_D])).
fof(f5699_D,plain,(
  ( ! [X19,X17,X18] : (sP7(X17,'T_0',X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f5662,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP7(X21,'T_0',X22,X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f118,f168])).
fof(f5661,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP7(X17,'T_0',X18,X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f118,f169])).
fof(f5695,plain,(
  ( ! [X26,X27,X25] : (sP7('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) | $spl116),
  inference(cnf_transformation,[],[f5695_D])).
fof(f5695_D,plain,(
  ( ! [X26,X27,X25] : (sP7('F_0',X25,X26,X27) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X26)) | p(s(bool,X25))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5655,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | sP7('F_0',X29,X30,X31) | p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X31),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f118,f168])).
fof(f5654,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X24)) | sP7('F_0',X25,X26,X27) | p(s(bool,X25)) | p(s(bool,X26)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X27),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f118,f169])).
fof(f5693,plain,(
  ( ! [X19,X17,X18] : (sP7('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) | $spl112),
  inference(cnf_transformation,[],[f5693_D])).
fof(f5693_D,plain,(
  ( ! [X19,X17,X18] : (sP7('T_0',X17,X18,X19) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | p(s(bool,X18)) | p(s(bool,X17))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f5653,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X20)) | sP7('T_0',X21,X22,X23) | p(s(bool,X21)) | p(s(bool,X22)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X23),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f118,f168])).
fof(f5652,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X16)) | sP7('T_0',X17,X18,X19) | p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X19),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f118,f169])).
fof(f5183,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5179,f235])).
fof(f5179,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,sK9)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f2430,f153])).
fof(f2430,plain,(
  ( ! [X134,X132,X133] : (~sP0(sK9,X132,X133,X134)) ) | $spl34),
  inference(subsumption_resolution,[],[f2393,f170])).
fof(f2393,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,'F_0')) | ~sP0(sK9,X132,X133,X134)) ) | $spl34),
  inference(superposition,[],[f152,f301])).
fof(f5166,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5162,f235])).
fof(f5162,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,sK9)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f2429,f153])).
fof(f5149,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl22 | $spl34)),
  inference(subsumption_resolution,[],[f5145,f235])).
fof(f5145,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X2),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl34),
  inference(resolution,[],[f2428,f153])).
fof(f2428,plain,(
  ( ! [X127,X128,X126] : (~sP0(X126,X127,sK9,X128)) ) | $spl34),
  inference(subsumption_resolution,[],[f2391,f170])).
fof(f2391,plain,(
  ( ! [X127,X128,X126] : (p(s(bool,'F_0')) | ~sP0(X126,X127,sK9,X128)) ) | $spl34),
  inference(superposition,[],[f150,f301])).
fof(f1394,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl76),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4679,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl34 | $spl46)),
  inference(superposition,[],[f4589,f168])).
fof(f4589,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl0 | $spl34 | $spl46)),
  inference(forward_demodulation,[],[f2916,f313])).
fof(f2916,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | ($spl0 | $spl34)),
  inference(backward_demodulation,[],[f301,f203])).
fof(f4678,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl34 | $spl46)),
  inference(superposition,[],[f4589,f168])).
fof(f4677,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl34 | $spl46)),
  inference(superposition,[],[f4589,f169])).
fof(f4676,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl34 | $spl46)),
  inference(superposition,[],[f4589,f169])).
fof(f883,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl70),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4474,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f169])).
fof(f4399,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | ($spl0 | $spl46)),
  inference(forward_demodulation,[],[f203,f313])).
fof(f4476,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f168])).
fof(f4570,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl30 | $spl46)),
  inference(superposition,[],[f4409,f169])).
fof(f4409,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | ($spl0 | $spl30 | $spl46)),
  inference(forward_demodulation,[],[f2069,f313])).
fof(f4572,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl30 | $spl46)),
  inference(superposition,[],[f4409,f168])).
fof(f4617,plain,(
  $false | ($spl22 | $spl39)),
  inference(subsumption_resolution,[],[f4616,f304])).
fof(f4616,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl22),
  inference(subsumption_resolution,[],[f4613,f170])).
fof(f4613,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK9) | $spl22),
  inference(superposition,[],[f235,f168])).
fof(f4615,plain,(
  $false | ($spl22 | $spl39)),
  inference(subsumption_resolution,[],[f4614,f304])).
fof(f4614,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl22),
  inference(subsumption_resolution,[],[f4611,f170])).
fof(f4611,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK9) | $spl22),
  inference(superposition,[],[f235,f169])).
fof(f4571,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl30 | $spl46)),
  inference(superposition,[],[f4409,f168])).
fof(f4569,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl30 | $spl46)),
  inference(superposition,[],[f4409,f169])).
fof(f4475,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f168])).
fof(f4473,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f169])).
fof(f4470,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f168])).
fof(f490,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | $spl54),
  inference(cnf_transformation,[],[f490_D])).
fof(f490_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4469,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f168])).
fof(f4468,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f169])).
fof(f4467,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(superposition,[],[f4399,f169])).
fof(f4405,plain,(
  ~p(s(bool,'F_0')) | $spl111),
  inference(cnf_transformation,[],[f4405_D])).
fof(f4405_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f157,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : ((p(s(bool,X10)) | p(s(bool,'F_0'))) <=> p(s(bool,X10)))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.OR_CLAUSES|split|3')).
fof(f159,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) | p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.OR_CLAUSES|split|2')).
fof(f163,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10] : ((p(s(bool,X10)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.AND_CLAUSES|split|3')).
fof(f167,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) & p(s(bool,X10))) <=> p(s(bool,'F_0')))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f4401,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(forward_demodulation,[],[f3324,f313])).
fof(f4400,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl0 | $spl46)),
  inference(forward_demodulation,[],[f3326,f313])).
fof(f4398,plain,(
  p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f4397,f200])).
fof(f4397,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl46)),
  inference(forward_demodulation,[],[f2043,f313])).
fof(f4395,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f4394,f200])).
fof(f4394,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK10)) | p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl3 | $spl46)),
  inference(forward_demodulation,[],[f1979,f313])).
fof(f4377,plain,(
  $false | ($spl34 | $spl39 | $spl42 | $spl46)),
  inference(subsumption_resolution,[],[f4319,f304])).
fof(f4319,plain,(
  s(bool,'T_0') = s(bool,sK9) | ($spl34 | $spl42 | $spl46)),
  inference(backward_demodulation,[],[f4001,f301])).
fof(f4001,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl42 | $spl46)),
  inference(backward_demodulation,[],[f313,f309])).
fof(f4376,plain,(
  $false | ($spl26 | $spl31 | $spl42 | $spl46)),
  inference(subsumption_resolution,[],[f4318,f296])).
fof(f296,plain,(
  s(bool,'T_0') != s(bool,sK10) | $spl31),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  s(bool,'T_0') != s(bool,sK10) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f4318,plain,(
  s(bool,'T_0') = s(bool,sK10) | ($spl26 | $spl42 | $spl46)),
  inference(backward_demodulation,[],[f4001,f293])).
fof(f4375,plain,(
  $false | ($spl42 | $spl46)),
  inference(subsumption_resolution,[],[f4364,f190])).
fof(f4364,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl46)),
  inference(duplicate_literal_removal,[],[f4317])).
fof(f4317,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl46)),
  inference(backward_demodulation,[],[f4001,f174])).
fof(f174,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f4374,plain,(
  $false | ($spl42 | $spl46)),
  inference(subsumption_resolution,[],[f4366,f190])).
fof(f4366,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl46)),
  inference(duplicate_literal_removal,[],[f4315])).
fof(f4315,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl46)),
  inference(backward_demodulation,[],[f4001,f172])).
fof(f172,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f4373,plain,(
  $false | ($spl42 | $spl46)),
  inference(subsumption_resolution,[],[f4313,f190])).
fof(f4313,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl46)),
  inference(backward_demodulation,[],[f4001,f170])).
fof(f4109,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4108,f1870])).
fof(f1870,plain,(
  ( ! [X64,X65,X63] : (~sP0(X63,X64,sK8,X65)) ) | $spl21),
  inference(resolution,[],[f231,f150])).
fof(f4108,plain,(
  ( ! [X142,X143,X141] : (sP0(X141,X142,sK8,X143) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4051,f200])).
fof(f4051,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,'T_0')) | sP0(X141,X142,sK8,X143) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f153,f313])).
fof(f4107,plain,(
  ( ! [X140,X138,X139] : (~p(s(bool,X138)) | ~p(s(bool,X139)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X140),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4106,f1871])).
fof(f1871,plain,(
  ( ! [X68,X66,X67] : (~sP0(X66,sK8,X67,X68)) ) | $spl21),
  inference(resolution,[],[f231,f151])).
fof(f4106,plain,(
  ( ! [X140,X138,X139] : (sP0(X138,sK8,X139,X140) | ~p(s(bool,X138)) | ~p(s(bool,X139)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X140),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4050,f200])).
fof(f4050,plain,(
  ( ! [X140,X138,X139] : (~p(s(bool,'T_0')) | sP0(X138,sK8,X139,X140) | ~p(s(bool,X138)) | ~p(s(bool,X139)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X140),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f153,f313])).
fof(f4105,plain,(
  ( ! [X136,X135,X137] : (~p(s(bool,X135)) | ~p(s(bool,X136)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X137),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4104,f1872])).
fof(f1872,plain,(
  ( ! [X70,X71,X69] : (~sP0(sK8,X69,X70,X71)) ) | $spl21),
  inference(resolution,[],[f231,f152])).
fof(f4104,plain,(
  ( ! [X136,X135,X137] : (sP0(sK8,X135,X136,X137) | ~p(s(bool,X135)) | ~p(s(bool,X136)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X137),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4049,f200])).
fof(f4049,plain,(
  ( ! [X136,X135,X137] : (~p(s(bool,'T_0')) | sP0(sK8,X135,X136,X137) | ~p(s(bool,X135)) | ~p(s(bool,X136)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X137),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f153,f313])).
fof(f4103,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4102,f1867])).
fof(f1867,plain,(
  ( ! [X54,X56,X55] : (~sP1(X54,sK8,X55,X56)) ) | $spl21),
  inference(resolution,[],[f231,f146])).
fof(f4102,plain,(
  ( ! [X125,X123,X124] : (sP1(X123,sK8,X124,X125) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4045,f200])).
fof(f4045,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,'T_0')) | sP1(X123,sK8,X124,X125) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(superposition,[],[f148,f313])).
fof(f4101,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4100,f1868])).
fof(f1868,plain,(
  ( ! [X59,X57,X58] : (~sP1(sK8,X57,X58,X59)) ) | $spl21),
  inference(resolution,[],[f231,f147])).
fof(f4100,plain,(
  ( ! [X121,X122,X120] : (sP1(sK8,X120,X121,X122) | ~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4044,f200])).
fof(f4044,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,'T_0')) | sP1(sK8,X120,X121,X122) | ~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(superposition,[],[f148,f313])).
fof(f4098,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4097,f1864])).
fof(f1864,plain,(
  ( ! [X47,X45,X46] : (~sP2(X45,X46,sK8,X47)) ) | $spl21),
  inference(resolution,[],[f231,f140])).
fof(f4097,plain,(
  ( ! [X107,X105,X106] : (sP2(X105,X106,sK8,X107) | ~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4039,f200])).
fof(f4039,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,'T_0')) | sP2(X105,X106,sK8,X107) | ~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f143,f313])).
fof(f4096,plain,(
  ( ! [X103,X102,X104] : (p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4095,f1865])).
fof(f1865,plain,(
  ( ! [X50,X48,X49] : (~sP2(sK8,X48,X49,X50)) ) | $spl21),
  inference(resolution,[],[f231,f142])).
fof(f4095,plain,(
  ( ! [X103,X102,X104] : (sP2(sK8,X102,X103,X104) | p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4038,f200])).
fof(f4038,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,'T_0')) | sP2(sK8,X102,X103,X104) | p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f143,f313])).
fof(f4093,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4092,f1861])).
fof(f1861,plain,(
  ( ! [X37,X38,X36] : (~sP3(sK8,X36,X37,X38)) ) | $spl21),
  inference(resolution,[],[f231,f137])).
fof(f4092,plain,(
  ( ! [X88,X87,X89] : (sP3(sK8,X87,X88,X89) | p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4033,f200])).
fof(f4033,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP3(sK8,X87,X88,X89) | p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(superposition,[],[f138,f313])).
fof(f4089,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4088,f1858])).
fof(f1858,plain,(
  ( ! [X28,X29,X27] : (~sP4(X27,X28,sK8,X29)) ) | $spl21),
  inference(resolution,[],[f231,f130])).
fof(f4088,plain,(
  ( ! [X70,X71,X69] : (sP4(X69,X70,sK8,X71) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4027,f200])).
fof(f4027,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | sP4(X69,X70,sK8,X71) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f133,f313])).
fof(f4087,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4086,f1859])).
fof(f1859,plain,(
  ( ! [X30,X31,X32] : (~sP4(X30,sK8,X31,X32)) ) | $spl21),
  inference(resolution,[],[f231,f131])).
fof(f4086,plain,(
  ( ! [X68,X66,X67] : (sP4(X66,sK8,X67,X68) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4026,f200])).
fof(f4026,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | sP4(X66,sK8,X67,X68) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f133,f313])).
fof(f4084,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4083,f1855])).
fof(f1855,plain,(
  ( ! [X19,X20,X18] : (~sP5(X18,sK8,X19,X20)) ) | $spl21),
  inference(resolution,[],[f231,f126])).
fof(f4083,plain,(
  ( ! [X52,X53,X51] : (sP5(X51,sK8,X52,X53) | p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4021,f200])).
fof(f4021,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | sP5(X51,sK8,X52,X53) | p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl46),
  inference(superposition,[],[f128,f313])).
fof(f4080,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl21 | $spl46)),
  inference(subsumption_resolution,[],[f4079,f1852])).
fof(f1852,plain,(
  ( ! [X10,X11,X9] : (~sP6(X9,X10,sK8,X11)) ) | $spl21),
  inference(resolution,[],[f231,f120])).
fof(f4079,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,sK8,X35) | p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(subsumption_resolution,[],[f4015,f200])).
fof(f4015,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | sP6(X33,X34,sK8,X35) | p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl46),
  inference(superposition,[],[f123,f313])).
fof(f3837,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl34)),
  inference(superposition,[],[f3177,f168])).
fof(f3177,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | ($spl0 | $spl26 | $spl34)),
  inference(backward_demodulation,[],[f293,f2916])).
fof(f3836,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl34)),
  inference(superposition,[],[f3177,f168])).
fof(f3835,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl34)),
  inference(superposition,[],[f3177,f169])).
fof(f3834,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26 | $spl34)),
  inference(superposition,[],[f3177,f169])).
fof(f3786,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f3785,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f3784,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f3783,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f3780,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f3779,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f168])).
fof(f3778,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f3777,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl26)),
  inference(superposition,[],[f3548,f169])).
fof(f3487,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f3482,f296])).
fof(f3482,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl27),
  inference(trivial_inequality_removal,[],[f3481])).
fof(f3481,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10) | $spl27),
  inference(superposition,[],[f292,f168])).
fof(f292,plain,(
  s(bool,'F_0') != s(bool,sK10) | $spl27),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  s(bool,'F_0') != s(bool,sK10) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f3485,plain,(
  $false | ($spl27 | $spl31)),
  inference(subsumption_resolution,[],[f3483,f296])).
fof(f3483,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl27),
  inference(trivial_inequality_removal,[],[f3479])).
fof(f3479,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10) | $spl27),
  inference(superposition,[],[f292,f169])).
fof(f3425,plain,(
  $false | ($spl24 | $spl31)),
  inference(subsumption_resolution,[],[f3424,f296])).
fof(f3424,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl24),
  inference(subsumption_resolution,[],[f3412,f170])).
fof(f3412,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK10) | $spl24),
  inference(superposition,[],[f240,f168])).
fof(f3423,plain,(
  $false | ($spl24 | $spl31)),
  inference(subsumption_resolution,[],[f3422,f296])).
fof(f3422,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl24),
  inference(subsumption_resolution,[],[f3410,f170])).
fof(f3410,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK10) | $spl24),
  inference(superposition,[],[f240,f169])).
fof(f3338,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3337,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3336,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f3335,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f306,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | $spl40),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3332,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3331,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3330,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f3329,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f3325,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f3323,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f3240,plain,(
  $false | ($spl24 | $spl26)),
  inference(subsumption_resolution,[],[f3178,f170])).
fof(f3178,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl26)),
  inference(backward_demodulation,[],[f293,f240])).
fof(f885,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl72),
  inference(cnf_transformation,[],[f885_D])).
fof(f885_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3104,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f168])).
fof(f3103,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f168])).
fof(f3102,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f169])).
fof(f3101,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f169])).
fof(f488,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | $spl52),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3089,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f168])).
fof(f3087,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | ($spl0 | $spl34)),
  inference(superposition,[],[f2916,f169])).
fof(f2901,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2900,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2899,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2898,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2895,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2894,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2893,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2892,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2889,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2888,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f2887,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2886,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f2774,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2773,f1966])).
fof(f1966,plain,(
  ( ! [X64,X65,X63] : (~sP0(X63,X64,sK9,X65)) ) | $spl23),
  inference(resolution,[],[f236,f150])).
fof(f2773,plain,(
  ( ! [X142,X143,X141] : (sP0(X141,X142,sK9,X143) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2712,f200])).
fof(f2712,plain,(
  ( ! [X142,X143,X141] : (~p(s(bool,'T_0')) | sP0(X141,X142,sK9,X143) | ~p(s(bool,X141)) | ~p(s(bool,X142)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X143),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f153,f305])).
fof(f2770,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2769,f1963])).
fof(f1963,plain,(
  ( ! [X54,X56,X55] : (~sP1(X54,sK9,X55,X56)) ) | $spl23),
  inference(resolution,[],[f236,f146])).
fof(f2769,plain,(
  ( ! [X125,X123,X124] : (sP1(X123,sK9,X124,X125) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2706,f200])).
fof(f2706,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,'T_0')) | sP1(X123,sK9,X124,X125) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X125),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(superposition,[],[f148,f305])).
fof(f2767,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2766,f1964])).
fof(f1964,plain,(
  ( ! [X59,X57,X58] : (~sP1(sK9,X57,X58,X59)) ) | $spl23),
  inference(resolution,[],[f236,f147])).
fof(f2766,plain,(
  ( ! [X121,X122,X120] : (sP1(sK9,X120,X121,X122) | ~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2705,f200])).
fof(f2705,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,'T_0')) | sP1(sK9,X120,X121,X122) | ~p(s(bool,X120)) | p(s(bool,X121)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X122),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(superposition,[],[f148,f305])).
fof(f2764,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2763,f1960])).
fof(f1960,plain,(
  ( ! [X47,X45,X46] : (~sP2(X45,X46,sK9,X47)) ) | $spl23),
  inference(resolution,[],[f236,f140])).
fof(f2763,plain,(
  ( ! [X107,X105,X106] : (sP2(X105,X106,sK9,X107) | ~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2700,f200])).
fof(f2700,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,'T_0')) | sP2(X105,X106,sK9,X107) | ~p(s(bool,X105)) | p(s(bool,X106)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X107),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f143,f305])).
fof(f2761,plain,(
  ( ! [X103,X102,X104] : (p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2760,f1961])).
fof(f1961,plain,(
  ( ! [X50,X48,X49] : (~sP2(sK9,X48,X49,X50)) ) | $spl23),
  inference(resolution,[],[f236,f142])).
fof(f2760,plain,(
  ( ! [X103,X102,X104] : (sP2(sK9,X102,X103,X104) | p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2699,f200])).
fof(f2699,plain,(
  ( ! [X103,X102,X104] : (~p(s(bool,'T_0')) | sP2(sK9,X102,X103,X104) | p(s(bool,X102)) | ~p(s(bool,X103)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X104),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f143,f305])).
fof(f2757,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2756,f1957])).
fof(f1957,plain,(
  ( ! [X37,X38,X36] : (~sP3(sK9,X36,X37,X38)) ) | $spl23),
  inference(resolution,[],[f236,f137])).
fof(f2756,plain,(
  ( ! [X88,X87,X89] : (sP3(sK9,X87,X88,X89) | p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2694,f200])).
fof(f2694,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP3(sK9,X87,X88,X89) | p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X89),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(superposition,[],[f138,f305])).
fof(f2753,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2752,f1954])).
fof(f1954,plain,(
  ( ! [X28,X29,X27] : (~sP4(X27,X28,sK9,X29)) ) | $spl23),
  inference(resolution,[],[f236,f130])).
fof(f2752,plain,(
  ( ! [X70,X71,X69] : (sP4(X69,X70,sK9,X71) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2688,f200])).
fof(f2688,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | sP4(X69,X70,sK9,X71) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X71),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f133,f305])).
fof(f2749,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2748,f1955])).
fof(f1955,plain,(
  ( ! [X30,X31,X32] : (~sP4(X30,sK9,X31,X32)) ) | $spl23),
  inference(resolution,[],[f236,f131])).
fof(f2748,plain,(
  ( ! [X68,X66,X67] : (sP4(X66,sK9,X67,X68) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2687,f200])).
fof(f2687,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | sP4(X66,sK9,X67,X68) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X68),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f133,f305])).
fof(f2745,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2744,f1951])).
fof(f1951,plain,(
  ( ! [X19,X20,X18] : (~sP5(X18,sK9,X19,X20)) ) | $spl23),
  inference(resolution,[],[f236,f126])).
fof(f2744,plain,(
  ( ! [X52,X53,X51] : (sP5(X51,sK9,X52,X53) | p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2682,f200])).
fof(f2682,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | sP5(X51,sK9,X52,X53) | p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X53),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0'))))) ) | $spl38),
  inference(superposition,[],[f128,f305])).
fof(f2739,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | ($spl23 | $spl38)),
  inference(subsumption_resolution,[],[f2738,f1948])).
fof(f1948,plain,(
  ( ! [X10,X11,X9] : (~sP6(X9,X10,sK9,X11)) ) | $spl23),
  inference(resolution,[],[f236,f120])).
fof(f2738,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,sK9,X35) | p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(subsumption_resolution,[],[f2676,f200])).
fof(f2676,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | sP6(X33,X34,sK9,X35) | p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X35),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0'))))) ) | $spl38),
  inference(superposition,[],[f123,f305])).
fof(f2554,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f168])).
fof(f2553,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f168])).
fof(f2552,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f169])).
fof(f2551,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30 | $spl34)),
  inference(superposition,[],[f2345,f169])).
fof(f492,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | $spl56),
  inference(cnf_transformation,[],[f492_D])).
fof(f492_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2260,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f168])).
fof(f2259,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f168])).
fof(f2258,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f169])).
fof(f2257,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f169])).
fof(f2254,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f168])).
fof(f2253,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f168])).
fof(f2252,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f169])).
fof(f2251,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | ($spl0 | $spl30)),
  inference(superposition,[],[f2069,f169])).
fof(f2130,plain,(
  $false | ($spl25 | $spl30)),
  inference(subsumption_resolution,[],[f2068,f190])).
fof(f2068,plain,(
  ~p(s(bool,'T_0')) | ($spl25 | $spl30)),
  inference(backward_demodulation,[],[f297,f241])).
fof(f1992,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl21 | $spl23)),
  inference(subsumption_resolution,[],[f1991,f231])).
fof(f1991,plain,(
  ~p(s(bool,sK10)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl3 | $spl23)),
  inference(subsumption_resolution,[],[f1986,f236])).
fof(f1940,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,sK9)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) | ($spl3 | $spl21)),
  inference(subsumption_resolution,[],[f1935,f231])).
fof(f1928,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1927,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1926,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1925,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1922,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1921,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1920,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1919,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1916,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1915,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f168])).
fof(f1914,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1913,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl0),
  inference(superposition,[],[f203,f169])).
fof(f1840,plain,(
  ~p(s(bool,sK10)) | ~p(s(bool,sK9)) | ~p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | $spl5),
  inference(resolution,[],[f211,f153])).
fof(f1839,plain,(
  p(s(bool,sK10)) | p(s(bool,sK9)) | p(s(bool,sK8)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl3),
  inference(resolution,[],[f206,f118])).
fof(f1822,plain,(
  $false | ($spl1 | $spl4)),
  inference(resolution,[],[f1772,f212])).
fof(f1772,plain,(
  ( ! [X2,X0,X1] : (~sP0(X0,X1,X2,sK11)) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f1756,f149])).
fof(f149,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f79])).
fof(f1756,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f1755,f423])).
fof(f423,plain,(
  s(bool,'T_0') = s(bool,sK10) | $spl4),
  inference(subsumption_resolution,[],[f420,f170])).
fof(f420,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK10) | $spl4),
  inference(superposition,[],[f321,f169])).
fof(f321,plain,(
  p(s(bool,sK10)) | $spl4),
  inference(resolution,[],[f212,f152])).
fof(f1755,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f1168,f389])).
fof(f389,plain,(
  s(bool,'T_0') = s(bool,sK9) | $spl4),
  inference(subsumption_resolution,[],[f386,f170])).
fof(f386,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK9) | $spl4),
  inference(superposition,[],[f320,f169])).
fof(f320,plain,(
  p(s(bool,sK9)) | $spl4),
  inference(resolution,[],[f212,f151])).
fof(f1168,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f355,f204])).
fof(f355,plain,(
  s(bool,'T_0') = s(bool,sK8) | $spl4),
  inference(subsumption_resolution,[],[f352,f170])).
fof(f352,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK8) | $spl4),
  inference(superposition,[],[f319,f169])).
fof(f319,plain,(
  p(s(bool,sK8)) | $spl4),
  inference(resolution,[],[f212,f150])).
fof(f1753,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1745,f212])).
fof(f1745,plain,(
  ~sP0(sK10,sK9,sK8,sK11) | $spl3),
  inference(resolution,[],[f206,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,X1,X2,X3) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f1731,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,X5,sK11) | ~sP7(X3,X4,X5,sK11) | sP0(X3,X4,X5,sK11) | sP1(X3,X4,X5,sK11) | sP3(X3,X4,X5,sK11) | sP4(X3,X4,X5,sK11) | sP5(X3,X4,X5,sK11)) ) | $spl78),
  inference(cnf_transformation,[],[f1731_D])).
fof(f1731_D,plain,(
  ( ! [X4,X5,X3] : (sP6(X3,X4,X5,sK11) | ~sP7(X3,X4,X5,sK11) | sP0(X3,X4,X5,sK11) | sP1(X3,X4,X5,sK11) | sP3(X3,X4,X5,sK11) | sP4(X3,X4,X5,sK11) | sP5(X3,X4,X5,sK11)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1727,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(X3,X4,X5,sK11) | sP5(X3,X4,X5,sK11) | sP4(X3,X4,X5,sK11) | sP3(X3,X4,X5,sK11) | sP1(X3,X4,X5,sK11) | sP0(X3,X4,X5,sK11) | ~sP7(X3,X4,X5,sK11)) ) | ($spl4 | $spl41)),
  inference(resolution,[],[f1675,f107])).
fof(f1675,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2,sK11)) ) | ($spl4 | $spl41)),
  inference(resolution,[],[f1398,f139])).
fof(f1398,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl4 | $spl41)),
  inference(backward_demodulation,[],[f423,f1163])).
fof(f1163,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | ($spl4 | $spl41)),
  inference(backward_demodulation,[],[f355,f307])).
fof(f1373,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl4 | $spl41)),
  inference(superposition,[],[f1163,f168])).
fof(f1372,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl4 | $spl41)),
  inference(superposition,[],[f1163,f168])).
fof(f1371,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | ($spl4 | $spl41)),
  inference(superposition,[],[f1163,f169])).
fof(f1370,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | ($spl4 | $spl41)),
  inference(superposition,[],[f1163,f169])).
fof(f886,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | $spl73),
  inference(cnf_transformation,[],[f886_D])).
fof(f886_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1130,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | $spl49),
  inference(superposition,[],[f315,f168])).
fof(f1129,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | $spl49),
  inference(superposition,[],[f315,f168])).
fof(f1128,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9) | $spl49),
  inference(superposition,[],[f315,f169])).
fof(f1127,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9) | $spl49),
  inference(superposition,[],[f315,f169])).
fof(f1124,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10) | $spl49),
  inference(superposition,[],[f315,f168])).
fof(f493,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | $spl57),
  inference(cnf_transformation,[],[f493_D])).
fof(f493_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f1123,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK10) | $spl49),
  inference(superposition,[],[f315,f168])).
fof(f1122,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10) | $spl49),
  inference(superposition,[],[f315,f169])).
fof(f1121,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK10) | $spl49),
  inference(superposition,[],[f315,f169])).
fof(f884,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | $spl71),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f999,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | $spl45),
  inference(superposition,[],[f311,f168])).
fof(f998,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | $spl45),
  inference(superposition,[],[f311,f168])).
fof(f997,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK9) | $spl45),
  inference(superposition,[],[f311,f169])).
fof(f996,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9) | $spl45),
  inference(superposition,[],[f311,f169])).
fof(f993,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | $spl45),
  inference(superposition,[],[f311,f168])).
fof(f491,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | $spl55),
  inference(cnf_transformation,[],[f491_D])).
fof(f491_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f992,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | $spl45),
  inference(superposition,[],[f311,f168])).
fof(f991,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK10) | $spl45),
  inference(superposition,[],[f311,f169])).
fof(f990,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10) | $spl45),
  inference(superposition,[],[f311,f169])).
fof(f878,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl41),
  inference(superposition,[],[f307,f168])).
fof(f877,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl41),
  inference(superposition,[],[f307,f168])).
fof(f876,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl41),
  inference(superposition,[],[f307,f169])).
fof(f875,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl41),
  inference(superposition,[],[f307,f169])).
fof(f864,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl41),
  inference(superposition,[],[f307,f168])).
fof(f489,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | $spl53),
  inference(cnf_transformation,[],[f489_D])).
fof(f489_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f863,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl41),
  inference(superposition,[],[f307,f168])).
fof(f862,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl41),
  inference(superposition,[],[f307,f169])).
fof(f861,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl41),
  inference(superposition,[],[f307,f169])).
fof(f748,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl37),
  inference(superposition,[],[f303,f168])).
fof(f747,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl37),
  inference(superposition,[],[f303,f168])).
fof(f746,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl37),
  inference(superposition,[],[f303,f169])).
fof(f745,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl37),
  inference(superposition,[],[f303,f169])).
fof(f736,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl37),
  inference(superposition,[],[f303,f168])).
fof(f735,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl37),
  inference(superposition,[],[f303,f168])).
fof(f734,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl37),
  inference(superposition,[],[f303,f169])).
fof(f733,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl37),
  inference(superposition,[],[f303,f169])).
fof(f614,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl33),
  inference(superposition,[],[f299,f168])).
fof(f613,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl33),
  inference(superposition,[],[f299,f168])).
fof(f612,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl33),
  inference(superposition,[],[f299,f169])).
fof(f611,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl33),
  inference(superposition,[],[f299,f169])).
fof(f608,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl33),
  inference(superposition,[],[f299,f168])).
fof(f607,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl33),
  inference(superposition,[],[f299,f168])).
fof(f606,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl33),
  inference(superposition,[],[f299,f169])).
fof(f605,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl33),
  inference(superposition,[],[f299,f169])).
fof(f481,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl29),
  inference(superposition,[],[f295,f168])).
fof(f480,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl29),
  inference(superposition,[],[f295,f168])).
fof(f479,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl29),
  inference(superposition,[],[f295,f169])).
fof(f478,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl29),
  inference(superposition,[],[f295,f169])).
fof(f475,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl29),
  inference(superposition,[],[f295,f168])).
fof(f474,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl29),
  inference(superposition,[],[f295,f168])).
fof(f473,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl29),
  inference(superposition,[],[f295,f169])).
fof(f472,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl29),
  inference(superposition,[],[f295,f169])).
fof(f315,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | $spl49),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f287,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f311,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | $spl45),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f286,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f285,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f284,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f307,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | $spl41),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f281,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f303,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | $spl37),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f280,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f279,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'F_0'))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK9) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f278,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,'T_0'))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK9) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f299,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | $spl33),
  inference(cnf_transformation,[],[f299_D])).
fof(f299_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f275,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f295,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | $spl29),
  inference(cnf_transformation,[],[f295_D])).
fof(f295_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f274,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl1),
  inference(superposition,[],[f204,f168])).
fof(f273,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'T_0') = s(bool,sK10) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f272,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'T_0'))),s(bool,sK9))),s(bool,sK8)))) | s(bool,'F_0') = s(bool,sK10) | $spl1),
  inference(superposition,[],[f204,f169])).
fof(f212,plain,(
  sP0(sK10,sK9,sK8,sK11) | $spl4),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  sP0(sK10,sK9,sK8,sK11) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f239,plain,(
  ~p(s(bool,sK10)) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f238,f122])).
fof(f238,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f237,f132])).
fof(f237,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f210,f127])).
fof(f210,plain,(
  ~p(s(bool,sK10)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f110])).
fof(f234,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f233,f141])).
fof(f233,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f232,f136])).
fof(f232,plain,(
  ~p(s(bool,sK9)) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f209,f121])).
fof(f209,plain,(
  ~p(s(bool,sK9)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f109])).
fof(f229,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f228,f145])).
fof(f228,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f227,f135])).
fof(f227,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(subsumption_resolution,[],[f208,f125])).
fof(f208,plain,(
  ~p(s(bool,sK8)) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f108])).
fof(f207,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) | sP6(sK10,sK9,sK8,sK11) | sP5(sK10,sK9,sK8,sK11) | sP4(sK10,sK9,sK8,sK11) | sP3(sK10,sK9,sK8,sK11) | sP2(sK10,sK9,sK8,sK11) | sP1(sK10,sK9,sK8,sK11) | sP0(sK10,sK9,sK8,sK11) | $spl2),
  inference(resolution,[],[f205,f107])).
fof(f203,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) | $spl0),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f205,plain,(
  sP7(sK10,sK9,sK8,sK11) | $spl2),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  sP7(sK10,sK9,sK8,sK11) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f154,plain,(
  sP7(sK10,sK9,sK8,sK11) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8))))),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) | sP7(sK10,sK9,sK8,sK11)) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) | ~sP7(sK10,sK9,sK8,sK11))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f80])).
fof(f80,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) | sP7(X2,X1,X0,X3)) & (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) | ~sP7(X2,X1,X0,X3)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) <~> sP7(X2,X1,X0,X3))),
  inference(definition_folding,[],[f45,f54,f53,f52,f51,f50,f49,f48,f47])).
fof(f45,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) <~> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,X2))),s(bool,X1))),s(bool,X0)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X3),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X0)) & ~p(s(bool,X1)) & ~p(s(bool,X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,X8))),s(bool,X7))),s(bool,X6)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)) & ~p(s(bool,X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,X8))),s(bool,X7))),s(bool,X6)))) <=> ((p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)) & p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'T_0')))) & p(s(bool,X6)) & ~p(s(bool,X7)) & ~p(s(bool,X8))) | (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),X9),s(bool,'F_0'))),s(bool,'F_0'))),s(bool,'F_0')))) & ~p(s(bool,X6)) & ~p(s(bool,X7)) & ~p(s(bool,X8)))))),
  file('thm.clos_relationProps.clos_relationProps_45_prove_ltb','thm.clos_relationProps.clos_relationProps_45_prove_ltb')).
fof(f204,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) | $spl1),
  inference(cnf_transformation,[],[f204_D])).
fof(f204_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f155,plain,(
  ~sP7(sK10,sK9,sK8,sK11) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),sK11),s(bool,sK10))),s(bool,sK9))),s(bool,sK8))))),
  inference(cnf_transformation,[],[f81])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 513
% Passive clauses: 5285
% Generated clauses: 13827
% Final active clauses: 103
% Final passive clauses: 1476
% Input formulas: 21
% Initial clauses: 93
% 
% Selected by SInE selection: 21
% SInE iterations: 1
% 
% Duplicate literals: 12
% Trivial inequalities: 31
% Fw subsumption resolutions: 4626
% Fw demodulations: 321
% Bw demodulations: 84
% 
% Simple tautologies: 202
% Equational tautologies: 1
% Forward subsumptions: 2674
% 
% Binary resolution: 4648
% Factoring: 18
% Forward superposition: 1609
% Backward superposition: 1368
% 
% Split clauses: 571
% Split components: 1360
% SAT solver clauses: 661
% SAT solver unit clauses: 4
% SAT solver binary clauses: 142
% 
% Sat splits: 571
% Sat splitting refutations: 52
% 
% Memory used [KB]: 4349
% Time elapsed: 0.494 s
% ------------------------------
% ------------------------------
% Success in time 12.26 s
