Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSoundInvariants
% (15419)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 3145
% Generated clauses: 16843
% Final active clauses: 139
% Final passive clauses: 2823
% Input formulas: 69
% Initial clauses: 102
% 
% Duplicate literals: 506
% Trivial inequalities: 280
% Fw subsumption resolutions: 1398
% Fw demodulations: 2368
% Bw demodulations: 17
% 
% Simple tautologies: 248
% Equational tautologies: 42
% Forward subsumptions: 8415
% 
% Binary resolution: 501
% Factoring: 1
% Forward superposition: 6035
% Backward superposition: 5821
% Self superposition: 152
% Equality factoring: 68
% Equality resolution: 26
% 
% Split clauses: 228
% Split components: 458
% SAT solver clauses: 607
% SAT solver unit clauses: 9
% SAT solver binary clauses: 148
% 
% Sat splits: 228
% 
% TWLsolver clauses: 607
% TWLsolver calls for satisfiability: 167
% 
% Memory used [KB]: 18038
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.typeSoundInvariants
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSoundInvariants
% SZS output start Proof for thm.typeSoundInvariants
fof(f9069,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9014,f2682,f298,f382,f8396,f217,f324,f329,f331,f7851,f7853,f7852,f237,f343,f276,f362,f7300,f7519,f216,f322,f277,f363,f8675,f2615,f8727])).
fof(f8727,plain,(
  $false | ($spl2 | $spl426)),
  inference(resolution,[],[f2615,f324])).
fof(f2615,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl426),
  inference(cnf_transformation,[],[f2615_D])).
fof(f2615_D,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f8675,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | ($spl96 | $spl1534)),
  inference(superposition,[],[f7852,f382])).
fof(f363,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X4),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) ) | $spl58),
  inference(cnf_transformation,[],[f363_D])).
fof(f363_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X4),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X4),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.ast.t',X0),X4),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.ast.t',X0),X4),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X16] : ! [X7,X20,X30,X31] : s(X16,happ(s(fun('type.ast.t',X16),X31),s('type.ast.t',X7))) = s(X16,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X7))),s(fun('type.ast.t',X16),X20),s(X16,X30),s(fun('type.ast.t',X16),X31)))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','thm.typeSoundInvariants.store_t_case_def|split|2')).
fof(f322,plain,(
  s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',sK0))) | $spl0),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f216,plain,(
  s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',sK0)))),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f171])).
fof(f171,plain,(
  ? [X0,X1] : s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X0)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))) != s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ~s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ~s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X7))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ~s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X7))) = s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Varray_t_1'(s('type.ast.t',X6)))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','thm.typeSoundInvariants.store_t_distinct|split|1')).
fof(f7519,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X1),s('type.ast.t',sK0))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X1)))) ) | $spl1454),
  inference(cnf_transformation,[],[f7519_D])).
fof(f7519_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X1),s('type.ast.t',sK0))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X1)))) ) <=> ~$spl1454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1454])])).
fof(f7300,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X1),s('type.ast.t',sK0))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X1)))) ) | ($spl0 | $spl58)),
  inference(superposition,[],[f363,f322])).
fof(f362,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X2),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) ) | $spl56),
  inference(cnf_transformation,[],[f362_D])).
fof(f362_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X2),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f276,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.t',X0),X2),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.ast.t',X0),X2),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.ast.t',X0),X2),s('type.ast.t',X1))) = s(X0,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X1))),s(fun('type.ast.t',X0),X2),s(X0,X3),s(fun('type.ast.t',X0),X4)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X16] : ! [X7,X20,X30,X31] : s(X16,happ(s(fun('type.ast.t',X16),X20),s('type.ast.t',X7))) = s(X16,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',X7))),s(fun('type.ast.t',X16),X20),s(X16,X30),s(fun('type.ast.t',X16),X31)))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','thm.typeSoundInvariants.store_t_case_def|split|0')).
fof(f343,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f237,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) <=> p(s(bool,X28)))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f7852,plain,(
  ( ! [X122] : (~p(s(bool,happ(s(fun('type.ast.t',bool),X122),s('type.ast.t',sK0))))) ) | $spl1534),
  inference(cnf_transformation,[],[f7852_D])).
fof(f7852_D,plain,(
  ( ! [X122] : (~p(s(bool,happ(s(fun('type.ast.t',bool),X122),s('type.ast.t',sK0))))) ) <=> ~$spl1534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1534])])).
fof(f7853,plain,(
  ( ! [X120] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ast.t',bool),X120),s('type.ast.t',sK1)))) ) | $spl1536),
  inference(cnf_transformation,[],[f7853_D])).
fof(f7853_D,plain,(
  ( ! [X120] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ast.t',bool),X120),s('type.ast.t',sK1)))) ) <=> ~$spl1536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1536])])).
fof(f7851,plain,(
  ( ! [X122,X120] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ast.t',bool),X120),s('type.ast.t',sK1))) | ~p(s(bool,happ(s(fun('type.ast.t',bool),X122),s('type.ast.t',sK0))))) ) | ($spl26 | $spl56 | $spl1454)),
  inference(forward_demodulation,[],[f7717,f362])).
fof(f7717,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,happ(s(fun('type.ast.t',bool),X122),s('type.ast.t',sK0)))) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.store_t_CASE_4'(s('type.typeSoundInvariants.store_t','const.typeSoundInvariants.Ref_t_1'(s('type.ast.t',sK1))),s(fun('type.ast.t',bool),X120),s(bool,X121),s(fun('type.ast.t',bool),X122)))) ) | ($spl26 | $spl1454)),
  inference(superposition,[],[f343,f7519])).
fof(f331,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f331_D])).
fof(f331_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f329,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f222,f217])).
fof(f222,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f40])).
fof(f40,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f324,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f217,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1',aHL_TRUTH)).
fof(f8396,plain,(
  $false | ($spl2 | $spl9 | $spl428)),
  inference(subsumption_resolution,[],[f7973,f324])).
fof(f7973,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl428)),
  inference(backward_demodulation,[],[f2682,f331])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) ) | $spl96),
  inference(cnf_transformation,[],[f382_D])).
fof(f382_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f298,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X10,X8] : ! [X11,X13] : s(X8,X11) = s(X8,happ(s(fun(X10,X8),'_dst_x0_1'(s(X8,X11))),s(X10,X13)))),
  file('thm.typeSoundInvariants.store_t_distinct|split|1','_dst_x0_ax')).
fof(f2682,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl428),
  inference(cnf_transformation,[],[f2682_D])).
fof(f2682_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f9014,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,X16)) ) | ($spl96 | $spl1536)),
  inference(superposition,[],[f7853,f382])).
% SZS output end Proof for thm.typeSoundInvariants
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 115
% Passive clauses: 1484
% Generated clauses: 9629
% Final active clauses: 68
% Final passive clauses: 511
% Input formulas: 69
% Initial clauses: 102
% 
% Duplicate literals: 222
% Trivial inequalities: 185
% Fw subsumption resolutions: 565
% Fw demodulations: 1105
% Bw demodulations: 463
% Condensations: 524
% Global subsumptions: 8
% 
% Simple tautologies: 106
% Equational tautologies: 35
% Forward subsumptions: 3850
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 426
% 
% Binary resolution: 177
% Factoring: 1
% Forward superposition: 2354
% Backward superposition: 2070
% Self superposition: 82
% Equality factoring: 52
% Equality resolution: 18
% 
% Split clauses: 73
% Split components: 146
% SAT solver clauses: 1515
% SAT solver unit clauses: 152
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 17
% 
% Sat splits: 73
% Sat splitting refutations: 23
% 
% TWLsolver clauses: 1481
% TWLsolver calls for satisfiability: 377
% 
% Memory used [KB]: 9978
% Time elapsed: 0.497 s
% ------------------------------
% ------------------------------
% Success in time 1.14 s
