Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f2355,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2354,f291,f442,f440,f451,f292,f443,f441])).
fof(f441,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) | $spl1),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f443,plain,(
  s('type.ast.word_size','const.ast.W64_0') != s('type.ast.word_size',sK0) | $spl3),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  s('type.ast.word_size','const.ast.W64_0') != s('type.ast.word_size',sK0) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f292,plain,(
  s('type.ast.word_size','const.ast.W64_0') != s('type.ast.word_size',sK0) | s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0)))),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  (s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',sK0)) & (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) | s('type.ast.word_size','const.ast.W64_0') != s('type.ast.word_size',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f262])).
fof(f262,plain,(
  ? [X0] : ((s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X0)) & (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) | s('type.ast.word_size','const.ast.W64_0') != s('type.ast.word_size',X0)))),
  inference(nnf_transformation,[],[f257])).
fof(f257,plain,(
  ? [X0] : (s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) <~> s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X0))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ~! [X0] : (s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) <=> s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X6))) <=> s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X6))) <=> s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X6))),
  file('thm.evalProps.Tword_simp|split|15','thm.evalProps.Tword_simp|split|15')).
fof(f451,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f297,f450])).
fof(f450,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f441,f442])).
fof(f297,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0')))),
  inference(cnf_transformation,[],[f101])).
fof(f101,axiom,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0')))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.Tword64_def')).
fof(f440,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) | $spl0),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f442,plain,(
  s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',sK0) | $spl2),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',sK0) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f291,plain,(
  s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',sK0) | s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',sK0)))),
  inference(cnf_transformation,[],[f263])).
fof(f2354,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2352,f352])).
fof(f352,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor','const.ast.TC_word64_0')),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor','const.ast.TC_word64_0')),
  inference(flattening,[],[f60])).
fof(f60,axiom,(
  ~s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word64_0')),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.tctor_distinct|split|50')).
fof(f2352,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word64_0') | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f2349])).
fof(f2349,plain,(
  s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tword64_0') | s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word64_0') | ($spl0 | $spl3)),
  inference(superposition,[],[f1832,f575])).
fof(f575,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0')))),
  inference(forward_demodulation,[],[f569,f297])).
fof(f569,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W64_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word64_0')))),
  inference(superposition,[],[f373,f296])).
fof(f296,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W64_0')))),
  inference(cnf_transformation,[],[f98])).
fof(f98,axiom,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W64_0')))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.TC_word_def|split|1')).
fof(f373,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X17] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X17))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X17)))))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.Tword_def')).
fof(f1832,plain,(
  ( ! [X4,X5] : (s('type.ast.t','const.ast.Tword64_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X4),s('type.ast.tctor',X5))) | s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X5)) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f436,f574])).
fof(f574,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f568,f470])).
fof(f470,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f466,f443])).
fof(f466,plain,(
  s('type.ast.t','const.ast.Tword64_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',sK0) | $spl0),
  inference(superposition,[],[f440,f374])).
fof(f374,plain,(
  ( ! [X0] : (s('type.ast.word_size','const.ast.W8_0') = s('type.ast.word_size',X0) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : (s('type.ast.word_size','const.ast.W8_0') = s('type.ast.word_size',X0) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X0))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : (s('type.ast.word_size','const.ast.W8_0') = s('type.ast.word_size',X9) | s('type.ast.word_size','const.ast.W64_0') = s('type.ast.word_size',X9))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.word_size_nchotomy')).
fof(f568,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0')))),
  inference(superposition,[],[f373,f295])).
fof(f295,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W8_0')))),
  inference(cnf_transformation,[],[f97])).
fof(f97,axiom,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W8_0')))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.TC_word_def|split|0')).
fof(f436,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | s('type.ast.tctor',X1) = s('type.ast.tctor',X3)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & (s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & ((s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3)) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) <=> (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X13,X14,X15,X16] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X13),s('type.ast.tctor',X14))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X15),s('type.ast.tctor',X16))) <=> (s('type.list.list'('type.ast.t'),X13) = s('type.list.list'('type.ast.t'),X15) & s('type.ast.tctor',X14) = s('type.ast.tctor',X16)))),
  file('thm.evalProps.Tword_simp|split|15','thm.ast.t_|49||49||split|2')).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 169
% Passive clauses: 414
% Generated clauses: 1949
% Final active clauses: 146
% Final passive clauses: 192
% Input formulas: 120
% Initial clauses: 147
% 
% Duplicate literals: 208
% Trivial inequalities: 33
% Fw subsumption resolutions: 139
% Fw demodulations: 43
% Bw demodulations: 16
% 
% Simple tautologies: 96
% Equational tautologies: 52
% Forward subsumptions: 911
% 
% Binary resolution: 35
% Factoring: 1
% Forward superposition: 528
% Backward superposition: 582
% Self superposition: 84
% Equality factoring: 86
% Equality resolution: 12
% 
% Split clauses: 51
% Split components: 102
% SAT solver clauses: 95
% SAT solver unit clauses: 18
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 51
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 93
% TWLsolver calls for satisfiability: 22
% 
% Memory used [KB]: 1663
% Time elapsed: 0.072 s
% ------------------------------
% ------------------------------
% Success in time 0.09 s
