Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.exh_to_patProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exh_to_patProof
% SZS output start Proof for thm.exh_to_patProof
fof(f6230,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6227,f190,f322])).
fof(f322,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f190,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.exh_to_patProof.exh_to_patProof_74_prove_ltb',aHL_FALSITY)).
fof(f6227,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f6226,f322])).
fof(f6226,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f6197,f3583])).
fof(f3583,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10)))))),
  inference(resolution,[],[f183,f202])).
fof(f202,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.exh_to_patProof.exh_to_patProof_74_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f183,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10))))))),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK12)))))) & p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK11)))))) & p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK11))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10)))))) & ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f81])).
fof(f81,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2)))))) & p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1)))))) & p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))) & ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2)))))) & (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1)))))) & (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))) & ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X1))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))) => p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X2))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X3),s('type.list.list'('type.patSem.v'),X0)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X8))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X7)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X7))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X6)))))) => p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X8))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X6)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X8))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X7)))))) => (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X7))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X6)))))) => p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X8))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',X9),s('type.list.list'('type.patSem.v'),X6)))))))))),
  file('thm.exh_to_patProof.exh_to_patProof_74_prove_ltb','thm.exh_to_patProof.exh_to_patProof_74_prove_ltb')).
fof(f6197,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10))))))),
  inference(resolution,[],[f3631,f182])).
fof(f182,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK11))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK10))))))),
  inference(cnf_transformation,[],[f117])).
fof(f3631,plain,(
  ( ! [X0] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK11))),s('type.patSem.v',X0)))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v',X0))))) )),
  inference(resolution,[],[f230,f181])).
fof(f181,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK12))),s('type.patSem.v','const.list.EL_2'(s('type.num.num',sK13),s('type.list.list'('type.patSem.v'),sK11))))))),
  inference(cnf_transformation,[],[f117])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X1),s('type.patSem.v',X2)))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X2))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | ! [X2] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X1),s('type.patSem.v',X2)))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X2))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) => ! [X2] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X1),s('type.patSem.v',X2)))) => p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X2))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X16,X17] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X16),s('type.patSem.v',X17)))) => ! [X18] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X17),s('type.patSem.v',X18)))) => p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X16),s('type.patSem.v',X18))))))),
  file('thm.exh_to_patProof.exh_to_patProof_74_prove_ltb','thm.exh_to_patProof.v_rel_trans')).
% SZS output end Proof for thm.exh_to_patProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 415
% Passive clauses: 1746
% Generated clauses: 6014
% Final active clauses: 413
% Final passive clauses: 1120
% Input formulas: 32
% Initial clauses: 136
% 
% Duplicate literals: 194
% Trivial inequalities: 52
% Fw subsumption resolutions: 630
% Fw demodulations: 26
% Bw demodulations: 2
% 
% Simple tautologies: 142
% Equational tautologies: 54
% Forward subsumptions: 3039
% 
% Binary resolution: 3005
% Factoring: 838
% Forward superposition: 258
% Backward superposition: 748
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 193
% Split components: 387
% SAT solver clauses: 350
% SAT solver unit clauses: 14
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 1
% 
% Sat splits: 193
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 349
% TWLsolver calls for satisfiability: 214
% 
% Memory used [KB]: 2942
% Time elapsed: 0.209 s
% ------------------------------
% ------------------------------
% Success in time 0.202 s
