Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.PairingHeap
% (71871)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 739
% Passive clauses: 4783
% Generated clauses: 22241
% Final active clauses: 596
% Final passive clauses: 3667
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 137
% Trivial inequalities: 41
% Fw subsumption resolutions: 1224
% Fw demodulations: 2106
% Bw demodulations: 135
% 
% Simple tautologies: 266
% Equational tautologies: 23
% Forward subsumptions: 13668
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 4133
% Factoring: 342
% Forward superposition: 6724
% Backward superposition: 6731
% Self superposition: 47
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 230
% Split components: 462
% SAT solver clauses: 360
% SAT solver unit clauses: 15
% SAT solver binary clauses: 74
% 
% Sat splits: 230
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 360
% TWLsolver calls for satisfiability: 239
% 
% Memory used [KB]: 12281
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.PairingHeap
% (71950)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 566
% Passive clauses: 2430
% Generated clauses: 21631
% Final active clauses: 467
% Final passive clauses: 1664
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 104
% Trivial inequalities: 25
% Fw subsumption resolutions: 689
% Fw demodulations: 2534
% Bw demodulations: 264
% Condensations: 94
% Global subsumptions: 24
% 
% Simple tautologies: 249
% Equational tautologies: 9
% Forward subsumptions: 12607
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 2154
% Factoring: 236
% Forward superposition: 6832
% Backward superposition: 5617
% Self superposition: 73
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 3387
% SAT solver unit clauses: 645
% SAT solver binary clauses: 126
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 15
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3291
% TWLsolver calls for satisfiability: 3255
% 
% Memory used [KB]: 27121
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.PairingHeap
% (72112)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1028
% Passive clauses: 8494
% Generated clauses: 89432
% Final active clauses: 819
% Final passive clauses: 4591
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 9
% Trivial inequalities: 58
% Fw subsumption resolutions: 1467
% Fw demodulations: 26707
% Bw demodulations: 1741
% Condensations: 31
% 
% Simple tautologies: 115
% Equational tautologies: 46
% Forward subsumptions: 50437
% Fw demodulations to eq. taut.: 351
% Bw demodulations to eq. taut.: 42
% 
% Unit resulting resolution: 23071
% Factoring: 1
% Forward superposition: 16845
% Backward superposition: 17905
% Self superposition: 195
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 10238
% SAT solver unit clauses: 3501
% SAT solver binary clauses: 3981
% 
% Sat splits: 25
% Sat splitting refutations: 1692
% 
% Memory used [KB]: 226179
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 25 next slice time: 25
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.PairingHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.PairingHeap
% SZS output start Proof for thm.PairingHeap
fof(f1012,plain,(
  $false),
  inference(subsumption_resolution,[],[f1010,f779])).
fof(f779,plain,(
  s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK14),sK17)),
  inference(subsumption_resolution,[],[f535,f768])).
fof(f768,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4) | ~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,sK79(X4,X3,X2,X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X1)),sK80(X4,X3,X2,X1,X0)))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK79,sK80])],[f533])).
fof(f533,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X7,X8] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X7),s('type.list.list'('type.PairingHeap.heap'(X1)),X8))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(rectify,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & (! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(flattening,[],[f531])).
fof(f531,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))) & ((! [X5,X6] : s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4)) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4))))))),
  inference(nnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4)))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> (? [X5,X6] : s('type.PairingHeap.heap'(X1),X4) = s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6))) & ~s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X12,X13] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X7,X6),X11),s(fun(X6,fun(X6,bool)),X12),s('type.PairingHeap.heap'(X7),X13)))) <=> (? [X14,X15] : s('type.PairingHeap.heap'(X7),X13) = s('type.PairingHeap.heap'(X7),'const.PairingHeap.Tree_2'(s(X7,X14),s('type.list.list'('type.PairingHeap.heap'(X7)),X15))) & ~s('type.PairingHeap.heap'(X7),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X7),X13)))),
  file('thm.PairingHeap.PairingHeap_95_prove_ltb','thm.PairingHeap.delete_min_side_def')).
fof(f535,plain,(
  s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK14),sK17) | p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.PairingHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.PairingHeap.heap'(sK14),sK17)))) | s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(sK14),sK17)) & (~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.PairingHeap.heap'(sK14),sK17)))) | s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK14),sK17))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f419])).
fof(f419,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4)) & (~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4)))),
  inference(nnf_transformation,[],[f352])).
fof(f352,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <~> s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') != s('type.PairingHeap.heap'(X1),X4))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) <=> ~s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.PairingHeap.heap'(X7),X10)))) <=> ~s('type.PairingHeap.heap'(X7),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X7),X10))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.PairingHeap.heap'(X7),X10)))) <=> ~s('type.PairingHeap.heap'(X7),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X7),X10))),
  file('thm.PairingHeap.PairingHeap_95_prove_ltb','thm.PairingHeap.PairingHeap_95_prove_ltb')).
fof(f1010,plain,(
  s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK14),sK17)),
  inference(equality_resolution,[],[f994])).
fof(f994,plain,(
  ( ! [X3] : (s('type.PairingHeap.heap'(sK14),sK17) != s('type.PairingHeap.heap'(sK14),X3) | s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK14),X3)) )),
  inference(superposition,[],[f915,f593])).
fof(f593,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK71(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK72(X1,X0)))) | s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,sK71(X1,X0)),s('type.list.list'('type.PairingHeap.heap'(X0)),sK72(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK71,sK72])],[f294])).
fof(f294,plain,(
  ! [X0,X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0] : ! [X1] : (s('type.PairingHeap.heap'(X0),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X0),X1) | ? [X2,X3] : s('type.PairingHeap.heap'(X0),X1) = s('type.PairingHeap.heap'(X0),'const.PairingHeap.Tree_2'(s(X0,X2),s('type.list.list'('type.PairingHeap.heap'(X0)),X3))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ! [X26] : (s('type.PairingHeap.heap'(X7),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X7),X26) | ? [X27,X28] : s('type.PairingHeap.heap'(X7),X26) = s('type.PairingHeap.heap'(X7),'const.PairingHeap.Tree_2'(s(X7,X27),s('type.list.list'('type.PairingHeap.heap'(X7)),X28))))),
  file('thm.PairingHeap.PairingHeap_95_prove_ltb','thm.PairingHeap.heap_nchotomy')).
fof(f915,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK14),sK17) != s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Tree_2'(s(sK14,X0),s('type.list.list'('type.PairingHeap.heap'(sK14)),X1)))) )),
  inference(subsumption_resolution,[],[f790,f779])).
fof(f790,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK14),sK17) | s('type.PairingHeap.heap'(sK14),sK17) != s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Tree_2'(s(sK14,X0),s('type.list.list'('type.PairingHeap.heap'(sK14)),X1)))) )),
  inference(resolution,[],[f778,f769])).
fof(f769,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.PairingHeap.heap'(X1),X4)))) | s('type.PairingHeap.heap'(X1),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(X1),X4) | s('type.PairingHeap.heap'(X1),X4) != s('type.PairingHeap.heap'(X1),'const.PairingHeap.Tree_2'(s(X1,X5),s('type.list.list'('type.PairingHeap.heap'(X1)),X6)))) )),
  inference(cnf_transformation,[],[f534])).
fof(f778,plain,(
  ~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.PairingHeap.heap'(sK14),sK17))))),
  inference(subsumption_resolution,[],[f536,f768])).
fof(f536,plain,(
  s('type.PairingHeap.heap'(sK14),'const.PairingHeap.Empty_0') = s('type.PairingHeap.heap'(sK14),sK17) | ~p(s(bool,'const.PairingHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.PairingHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f420])).
% SZS output end Proof for thm.PairingHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 154
% Passive clauses: 121
% Generated clauses: 227
% Final active clauses: 154
% Final passive clauses: 115
% Input formulas: 261
% Initial clauses: 235
% 
% Selected by SInE selection: 53
% SInE iterations: 1
% 
% Duplicate literals: 19
% Trivial inequalities: 3
% Fw subsumption resolutions: 22
% 
% Simple tautologies: 96
% Equational tautologies: 2
% Forward subsumptions: 43
% 
% Binary resolution: 103
% Forward superposition: 79
% Equality resolution: 2
% 
% Split clauses: 13
% Split components: 29
% SAT solver clauses: 14
% SAT solver binary clauses: 10
% 
% Sat splits: 13
% 
% Memory used [KB]: 2046
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
% Success in time 12.455 s
