Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% (125526)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 178
% Passive clauses: 1041
% Generated clauses: 14422
% Final active clauses: 162
% Final passive clauses: 707
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 679
% Trivial inequalities: 288
% Fw subsumption resolutions: 2462
% Fw demodulations: 2393
% Bw demodulations: 10
% 
% Simple tautologies: 355
% Equational tautologies: 23
% Forward subsumptions: 6167
% Fw demodulations to eq. taut.: 26
% 
% Binary resolution: 195
% Forward superposition: 4497
% Backward superposition: 3467
% Self superposition: 109
% Equality factoring: 55
% Equality resolution: 50
% 
% Split clauses: 502
% Split components: 1016
% SAT solver clauses: 1140
% SAT solver unit clauses: 7
% SAT solver binary clauses: 77
% 
% Sat splits: 502
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1140
% TWLsolver calls for satisfiability: 176
% 
% Memory used [KB]: 6652
% Time elapsed: 0.629 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_knownProof
% (125547)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 196
% Passive clauses: 595
% Generated clauses: 18445
% Final active clauses: 194
% Final passive clauses: 387
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 1203
% Trivial inequalities: 342
% Fw subsumption resolutions: 1970
% Fw demodulations: 3031
% Bw demodulations: 13
% Condensations: 1106
% 
% Simple tautologies: 338
% Equational tautologies: 19
% Forward subsumptions: 8893
% Fw demodulations to eq. taut.: 32
% 
% Binary resolution: 69
% Forward superposition: 6316
% Backward superposition: 3572
% Self superposition: 89
% Equality factoring: 34
% Equality resolution: 51
% 
% Split clauses: 68
% Split components: 145
% SAT solver clauses: 879
% SAT solver unit clauses: 72
% SAT solver binary clauses: 18
% 
% Sat splits: 68
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 879
% TWLsolver calls for satisfiability: 299
% 
% Memory used [KB]: 11513
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_knownProof
% (125596)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 382
% Passive clauses: 2720
% Generated clauses: 142704
% Final active clauses: 199
% Final passive clauses: 625
% Input formulas: 36
% Initial clauses: 53
% 
% Duplicate literals: 7733
% Trivial inequalities: 918
% Fw subsumption resolutions: 14773
% Fw demodulations: 9291
% Bw demodulations: 1390
% Condensations: 20169
% Global subsumptions: 765
% 
% Simple tautologies: 1294
% Equational tautologies: 1371
% Forward subsumptions: 82400
% Fw demodulations to eq. taut.: 113
% Bw demodulations to eq. taut.: 77
% 
% Unit resulting resolution: 5265
% Forward superposition: 35601
% Backward superposition: 33898
% Self superposition: 932
% Equality factoring: 493
% Equality resolution: 94
% 
% Split clauses: 238
% Split components: 496
% SAT solver clauses: 5036
% SAT solver unit clauses: 118
% SAT solver binary clauses: 193
% 
% Sat splits: 238
% Sat splitting refutations: 1073
% 
% Memory used [KB]: 69593
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f3843,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f121,f122,f141,f198,f202,f204,f207,f212,f214,f218,f221,f227,f230,f231,f236,f238,f242,f246,f247,f303,f305,f307,f311,f313,f317,f322,f324,f325,f328,f199,f331,f335,f345,f364,f376,f377,f394,f395,f410,f409,f408,f431,f430,f429,f447,f457,f458,f461,f474,f475,f476,f477,f478,f479,f498,f481,f482,f496,f483,f205,f511,f512,f514,f200,f617,f619,f620,f621,f624,f625,f658,f660,f661,f662,f665,f666,f709,f711,f712,f713,f716,f717,f731,f732,f763,f765,f760,f766,f910,f911,f913,f915,f916,f917,f918,f920,f922,f924,f925,f926,f930,f932,f934,f935,f937,f939,f941,f942,f945,f946,f948,f949,f952,f965,f967,f968,f969,f972,f973,f975,f977,f978,f980,f982,f983,f1011,f1013,f1014,f1015,f1018,f1019,f1031,f1032,f1046,f1048,f1049,f1050,f1053,f1054,f1089,f1090,f1087,f1091,f1095,f1097,f1098,f1099,f1102,f1103,f1107,f1109,f1110,f1112,f1114,f1115,f1162,f1164,f1165,f1166,f1169,f1170,f1183,f1184,f1188,f1190,f1191,f1193,f1195,f1196,f1267,f1269,f1270,f1271,f1274,f1275,f1289,f1290,f1334,f1336,f1337,f1338,f1341,f1342,f1356,f1357,f1382,f1383,f1422,f1424,f1425,f1426,f1429,f1430,f1436,f1563,f1564,f1565,f1570,f1572,f1574,f1566,f1587,f1588,f1592,f1366,f1595,f1596,f1597,f1299,f1598,f1599,f1600,f1240,f1226,f1602,f1603,f1072,f1041,f1606,f1607,f741,f1609,f1611,f1612,f1591,f1709,f1241,f1239,f1073,f1071,f123,f1765,f1767,f1768,f1770,f1772,f1773,f1782,f1783,f1949,f1957,f1958,f2019,f2021,f2022,f2024,f2026,f2027,f304,f2042,f2044,f2045,f2046,f2049,f2050,f638,f2058,f126,f124,f2059,f130,f128,f2070,f2069,f2177,f2178,f2179,f2181,f2182,f2183,f2180,f2185,f2187,f2188,f2189,f2192,f2193,f678,f2215,f2217,f2218,f2219,f2222,f2223,f729,f2345,f1608,f2346,f1605,f2347,f1604,f1601,f2363,f2364,f2366,f2367,f2368,f2369,f2370,f2371,f2373,f2374,f2375,f2378,f2379,f994,f2381,f2383,f2384,f2386,f2388,f2389,f997,f2419,f2421,f2422,f2423,f2426,f2427,f1029,f2439,f2441,f2442,f2443,f2446,f2447,f1061,f2466,f2468,f2469,f2470,f2473,f2474,f1181,f2495,f2497,f2498,f2499,f2502,f2503,f1131,f2507,f2509,f2510,f2512,f2514,f2515,f1135,f2543,f2546,f2556,f2557,f2560,f2562,f2563,f2565,f2567,f2568,f1213,f2587,f2589,f2590,f2591,f2594,f2595,f1287,f2636,f2545,f2637,f497,f2548,f2638,f2640,f2641,f2642,f2645,f2646,f1354,f2687,f1594,f2707,f2709,f2710,f2711,f2714,f2715,f2720,f2724,f2726,f2727,f2728,f2731,f2732,f2737,f2829,f2831,f2834,f2837,f2838,f2839,f2840,f2842,f2845,f2848,f2849,f2850,f2851,f2854,f2855,f2856,f2858,f2830,f2859,f2860,f2861,f2857,f2864,f2866,f2867,f2868,f2871,f2872,f2832,f2886,f2887,f2910,f2911,f2933,f2934,f2976,f2978,f2979,f2980,f2983,f2984,f2997,f2999,f3000,f3018,f3020,f3021,f3022,f3025,f3026,f3032,f3039,f3041,f3042,f3044,f3046,f3047,f3061,f3067,f3069,f3070,f3071,f3074,f3075,f3081,f3092,f3093,f3218,f3220,f3223,f3226,f3227,f3228,f3229,f3231,f3234,f3237,f3238,f3239,f3240,f3243,f3245,f3247,f1589,f3248,f3249,f3250,f3252,f3253,f3254,f3255,f3256,f3258,f3244,f3259,f3260,f3219,f3268,f3270,f3271,f3272,f3275,f3276,f3221,f3296,f3302,f3303,f3297,f3319,f3321,f3322,f3323,f3326,f3327,f3333,f3347,f3348,f3353,f306,f3354,f3374,f3375,f499,f3424,f3425,f3426,f3427,f3428,f3429,f3430,f3417,f3433,f3419,f3435,f3420,f3431,f3436,f3454,f3455,f206,f3456,f3457,f3459,f3460,f3462,f3463,f3464,f3466,f3449,f3467,f3451,f3468,f3452,f3465,f3469,f3512,f3246,f3527,f3530,f3531,f3528,f3304,f3558,f3560,f3561,f3562,f3565,f3566,f3571,f3576,f3578,f3579,f3580,f3583,f3584,f3591,f3598,f3600,f3601,f3602,f3605,f3606,f3621,f3623,f3624,f3648,f3650,f3651,f203,f3652,f228,f3655,f215,f3656,f229,f3662,f3810,f3812,f764,f3816,f1571,f3817,f1573,f3818,f3826,f3811,f3827,f3828,f1575,f3836,f3257,f3813,f3842])).
fof(f3842,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3841,f118])).
fof(f118,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd',aHL_TRUTH)).
fof(f3841,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3840,f125])).
fof(f125,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl2),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3840,plain,(
  ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3839,f2455])).
fof(f2455,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | $spl0),
  inference(forward_demodulation,[],[f2454,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X1),X2))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X1),X2))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X1),X2))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.closSem.state'(X1),X2)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X19,X18] : ! [X8,X10] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X18),X8))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X19),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X18),'type.ffi.ffi_state'(X19)),X10),s('type.closSem.state'(X18),X8)))))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.closSem.closSem_64_prove_ltb')).
fof(f2454,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | $spl0),
  inference(subsumption_resolution,[],[f2450,f116])).
fof(f116,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd',aHL_FALSITY)).
fof(f2450,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | $spl0),
  inference(superposition,[],[f105,f124])).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X6,X7] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X7,X6,X2,X1,X0)))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X6,X2,X1,X0)),s('type.closSem.v',X7))))))) & ((s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',sK7(X2,X1,X0))))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X6,X7] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X8] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X8))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X8),s('type.closSem.v',X7))))))) & (? [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(rectify,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4))))))) & (? [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) & ! [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) => ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) <=> ! [X3,X4] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X4))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) => ? [X5] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',X4)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X8,X9] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X8),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) <=> ! [X20,X21] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X21))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X20),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),X8))))) => ? [X22] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X22))) = s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X20),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) & p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X22),s('type.closSem.v',X21)))))))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.clos_knownProof.state_globals_approx_def')).
fof(f3839,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f3838])).
fof(f3838,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f3661,f103])).
fof(f103,plain,(
  ( ! [X6,X2,X0,X7,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X6),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X7,X6,X2,X1,X0)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f3661,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3660,f116])).
fof(f3660,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3647,f124])).
fof(f3647,plain,(
  p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)))) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f106])).
fof(f106,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X5),s('type.closSem.v',sK7(X2,X1,X0))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X5))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f3646,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) | ($spl0 | $spl2)),
  inference(equality_resolution,[],[f2696])).
fof(f2696,plain,(
  ( ! [X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X1))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X1,sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',X1))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f2536,f2455])).
fof(f2536,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X1),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X0,X1,sK3,sK2,sK1)),s('type.closSem.v',X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2527,f120])).
fof(f120,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f113])).
fof(f113,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.bool.REFL_CLAUSE')).
fof(f2527,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X0,X1,sK3,sK2,sK1)),s('type.closSem.v',X0)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X1),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2)))))) ) | $spl2),
  inference(superposition,[],[f104,f125])).
fof(f104,plain,(
  ( ! [X6,X2,X0,X7,X1] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X6,X2,X1,X0)),s('type.closSem.v',X7)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X6),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f3813,plain,(
  ( ! [X132,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X132) | s(bool,X131) != s(bool,X132)) ) | $spl230),
  inference(cnf_transformation,[],[f3813_D])).
fof(f3813_D,plain,(
  ( ! [X132,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X132) | s(bool,X131) != s(bool,X132)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f3257,plain,(
  ( ! [X596] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X596)))) ) | $spl172),
  inference(cnf_transformation,[],[f3257_D])).
fof(f3257_D,plain,(
  ( ! [X596] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X596)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f3836,plain,(
  ( ! [X280,X279,X278] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X278))) | s(bool,'T_0') = s(bool,X279) | s(bool,X280) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X279) != s(bool,X280)) ) | ($spl0 | $spl2 | $spl24 | $spl190)),
  inference(subsumption_resolution,[],[f3766,f2733])).
fof(f2733,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X41,sK2,sK1),sK6(X41,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X41,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X41)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f472])).
fof(f472,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(trivial_inequality_removal,[],[f466])).
fof(f466,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f304,f115])).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd',aHL_BOOL_CASES)).
fof(f2716,plain,(
  ( ! [X36] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X36,sK2,sK1),sK6(X36,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X36,sK2,sK1))))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2705,f472])).
fof(f2705,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X0,sK2,sK1),sK6(X0,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X0,sK2,sK1))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) ) | $spl2),
  inference(equality_resolution,[],[f2550])).
fof(f2550,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X1))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X0,sK2,sK1)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X1,sK6(X0,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) ) | $spl2),
  inference(superposition,[],[f2536,f105])).
fof(f3766,plain,(
  ( ! [X280,X279,X278] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X278,sK2,sK1),sK6(X278,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X278,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X278))) | s(bool,'T_0') = s(bool,X279) | s(bool,X280) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X279) != s(bool,X280)) ) | ($spl0 | $spl2 | $spl24 | $spl190)),
  inference(superposition,[],[f3427,f3657])).
fof(f3657,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | ($spl0 | $spl2 | $spl24)),
  inference(resolution,[],[f3646,f472])).
fof(f1575,plain,(
  ( ! [X198,X196,X194,X197] : (s(bool,X196) != s(bool,X197) | s(bool,X197) != s(bool,X198) | s(bool,'F_0') = s(bool,X198) | s(bool,X194) != s(bool,X196)) ) | $spl94),
  inference(cnf_transformation,[],[f1575_D])).
fof(f1575_D,plain,(
  ( ! [X198,X196,X194,X197] : (s(bool,X196) != s(bool,X197) | s(bool,X197) != s(bool,X198) | s(bool,'F_0') = s(bool,X198) | s(bool,X194) != s(bool,X196)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3828,plain,(
  ( ! [X227,X225,X226] : (s(bool,X225) != s(bool,X226) | s(bool,'T_0') = s(bool,X227) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X225) != s(bool,X227)) ) | $spl232),
  inference(cnf_transformation,[],[f3828_D])).
fof(f3828_D,plain,(
  ( ! [X227,X225,X226] : (s(bool,X225) != s(bool,X226) | s(bool,'T_0') = s(bool,X227) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X225) != s(bool,X227)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f3827,plain,(
  ( ! [X229,X227,X225,X230,X228,X226,X224] : (s(bool,X225) != s(bool,X226) | s(bool,X225) != s(bool,X227) | s(bool,X224) != s(bool,X228) | s(bool,X229) != s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,X228) != s(bool,X229) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X227)) ) | ($spl0 | $spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3755,f1020])).
fof(f1020,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) != s(bool,X65) | s(bool,X62) != s(bool,X64) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,X63)) ) | ($spl24 | $spl40)),
  inference(resolution,[],[f729,f472])).
fof(f3755,plain,(
  ( ! [X229,X227,X225,X230,X228,X226,X224] : (s(bool,'T_0') != s(bool,X224) | s(bool,X225) != s(bool,X226) | s(bool,X225) != s(bool,X227) | s(bool,X224) != s(bool,X228) | s(bool,X229) != s(bool,X230) | s(bool,'F_0') = s(bool,X230) | s(bool,X228) != s(bool,X229) | s(bool,X226) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X227)) ) | ($spl0 | $spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f2985,f3657])).
fof(f2985,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X98,X104] : (s(bool,X101) != s(bool,X102) | s(bool,X98) != s(bool,X100) | s(bool,X98) != s(bool,X99) | s(bool,X102) != s(bool,X103) | s(bool,X104) != s(bool,X105) | s(bool,'F_0') = s(bool,X105) | s(bool,X103) != s(bool,X104) | s(bool,X100) != s(bool,X101) | s(bool,'T_0') = s(bool,X99)) ) | ($spl24 | $spl78)),
  inference(resolution,[],[f1354,f472])).
fof(f3811,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X122)) ) | $spl228),
  inference(cnf_transformation,[],[f3811_D])).
fof(f3811_D,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'T_0') = s(bool,X122)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f3826,plain,(
  ( ! [X222,X223] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X223) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2 | $spl24 | $spl140)),
  inference(subsumption_resolution,[],[f3754,f2733])).
fof(f3754,plain,(
  ( ! [X222,X223] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X222,sK2,sK1),sK6(X222,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X222,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X223) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2 | $spl24 | $spl140)),
  inference(superposition,[],[f2857,f3657])).
fof(f3818,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X183) != s(bool,X185) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X187) | s(bool,X186) != s(bool,X187) | s(bool,X184) != s(bool,X188) | s(bool,'T_0') = s(bool,X188)) ) | ($spl0 | $spl2 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f3739,f1020])).
fof(f3739,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'T_0') != s(bool,X183) | s(bool,X184) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X183) != s(bool,X185) | s(bool,X185) != s(bool,X186) | s(bool,'F_0') = s(bool,X187) | s(bool,X186) != s(bool,X187) | s(bool,X184) != s(bool,X188) | s(bool,'T_0') = s(bool,X188)) ) | ($spl0 | $spl2 | $spl24 | $spl74)),
  inference(superposition,[],[f1343,f3657])).
fof(f1343,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (s(bool,X91) != s(bool,X92) | s(bool,X89) != s(bool,X91) | s(bool,X92) != s(bool,X93) | s(bool,X93) != s(bool,X94) | s(bool,'F_0') = s(bool,X95) | s(bool,X94) != s(bool,X95) | s(bool,X89) != s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | ($spl24 | $spl74)),
  inference(resolution,[],[f1287,f472])).
fof(f1573,plain,(
  ( ! [X173,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'F_0') = s(bool,X173) | s(bool,X172) != s(bool,X173)) ) | $spl92),
  inference(cnf_transformation,[],[f1573_D])).
fof(f1573_D,plain,(
  ( ! [X173,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'F_0') = s(bool,X173) | s(bool,X172) != s(bool,X173)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3817,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,X164) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X163) != s(bool,X165) | s(bool,X165) != s(bool,X166) | s(bool,'F_0') = s(bool,X166) | s(bool,X164) != s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | ($spl0 | $spl2 | $spl24 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f3735,f718])).
fof(f718,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) != s(bool,X55) | s(bool,X53) != s(bool,X54) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X54)) ) | ($spl24 | $spl38)),
  inference(resolution,[],[f678,f472])).
fof(f3735,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,'T_0') != s(bool,X163) | s(bool,X164) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X163) != s(bool,X165) | s(bool,X165) != s(bool,X166) | s(bool,'F_0') = s(bool,X166) | s(bool,X164) != s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | ($spl0 | $spl2 | $spl24 | $spl68)),
  inference(superposition,[],[f1276,f3657])).
fof(f1276,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X82) != s(bool,X83) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | ($spl24 | $spl68)),
  inference(resolution,[],[f1181,f472])).
fof(f1571,plain,(
  ( ! [X151,X153] : (s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153)) ) | $spl90),
  inference(cnf_transformation,[],[f1571_D])).
fof(f1571_D,plain,(
  ( ! [X151,X153] : (s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3816,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X148) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X147) != s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) != s(bool,X150) | s(bool,'T_0') = s(bool,X150)) ) | ($spl0 | $spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f3731,f655])).
fof(f655,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(subsumption_resolution,[],[f647,f120])).
fof(f647,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl24),
  inference(superposition,[],[f462,f115])).
fof(f462,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f304,f115])).
fof(f3731,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,'T_0') != s(bool,X147) | s(bool,X148) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X147) != s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X148) != s(bool,X150) | s(bool,'T_0') = s(bool,X150)) ) | ($spl0 | $spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1171,f3657])).
fof(f1171,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,X71) != s(bool,X73) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X71) != s(bool,X72) | s(bool,'T_0') = s(bool,X72)) ) | ($spl24 | $spl54)),
  inference(resolution,[],[f1029,f472])).
fof(f764,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) | $spl44),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3812,plain,(
  ( ! [X132,X130,X131] : (s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,X130) | s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X132)) ) | ($spl0 | $spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f3725,f115])).
fof(f3725,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') != s(bool,X130) | s(bool,X131) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,X130) | s(bool,X131) != s(bool,X132) | s(bool,'T_0') = s(bool,X132)) ) | ($spl0 | $spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1020,f3657])).
fof(f3810,plain,(
  ( ! [X121,X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl0 | $spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f3720,f115])).
fof(f3720,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,X121) | s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl0 | $spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f718,f3657])).
fof(f3662,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | $spl226),
  inference(cnf_transformation,[],[f3662_D])).
fof(f3662_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f229,plain,(
  ( ! [X80] : (p(s(bool,X80))) ) | $spl22),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ( ! [X80] : (p(s(bool,X80))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3656,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f88])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X16,X12,X17] : ((s(bool,X11) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X12) = s(bool,X17))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.bool.IMP_CONG')).
fof(f215,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | ~p(s(bool,X45))) ) | $spl18),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | ~p(s(bool,X45))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3655,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X22) != s(bool,X23)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f228,plain,(
  ( ! [X81] : (~p(s(bool,X81))) ) | $spl20),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X81] : (~p(s(bool,X81))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3652,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f90])).
fof(f90,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f203,plain,(
  ( ! [X23,X25] : (s(bool,X23) != s(bool,X25) | p(s(bool,X23))) ) | $spl12),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  ( ! [X23,X25] : (s(bool,X23) != s(bool,X25) | p(s(bool,X23))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3651,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f89])).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f3650,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,X7) != s(bool,X8)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f3648,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK3,sK2,sK1)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f3646,f85])).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f75])).
fof(f3624,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(bool,'F_0') != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,X32) != s(bool,X33) | s(bool,X33) != s(bool,X34) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) != s(bool,X35) | s(bool,X31) != s(bool,X32) | s(bool,X28) != s(bool,X29)) ) | $spl224),
  inference(cnf_transformation,[],[f3624_D])).
fof(f3624_D,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(bool,'F_0') != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,X32) != s(bool,X33) | s(bool,X33) != s(bool,X34) | s(bool,'F_0') = s(bool,X35) | s(bool,X34) != s(bool,X35) | s(bool,X31) != s(bool,X32) | s(bool,X28) != s(bool,X29)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f3623,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,X32) != s(bool,X33) | s(bool,X29) != s(bool,X30)) ) | ($spl10 | $spl148)),
  inference(subsumption_resolution,[],[f3616,f200])).
fof(f3616,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | s(bool,X32) != s(bool,X33) | s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,X27)) ) | $spl148),
  inference(superposition,[],[f2997,f115])).
fof(f3621,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X10) | p(s(bool,X10)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X8) | s(bool,X7) != s(bool,X8) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl222),
  inference(cnf_transformation,[],[f3621_D])).
fof(f3621_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X10) | p(s(bool,X10)) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X8) | s(bool,X7) != s(bool,X8) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f3606,plain,(
  ( ! [X103,X101,X99,X97,X105,X102,X100,X98,X96,X106,X104] : (s(bool,X96) != s(bool,X97) | s(bool,X96) != s(bool,X98) | s(bool,X99) != s(bool,X100) | s(bool,X100) != s(bool,X101) | s(bool,X102) != s(bool,X103) | s(bool,X103) != s(bool,X104) | s(bool,'F_0') = s(bool,X104) | s(bool,X101) != s(bool,X102) | s(bool,X98) != s(bool,X99) | p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X97) != s(bool,X106)) ) | $spl148),
  inference(resolution,[],[f2997,f119])).
fof(f3605,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,X84) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,X90) != s(bool,X91) | s(bool,X91) != s(bool,X92) | s(bool,'F_0') = s(bool,X92) | s(bool,X89) != s(bool,X90) | s(bool,X86) != s(bool,X87) | p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,X85) != s(bool,X93) | s(bool,X94) != s(bool,X95)) ) | $spl148),
  inference(resolution,[],[f2997,f91])).
fof(f3602,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,X53) != s(bool,X54) | s(bool,X50) != s(bool,X51) | p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X49) != s(bool,X59)) ) | $spl148),
  inference(resolution,[],[f2997,f90])).
fof(f3601,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,X42) != s(bool,X43) | s(bool,X43) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | s(bool,X41) != s(bool,X42) | s(bool,X38) != s(bool,X39) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X37) != s(bool,X46)) ) | $spl148),
  inference(resolution,[],[f2997,f89])).
fof(f3600,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | s(bool,X26) != s(bool,X27) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X25) != s(bool,X33) | s(bool,X34) != s(bool,X35)) ) | $spl148),
  inference(resolution,[],[f2997,f87])).
fof(f3598,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,X7) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X1) != s(bool,X10)) ) | $spl148),
  inference(resolution,[],[f2997,f85])).
fof(f3591,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X24))))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24))),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26))) ) | $spl220),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X24))))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24))),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f3584,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X84,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X85),'type.ffi.ffi_state'(X86)),X87),s('type.closSem.state'(X85),X88)),X86)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X84,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X85),'type.ffi.ffi_state'(X86)),X87),s('type.closSem.state'(X85),X88)),X86)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X85),X88))))) | p(s(bool,X89)) | p(s(bool,X90)) | s(bool,X90) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X86),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X85),'type.ffi.ffi_state'(X86)),X87),s('type.closSem.state'(X85),X88))),s('type.sptree.spt'('type.clos_known.val_approx'),X84)))) )),
  inference(resolution,[],[f1060,f119])).
fof(f1060,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.closSem.state'(X1),X3))),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X4,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.closSem.state'(X1),X3)),X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X4,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.closSem.state'(X1),X3)),X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X1),X3)))))) )),
  inference(superposition,[],[f105,f97])).
fof(f3583,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X76,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X77),'type.ffi.ffi_state'(X78)),X79),s('type.closSem.state'(X77),X80)),X78)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X76,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X77),'type.ffi.ffi_state'(X78)),X79),s('type.closSem.state'(X77),X80)),X78)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X77),X80))))) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X78),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X77),'type.ffi.ffi_state'(X78)),X79),s('type.closSem.state'(X77),X80))),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) | s(bool,X82) != s(bool,X83)) )),
  inference(resolution,[],[f1060,f91])).
fof(f3580,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X52,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X53),'type.ffi.ffi_state'(X54)),X55),s('type.closSem.state'(X53),X56)),X54)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X52,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X53),'type.ffi.ffi_state'(X54)),X55),s('type.closSem.state'(X53),X56)),X54)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X53),X56))))) | p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X59) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X54),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X53),'type.ffi.ffi_state'(X54)),X55),s('type.closSem.state'(X53),X56))),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) )),
  inference(resolution,[],[f1060,f90])).
fof(f3579,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X44,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X45),'type.ffi.ffi_state'(X46)),X47),s('type.closSem.state'(X45),X48)),X46)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X44,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X45),'type.ffi.ffi_state'(X46)),X47),s('type.closSem.state'(X45),X48)),X46)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X45),X48))))) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X45),'type.ffi.ffi_state'(X46)),X47),s('type.closSem.state'(X45),X48))),s('type.sptree.spt'('type.clos_known.val_approx'),X44)))) )),
  inference(resolution,[],[f1060,f89])).
fof(f3578,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X36,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X37),'type.ffi.ffi_state'(X38)),X39),s('type.closSem.state'(X37),X40)),X38)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X36,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X37),'type.ffi.ffi_state'(X38)),X39),s('type.closSem.state'(X37),X40)),X38)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X37),X40))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X38),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X37),'type.ffi.ffi_state'(X38)),X39),s('type.closSem.state'(X37),X40))),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) | s(bool,X42) != s(bool,X43)) )),
  inference(resolution,[],[f1060,f87])).
fof(f3576,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X20,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24)),X22)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X21),X24))))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X21),'type.ffi.ffi_state'(X22)),X23),s('type.closSem.state'(X21),X24))),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) )),
  inference(resolution,[],[f1060,f85])).
fof(f3571,plain,(
  ( ! [X14,X12,X11] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X11,sK2,sK1),sK6(X11,sK2,sK1),X12,sK2,sK1)),s('type.closSem.v',sK7(X11,sK2,sK1)))) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X11))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X12,sK2,sK1),sK6(X12,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X12,sK2,sK1))))) ) | $spl218),
  inference(cnf_transformation,[],[f3571_D])).
fof(f3571_D,plain,(
  ( ! [X14,X12,X11] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X11,sK2,sK1),sK6(X11,sK2,sK1),X12,sK2,sK1)),s('type.closSem.v',sK7(X11,sK2,sK1)))) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X11))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X12,sK2,sK1),sK6(X12,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X12,sK2,sK1))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f3566,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X51,sK2,sK1),sK6(X51,sK2,sK1),X52,sK2,sK1)),s('type.closSem.v',sK7(X51,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X52,sK2,sK1),sK6(X52,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X52,sK2,sK1)))) | p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X51)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f119])).
fof(f3038,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X9,sK2,sK1),sK6(X9,sK2,sK1),X10,sK2,sK1)),s('type.closSem.v',sK7(X9,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X10,sK2,sK1),sK6(X10,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X10,sK2,sK1))))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3035,f2716])).
fof(f3035,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X5,X6,X7),sK6(X5,X6,X7),X8,X6,X7)),s('type.closSem.v',sK7(X5,X6,X7))))) )),
  inference(subsumption_resolution,[],[f3030,f116])).
fof(f3030,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'F_0')) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X5,X6,X7),sK6(X5,X6,X7),X8,X6,X7)),s('type.closSem.v',sK7(X5,X6,X7))))) )),
  inference(superposition,[],[f3016,f115])).
fof(f3016,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X0,X1,X2),sK6(X0,X1,X2),X3,X1,X2)),s('type.closSem.v',sK7(X0,X1,X2))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))))) )),
  inference(equality_resolution,[],[f1126])).
fof(f1126,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X5))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X3,X4)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X5,sK6(X2,X3,X4),X6,X3,X4)),s('type.closSem.v',X5)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X3),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X3),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) )),
  inference(superposition,[],[f104,f105])).
fof(f3565,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),X47,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X47,sK2,sK1),sK6(X47,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X47,sK2,sK1)))) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,X49) != s(bool,X50)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f91])).
fof(f3562,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X31,sK2,sK1),sK6(X31,sK2,sK1),X32,sK2,sK1)),s('type.closSem.v',sK7(X31,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X32,sK2,sK1),sK6(X32,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X32,sK2,sK1)))) | p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f90])).
fof(f3561,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,sK2,sK1),sK6(X26,sK2,sK1),X27,sK2,sK1)),s('type.closSem.v',sK7(X26,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,sK2,sK1),sK6(X27,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X27,sK2,sK1)))) | p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f89])).
fof(f3560,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X21,sK2,sK1),sK6(X21,sK2,sK1),X22,sK2,sK1)),s('type.closSem.v',sK7(X21,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | s(bool,X24) != s(bool,X25)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f87])).
fof(f3558,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X11,sK2,sK1),sK6(X11,sK2,sK1),X12,sK2,sK1)),s('type.closSem.v',sK7(X11,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X12,sK2,sK1),sK6(X12,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X12,sK2,sK1)))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f3038,f85])).
fof(f3304,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | $spl178),
  inference(cnf_transformation,[],[f3304_D])).
fof(f3304_D,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f3528,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1))))) ) | $spl216),
  inference(cnf_transformation,[],[f3528_D])).
fof(f3528_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f3531,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,X36,X35),sK6(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK7(X38,X36,X35)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl14 | $spl24 | $spl34 | $spl186)),
  inference(subsumption_resolution,[],[f3521,f3375])).
fof(f3521,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,X36,X35),sK6(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK7(X38,X36,X35)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X31,X32,X33),sK6(X31,X32,X33),X34,X32,X33)),s('type.closSem.v',sK7(X31,X32,X33)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl14 | $spl24 | $spl34)),
  inference(superposition,[],[f3328,f3076])).
fof(f3076,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X79,X80,X81),sK6(X79,X80,X81),X82,X80,X81)),s('type.closSem.v',sK7(X79,X80,X81)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X81),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X82))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X81),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X79)))) ) | ($spl24 | $spl34)),
  inference(resolution,[],[f3049,f472])).
fof(f3049,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X81)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X81,X80,X79),sK6(X81,X80,X79),X82,X80,X79)),s('type.closSem.v',sK7(X81,X80,X79)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X80),s('type.sptree.spt'('type.clos_known.val_approx'),X82)))) ) | $spl34),
  inference(resolution,[],[f3035,f512])).
fof(f3328,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X95) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X93),X92),s('type.sptree.spt'('type.clos_known.val_approx'),X94))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X91,X92,X93),sK6(X91,X92,X93),X94,X92,X93)),s('type.closSem.v',sK7(X91,X92,X93)))) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X93),X92),s('type.sptree.spt'('type.clos_known.val_approx'),X91)))) ) | ($spl14 | $spl24)),
  inference(resolution,[],[f3050,f472])).
fof(f3050,plain,(
  ( ! [X87,X85,X83,X86,X84] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X84),s('type.sptree.spt'('type.clos_known.val_approx'),X85)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X85,X84,X83),sK6(X85,X84,X83),X86,X84,X83)),s('type.closSem.v',sK7(X85,X84,X83)))) | s(bool,X87) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X83),X84),s('type.sptree.spt'('type.clos_known.val_approx'),X86))) | s(bool,'F_0') = s(bool,X87)) ) | $spl14),
  inference(resolution,[],[f3035,f516])).
fof(f516,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f507,f120])).
fof(f507,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(superposition,[],[f205,f115])).
fof(f3530,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X30,X28,X27),sK6(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK7(X30,X28,X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl14 | $spl24 | $spl146)),
  inference(subsumption_resolution,[],[f3520,f2934])).
fof(f3520,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X30,X28,X27),sK6(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK7(X30,X28,X27)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,sK2,sK1),sK6(X26,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X26,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl14 | $spl24)),
  inference(superposition,[],[f3328,f2733])).
fof(f3527,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) ) | ($spl14 | $spl24)),
  inference(duplicate_literal_removal,[],[f3516])).
fof(f3516,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f3328,f115])).
fof(f3246,plain,(
  ( ! [X397,X399,X403,X398,X400,X402] : (s(bool,X402) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X397,X398,X399),sK6(X397,X398,X399),X400,X398,X399)),s('type.closSem.v',sK7(X397,X398,X399)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X397))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X400))) | s(bool,'T_0') = s(bool,X403) | s(bool,X402) != s(bool,X403)) ) | $spl166),
  inference(cnf_transformation,[],[f3246_D])).
fof(f3246_D,plain,(
  ( ! [X397,X399,X403,X398,X400,X402] : (s(bool,X402) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X397,X398,X399),sK6(X397,X398,X399),X400,X398,X399)),s('type.closSem.v',sK7(X397,X398,X399)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X397))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X400))) | s(bool,'T_0') = s(bool,X403) | s(bool,X402) != s(bool,X403)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f3512,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X18,X19,X20),sK6(X18,X19,X20),X21,X19,X20)),s('type.closSem.v',sK7(X18,X19,X20)))) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl2 | $spl24 | $spl34 | $spl190)),
  inference(subsumption_resolution,[],[f3504,f2733])).
fof(f3504,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X18,X19,X20),sK6(X18,X19,X20),X21,X19,X20)),s('type.closSem.v',sK7(X18,X19,X20)))) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl24 | $spl34 | $spl190)),
  inference(superposition,[],[f3427,f3076])).
fof(f3469,plain,(
  ( ! [X28,X26,X31,X29,X27] : (s(bool,X31) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'F_0') = s(bool,X31)) ) | $spl214),
  inference(cnf_transformation,[],[f3469_D])).
fof(f3469_D,plain,(
  ( ! [X28,X26,X31,X29,X27] : (s(bool,X31) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'F_0') = s(bool,X31)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f3465,plain,(
  ( ! [X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X3)) ) | $spl208),
  inference(cnf_transformation,[],[f3465_D])).
fof(f3465_D,plain,(
  ( ! [X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X3)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f3452,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,X31) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl24 | $spl34 | $spl94)),
  inference(superposition,[],[f1575,f3076])).
fof(f3468,plain,(
  ( ! [X24,X22] : (s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'F_0') = s(bool,X24)) ) | $spl212),
  inference(cnf_transformation,[],[f3468_D])).
fof(f3468_D,plain,(
  ( ! [X24,X22] : (s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f3451,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X24) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X25) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24 | $spl94)),
  inference(superposition,[],[f1575,f2733])).
fof(f3467,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X17)) ) | $spl210),
  inference(cnf_transformation,[],[f3467_D])).
fof(f3467_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f3449,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X16) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X18)) ) | ($spl2 | $spl94)),
  inference(superposition,[],[f1575,f125])).
fof(f3466,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | $spl94),
  inference(subsumption_resolution,[],[f3447,f115])).
fof(f3447,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | $spl94),
  inference(superposition,[],[f1575,f115])).
fof(f3464,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl38 | $spl94)),
  inference(subsumption_resolution,[],[f3445,f719])).
fof(f719,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X56) | s(bool,X56) != s(bool,X57) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(resolution,[],[f678,f116])).
fof(f3445,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | $spl94),
  inference(superposition,[],[f1575,f115])).
fof(f3463,plain,(
  ( ! [X28,X26,X29,X27,X32] : (s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) ) | $spl206),
  inference(cnf_transformation,[],[f3463_D])).
fof(f3463_D,plain,(
  ( ! [X28,X26,X29,X27,X32] : (s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f3462,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) ) | ($spl24 | $spl34 | $spl94)),
  inference(subsumption_resolution,[],[f3461,f655])).
fof(f3461,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) ) | ($spl24 | $spl34 | $spl94)),
  inference(subsumption_resolution,[],[f3444,f1575])).
fof(f3444,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X32) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,X27,X28),sK6(X26,X27,X28),X29,X27,X28)),s('type.closSem.v',sK7(X26,X27,X28)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X28),X27),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl24 | $spl34 | $spl94)),
  inference(superposition,[],[f1575,f3076])).
fof(f3460,plain,(
  ( ! [X25,X22] : (s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1))))) ) | $spl204),
  inference(cnf_transformation,[],[f3460_D])).
fof(f3460_D,plain,(
  ( ! [X25,X22] : (s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f3459,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1))))) ) | ($spl2 | $spl24 | $spl94)),
  inference(subsumption_resolution,[],[f3458,f655])).
fof(f3458,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1))))) ) | ($spl2 | $spl24 | $spl94)),
  inference(subsumption_resolution,[],[f3443,f1575])).
fof(f3443,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24 | $spl94)),
  inference(superposition,[],[f1575,f2733])).
fof(f3457,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl202),
  inference(cnf_transformation,[],[f3457_D])).
fof(f3457_D,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f3456,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | ($spl2 | $spl24 | $spl94)),
  inference(subsumption_resolution,[],[f3441,f655])).
fof(f3441,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | ($spl2 | $spl94)),
  inference(superposition,[],[f1575,f125])).
fof(f206,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) | $spl16),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3455,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') != s(bool,X11)) ) | $spl94),
  inference(subsumption_resolution,[],[f3439,f1575])).
fof(f3439,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') = s(bool,X8)) ) | $spl94),
  inference(superposition,[],[f1575,f115])).
fof(f3454,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl94)),
  inference(subsumption_resolution,[],[f3437,f655])).
fof(f3437,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | $spl94),
  inference(superposition,[],[f1575,f115])).
fof(f3436,plain,(
  ( ! [X57,X64,X54,X62,X56,X55,X63] : (s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,X62) != s(bool,X63) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) | $spl200),
  inference(cnf_transformation,[],[f3436_D])).
fof(f3436_D,plain,(
  ( ! [X57,X64,X54,X62,X56,X55,X63] : (s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,X62) != s(bool,X63) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f3431,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X2) != s(bool,X3)) ) | $spl194),
  inference(cnf_transformation,[],[f3431_D])).
fof(f3431_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X2) != s(bool,X3)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f3420,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,'T_0') != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X62) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X58) != s(bool,X60) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl78)),
  inference(superposition,[],[f2985,f3076])).
fof(f3435,plain,(
  ( ! [X52,X46,X53,X51] : (s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53)) ) | $spl198),
  inference(cnf_transformation,[],[f3435_D])).
fof(f3435_D,plain,(
  ( ! [X52,X46,X53,X51] : (s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,X51) != s(bool,X52) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f3419,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,'T_0') != s(bool,X47) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X51) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X47) != s(bool,X49) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f2985,f2733])).
fof(f3433,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) | $spl196),
  inference(cnf_transformation,[],[f3433_D])).
fof(f3433_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f3417,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X36) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X32) != s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f2985,f125])).
fof(f3430,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl54 | $spl78)),
  inference(subsumption_resolution,[],[f3413,f1172])).
fof(f1172,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) != s(bool,X78)) ) | $spl54),
  inference(resolution,[],[f1029,f116])).
fof(f3413,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X2) != s(bool,X4) | s(bool,X0) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X1) != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl78)),
  inference(superposition,[],[f2985,f115])).
fof(f3429,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (s(bool,X59) != s(bool,X60) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,X59) != s(bool,X61)) ) | $spl192),
  inference(cnf_transformation,[],[f3429_D])).
fof(f3429_D,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (s(bool,X59) != s(bool,X60) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,X59) != s(bool,X61)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f3428,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X58) != s(bool,X62) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3412,f1020])).
fof(f3412,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X63] : (s(bool,'T_0') != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,X59) != s(bool,X61) | s(bool,X58) != s(bool,X62) | s(bool,X63) != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X63) | s(bool,X60) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X54,X55,X56),sK6(X54,X55,X56),X57,X55,X56)),s('type.closSem.v',sK7(X54,X55,X56)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X55),s('type.sptree.spt'('type.clos_known.val_approx'),X54)))) ) | ($spl24 | $spl34 | $spl78)),
  inference(superposition,[],[f2985,f3076])).
fof(f3427,plain,(
  ( ! [X50,X48,X46,X49] : (s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50)) ) | $spl190),
  inference(cnf_transformation,[],[f3427_D])).
fof(f3427_D,plain,(
  ( ! [X50,X48,X46,X49] : (s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f3426,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X47) != s(bool,X51) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3411,f1020])).
fof(f3411,plain,(
  ( ! [X47,X52,X50,X48,X46,X53,X51,X49] : (s(bool,'T_0') != s(bool,X47) | s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,X47) != s(bool,X51) | s(bool,X52) != s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X46,sK2,sK1),sK6(X46,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X46,sK2,sK1)))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X46)))) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f2985,f2733])).
fof(f3425,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) | $spl188),
  inference(cnf_transformation,[],[f3425_D])).
fof(f3425_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f3424,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X32) != s(bool,X36) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl40 | $spl78)),
  inference(subsumption_resolution,[],[f3409,f1020])).
fof(f3409,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X32) != s(bool,X36) | s(bool,X37) != s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X36) != s(bool,X37) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl2 | $spl24 | $spl78)),
  inference(superposition,[],[f2985,f125])).
fof(f499,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl32),
  inference(cnf_transformation,[],[f499_D])).
fof(f499_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3375,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) | $spl186),
  inference(cnf_transformation,[],[f3375_D])).
fof(f3375_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f3374,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl162)),
  inference(subsumption_resolution,[],[f3363,f472])).
fof(f3363,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5)) ) | $spl162),
  inference(superposition,[],[f3221,f115])).
fof(f3354,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),X3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | $spl184),
  inference(cnf_transformation,[],[f3354_D])).
fof(f3354_D,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),X3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f306,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26)) ) | $spl26),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3353,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),X3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl0 | $spl24 | $spl156)),
  inference(subsumption_resolution,[],[f3343,f304])).
fof(f3343,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0),X3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl0 | $spl156)),
  inference(superposition,[],[f3081,f124])).
fof(f3348,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X8,X7,X6),sK6(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK7(X8,X7,X6))))) ) | $spl182),
  inference(cnf_transformation,[],[f3348_D])).
fof(f3348_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X8,X7,X6),sK6(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK7(X8,X7,X6))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f3347,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X8,X7,X6),sK6(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK7(X8,X7,X6)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl24 | $spl156)),
  inference(subsumption_resolution,[],[f3337,f472])).
fof(f3337,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X8))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X8,X7,X6),sK6(X8,X7,X6),X9,X7,X6)),s('type.closSem.v',sK7(X8,X7,X6)))) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | s(bool,'T_0') = s(bool,X5)) ) | $spl156),
  inference(superposition,[],[f3081,f115])).
fof(f3333,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X20,X21,X22),sK6(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK7(X20,X21,X22)))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) ) | $spl180),
  inference(cnf_transformation,[],[f3333_D])).
fof(f3333_D,plain,(
  ( ! [X26,X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X20,X21,X22),sK6(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK7(X20,X21,X22)))) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,X24) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f3327,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X84,X85,X86),sK6(X84,X85,X86),X87,X85,X86)),s('type.closSem.v',sK7(X84,X85,X86)))) | s(bool,X88) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X86),X85),s('type.sptree.spt'('type.clos_known.val_approx'),X87))) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X89)) | p(s(bool,X90)) | s(bool,X90) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X86),X85),s('type.sptree.spt'('type.clos_known.val_approx'),X84)))) ) | $spl14),
  inference(resolution,[],[f3050,f119])).
fof(f3326,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X76,X77,X78),sK6(X76,X77,X78),X79,X77,X78)),s('type.closSem.v',sK7(X76,X77,X78)))) | s(bool,X80) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X78),X77),s('type.sptree.spt'('type.clos_known.val_approx'),X79))) | s(bool,'F_0') = s(bool,X80) | p(s(bool,X81)) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X78),X77),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) | s(bool,X82) != s(bool,X83)) ) | $spl14),
  inference(resolution,[],[f3050,f91])).
fof(f3323,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X52,X53,X54),sK6(X52,X53,X54),X55,X53,X54)),s('type.closSem.v',sK7(X52,X53,X54)))) | s(bool,X56) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X54),X53),s('type.sptree.spt'('type.clos_known.val_approx'),X55))) | s(bool,'F_0') = s(bool,X56) | p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X59) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X54),X53),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) ) | $spl14),
  inference(resolution,[],[f3050,f90])).
fof(f3322,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X44,X45,X46),sK6(X44,X45,X46),X47,X45,X46)),s('type.closSem.v',sK7(X44,X45,X46)))) | s(bool,X48) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X45),s('type.sptree.spt'('type.clos_known.val_approx'),X47))) | s(bool,'F_0') = s(bool,X48) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X45),s('type.sptree.spt'('type.clos_known.val_approx'),X44)))) ) | $spl14),
  inference(resolution,[],[f3050,f89])).
fof(f3321,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X36,X37,X38),sK6(X36,X37,X38),X39,X37,X38)),s('type.closSem.v',sK7(X36,X37,X38)))) | s(bool,X40) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X38),X37),s('type.sptree.spt'('type.clos_known.val_approx'),X39))) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X38),X37),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) | s(bool,X42) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f3050,f87])).
fof(f3319,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X20,X21,X22),sK6(X20,X21,X22),X23,X21,X22)),s('type.closSem.v',sK7(X20,X21,X22)))) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X22),X21),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) | $spl14),
  inference(resolution,[],[f3050,f85])).
fof(f3297,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))))) ) | $spl176),
  inference(cnf_transformation,[],[f3297_D])).
fof(f3297_D,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f3303,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,X36,X35),sK6(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK7(X38,X36,X35)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl10 | $spl24 | $spl34 | $spl154)),
  inference(subsumption_resolution,[],[f3290,f3207])).
fof(f3207,plain,(
  ( ! [X313,X312,X314,X311] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X311,X312,X313),sK6(X311,X312,X313),X314,X312,X313)),s('type.closSem.v',sK7(X311,X312,X313))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X313),X312),s('type.sptree.spt'('type.clos_known.val_approx'),X314))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X313),X312),s('type.sptree.spt'('type.clos_known.val_approx'),X311)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(trivial_inequality_removal,[],[f3157])).
fof(f3157,plain,(
  ( ! [X313,X312,X314,X311] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X311,X312,X313),sK6(X311,X312,X313),X314,X312,X313)),s('type.closSem.v',sK7(X311,X312,X313))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X313),X312),s('type.sptree.spt'('type.clos_known.val_approx'),X314))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X313),X312),s('type.sptree.spt'('type.clos_known.val_approx'),X311)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(superposition,[],[f200,f3076])).
fof(f3290,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X31,X32,X33),sK6(X31,X32,X33),X34,X32,X33)),s('type.closSem.v',sK7(X31,X32,X33))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,X36,X35),sK6(X38,X36,X35),X37,X36,X35)),s('type.closSem.v',sK7(X38,X36,X35)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31)))) ) | ($spl24 | $spl34 | $spl154)),
  inference(superposition,[],[f3061,f3076])).
fof(f3302,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X30,X28,X27),sK6(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK7(X30,X28,X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl10 | $spl24 | $spl154)),
  inference(subsumption_resolution,[],[f3289,f2825])).
fof(f2825,plain,(
  ( ! [X156] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X156,sK2,sK1),sK6(X156,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X156,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X156)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(trivial_inequality_removal,[],[f2783])).
fof(f2783,plain,(
  ( ! [X156] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X156,sK2,sK1),sK6(X156,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X156,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X156)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(superposition,[],[f200,f2733])).
fof(f3289,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X27),X28),s('type.sptree.spt'('type.clos_known.val_approx'),X30)))) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X26,sK2,sK1),sK6(X26,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X26,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X30,X28,X27),sK6(X30,X28,X27),X29,X28,X27)),s('type.closSem.v',sK7(X30,X28,X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X26)))) ) | ($spl2 | $spl24 | $spl154)),
  inference(superposition,[],[f3061,f2733])).
fof(f3296,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl34 | $spl154)),
  inference(subsumption_resolution,[],[f3285,f512])).
fof(f3285,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X4)))) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X4,X2,X1),sK6(X4,X2,X1),X3,X2,X1)),s('type.closSem.v',sK7(X4,X2,X1)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl154),
  inference(superposition,[],[f3061,f115])).
fof(f3221,plain,(
  ( ! [X80,X78,X79,X77,X82] : (s(bool,X82) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X77,X78,X79),sK6(X77,X78,X79),X80,X78,X79)),s('type.closSem.v',sK7(X77,X78,X79)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X80))) | p(s(bool,X82))) ) | $spl162),
  inference(cnf_transformation,[],[f3221_D])).
fof(f3221_D,plain,(
  ( ! [X80,X78,X79,X77,X82] : (s(bool,X82) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X77,X78,X79),sK6(X77,X78,X79),X80,X78,X79)),s('type.closSem.v',sK7(X77,X78,X79)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X80))) | p(s(bool,X82))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f3276,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63))) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X63,X61,X60),sK6(X63,X61,X60),X62,X61,X60)),s('type.closSem.v',sK7(X63,X61,X60))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f119])).
fof(f3275,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X55))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X56))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X56,X54,X53),sK6(X56,X54,X53),X55,X54,X53)),s('type.closSem.v',sK7(X56,X54,X53)))) | s(bool,X58) != s(bool,X59)) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f91])).
fof(f3272,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X33,X32),sK6(X35,X33,X32),X34,X33,X32)),s('type.closSem.v',sK7(X35,X33,X32))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f90])).
fof(f3271,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X28,X26,X25),sK6(X28,X26,X25),X27,X26,X25)),s('type.closSem.v',sK7(X28,X26,X25))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f89])).
fof(f3270,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X21,X19,X18),sK6(X21,X19,X18),X20,X19,X18)),s('type.closSem.v',sK7(X21,X19,X18)))) | s(bool,X23) != s(bool,X24)) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f87])).
fof(f3268,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X7,X5,X4),sK6(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK7(X7,X5,X4))))) ) | ($spl10 | $spl24 | $spl34)),
  inference(resolution,[],[f3207,f85])).
fof(f3219,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X36,X37),sK6(X35,X36,X37),X38,X36,X37)),s('type.closSem.v',sK7(X35,X36,X37))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X35))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) | $spl160),
  inference(cnf_transformation,[],[f3219_D])).
fof(f3219_D,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X36,X37),sK6(X35,X36,X37),X38,X36,X37)),s('type.closSem.v',sK7(X35,X36,X37))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X35))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f3260,plain,(
  ( ! [X617,X619,X620,X618] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X620)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X617))))) ) | $spl174),
  inference(cnf_transformation,[],[f3260_D])).
fof(f3260_D,plain,(
  ( ! [X617,X619,X620,X618] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X620)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X617))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f3259,plain,(
  ( ! [X617,X619,X613,X615,X620,X616,X618,X614] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X620)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X613,X614,X615),sK6(X613,X614,X615),X616,X614,X615)),s('type.closSem.v',sK7(X613,X614,X615))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X617)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X615),X614),s('type.sptree.spt'('type.clos_known.val_approx'),X616))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X615),X614),s('type.sptree.spt'('type.clos_known.val_approx'),X613)))) ) | ($spl24 | $spl34 | $spl152)),
  inference(subsumption_resolution,[],[f3204,f3035])).
fof(f3204,plain,(
  ( ! [X617,X619,X613,X615,X620,X616,X618,X614] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X617,X618,X619),sK6(X617,X618,X619),X620,X618,X619)),s('type.closSem.v',sK7(X617,X618,X619)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X620)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X613,X614,X615),sK6(X613,X614,X615),X616,X614,X615)),s('type.closSem.v',sK7(X613,X614,X615))))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X619),X618),s('type.sptree.spt'('type.clos_known.val_approx'),X617)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X615),X614),s('type.sptree.spt'('type.clos_known.val_approx'),X616))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X615),X614),s('type.sptree.spt'('type.clos_known.val_approx'),X613)))) ) | ($spl24 | $spl34 | $spl152)),
  inference(superposition,[],[f3032,f3076])).
fof(f3244,plain,(
  ( ! [X368,X370,X373,X369,X371] : (s(bool,X373) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X368,X369,X370),sK6(X368,X369,X370),X371,X369,X370)),s('type.closSem.v',sK7(X368,X369,X370)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X368))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X371))) | s(bool,'T_0') = s(bool,X373)) ) | $spl164),
  inference(cnf_transformation,[],[f3244_D])).
fof(f3244_D,plain,(
  ( ! [X368,X370,X373,X369,X371] : (s(bool,X373) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X368,X369,X370),sK6(X368,X369,X370),X371,X369,X370)),s('type.closSem.v',sK7(X368,X369,X370)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X368))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X371))) | s(bool,'T_0') = s(bool,X373)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f3258,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X601))) | s(bool,'T_0') = s(bool,X602) | s(bool,X602) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X597,X598,X599),sK6(X597,X598,X599),X600,X598,X599)),s('type.closSem.v',sK7(X597,X598,X599)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X599),X598),s('type.sptree.spt'('type.clos_known.val_approx'),X600))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X599),X598),s('type.sptree.spt'('type.clos_known.val_approx'),X597)))) ) | ($spl2 | $spl24 | $spl34 | $spl140)),
  inference(subsumption_resolution,[],[f3202,f2733])).
fof(f3202,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X601,sK2,sK1),sK6(X601,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X601,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X601))) | s(bool,'T_0') = s(bool,X602) | s(bool,X602) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X597,X598,X599),sK6(X597,X598,X599),X600,X598,X599)),s('type.closSem.v',sK7(X597,X598,X599)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X599),X598),s('type.sptree.spt'('type.clos_known.val_approx'),X600))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X599),X598),s('type.sptree.spt'('type.clos_known.val_approx'),X597)))) ) | ($spl24 | $spl34 | $spl140)),
  inference(superposition,[],[f2857,f3076])).
fof(f3256,plain,(
  ( ! [X596,X592,X594,X593,X595] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X596))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X592,X593,X594),sK6(X592,X593,X594),X595,X593,X594)),s('type.closSem.v',sK7(X592,X593,X594))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X594),X593),s('type.sptree.spt'('type.clos_known.val_approx'),X595))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X594),X593),s('type.sptree.spt'('type.clos_known.val_approx'),X592)))) ) | ($spl2 | $spl24 | $spl34 | $spl136)),
  inference(subsumption_resolution,[],[f3201,f2733])).
fof(f3201,plain,(
  ( ! [X596,X592,X594,X593,X595] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X596,sK2,sK1),sK6(X596,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X596,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X596))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X592,X593,X594),sK6(X592,X593,X594),X595,X593,X594)),s('type.closSem.v',sK7(X592,X593,X594))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X594),X593),s('type.sptree.spt'('type.clos_known.val_approx'),X595))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X594),X593),s('type.sptree.spt'('type.clos_known.val_approx'),X592)))) ) | ($spl24 | $spl34 | $spl136)),
  inference(superposition,[],[f2832,f3076])).
fof(f3255,plain,(
  ( ! [X586] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X586,sK2,sK1),sK6(X586,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X586,sK2,sK1))))) ) | $spl170),
  inference(cnf_transformation,[],[f3255_D])).
fof(f3255_D,plain,(
  ( ! [X586] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X586,sK2,sK1),sK6(X586,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X586,sK2,sK1))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f3254,plain,(
  ( ! [X584,X586,X582,X585,X583] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X582,X583,X584),sK6(X582,X583,X584),X585,X583,X584)),s('type.closSem.v',sK7(X582,X583,X584))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X586,sK2,sK1),sK6(X586,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X586,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X584),X583),s('type.sptree.spt'('type.clos_known.val_approx'),X585))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X584),X583),s('type.sptree.spt'('type.clos_known.val_approx'),X582)))) ) | ($spl2 | $spl24 | $spl34 | $spl132)),
  inference(subsumption_resolution,[],[f3199,f2733])).
fof(f3199,plain,(
  ( ! [X584,X586,X582,X585,X583] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X586))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X582,X583,X584),sK6(X582,X583,X584),X585,X583,X584)),s('type.closSem.v',sK7(X582,X583,X584))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X586,sK2,sK1),sK6(X586,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X586,sK2,sK1)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X584),X583),s('type.sptree.spt'('type.clos_known.val_approx'),X585))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X584),X583),s('type.sptree.spt'('type.clos_known.val_approx'),X582)))) ) | ($spl24 | $spl34 | $spl132)),
  inference(superposition,[],[f2737,f3076])).
fof(f3253,plain,(
  ( ! [X581] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X581))))) ) | $spl168),
  inference(cnf_transformation,[],[f3253_D])).
fof(f3253_D,plain,(
  ( ! [X581] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X581))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f3252,plain,(
  ( ! [X580,X578,X581,X577,X579] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X581)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X577,X578,X579),sK6(X577,X578,X579),X580,X578,X579)),s('type.closSem.v',sK7(X577,X578,X579))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X579),X578),s('type.sptree.spt'('type.clos_known.val_approx'),X580))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X579),X578),s('type.sptree.spt'('type.clos_known.val_approx'),X577)))) ) | ($spl2 | $spl24 | $spl34 | $spl130)),
  inference(subsumption_resolution,[],[f3198,f2716])).
fof(f3198,plain,(
  ( ! [X580,X578,X581,X577,X579] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X581,sK2,sK1),sK6(X581,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X581,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X581)))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X577,X578,X579),sK6(X577,X578,X579),X580,X578,X579)),s('type.closSem.v',sK7(X577,X578,X579))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X579),X578),s('type.sptree.spt'('type.clos_known.val_approx'),X580))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X579),X578),s('type.sptree.spt'('type.clos_known.val_approx'),X577)))) ) | ($spl24 | $spl34 | $spl130)),
  inference(superposition,[],[f2720,f3076])).
fof(f3250,plain,(
  ( ! [X509,X511,X507,X512,X514,X508,X510,X506,X513,X515] : (s(bool,X511) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X506,X507,X508),sK6(X506,X507,X508),X509,X507,X508)),s('type.closSem.v',sK7(X506,X507,X508)))) | s(bool,X510) != s(bool,X512) | s(bool,X512) != s(bool,X513) | s(bool,'F_0') = s(bool,X514) | s(bool,X513) != s(bool,X514) | s(bool,X511) != s(bool,X515) | s(bool,'T_0') = s(bool,X515) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X508),X507),s('type.sptree.spt'('type.clos_known.val_approx'),X509))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X508),X507),s('type.sptree.spt'('type.clos_known.val_approx'),X506)))) ) | ($spl24 | $spl34 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f3188,f1020])).
fof(f3188,plain,(
  ( ! [X509,X511,X507,X512,X514,X508,X510,X506,X513,X515] : (s(bool,'T_0') != s(bool,X510) | s(bool,X511) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X506,X507,X508),sK6(X506,X507,X508),X509,X507,X508)),s('type.closSem.v',sK7(X506,X507,X508)))) | s(bool,X510) != s(bool,X512) | s(bool,X512) != s(bool,X513) | s(bool,'F_0') = s(bool,X514) | s(bool,X513) != s(bool,X514) | s(bool,X511) != s(bool,X515) | s(bool,'T_0') = s(bool,X515) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X508),X507),s('type.sptree.spt'('type.clos_known.val_approx'),X509))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X508),X507),s('type.sptree.spt'('type.clos_known.val_approx'),X506)))) ) | ($spl24 | $spl34 | $spl74)),
  inference(superposition,[],[f1343,f3076])).
fof(f3249,plain,(
  ( ! [X476,X478,X472,X474,X471,X475,X477,X473,X470] : (s(bool,X475) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X470,X471,X472),sK6(X470,X471,X472),X473,X471,X472)),s('type.closSem.v',sK7(X470,X471,X472)))) | s(bool,X474) != s(bool,X476) | s(bool,X476) != s(bool,X477) | s(bool,'F_0') = s(bool,X477) | s(bool,X475) != s(bool,X478) | s(bool,'T_0') = s(bool,X478) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X472),X471),s('type.sptree.spt'('type.clos_known.val_approx'),X473))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X472),X471),s('type.sptree.spt'('type.clos_known.val_approx'),X470)))) ) | ($spl24 | $spl34 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f3184,f718])).
fof(f3184,plain,(
  ( ! [X476,X478,X472,X474,X471,X475,X477,X473,X470] : (s(bool,'T_0') != s(bool,X474) | s(bool,X475) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X470,X471,X472),sK6(X470,X471,X472),X473,X471,X472)),s('type.closSem.v',sK7(X470,X471,X472)))) | s(bool,X474) != s(bool,X476) | s(bool,X476) != s(bool,X477) | s(bool,'F_0') = s(bool,X477) | s(bool,X475) != s(bool,X478) | s(bool,'T_0') = s(bool,X478) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X472),X471),s('type.sptree.spt'('type.clos_known.val_approx'),X473))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X472),X471),s('type.sptree.spt'('type.clos_known.val_approx'),X470)))) ) | ($spl24 | $spl34 | $spl68)),
  inference(superposition,[],[f1276,f3076])).
fof(f3248,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X439] : (s(bool,X443) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X438,X439,X440),sK6(X438,X439,X440),X441,X439,X440)),s('type.closSem.v',sK7(X438,X439,X440)))) | s(bool,X442) != s(bool,X444) | s(bool,'F_0') = s(bool,X444) | s(bool,X443) != s(bool,X445) | s(bool,'T_0') = s(bool,X445) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X440),X439),s('type.sptree.spt'('type.clos_known.val_approx'),X441))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X440),X439),s('type.sptree.spt'('type.clos_known.val_approx'),X438)))) ) | ($spl24 | $spl34 | $spl54)),
  inference(subsumption_resolution,[],[f3180,f655])).
fof(f3180,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X439] : (s(bool,'T_0') != s(bool,X442) | s(bool,X443) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X438,X439,X440),sK6(X438,X439,X440),X441,X439,X440)),s('type.closSem.v',sK7(X438,X439,X440)))) | s(bool,X442) != s(bool,X444) | s(bool,'F_0') = s(bool,X444) | s(bool,X443) != s(bool,X445) | s(bool,'T_0') = s(bool,X445) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X440),X439),s('type.sptree.spt'('type.clos_known.val_approx'),X441))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X440),X439),s('type.sptree.spt'('type.clos_known.val_approx'),X438)))) ) | ($spl24 | $spl34 | $spl54)),
  inference(superposition,[],[f1171,f3076])).
fof(f1589,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) | $spl96),
  inference(cnf_transformation,[],[f1589_D])).
fof(f1589_D,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3247,plain,(
  ( ! [X428,X430,X426,X429,X431,X425,X427] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X431,X430,X429)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X431,X430,X429)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X429),X430))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X425,X426,X427),sK6(X425,X426,X427),X428,X426,X427)),s('type.closSem.v',sK7(X425,X426,X427))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X427),X426),s('type.sptree.spt'('type.clos_known.val_approx'),X428))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X427),X426),s('type.sptree.spt'('type.clos_known.val_approx'),X425)))) ) | ($spl24 | $spl34 | $spl58)),
  inference(subsumption_resolution,[],[f3178,f1062])).
fof(f1062,plain,(
  ( ! [X4,X5,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X5,X4,X3)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X5,X4,X3)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X4),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) )),
  inference(subsumption_resolution,[],[f1059,f116])).
fof(f1059,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X5,X4,X3)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X5,X4,X3)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X3),X4))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X4),s('type.sptree.spt'('type.clos_known.val_approx'),X5)))) )),
  inference(superposition,[],[f105,f115])).
fof(f3178,plain,(
  ( ! [X428,X430,X426,X429,X431,X425,X427] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X429),X430),s('type.sptree.spt'('type.clos_known.val_approx'),X431))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X431,X430,X429)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X431,X430,X429)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X429),X430))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X425,X426,X427),sK6(X425,X426,X427),X428,X426,X427)),s('type.closSem.v',sK7(X425,X426,X427))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X427),X426),s('type.sptree.spt'('type.clos_known.val_approx'),X428))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X427),X426),s('type.sptree.spt'('type.clos_known.val_approx'),X425)))) ) | ($spl24 | $spl34 | $spl58)),
  inference(superposition,[],[f1061,f3076])).
fof(f3245,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (s(bool,X402) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X397,X398,X399),sK6(X397,X398,X399),X400,X398,X399)),s('type.closSem.v',sK7(X397,X398,X399)))) | s(bool,'F_0') = s(bool,X401) | s(bool,X402) != s(bool,X403) | s(bool,'T_0') = s(bool,X403) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X400))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X397)))) ) | ($spl24 | $spl34 | $spl40)),
  inference(subsumption_resolution,[],[f3174,f115])).
fof(f3174,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (s(bool,'T_0') != s(bool,X401) | s(bool,X402) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X397,X398,X399),sK6(X397,X398,X399),X400,X398,X399)),s('type.closSem.v',sK7(X397,X398,X399)))) | s(bool,'F_0') = s(bool,X401) | s(bool,X402) != s(bool,X403) | s(bool,'T_0') = s(bool,X403) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X400))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X399),X398),s('type.sptree.spt'('type.clos_known.val_approx'),X397)))) ) | ($spl24 | $spl34 | $spl40)),
  inference(superposition,[],[f1020,f3076])).
fof(f3243,plain,(
  ( ! [X372,X368,X370,X373,X369,X371] : (s(bool,X373) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X368,X369,X370),sK6(X368,X369,X370),X371,X369,X370)),s('type.closSem.v',sK7(X368,X369,X370)))) | s(bool,'F_0') = s(bool,X372) | s(bool,'T_0') = s(bool,X373) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X371))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X368)))) ) | ($spl24 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f3169,f115])).
fof(f3169,plain,(
  ( ! [X372,X368,X370,X373,X369,X371] : (s(bool,'T_0') != s(bool,X372) | s(bool,X373) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X368,X369,X370),sK6(X368,X369,X370),X371,X369,X370)),s('type.closSem.v',sK7(X368,X369,X370)))) | s(bool,'F_0') = s(bool,X372) | s(bool,'T_0') = s(bool,X373) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X371))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X370),X369),s('type.sptree.spt'('type.clos_known.val_approx'),X368)))) ) | ($spl24 | $spl34 | $spl38)),
  inference(superposition,[],[f718,f3076])).
fof(f3240,plain,(
  ( ! [X315,X316,X318,X317,X319] : (s(bool,'T_0') != s(bool,X319) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X315,X316,X317),sK6(X315,X316,X317),X318,X316,X317)),s('type.closSem.v',sK7(X315,X316,X317))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X318))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X315)))) ) | ($spl10 | $spl14 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3158,f200])).
fof(f3158,plain,(
  ( ! [X315,X316,X318,X317,X319] : (s(bool,'T_0') != s(bool,X319) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X315,X316,X317),sK6(X315,X316,X317),X318,X316,X317)),s('type.closSem.v',sK7(X315,X316,X317))))) | ~p(s(bool,X319)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X318))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X317),X316),s('type.sptree.spt'('type.clos_known.val_approx'),X315)))) ) | ($spl14 | $spl24 | $spl34)),
  inference(superposition,[],[f205,f3076])).
fof(f3239,plain,(
  ( ! [X308,X310,X306,X309,X305,X307] : (p(s(bool,X309)) | p(s(bool,X310)) | s(bool,X310) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X305,X306,X307),sK6(X305,X306,X307),X308,X306,X307)),s('type.closSem.v',sK7(X305,X306,X307)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X307),X306),s('type.sptree.spt'('type.clos_known.val_approx'),X308))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X307),X306),s('type.sptree.spt'('type.clos_known.val_approx'),X305)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3156,f120])).
fof(f3156,plain,(
  ( ! [X308,X310,X306,X309,X305,X307] : (~p(s(bool,'T_0')) | p(s(bool,X309)) | p(s(bool,X310)) | s(bool,X310) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X305,X306,X307),sK6(X305,X306,X307),X308,X306,X307)),s('type.closSem.v',sK7(X305,X306,X307)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X307),X306),s('type.sptree.spt'('type.clos_known.val_approx'),X308))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X307),X306),s('type.sptree.spt'('type.clos_known.val_approx'),X305)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f119,f3076])).
fof(f3238,plain,(
  ( ! [X292,X288,X290,X287,X289,X291] : (s(bool,'T_0') != s(bool,X291) | p(s(bool,X292)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X287,X288,X289),sK6(X287,X288,X289),X290,X288,X289)),s('type.closSem.v',sK7(X287,X288,X289))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X289),X288),s('type.sptree.spt'('type.clos_known.val_approx'),X290))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X289),X288),s('type.sptree.spt'('type.clos_known.val_approx'),X287)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3153,f200])).
fof(f3153,plain,(
  ( ! [X292,X288,X290,X287,X289,X291] : (s(bool,'T_0') != s(bool,X291) | p(s(bool,X292)) | ~p(s(bool,X291)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X287,X288,X289),sK6(X287,X288,X289),X290,X288,X289)),s('type.closSem.v',sK7(X287,X288,X289))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X289),X288),s('type.sptree.spt'('type.clos_known.val_approx'),X290))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X289),X288),s('type.sptree.spt'('type.clos_known.val_approx'),X287)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f119,f3076])).
fof(f3237,plain,(
  ( ! [X269,X271,X267,X268,X270,X266,X272] : (p(s(bool,X270)) | ~p(s(bool,X271)) | s(bool,X270) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X266,X267,X268),sK6(X266,X267,X268),X269,X267,X268)),s('type.closSem.v',sK7(X266,X267,X268)))) | s(bool,X271) != s(bool,X272) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X268),X267),s('type.sptree.spt'('type.clos_known.val_approx'),X269))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X268),X267),s('type.sptree.spt'('type.clos_known.val_approx'),X266)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3150,f120])).
fof(f3150,plain,(
  ( ! [X269,X271,X267,X268,X270,X266,X272] : (~p(s(bool,'T_0')) | p(s(bool,X270)) | ~p(s(bool,X271)) | s(bool,X270) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X266,X267,X268),sK6(X266,X267,X268),X269,X267,X268)),s('type.closSem.v',sK7(X266,X267,X268)))) | s(bool,X271) != s(bool,X272) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X268),X267),s('type.sptree.spt'('type.clos_known.val_approx'),X269))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X268),X267),s('type.sptree.spt'('type.clos_known.val_approx'),X266)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f91,f3076])).
fof(f3234,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X238] : (s(bool,'T_0') != s(bool,X242) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X238,X239,X240),sK6(X238,X239,X240),X241,X239,X240)),s('type.closSem.v',sK7(X238,X239,X240))))) | ~p(s(bool,X243)) | s(bool,X243) != s(bool,X244) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X240),X239),s('type.sptree.spt'('type.clos_known.val_approx'),X241))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X240),X239),s('type.sptree.spt'('type.clos_known.val_approx'),X238)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3146,f200])).
fof(f3146,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X238] : (s(bool,'T_0') != s(bool,X242) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X238,X239,X240),sK6(X238,X239,X240),X241,X239,X240)),s('type.closSem.v',sK7(X238,X239,X240))))) | ~p(s(bool,X243)) | ~p(s(bool,X242)) | s(bool,X243) != s(bool,X244) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X240),X239),s('type.sptree.spt'('type.clos_known.val_approx'),X241))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X240),X239),s('type.sptree.spt'('type.clos_known.val_approx'),X238)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f91,f3076])).
fof(f3231,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (p(s(bool,X214)) | ~p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X210,X211,X212),sK6(X210,X211,X212),X213,X211,X212)),s('type.closSem.v',sK7(X210,X211,X212)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X212),X211),s('type.sptree.spt'('type.clos_known.val_approx'),X213))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X212),X211),s('type.sptree.spt'('type.clos_known.val_approx'),X210)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3142,f120])).
fof(f3142,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X214] : (~p(s(bool,'T_0')) | p(s(bool,X214)) | ~p(s(bool,X215)) | p(s(bool,X216)) | s(bool,X216) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X210,X211,X212),sK6(X210,X211,X212),X213,X211,X212)),s('type.closSem.v',sK7(X210,X211,X212)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X212),X211),s('type.sptree.spt'('type.clos_known.val_approx'),X213))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X212),X211),s('type.sptree.spt'('type.clos_known.val_approx'),X210)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f90,f3076])).
fof(f3229,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X194)) | ~p(s(bool,X195)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X189,X190,X191),sK6(X189,X190,X191),X192,X190,X191)),s('type.closSem.v',sK7(X189,X190,X191))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X191),X190),s('type.sptree.spt'('type.clos_known.val_approx'),X192))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X191),X190),s('type.sptree.spt'('type.clos_known.val_approx'),X189)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3139,f200])).
fof(f3139,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (s(bool,'T_0') != s(bool,X193) | p(s(bool,X194)) | ~p(s(bool,X193)) | ~p(s(bool,X195)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X189,X190,X191),sK6(X189,X190,X191),X192,X190,X191)),s('type.closSem.v',sK7(X189,X190,X191))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X191),X190),s('type.sptree.spt'('type.clos_known.val_approx'),X192))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X191),X190),s('type.sptree.spt'('type.clos_known.val_approx'),X189)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f90,f3076])).
fof(f3228,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X179) != s(bool,X181) | s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X175,X176,X177),sK6(X175,X176,X177),X178,X176,X177)),s('type.closSem.v',sK7(X175,X176,X177)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X177),X176),s('type.sptree.spt'('type.clos_known.val_approx'),X178))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X177),X176),s('type.sptree.spt'('type.clos_known.val_approx'),X175)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3137,f120])).
fof(f3137,plain,(
  ( ! [X177,X175,X180,X178,X176,X181,X179] : (~p(s(bool,'T_0')) | p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X179) != s(bool,X181) | s(bool,X180) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X175,X176,X177),sK6(X175,X176,X177),X178,X176,X177)),s('type.closSem.v',sK7(X175,X176,X177)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X177),X176),s('type.sptree.spt'('type.clos_known.val_approx'),X178))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X177),X176),s('type.sptree.spt'('type.clos_known.val_approx'),X175)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f89,f3076])).
fof(f3227,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (s(bool,'T_0') != s(bool,X144) | p(s(bool,X145)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X140,X141,X142),sK6(X140,X141,X142),X143,X141,X142)),s('type.closSem.v',sK7(X140,X141,X142))))) | s(bool,X145) != s(bool,X146) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X142),X141),s('type.sptree.spt'('type.clos_known.val_approx'),X143))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X142),X141),s('type.sptree.spt'('type.clos_known.val_approx'),X140)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3132,f200])).
fof(f3132,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (s(bool,'T_0') != s(bool,X144) | p(s(bool,X145)) | ~p(s(bool,X144)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X140,X141,X142),sK6(X140,X141,X142),X143,X141,X142)),s('type.closSem.v',sK7(X140,X141,X142))))) | s(bool,X145) != s(bool,X146) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X142),X141),s('type.sptree.spt'('type.clos_known.val_approx'),X143))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X142),X141),s('type.sptree.spt'('type.clos_known.val_approx'),X140)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f89,f3076])).
fof(f3226,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X126,X127,X128),sK6(X126,X127,X128),X129,X127,X128)),s('type.closSem.v',sK7(X126,X127,X128)))) | s(bool,X131) != s(bool,X132) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X128),X127),s('type.sptree.spt'('type.clos_known.val_approx'),X129))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X128),X127),s('type.sptree.spt'('type.clos_known.val_approx'),X126)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3130,f120])).
fof(f3130,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X126,X127,X128),sK6(X126,X127,X128),X129,X127,X128)),s('type.closSem.v',sK7(X126,X127,X128)))) | s(bool,X131) != s(bool,X132) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X128),X127),s('type.sptree.spt'('type.clos_known.val_approx'),X129))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X128),X127),s('type.sptree.spt'('type.clos_known.val_approx'),X126)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f87,f3076])).
fof(f3223,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X105,X106,X107),sK6(X105,X106,X107),X108,X106,X107)),s('type.closSem.v',sK7(X105,X106,X107))))) | ~p(s(bool,X110)) | s(bool,X110) != s(bool,X111) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X107),X106),s('type.sptree.spt'('type.clos_known.val_approx'),X108))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X107),X106),s('type.sptree.spt'('type.clos_known.val_approx'),X105)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3127,f200])).
fof(f3127,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X105,X106,X107),sK6(X105,X106,X107),X108,X106,X107)),s('type.closSem.v',sK7(X105,X106,X107))))) | ~p(s(bool,X110)) | ~p(s(bool,X109)) | s(bool,X110) != s(bool,X111) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X107),X106),s('type.sptree.spt'('type.clos_known.val_approx'),X108))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X107),X106),s('type.sptree.spt'('type.clos_known.val_approx'),X105)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f87,f3076])).
fof(f3220,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X77,X78,X79),sK6(X77,X78,X79),X80,X78,X79)),s('type.closSem.v',sK7(X77,X78,X79)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X80))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) ) | ($spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3123,f120])).
fof(f3123,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,X82) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X77,X78,X79),sK6(X77,X78,X79),X80,X78,X79)),s('type.closSem.v',sK7(X77,X78,X79)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X80))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f85,f3076])).
fof(f3218,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X40)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X36,X37),sK6(X35,X36,X37),X38,X36,X37)),s('type.closSem.v',sK7(X35,X36,X37))))) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) ) | ($spl10 | $spl24 | $spl34)),
  inference(subsumption_resolution,[],[f3117,f200])).
fof(f3117,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X40)) | ~p(s(bool,X39)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X36,X37),sK6(X35,X36,X37),X38,X36,X37)),s('type.closSem.v',sK7(X35,X36,X37))))) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X38))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X37),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) ) | ($spl24 | $spl34)),
  inference(superposition,[],[f85,f3076])).
fof(f3093,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))))) ) | $spl158),
  inference(cnf_transformation,[],[f3093_D])).
fof(f3093_D,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f3092,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6))))) ) | ($spl10 | $spl152)),
  inference(subsumption_resolution,[],[f3085,f200])).
fof(f3085,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,X7,X8),sK6(X6,X7,X8),X9,X7,X8)),s('type.closSem.v',sK7(X6,X7,X8)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X8),X7),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl152),
  inference(superposition,[],[f3032,f115])).
fof(f3081,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X17,X18,X19),sK6(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK7(X17,X18,X19)))) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) | $spl156),
  inference(cnf_transformation,[],[f3081_D])).
fof(f3081_D,plain,(
  ( ! [X19,X17,X22,X20,X18] : (s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X17,X18,X19),sK6(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK7(X17,X18,X19)))) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f3075,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X73,X74,X75),sK6(X73,X74,X75),X76,X74,X75)),s('type.closSem.v',sK7(X73,X74,X75)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X75),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) | p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X78) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X75),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X73)))) ) | $spl34),
  inference(resolution,[],[f3049,f119])).
fof(f3074,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X66,X67,X68),sK6(X66,X67,X68),X69,X67,X68)),s('type.closSem.v',sK7(X66,X67,X68)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X68),X67),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X68),X67),s('type.sptree.spt'('type.clos_known.val_approx'),X66))) | s(bool,X71) != s(bool,X72)) ) | $spl34),
  inference(resolution,[],[f3049,f91])).
fof(f3071,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X45,X46,X47),sK6(X45,X46,X47),X48,X46,X47)),s('type.closSem.v',sK7(X45,X46,X47)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X47),X46),s('type.sptree.spt'('type.clos_known.val_approx'),X48))) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X47),X46),s('type.sptree.spt'('type.clos_known.val_approx'),X45)))) ) | $spl34),
  inference(resolution,[],[f3049,f90])).
fof(f3070,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,X39,X40),sK6(X38,X39,X40),X41,X39,X40)),s('type.closSem.v',sK7(X38,X39,X40)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X40),X39),s('type.sptree.spt'('type.clos_known.val_approx'),X41))) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X42) != s(bool,X44) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X40),X39),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) | $spl34),
  inference(resolution,[],[f3049,f89])).
fof(f3069,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X31,X32,X33),sK6(X31,X32,X33),X34,X32,X33)),s('type.closSem.v',sK7(X31,X32,X33)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X32),s('type.sptree.spt'('type.clos_known.val_approx'),X31))) | s(bool,X36) != s(bool,X37)) ) | $spl34),
  inference(resolution,[],[f3049,f87])).
fof(f3067,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X17,X18,X19),sK6(X17,X18,X19),X20,X18,X19)),s('type.closSem.v',sK7(X17,X18,X19)))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X20))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X19),X18),s('type.sptree.spt'('type.clos_known.val_approx'),X17)))) ) | $spl34),
  inference(resolution,[],[f3049,f85])).
fof(f3061,plain,(
  ( ! [X14,X12,X13,X11,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | ~p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X13,X12,X11),sK6(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK7(X13,X12,X11))))) ) | $spl154),
  inference(cnf_transformation,[],[f3061_D])).
fof(f3061_D,plain,(
  ( ! [X14,X12,X13,X11,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | ~p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X13,X12,X11),sK6(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK7(X13,X12,X11))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f3047,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X67),X68),s('type.sptree.spt'('type.clos_known.val_approx'),X69)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X69,X68,X67),sK6(X69,X68,X67),X70,X68,X67)),s('type.closSem.v',sK7(X69,X68,X67)))) | p(s(bool,X71)) | ~p(s(bool,X72)) | s(bool,X72) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X67),X68),s('type.sptree.spt'('type.clos_known.val_approx'),X70)))) )),
  inference(resolution,[],[f3035,f119])).
fof(f3046,plain,(
  ( ! [X61,X66,X64,X62,X60,X65,X63] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X62,X61,X60),sK6(X62,X61,X60),X63,X61,X60)),s('type.closSem.v',sK7(X62,X61,X60)))) | ~p(s(bool,X64)) | ~p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63))) | s(bool,X64) != s(bool,X66)) )),
  inference(resolution,[],[f3035,f91])).
fof(f3044,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X47),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X48,X47,X46),sK6(X48,X47,X46),X49,X47,X46)),s('type.closSem.v',sK7(X48,X47,X46)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X46),X47),s('type.sptree.spt'('type.clos_known.val_approx'),X49)))) )),
  inference(resolution,[],[f3035,f90])).
fof(f3042,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X34,X33,X32),sK6(X34,X33,X32),X35,X33,X32)),s('type.closSem.v',sK7(X34,X33,X32)))) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) )),
  inference(resolution,[],[f3035,f89])).
fof(f3041,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,X26,X25),sK6(X27,X26,X25),X28,X26,X25)),s('type.closSem.v',sK7(X27,X26,X25)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f3035,f87])).
fof(f3039,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X13,X12,X11),sK6(X13,X12,X11),X14,X12,X11)),s('type.closSem.v',sK7(X13,X12,X11)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X11),X12),s('type.sptree.spt'('type.clos_known.val_approx'),X14)))) )),
  inference(resolution,[],[f3035,f85])).
fof(f3032,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X7,X5,X4),sK6(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK7(X7,X5,X4)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,X9)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))))) ) | $spl152),
  inference(cnf_transformation,[],[f3032_D])).
fof(f3032_D,plain,(
  ( ! [X6,X4,X7,X5,X9] : (s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X7,X5,X4),sK6(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK7(X7,X5,X4)))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,X9)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f3026,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X62)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X60),X61),s('type.sptree.spt'('type.clos_known.val_approx'),X63)))) | p(s(bool,X64)) | p(s(bool,X65)) | s(bool,X65) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X63,X61,X60),sK6(X63,X61,X60),X62,X61,X60)),s('type.closSem.v',sK7(X63,X61,X60))))) )),
  inference(resolution,[],[f3016,f119])).
fof(f3025,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X55)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X53),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X56)))) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X56,X54,X53),sK6(X56,X54,X53),X55,X54,X53)),s('type.closSem.v',sK7(X56,X54,X53)))) | s(bool,X58) != s(bool,X59)) )),
  inference(resolution,[],[f3016,f91])).
fof(f3022,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X34)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X32),X33),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X35,X33,X32),sK6(X35,X33,X32),X34,X33,X32)),s('type.closSem.v',sK7(X35,X33,X32))))) )),
  inference(resolution,[],[f3016,f90])).
fof(f3021,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X28)))) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X28,X26,X25),sK6(X28,X26,X25),X27,X26,X25)),s('type.closSem.v',sK7(X28,X26,X25))))) )),
  inference(resolution,[],[f3016,f89])).
fof(f3020,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X20)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X18),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X21)))) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X21,X19,X18),sK6(X21,X19,X18),X20,X19,X18)),s('type.closSem.v',sK7(X21,X19,X18)))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f3016,f87])).
fof(f3018,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X4),X5),s('type.sptree.spt'('type.clos_known.val_approx'),X7)))) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X7,X5,X4),sK6(X7,X5,X4),X6,X5,X4)),s('type.closSem.v',sK7(X7,X5,X4))))) )),
  inference(resolution,[],[f3016,f85])).
fof(f3000,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X28) != s(bool,X29) | s(bool,X25) != s(bool,X26)) ) | $spl150),
  inference(cnf_transformation,[],[f3000_D])).
fof(f3000_D,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,X27) != s(bool,X28) | s(bool,X29) != s(bool,X30) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X28) != s(bool,X29) | s(bool,X25) != s(bool,X26)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f2999,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X24)) | s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) != s(bool,X30) | s(bool,X26) != s(bool,X27)) ) | ($spl10 | $spl78)),
  inference(subsumption_resolution,[],[f2993,f200])).
fof(f2993,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X24)) | s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) != s(bool,X30) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X24)) ) | $spl78),
  inference(superposition,[],[f1354,f115])).
fof(f2997,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1,X9] : (p(s(bool,X9)) | s(bool,X1) != s(bool,X9) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl148),
  inference(cnf_transformation,[],[f2997_D])).
fof(f2997_D,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1,X9] : (p(s(bool,X9)) | s(bool,X1) != s(bool,X9) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f2984,plain,(
  ( ! [X94,X92,X90,X88,X97,X95,X93,X91,X89,X96] : (s(bool,X88) != s(bool,X89) | s(bool,X88) != s(bool,X90) | s(bool,X91) != s(bool,X92) | s(bool,X92) != s(bool,X93) | s(bool,X94) != s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,X93) != s(bool,X94) | s(bool,X90) != s(bool,X91) | p(s(bool,X96)) | p(s(bool,X97)) | s(bool,X89) != s(bool,X97)) ) | $spl78),
  inference(resolution,[],[f1354,f119])).
fof(f2983,plain,(
  ( ! [X80,X78,X87,X85,X83,X81,X79,X77,X86,X84,X82] : (s(bool,X77) != s(bool,X78) | s(bool,X77) != s(bool,X79) | s(bool,X80) != s(bool,X81) | s(bool,X81) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,'F_0') = s(bool,X84) | s(bool,X82) != s(bool,X83) | s(bool,X79) != s(bool,X80) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X78) != s(bool,X85) | s(bool,X86) != s(bool,X87)) ) | $spl78),
  inference(resolution,[],[f1354,f91])).
fof(f2980,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X44,X53,X51,X49] : (s(bool,X44) != s(bool,X45) | s(bool,X44) != s(bool,X46) | s(bool,X47) != s(bool,X48) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,X51) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) != s(bool,X50) | s(bool,X46) != s(bool,X47) | p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X45) != s(bool,X54)) ) | $spl78),
  inference(resolution,[],[f1354,f90])).
fof(f2979,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X42,X40] : (s(bool,X33) != s(bool,X34) | s(bool,X33) != s(bool,X35) | s(bool,X36) != s(bool,X37) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X38) != s(bool,X39) | s(bool,X35) != s(bool,X36) | p(s(bool,X41)) | p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X34) != s(bool,X42)) ) | $spl78),
  inference(resolution,[],[f1354,f89])).
fof(f2978,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32,X22] : (s(bool,X22) != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,X25) != s(bool,X26) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,X27) != s(bool,X28) | s(bool,X24) != s(bool,X25) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X23) != s(bool,X30) | s(bool,X31) != s(bool,X32)) ) | $spl78),
  inference(resolution,[],[f1354,f87])).
fof(f2976,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X1) != s(bool,X9)) ) | $spl78),
  inference(resolution,[],[f1354,f85])).
fof(f2934,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | $spl146),
  inference(cnf_transformation,[],[f2934_D])).
fof(f2934_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f2933,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl136)),
  inference(subsumption_resolution,[],[f2926,f472])).
fof(f2926,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl136),
  inference(superposition,[],[f2832,f115])).
fof(f2911,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) ) | $spl144),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f2910,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) ) | ($spl10 | $spl132)),
  inference(subsumption_resolution,[],[f2901,f200])).
fof(f2901,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl132),
  inference(superposition,[],[f2737,f115])).
fof(f2887,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))))) ) | $spl142),
  inference(cnf_transformation,[],[f2887_D])).
fof(f2887_D,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f2886,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,X2))) ) | ($spl10 | $spl130)),
  inference(subsumption_resolution,[],[f2880,f200])).
fof(f2880,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl130),
  inference(superposition,[],[f2720,f115])).
fof(f2832,plain,(
  ( ! [X29,X27] : (s(bool,X29) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,sK2,sK1),sK6(X27,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X27,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X27))) | p(s(bool,X29))) ) | $spl136),
  inference(cnf_transformation,[],[f2832_D])).
fof(f2832_D,plain,(
  ( ! [X29,X27] : (s(bool,X29) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,sK2,sK1),sK6(X27,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X27,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X27))) | p(s(bool,X29))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f2872,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X33))) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X33,sK2,sK1),sK6(X33,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X33,sK2,sK1))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f119])).
fof(f2871,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X29,sK2,sK1),sK6(X29,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X29,sK2,sK1)))) | s(bool,X31) != s(bool,X32)) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f91])).
fof(f2868,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X17))) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X17,sK2,sK1),sK6(X17,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X17,sK2,sK1))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f90])).
fof(f2867,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X13,sK2,sK1),sK6(X13,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X13,sK2,sK1))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f89])).
fof(f2866,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X9,sK2,sK1),sK6(X9,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X9,sK2,sK1)))) | s(bool,X11) != s(bool,X12)) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f87])).
fof(f2864,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X1,sK2,sK1),sK6(X1,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X1,sK2,sK1))))) ) | ($spl2 | $spl10 | $spl24)),
  inference(resolution,[],[f2825,f85])).
fof(f2857,plain,(
  ( ! [X191,X194,X193] : (s(bool,X193) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X191,sK2,sK1),sK6(X191,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X191,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X191))) | s(bool,'T_0') = s(bool,X194) | s(bool,X193) != s(bool,X194)) ) | $spl140),
  inference(cnf_transformation,[],[f2857_D])).
fof(f2857_D,plain,(
  ( ! [X191,X194,X193] : (s(bool,X193) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X191,sK2,sK1),sK6(X191,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X191,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X191))) | s(bool,'T_0') = s(bool,X194) | s(bool,X193) != s(bool,X194)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2861,plain,(
  ( ! [X261,X263,X259,X260,X262,X258,X264] : (s(bool,X260) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X258,sK2,sK1),sK6(X258,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X258,sK2,sK1)))) | s(bool,X259) != s(bool,X261) | s(bool,X261) != s(bool,X262) | s(bool,'F_0') = s(bool,X263) | s(bool,X262) != s(bool,X263) | s(bool,X260) != s(bool,X264) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X258)))) ) | ($spl2 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f2814,f1020])).
fof(f2814,plain,(
  ( ! [X261,X263,X259,X260,X262,X258,X264] : (s(bool,'T_0') != s(bool,X259) | s(bool,X260) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X258,sK2,sK1),sK6(X258,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X258,sK2,sK1)))) | s(bool,X259) != s(bool,X261) | s(bool,X261) != s(bool,X262) | s(bool,'F_0') = s(bool,X263) | s(bool,X262) != s(bool,X263) | s(bool,X260) != s(bool,X264) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X258)))) ) | ($spl2 | $spl24 | $spl74)),
  inference(superposition,[],[f1343,f2733])).
fof(f2860,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (s(bool,X236) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X234,sK2,sK1),sK6(X234,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X234,sK2,sK1)))) | s(bool,X235) != s(bool,X237) | s(bool,X237) != s(bool,X238) | s(bool,'F_0') = s(bool,X238) | s(bool,X236) != s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X234)))) ) | ($spl2 | $spl24 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f2810,f718])).
fof(f2810,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (s(bool,'T_0') != s(bool,X235) | s(bool,X236) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X234,sK2,sK1),sK6(X234,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X234,sK2,sK1)))) | s(bool,X235) != s(bool,X237) | s(bool,X237) != s(bool,X238) | s(bool,'F_0') = s(bool,X238) | s(bool,X236) != s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X234)))) ) | ($spl2 | $spl24 | $spl68)),
  inference(superposition,[],[f1276,f2733])).
fof(f2859,plain,(
  ( ! [X218,X216,X215,X217,X214] : (s(bool,X216) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X214,sK2,sK1),sK6(X214,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X214,sK2,sK1)))) | s(bool,X215) != s(bool,X217) | s(bool,'F_0') = s(bool,X217) | s(bool,X216) != s(bool,X218) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X214)))) ) | ($spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f2806,f655])).
fof(f2806,plain,(
  ( ! [X218,X216,X215,X217,X214] : (s(bool,'T_0') != s(bool,X215) | s(bool,X216) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X214,sK2,sK1),sK6(X214,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X214,sK2,sK1)))) | s(bool,X215) != s(bool,X217) | s(bool,'F_0') = s(bool,X217) | s(bool,X216) != s(bool,X218) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X214)))) ) | ($spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1171,f2733])).
fof(f2830,plain,(
  ( ! [X3] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | $spl134),
  inference(cnf_transformation,[],[f2830_D])).
fof(f2830_D,plain,(
  ( ! [X3] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2858,plain,(
  ( ! [X210,X208,X209,X207] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X210,X209,X208)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X210,X209,X208)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X208),X209))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X207,sK2,sK1),sK6(X207,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X207,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X207)))) ) | ($spl2 | $spl24 | $spl58)),
  inference(subsumption_resolution,[],[f2804,f1062])).
fof(f2804,plain,(
  ( ! [X210,X208,X209,X207] : (s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X208),X209),s('type.sptree.spt'('type.clos_known.val_approx'),X210))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X210,X209,X208)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X210,X209,X208)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X208),X209))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X207,sK2,sK1),sK6(X207,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X207,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X207)))) ) | ($spl2 | $spl24 | $spl58)),
  inference(superposition,[],[f1061,f2733])).
fof(f2856,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,X193) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X191,sK2,sK1),sK6(X191,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X191,sK2,sK1)))) | s(bool,'F_0') = s(bool,X192) | s(bool,X193) != s(bool,X194) | s(bool,'T_0') = s(bool,X194) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X191)))) ) | ($spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f2800,f115])).
fof(f2800,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,'T_0') != s(bool,X192) | s(bool,X193) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X191,sK2,sK1),sK6(X191,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X191,sK2,sK1)))) | s(bool,'F_0') = s(bool,X192) | s(bool,X193) != s(bool,X194) | s(bool,'T_0') = s(bool,X194) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X191)))) ) | ($spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1020,f2733])).
fof(f2855,plain,(
  ( ! [X177,X179] : (s(bool,X179) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X177,sK2,sK1),sK6(X177,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X177,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X177))) | s(bool,'T_0') = s(bool,X179)) ) | $spl138),
  inference(cnf_transformation,[],[f2855_D])).
fof(f2855_D,plain,(
  ( ! [X177,X179] : (s(bool,X179) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X177,sK2,sK1),sK6(X177,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X177,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X177))) | s(bool,'T_0') = s(bool,X179)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2854,plain,(
  ( ! [X177,X178,X179] : (s(bool,X179) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X177,sK2,sK1),sK6(X177,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X177,sK2,sK1)))) | s(bool,'F_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X179) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X177)))) ) | ($spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f2795,f115])).
fof(f2795,plain,(
  ( ! [X177,X178,X179] : (s(bool,'T_0') != s(bool,X178) | s(bool,X179) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X177,sK2,sK1),sK6(X177,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X177,sK2,sK1)))) | s(bool,'F_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X179) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X177)))) ) | ($spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f718,f2733])).
fof(f2851,plain,(
  ( ! [X158,X157] : (s(bool,'T_0') != s(bool,X158) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X157,sK2,sK1),sK6(X157,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X157,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X157)))) ) | ($spl2 | $spl10 | $spl14 | $spl24)),
  inference(subsumption_resolution,[],[f2784,f200])).
fof(f2784,plain,(
  ( ! [X158,X157] : (s(bool,'T_0') != s(bool,X158) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X157,sK2,sK1),sK6(X157,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X157,sK2,sK1))))) | ~p(s(bool,X158)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X157)))) ) | ($spl2 | $spl14 | $spl24)),
  inference(superposition,[],[f205,f2733])).
fof(f2850,plain,(
  ( ! [X154,X155,X153] : (p(s(bool,X154)) | p(s(bool,X155)) | s(bool,X155) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X153,sK2,sK1),sK6(X153,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X153,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X153)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2782,f120])).
fof(f2782,plain,(
  ( ! [X154,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,X154)) | p(s(bool,X155)) | s(bool,X155) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X153,sK2,sK1),sK6(X153,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X153,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X153)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f119,f2733])).
fof(f2849,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') != s(bool,X145) | p(s(bool,X146)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X144,sK2,sK1),sK6(X144,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X144,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X144)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2779,f200])).
fof(f2779,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') != s(bool,X145) | p(s(bool,X146)) | ~p(s(bool,X145)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X144,sK2,sK1),sK6(X144,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X144,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X144)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f119,f2733])).
fof(f2848,plain,(
  ( ! [X138,X136,X135,X137] : (p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X135,sK2,sK1),sK6(X135,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X135,sK2,sK1)))) | s(bool,X137) != s(bool,X138) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X135)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2776,f120])).
fof(f2776,plain,(
  ( ! [X138,X136,X135,X137] : (~p(s(bool,'T_0')) | p(s(bool,X136)) | ~p(s(bool,X137)) | s(bool,X136) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X135,sK2,sK1),sK6(X135,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X135,sK2,sK1)))) | s(bool,X137) != s(bool,X138) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X135)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f91,f2733])).
fof(f2845,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,X120) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X119,sK2,sK1),sK6(X119,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X119,sK2,sK1))))) | ~p(s(bool,X121)) | s(bool,X121) != s(bool,X122) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X119)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2772,f200])).
fof(f2772,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,X120) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X119,sK2,sK1),sK6(X119,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X119,sK2,sK1))))) | ~p(s(bool,X121)) | ~p(s(bool,X120)) | s(bool,X121) != s(bool,X122) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X119)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f91,f2733])).
fof(f2842,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | ~p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X106) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X103,sK2,sK1),sK6(X103,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X103,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X103)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2768,f120])).
fof(f2768,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | ~p(s(bool,X105)) | p(s(bool,X106)) | s(bool,X106) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X103,sK2,sK1),sK6(X103,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X103,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X103)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f90,f2733])).
fof(f2840,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X94)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X91,sK2,sK1),sK6(X91,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X91,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X91)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2765,f200])).
fof(f2765,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,X93)) | ~p(s(bool,X92)) | ~p(s(bool,X94)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X91,sK2,sK1),sK6(X91,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X91,sK2,sK1))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X91)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f90,f2733])).
fof(f2839,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X84) != s(bool,X86) | s(bool,X85) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X83,sK2,sK1),sK6(X83,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X83,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X83)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2763,f120])).
fof(f2763,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X84) != s(bool,X86) | s(bool,X85) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X83,sK2,sK1),sK6(X83,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X83,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X83)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f89,f2733])).
fof(f2838,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | p(s(bool,X65)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X63,sK2,sK1),sK6(X63,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X63,sK2,sK1))))) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X63)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2758,f200])).
fof(f2758,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | p(s(bool,X65)) | ~p(s(bool,X64)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X63,sK2,sK1),sK6(X63,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X63,sK2,sK1))))) | s(bool,X65) != s(bool,X66) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X63)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f89,f2733])).
fof(f2837,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X55,sK2,sK1),sK6(X55,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X55,sK2,sK1)))) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X55)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2756,f120])).
fof(f2756,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X55,sK2,sK1),sK6(X55,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X55,sK2,sK1)))) | s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X55)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f87,f2733])).
fof(f2834,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X43,sK2,sK1),sK6(X43,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X43,sK2,sK1))))) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2753,f200])).
fof(f2753,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X43,sK2,sK1),sK6(X43,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X43,sK2,sK1))))) | ~p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f87,f2733])).
fof(f2831,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,sK2,sK1),sK6(X27,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X27,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f2749,f120])).
fof(f2749,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X27,sK2,sK1),sK6(X27,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X27,sK2,sK1)))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f85,f2733])).
fof(f2829,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) | s(bool,X5) != s(bool,X6) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl2 | $spl10 | $spl24)),
  inference(subsumption_resolution,[],[f2743,f200])).
fof(f2743,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X4)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X3,sK2,sK1),sK6(X3,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X3,sK2,sK1))))) | s(bool,X5) != s(bool,X6) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f85,f2733])).
fof(f2737,plain,(
  ( ! [X6,X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,sK2,sK1),sK6(X6,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X6,sK2,sK1))))) ) | $spl132),
  inference(cnf_transformation,[],[f2737_D])).
fof(f2737_D,plain,(
  ( ! [X6,X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X6))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,sK2,sK1),sK6(X6,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X6,sK2,sK1))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f2732,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X38,sK2,sK1),sK6(X38,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X38,sK2,sK1)))) | p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X38)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f119])).
fof(f2731,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X34,sK2,sK1),sK6(X34,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X34,sK2,sK1)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X34))) | s(bool,X36) != s(bool,X37)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f91])).
fof(f2728,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X22,sK2,sK1),sK6(X22,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X22,sK2,sK1)))) | p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f90])).
fof(f2727,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X18,sK2,sK1),sK6(X18,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X18,sK2,sK1)))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f89])).
fof(f2726,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X14,sK2,sK1),sK6(X14,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X14,sK2,sK1)))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X14))) | s(bool,X16) != s(bool,X17)) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f87])).
fof(f2724,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X6,sK2,sK1),sK6(X6,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X6,sK2,sK1)))) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X6)))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f2716,f85])).
fof(f2720,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X1,sK2,sK1),sK6(X1,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X1,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X3))) ) | $spl130),
  inference(cnf_transformation,[],[f2720_D])).
fof(f2720_D,plain,(
  ( ! [X3,X1] : (s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X1,sK2,sK1),sK6(X1,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X1,sK2,sK1)))) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X3))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2715,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X33)))) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X33,sK2,sK1),sK6(X33,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X33,sK2,sK1))))) ) | $spl2),
  inference(resolution,[],[f2705,f119])).
fof(f2714,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X29,sK2,sK1),sK6(X29,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X29,sK2,sK1)))) | s(bool,X31) != s(bool,X32)) ) | $spl2),
  inference(resolution,[],[f2705,f91])).
fof(f2711,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X17)))) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X17,sK2,sK1),sK6(X17,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X17,sK2,sK1))))) ) | $spl2),
  inference(resolution,[],[f2705,f90])).
fof(f2710,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X13,sK2,sK1),sK6(X13,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X13,sK2,sK1))))) ) | $spl2),
  inference(resolution,[],[f2705,f89])).
fof(f2709,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X9,sK2,sK1),sK6(X9,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X9,sK2,sK1)))) | s(bool,X11) != s(bool,X12)) ) | $spl2),
  inference(resolution,[],[f2705,f87])).
fof(f2707,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(X1,sK2,sK1),sK6(X1,sK2,sK1),sK3,sK2,sK1)),s('type.closSem.v',sK7(X1,sK2,sK1))))) ) | $spl2),
  inference(resolution,[],[f2705,f85])).
fof(f1594,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | $spl100),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2687,plain,(
  ( ! [X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | s(bool,'F_0') != s(bool,X1)) ) | ($spl0 | $spl24 | $spl58)),
  inference(forward_demodulation,[],[f2686,f97])).
fof(f2686,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)))))))) ) | ($spl0 | $spl24 | $spl58)),
  inference(subsumption_resolution,[],[f2680,f304])).
fof(f2680,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | p(s(bool,X1))) ) | ($spl0 | $spl58)),
  inference(superposition,[],[f1061,f124])).
fof(f1354,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,X8)) | s(bool,X1) != s(bool,X8) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f1354_D])).
fof(f1354_D,plain,(
  ( ! [X6,X4,X2,X0,X8,X5,X3,X1] : (p(s(bool,X8)) | s(bool,X1) != s(bool,X8) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2646,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X86) | s(bool,X85) != s(bool,X86) | s(bool,X82) != s(bool,X83) | p(s(bool,X87)) | p(s(bool,X88)) | s(bool,X81) != s(bool,X88)) ) | $spl74),
  inference(resolution,[],[f1287,f119])).
fof(f2645,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X70) != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X73) != s(bool,X74) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) != s(bool,X76) | s(bool,X72) != s(bool,X73) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X71) != s(bool,X77) | s(bool,X78) != s(bool,X79)) ) | $spl74),
  inference(resolution,[],[f1287,f91])).
fof(f2642,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) != s(bool,X46) | s(bool,X42) != s(bool,X43) | p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X41) != s(bool,X49)) ) | $spl74),
  inference(resolution,[],[f1287,f90])).
fof(f2641,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) != s(bool,X36) | s(bool,X32) != s(bool,X33) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X31) != s(bool,X38)) ) | $spl74),
  inference(resolution,[],[f1287,f89])).
fof(f2640,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X21) != s(bool,X27) | s(bool,X28) != s(bool,X29)) ) | $spl74),
  inference(resolution,[],[f1287,f87])).
fof(f2638,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X1) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f1287,f85])).
fof(f2548,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | $spl128),
  inference(cnf_transformation,[],[f2548_D])).
fof(f2548_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f497,plain,(
  ( ! [X4,X7] : (s(bool,X4) != s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) | $spl30),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  ( ! [X4,X7] : (s(bool,X4) != s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2637,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl24 | $spl88)),
  inference(subsumption_resolution,[],[f2634,f655])).
fof(f2634,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | $spl88),
  inference(superposition,[],[f1566,f115])).
fof(f2545,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | $spl127),
  inference(cnf_transformation,[],[f2545_D])).
fof(f2545_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f2636,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl48 | $spl88)),
  inference(subsumption_resolution,[],[f2632,f766])).
fof(f2632,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl88),
  inference(superposition,[],[f1566,f115])).
fof(f1287,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) | $spl74),
  inference(cnf_transformation,[],[f1287_D])).
fof(f1287_D,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,X7)) | s(bool,X1) != s(bool,X7) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2595,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,X72) != s(bool,X74) | s(bool,X75) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X74) != s(bool,X75) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X73) != s(bool,X79)) ) | $spl68),
  inference(resolution,[],[f1181,f119])).
fof(f2594,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl68),
  inference(resolution,[],[f1181,f91])).
fof(f2591,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl68),
  inference(resolution,[],[f1181,f90])).
fof(f2590,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl68),
  inference(resolution,[],[f1181,f89])).
fof(f2589,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl68),
  inference(resolution,[],[f1181,f87])).
fof(f2587,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl68),
  inference(resolution,[],[f1181,f85])).
fof(f1213,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | ~p(s(bool,X17)) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X13,X11,X12,X14,X15))))) ) | $spl72),
  inference(cnf_transformation,[],[f1213_D])).
fof(f1213_D,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | ~p(s(bool,X17)) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X13,X11,X12,X14,X15))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2568,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X75),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X77,X75,X76,X78,X79)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X77))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X75),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X76)))) )),
  inference(resolution,[],[f103,f119])).
fof(f2567,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X67),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X69,X67,X68,X70,X71)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X69))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X67),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f103,f91])).
fof(f2565,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X51),s('type.sptree.spt'('type.clos_known.val_approx'),X52))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X53,X51,X52,X54,X55)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X53))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X51),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) )),
  inference(resolution,[],[f103,f90])).
fof(f2563,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X35),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X37,X35,X36,X38,X39)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X37))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X35),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) )),
  inference(resolution,[],[f103,f89])).
fof(f2562,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X27),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X29,X27,X28,X30,X31)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X29))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X27),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f103,f87])).
fof(f2560,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X13,X11,X12,X14,X15)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12)))) )),
  inference(resolution,[],[f103,f85])).
fof(f2557,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | $spl86),
  inference(subsumption_resolution,[],[f2554,f115])).
fof(f2554,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | $spl86),
  inference(superposition,[],[f1564,f115])).
fof(f2556,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl86),
  inference(duplicate_literal_removal,[],[f2552])).
fof(f2552,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl86),
  inference(superposition,[],[f1564,f115])).
fof(f2546,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl24 | $spl84)),
  inference(subsumption_resolution,[],[f2541,f304])).
fof(f2541,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | $spl84),
  inference(superposition,[],[f1436,f115])).
fof(f2543,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | p(s(bool,X1))) ) | ($spl10 | $spl84)),
  inference(subsumption_resolution,[],[f2539,f200])).
fof(f2539,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl84),
  inference(superposition,[],[f1436,f115])).
fof(f1135,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | ~p(s(bool,X17)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X11,X12,X13,X14,X15)),s('type.closSem.v',X11))))) ) | $spl66),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ( ! [X14,X12,X17,X15,X13,X11] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13))) | ~p(s(bool,X17)) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X11,X12,X13,X14,X15)),s('type.closSem.v',X11))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2515,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X75,X76,X77,X78,X79)),s('type.closSem.v',X75)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X75))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X76),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) )),
  inference(resolution,[],[f104,f119])).
fof(f2514,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X67,X68,X69,X70,X71)),s('type.closSem.v',X67)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X67))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X68),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f104,f91])).
fof(f2512,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X51,X52,X53,X54,X55)),s('type.closSem.v',X51)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X51))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X52),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X53)))) )),
  inference(resolution,[],[f104,f90])).
fof(f2510,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X35,X36,X37,X38,X39)),s('type.closSem.v',X35)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X35))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X36),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) )),
  inference(resolution,[],[f104,f89])).
fof(f2509,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X27,X28,X29,X30,X31)),s('type.closSem.v',X27)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X27))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X28),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f104,f87])).
fof(f2507,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X11,X12,X13,X14,X15)),s('type.closSem.v',X11)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) )),
  inference(resolution,[],[f104,f85])).
fof(f1131,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11))))) ) | $spl64),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  ( ! [X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X8,X11,X10,X9)),s('type.closSem.v',X7))) | p(s(bool,X13)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2503,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X71))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X72),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X73),X74))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X73),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X75)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X71,X72,X75,X74,X73)),s('type.closSem.v',X71)))) )),
  inference(resolution,[],[f104,f119])).
fof(f2502,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X63))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X64),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X65),X66))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X65),X66),s('type.sptree.spt'('type.clos_known.val_approx'),X67)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X63,X64,X67,X66,X65)),s('type.closSem.v',X63))) | s(bool,X69) != s(bool,X70)) )),
  inference(resolution,[],[f104,f91])).
fof(f2499,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X39))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X40),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X41),X42))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X42),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X39,X40,X43,X42,X41)),s('type.closSem.v',X39)))) )),
  inference(resolution,[],[f104,f90])).
fof(f2498,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X31))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X32),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X34),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X31,X32,X35,X34,X33)),s('type.closSem.v',X31)))) )),
  inference(resolution,[],[f104,f89])).
fof(f2497,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X23))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X24),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X25),X26))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X23,X24,X27,X26,X25)),s('type.closSem.v',X23))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f104,f87])).
fof(f2495,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X8,X11,X10,X9)),s('type.closSem.v',X7)))) )),
  inference(resolution,[],[f104,f85])).
fof(f1181,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) | $spl68),
  inference(cnf_transformation,[],[f1181_D])).
fof(f1181_D,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,X6)) | s(bool,X1) != s(bool,X6) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2474,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,X64) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X67) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X65) != s(bool,X70)) ) | $spl54),
  inference(resolution,[],[f1029,f119])).
fof(f2473,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl54),
  inference(resolution,[],[f1029,f91])).
fof(f2470,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl54),
  inference(resolution,[],[f1029,f90])).
fof(f2469,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl54),
  inference(resolution,[],[f1029,f89])).
fof(f2468,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f1029,f87])).
fof(f2466,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl54),
  inference(resolution,[],[f1029,f85])).
fof(f1061,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X4))) ) | $spl58),
  inference(cnf_transformation,[],[f1061_D])).
fof(f1061_D,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X4))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2447,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X48,X49,X50)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X48,X49,X50)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X50),X49))))) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X50),X49),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) )),
  inference(resolution,[],[f105,f119])).
fof(f2446,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X42,X43,X44)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X42,X43,X44)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X44),X43))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X44),X43),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f105,f91])).
fof(f2443,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X24,X25,X26)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X24,X25,X26)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X26),X25))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X26),X25),s('type.sptree.spt'('type.clos_known.val_approx'),X24)))) )),
  inference(resolution,[],[f105,f90])).
fof(f2442,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X18,X19,X20)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X18,X19,X20)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X20),X19))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) )),
  inference(resolution,[],[f105,f89])).
fof(f2441,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X12,X13,X14)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X12,X13,X14)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X14),X13))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X13),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f105,f87])).
fof(f2439,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0)))) )),
  inference(resolution,[],[f105,f85])).
fof(f1029,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl54),
  inference(cnf_transformation,[],[f1029_D])).
fof(f1029_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,X5)) | s(bool,X1) != s(bool,X5) | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2427,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl40),
  inference(resolution,[],[f729,f119])).
fof(f2426,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f729,f91])).
fof(f2423,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f729,f90])).
fof(f2422,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f729,f89])).
fof(f2421,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f729,f87])).
fof(f2419,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f729,f85])).
fof(f997,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK7(X2,X1,X0)))) | ~p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) ) | $spl52),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK7(X2,X1,X0)))) | ~p(s(bool,X5)) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2389,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X57),s('type.sptree.spt'('type.clos_known.val_approx'),X58)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X59))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X58,X57,X56)),s('type.sptree.spt'('type.clos_known.val_approx'),X58))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X59),s('type.closSem.v',sK7(X58,X57,X56))))) )),
  inference(resolution,[],[f106,f119])).
fof(f2388,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X49),X50),s('type.sptree.spt'('type.clos_known.val_approx'),X51)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X52))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X51,X50,X49)),s('type.sptree.spt'('type.clos_known.val_approx'),X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X52),s('type.closSem.v',sK7(X51,X50,X49)))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f106,f91])).
fof(f2386,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X38))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X37,X36,X35)),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X38),s('type.closSem.v',sK7(X37,X36,X35))))) )),
  inference(resolution,[],[f106,f90])).
fof(f2384,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X21),X22),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X24))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X23,X22,X21)),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X24),s('type.closSem.v',sK7(X23,X22,X21))))) )),
  inference(resolution,[],[f106,f89])).
fof(f2383,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X15),s('type.sptree.spt'('type.clos_known.val_approx'),X16)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X17))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X16,X15,X14)),s('type.sptree.spt'('type.clos_known.val_approx'),X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X17),s('type.closSem.v',sK7(X16,X15,X14)))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f106,f87])).
fof(f2381,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK7(X2,X1,X0))))) )),
  inference(resolution,[],[f106,f85])).
fof(f994,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X5)) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK7(X1,X2,X3)))))) ) | $spl50),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X5)) | ~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK7(X1,X2,X3)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2379,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X56),s('type.closSem.v',sK7(X57,X58,X59))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X56))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X57,X58,X59)),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X59),X58),s('type.sptree.spt'('type.clos_known.val_approx'),X57)))) )),
  inference(resolution,[],[f106,f119])).
fof(f2378,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X49),s('type.closSem.v',sK7(X50,X51,X52))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X49))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X50,X51,X52)),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X52),X51),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f106,f91])).
fof(f2375,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X28),s('type.closSem.v',sK7(X29,X30,X31))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X28))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X29,X30,X31)),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) )),
  inference(resolution,[],[f106,f90])).
fof(f2374,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X21),s('type.closSem.v',sK7(X22,X23,X24))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X21))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X22,X23,X24)),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X24),X23),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) )),
  inference(resolution,[],[f106,f89])).
fof(f2373,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X14),s('type.closSem.v',sK7(X15,X16,X17))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X14))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X15,X16,X17)),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X17),X16),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f106,f87])).
fof(f2371,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK7(X1,X2,X3))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) )),
  inference(resolution,[],[f106,f85])).
fof(f2370,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl124),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2369,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl112)),
  inference(subsumption_resolution,[],[f2361,f115])).
fof(f2361,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl112)),
  inference(superposition,[],[f1604,f124])).
fof(f2368,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | $spl112),
  inference(subsumption_resolution,[],[f2359,f1604])).
fof(f2359,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl112),
  inference(superposition,[],[f1604,f115])).
fof(f2367,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl112),
  inference(subsumption_resolution,[],[f2358,f1604])).
fof(f2358,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | $spl112),
  inference(superposition,[],[f1604,f115])).
fof(f2366,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(subsumption_resolution,[],[f2357,f115])).
fof(f2357,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(superposition,[],[f1604,f115])).
fof(f2364,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) | $spl122),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2363,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(duplicate_literal_removal,[],[f2351])).
fof(f2351,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl112),
  inference(superposition,[],[f1604,f115])).
fof(f1601,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X16) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) | $spl110),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X16) != s(bool,X17) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1604,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4)) ) | $spl112),
  inference(cnf_transformation,[],[f1604_D])).
fof(f1604_D,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2347,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (s(bool,X168) != s(bool,X169) | s(bool,X170) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,X172) | s(bool,X171) != s(bool,X172) | s(bool,X169) != s(bool,X173) | s(bool,'T_0') = s(bool,X173)) ) | ($spl0 | $spl24 | $spl38 | $spl74)),
  inference(subsumption_resolution,[],[f2301,f718])).
fof(f2301,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (s(bool,'F_0') != s(bool,X168) | s(bool,X168) != s(bool,X169) | s(bool,X170) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X170) != s(bool,X171) | s(bool,'F_0') = s(bool,X172) | s(bool,X171) != s(bool,X172) | s(bool,X169) != s(bool,X173) | s(bool,'T_0') = s(bool,X173)) ) | ($spl0 | $spl24 | $spl74)),
  inference(superposition,[],[f1343,f124])).
fof(f1605,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X13)) ) | $spl114),
  inference(cnf_transformation,[],[f1605_D])).
fof(f1605_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2346,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X153) != s(bool,X154) | s(bool,X155) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,X156) | s(bool,X154) != s(bool,X157) | s(bool,'T_0') = s(bool,X157)) ) | ($spl0 | $spl24 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f2298,f718])).
fof(f2298,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'F_0') != s(bool,X153) | s(bool,X153) != s(bool,X154) | s(bool,X155) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X155) != s(bool,X156) | s(bool,'F_0') = s(bool,X156) | s(bool,X154) != s(bool,X157) | s(bool,'T_0') = s(bool,X157)) ) | ($spl0 | $spl24 | $spl68)),
  inference(superposition,[],[f1276,f124])).
fof(f1608,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X8)) ) | $spl116),
  inference(cnf_transformation,[],[f1608_D])).
fof(f1608_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2345,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X141) != s(bool,X142) | s(bool,X143) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X143) | s(bool,X142) != s(bool,X144) | s(bool,'T_0') = s(bool,X144)) ) | ($spl0 | $spl24 | $spl38 | $spl54)),
  inference(subsumption_resolution,[],[f2295,f718])).
fof(f2295,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'F_0') != s(bool,X141) | s(bool,X141) != s(bool,X142) | s(bool,X143) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X143) | s(bool,X142) != s(bool,X144) | s(bool,'T_0') = s(bool,X144)) ) | ($spl0 | $spl24 | $spl54)),
  inference(superposition,[],[f1171,f124])).
fof(f729,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f729_D])).
fof(f729_D,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,X4)) | s(bool,X1) != s(bool,X4) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2223,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl38),
  inference(resolution,[],[f678,f119])).
fof(f2222,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f678,f91])).
fof(f2219,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f678,f90])).
fof(f2218,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f678,f89])).
fof(f2217,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f678,f87])).
fof(f2215,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f678,f85])).
fof(f678,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,X3)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2193,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f516,f119])).
fof(f2192,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f516,f91])).
fof(f2189,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl14),
  inference(resolution,[],[f516,f90])).
fof(f2188,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f516,f89])).
fof(f2187,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f516,f87])).
fof(f2185,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f516,f85])).
fof(f2180,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) != s(bool,X126)) ) | $spl120),
  inference(cnf_transformation,[],[f2180_D])).
fof(f2180_D,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X126) | s(bool,X125) != s(bool,X126)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2183,plain,(
  ( ! [X163,X161,X159,X164,X162,X160] : (s(bool,X160) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X159) != s(bool,X161) | s(bool,X161) != s(bool,X162) | s(bool,'F_0') = s(bool,X163) | s(bool,X162) != s(bool,X163) | s(bool,X160) != s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | ($spl2 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f2137,f1020])).
fof(f2137,plain,(
  ( ! [X163,X161,X159,X164,X162,X160] : (s(bool,'T_0') != s(bool,X159) | s(bool,X160) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X159) != s(bool,X161) | s(bool,X161) != s(bool,X162) | s(bool,'F_0') = s(bool,X163) | s(bool,X162) != s(bool,X163) | s(bool,X160) != s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | ($spl2 | $spl24 | $spl74)),
  inference(superposition,[],[f1343,f125])).
fof(f2182,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,X146) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X145) != s(bool,X147) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,X148) | s(bool,X146) != s(bool,X149) | s(bool,'T_0') = s(bool,X149)) ) | ($spl2 | $spl24 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f2134,f718])).
fof(f2134,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'T_0') != s(bool,X145) | s(bool,X146) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X145) != s(bool,X147) | s(bool,X147) != s(bool,X148) | s(bool,'F_0') = s(bool,X148) | s(bool,X146) != s(bool,X149) | s(bool,'T_0') = s(bool,X149)) ) | ($spl2 | $spl24 | $spl68)),
  inference(superposition,[],[f1276,f125])).
fof(f2181,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X135) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,X136) | s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl2 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f2131,f655])).
fof(f2131,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X134) | s(bool,X135) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,X136) | s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl2 | $spl24 | $spl54)),
  inference(superposition,[],[f1171,f125])).
fof(f2179,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,X126)) ) | ($spl2 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f2127,f115])).
fof(f2127,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'T_0') = s(bool,X126)) ) | ($spl2 | $spl24 | $spl40)),
  inference(superposition,[],[f1020,f125])).
fof(f2178,plain,(
  ( ! [X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X119)) ) | $spl118),
  inference(cnf_transformation,[],[f2178_D])).
fof(f2178_D,plain,(
  ( ! [X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X119)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f2177,plain,(
  ( ! [X118,X119] : (s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X119)) ) | ($spl2 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f2123,f115])).
fof(f2123,plain,(
  ( ! [X118,X119] : (s(bool,'T_0') != s(bool,X118) | s(bool,X119) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X119)) ) | ($spl2 | $spl24 | $spl38)),
  inference(superposition,[],[f718,f125])).
fof(f2069,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f2061])).
fof(f2061,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(superposition,[],[f497,f115])).
fof(f2070,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(subsumption_resolution,[],[f2060,f497])).
fof(f2060,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f497,f115])).
fof(f128,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl4),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f130,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl7),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f2059,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f83,plain,(
  s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2,X3,X4] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X4),s('type.closSem.state'(X1),X2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3,X4] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X4),s('type.closSem.state'(X1),X2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X1),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X3))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X4),s('type.closSem.state'(X1),X2))),s('type.sptree.spt'('type.clos_known.val_approx'),X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X8),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X10),s('type.closSem.state'(X7),X8))),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X7),X8),s('type.sptree.spt'('type.clos_known.val_approx'),X9))) = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X6),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X6)),X10),s('type.closSem.state'(X7),X8))),s('type.sptree.spt'('type.clos_known.val_approx'),X9)))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.clos_knownProof.state_globals_approx_ffifupd')).
fof(f124,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl0),
  inference(cnf_transformation,[],[f124_D])).
fof(f124_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f126,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl3),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2058,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f638,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl36),
  inference(cnf_transformation,[],[f638_D])).
fof(f638_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2050,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f512,f119])).
fof(f2049,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f512,f91])).
fof(f2046,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl34),
  inference(resolution,[],[f512,f90])).
fof(f2045,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f512,f89])).
fof(f2044,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f512,f87])).
fof(f2042,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f512,f85])).
fof(f304,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | ~p(s(bool,X18))) ) | $spl24),
  inference(cnf_transformation,[],[f304_D])).
fof(f304_D,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | ~p(s(bool,X18))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2027,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f116,f119])).
fof(f2026,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f116,f91])).
fof(f2024,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f116,f90])).
fof(f2022,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f116,f89])).
fof(f2021,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f116,f87])).
fof(f2019,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f116,f85])).
fof(f1958,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f1785,f499])).
fof(f1785,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl32),
  inference(backward_demodulation,[],[f499,f83])).
fof(f1957,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f1784,f118])).
fof(f1784,plain,(
  ~p(s(bool,'T_0')) | $spl32),
  inference(backward_demodulation,[],[f499,f116])).
fof(f1949,plain,(
  $false | ($spl7 | $spl32)),
  inference(trivial_inequality_removal,[],[f1941])).
fof(f1941,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl7 | $spl32)),
  inference(superposition,[],[f142,f499])).
fof(f142,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(subsumption_resolution,[],[f138,f130])).
fof(f138,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl7),
  inference(superposition,[],[f130,f115])).
fof(f1783,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f1782,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f1773,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X25)) )),
  inference(resolution,[],[f116,f119])).
fof(f1772,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X22) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f116,f91])).
fof(f1770,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,X16)) )),
  inference(resolution,[],[f116,f90])).
fof(f1768,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,'F_0') != s(bool,X10)) )),
  inference(resolution,[],[f116,f89])).
fof(f1767,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f116,f87])).
fof(f1765,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') != s(bool,X1)) )),
  inference(resolution,[],[f116,f85])).
fof(f123,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl1),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1071,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X1)) ) | $spl46),
  inference(duplicate_literal_removal,[],[f1064])).
fof(f1064,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl46),
  inference(superposition,[],[f765,f115])).
fof(f1073,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl46),
  inference(subsumption_resolution,[],[f1068,f115])).
fof(f1068,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl46),
  inference(superposition,[],[f765,f115])).
fof(f1239,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl48 | $spl60)),
  inference(subsumption_resolution,[],[f1232,f766])).
fof(f1232,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl60),
  inference(superposition,[],[f1090,f115])).
fof(f1241,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | ($spl24 | $spl60)),
  inference(subsumption_resolution,[],[f1236,f655])).
fof(f1236,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl60),
  inference(superposition,[],[f1090,f115])).
fof(f1709,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f1615,f118])).
fof(f1615,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f1614,f116])).
fof(f1614,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f124,f128])).
fof(f1591,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl99),
  inference(cnf_transformation,[],[f1591_D])).
fof(f1591_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1612,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f465,f304])).
fof(f465,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f304,f115])).
fof(f1611,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f522,f200])).
fof(f522,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f200,f115])).
fof(f1609,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(subsumption_resolution,[],[f737,f719])).
fof(f737,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl38),
  inference(superposition,[],[f719,f115])).
fof(f741,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X8) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl4 | $spl38)),
  inference(superposition,[],[f719,f128])).
fof(f1607,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | $spl48),
  inference(subsumption_resolution,[],[f957,f766])).
fof(f957,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(superposition,[],[f766,f115])).
fof(f1606,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl40),
  inference(subsumption_resolution,[],[f1037,f1021])).
fof(f1021,plain,(
  ( ! [X68,X66,X67] : (s(bool,'F_0') != s(bool,X66) | s(bool,X66) != s(bool,X67) | s(bool,'F_0') = s(bool,X68) | s(bool,X67) != s(bool,X68)) ) | $spl40),
  inference(resolution,[],[f729,f116])).
fof(f1037,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl40),
  inference(superposition,[],[f1021,f115])).
fof(f1041,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X12) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13)) ) | ($spl4 | $spl40)),
  inference(superposition,[],[f1021,f128])).
fof(f1072,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f1071,f128])).
fof(f1603,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5)) ) | $spl62),
  inference(subsumption_resolution,[],[f1144,f1091])).
fof(f1144,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl62),
  inference(superposition,[],[f1091,f115])).
fof(f1602,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(subsumption_resolution,[],[f1222,f1172])).
fof(f1222,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(superposition,[],[f1172,f115])).
fof(f1226,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X16) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,X16) != s(bool,X17)) ) | ($spl4 | $spl54)),
  inference(superposition,[],[f1172,f128])).
fof(f1240,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3)) ) | ($spl4 | $spl48 | $spl60)),
  inference(forward_demodulation,[],[f1239,f128])).
fof(f1600,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) | $spl108),
  inference(cnf_transformation,[],[f1600_D])).
fof(f1600_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1599,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl68),
  inference(subsumption_resolution,[],[f1295,f1277])).
fof(f1277,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,'F_0') != s(bool,X86) | s(bool,X86) != s(bool,X87) | s(bool,X88) != s(bool,X89) | s(bool,X89) != s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,X87) != s(bool,X88)) ) | $spl68),
  inference(resolution,[],[f1181,f116])).
fof(f1295,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl68),
  inference(superposition,[],[f1277,f115])).
fof(f1598,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23)) ) | $spl106),
  inference(cnf_transformation,[],[f1598_D])).
fof(f1598_D,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1299,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21)) ) | ($spl4 | $spl68)),
  inference(superposition,[],[f1277,f128])).
fof(f1597,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl104),
  inference(cnf_transformation,[],[f1597_D])).
fof(f1597_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1596,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(subsumption_resolution,[],[f1362,f1344])).
fof(f1344,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(bool,'F_0') != s(bool,X96) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,X99) | s(bool,X99) != s(bool,X100) | s(bool,'F_0') = s(bool,X101) | s(bool,X100) != s(bool,X101) | s(bool,X97) != s(bool,X98)) ) | $spl74),
  inference(resolution,[],[f1287,f116])).
fof(f1362,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) != s(bool,X1) | s(bool,X2) != s(bool,X3) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl74),
  inference(superposition,[],[f1344,f115])).
fof(f1595,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27)) ) | $spl102),
  inference(cnf_transformation,[],[f1595_D])).
fof(f1595_D,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1366,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X25) != s(bool,X26) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X28) | s(bool,X27) != s(bool,X28) | s(bool,X24) != s(bool,X25)) ) | ($spl4 | $spl74)),
  inference(superposition,[],[f1344,f128])).
fof(f1592,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | s(bool,'T_0') != s(bool,'F_0') | ($spl4 | $spl82)),
  inference(forward_demodulation,[],[f1389,f97])).
fof(f1389,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | ($spl4 | $spl82)),
  inference(superposition,[],[f1383,f128])).
fof(f1588,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1)))))) ) | $spl82),
  inference(subsumption_resolution,[],[f1390,f1383])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X2,X1,X0)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) ) | $spl82),
  inference(superposition,[],[f1383,f115])).
fof(f1587,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1586,f118])).
fof(f1586,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1585,f128])).
fof(f1585,plain,(
  ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1584,f298])).
fof(f298,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | $spl6),
  inference(subsumption_resolution,[],[f256,f116])).
fof(f256,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | $spl6),
  inference(superposition,[],[f105,f129])).
fof(f129,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl6),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1584,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1583,f97])).
fof(f1583,plain,(
  s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f1582])).
fof(f1582,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f1435,f103])).
fof(f1435,plain,(
  s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1434,f116])).
fof(f1434,plain,(
  p(s(bool,'F_0')) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f1421,f129])).
fof(f1421,plain,(
  p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(sK3,sK2,sK1)),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) != s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f106])).
fof(f1420,plain,(
  p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) | ($spl4 | $spl6)),
  inference(equality_resolution,[],[f1395])).
fof(f1395,plain,(
  ( ! [X0] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X0))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(sK3,sK2,sK1)))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X0,sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',X0))))) ) | ($spl4 | $spl6)),
  inference(superposition,[],[f584,f298])).
fof(f584,plain,(
  ( ! [X2,X3] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK1),sK2))))) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X2,X3,sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f583,f97])).
fof(f583,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X2,X3,sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f532,f120])).
fof(f532,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X2,X3,sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',X2)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X2))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X3),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)))))))) ) | $spl4),
  inference(superposition,[],[f104,f128])).
fof(f1566,plain,(
  ( ! [X132,X133] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X133) | s(bool,X132) != s(bool,X133)) ) | $spl88),
  inference(cnf_transformation,[],[f1566_D])).
fof(f1566_D,plain,(
  ( ! [X132,X133] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X133) | s(bool,X132) != s(bool,X133)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1574,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (s(bool,X195) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X194) != s(bool,X196) | s(bool,X196) != s(bool,X197) | s(bool,'F_0') = s(bool,X198) | s(bool,X197) != s(bool,X198) | s(bool,X195) != s(bool,X199) | s(bool,'T_0') = s(bool,X199)) ) | ($spl4 | $spl6 | $spl24 | $spl40 | $spl74)),
  inference(subsumption_resolution,[],[f1520,f1020])).
fof(f1520,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (s(bool,'T_0') != s(bool,X194) | s(bool,X195) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X194) != s(bool,X196) | s(bool,X196) != s(bool,X197) | s(bool,'F_0') = s(bool,X198) | s(bool,X197) != s(bool,X198) | s(bool,X195) != s(bool,X199) | s(bool,'T_0') = s(bool,X199)) ) | ($spl4 | $spl6 | $spl24 | $spl74)),
  inference(superposition,[],[f1343,f1431])).
fof(f1431,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f1420,f472])).
fof(f1572,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(bool,X171) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X170) != s(bool,X172) | s(bool,X172) != s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,X171) != s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl4 | $spl6 | $spl24 | $spl38 | $spl68)),
  inference(subsumption_resolution,[],[f1515,f718])).
fof(f1515,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(bool,'T_0') != s(bool,X170) | s(bool,X171) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X170) != s(bool,X172) | s(bool,X172) != s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,X171) != s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl4 | $spl6 | $spl24 | $spl68)),
  inference(superposition,[],[f1276,f1431])).
fof(f1570,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X152) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153) | s(bool,X152) != s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | ($spl4 | $spl6 | $spl24 | $spl54)),
  inference(subsumption_resolution,[],[f1510,f655])).
fof(f1510,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,'T_0') != s(bool,X151) | s(bool,X152) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,X153) | s(bool,X152) != s(bool,X154) | s(bool,'T_0') = s(bool,X154)) ) | ($spl4 | $spl6 | $spl24 | $spl54)),
  inference(superposition,[],[f1171,f1431])).
fof(f1565,plain,(
  ( ! [X132,X133,X131] : (s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,X133) | s(bool,'T_0') = s(bool,X133)) ) | ($spl4 | $spl6 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f1502,f115])).
fof(f1502,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') != s(bool,X131) | s(bool,X132) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,X133) | s(bool,'T_0') = s(bool,X133)) ) | ($spl4 | $spl6 | $spl24 | $spl40)),
  inference(superposition,[],[f1020,f1431])).
fof(f1564,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X122)) ) | $spl86),
  inference(cnf_transformation,[],[f1564_D])).
fof(f1564_D,plain,(
  ( ! [X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'T_0') = s(bool,X122)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1563,plain,(
  ( ! [X121,X122] : (s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl4 | $spl6 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f1495,f115])).
fof(f1495,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,X121) | s(bool,X122) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X122)) ) | ($spl4 | $spl6 | $spl24 | $spl38)),
  inference(superposition,[],[f718,f1431])).
fof(f1436,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | p(s(bool,X1))) ) | $spl84),
  inference(cnf_transformation,[],[f1436_D])).
fof(f1436_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | p(s(bool,X1))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1430,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f119])).
fof(f1429,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X22) != s(bool,X23)) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f91])).
fof(f1426,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f90])).
fof(f1425,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f89])).
fof(f1424,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1)))) | s(bool,X7) != s(bool,X8)) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f87])).
fof(f1422,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(sK7(sK3,sK2,sK1),sK6(sK3,sK2,sK1),sK3,'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2)),sK0)),s('type.closSem.v',sK7(sK3,sK2,sK1))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f1420,f85])).
fof(f1383,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6)))))) ) | $spl82),
  inference(cnf_transformation,[],[f1383_D])).
fof(f1383_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1382,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6))))) | p(s(bool,X4))) ) | ($spl10 | $spl58)),
  inference(subsumption_resolution,[],[f1373,f200])).
fof(f1373,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X5),X6),s('type.sptree.spt'('type.clos_known.val_approx'),X7))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X7,X6,X5)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X7,X6,X5)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X5),X6))))) | p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl58),
  inference(superposition,[],[f1061,f115])).
fof(f1357,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) | $spl80),
  inference(cnf_transformation,[],[f1357_D])).
fof(f1357_D,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1356,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24)) ) | ($spl10 | $spl74)),
  inference(subsumption_resolution,[],[f1351,f200])).
fof(f1351,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) != s(bool,X27) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X21)) ) | $spl74),
  inference(superposition,[],[f1287,f115])).
fof(f1342,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X86,X84,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X80) != s(bool,X82) | s(bool,X83) != s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X86) | s(bool,X85) != s(bool,X86) | s(bool,X82) != s(bool,X83) | p(s(bool,X87)) | p(s(bool,X88)) | s(bool,X81) != s(bool,X88)) ) | $spl74),
  inference(resolution,[],[f1287,f119])).
fof(f1341,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(bool,X70) != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X73) != s(bool,X74) | s(bool,X74) != s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,X75) != s(bool,X76) | s(bool,X72) != s(bool,X73) | p(s(bool,X77)) | ~p(s(bool,X78)) | s(bool,X71) != s(bool,X77) | s(bool,X78) != s(bool,X79)) ) | $spl74),
  inference(resolution,[],[f1287,f91])).
fof(f1338,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(bool,X40) != s(bool,X41) | s(bool,X40) != s(bool,X42) | s(bool,X43) != s(bool,X44) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) != s(bool,X46) | s(bool,X42) != s(bool,X43) | p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X41) != s(bool,X49)) ) | $spl74),
  inference(resolution,[],[f1287,f90])).
fof(f1337,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X30) != s(bool,X32) | s(bool,X33) != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X36) | s(bool,X35) != s(bool,X36) | s(bool,X32) != s(bool,X33) | p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X31) != s(bool,X38)) ) | $spl74),
  inference(resolution,[],[f1287,f89])).
fof(f1336,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,X23) != s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X26) | s(bool,X25) != s(bool,X26) | s(bool,X22) != s(bool,X23) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X21) != s(bool,X27) | s(bool,X28) != s(bool,X29)) ) | $spl74),
  inference(resolution,[],[f1287,f87])).
fof(f1334,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X6) | s(bool,X5) != s(bool,X6) | s(bool,X2) != s(bool,X3) | p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X1) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f1287,f85])).
fof(f1290,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23) | s(bool,X19) != s(bool,X20)) ) | $spl76),
  inference(cnf_transformation,[],[f1290_D])).
fof(f1290_D,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) != s(bool,X23) | s(bool,X19) != s(bool,X20)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1289,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21)) ) | ($spl10 | $spl68)),
  inference(subsumption_resolution,[],[f1284,f200])).
fof(f1284,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X18)) ) | $spl68),
  inference(superposition,[],[f1181,f115])).
fof(f1275,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s(bool,X72) != s(bool,X73) | s(bool,X72) != s(bool,X74) | s(bool,X75) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X74) != s(bool,X75) | p(s(bool,X78)) | p(s(bool,X79)) | s(bool,X73) != s(bool,X79)) ) | $spl68),
  inference(resolution,[],[f1181,f119])).
fof(f1274,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s(bool,X63) != s(bool,X64) | s(bool,X63) != s(bool,X65) | s(bool,X66) != s(bool,X67) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X65) != s(bool,X66) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X64) != s(bool,X69) | s(bool,X70) != s(bool,X71)) ) | $spl68),
  inference(resolution,[],[f1181,f91])).
fof(f1271,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X44,X42,X40] : (s(bool,X36) != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,X38) != s(bool,X39) | p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X37) != s(bool,X44)) ) | $spl68),
  inference(resolution,[],[f1181,f90])).
fof(f1270,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,X27) != s(bool,X28) | s(bool,X27) != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | s(bool,X29) != s(bool,X30) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X28) != s(bool,X34)) ) | $spl68),
  inference(resolution,[],[f1181,f89])).
fof(f1269,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,X21) != s(bool,X22) | s(bool,X22) != s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,X20) != s(bool,X21) | p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X19) != s(bool,X24) | s(bool,X25) != s(bool,X26)) ) | $spl68),
  inference(resolution,[],[f1181,f87])).
fof(f1267,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X2) != s(bool,X3) | p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X1) != s(bool,X7)) ) | $spl68),
  inference(resolution,[],[f1181,f85])).
fof(f1196,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X75),s('type.sptree.spt'('type.clos_known.val_approx'),X76))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X77,X75,X76,X78,X79)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X77))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X75),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X76)))) )),
  inference(resolution,[],[f103,f119])).
fof(f1195,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X67),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X69,X67,X68,X70,X71)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X69))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X67),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X68))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f103,f91])).
fof(f1193,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X51),s('type.sptree.spt'('type.clos_known.val_approx'),X52))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X53,X51,X52,X54,X55)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X53))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X51),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X52)))) )),
  inference(resolution,[],[f103,f90])).
fof(f1191,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X35),s('type.sptree.spt'('type.clos_known.val_approx'),X36))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X37,X35,X36,X38,X39)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X37))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X35),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X36)))) )),
  inference(resolution,[],[f103,f89])).
fof(f1190,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X27),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X29,X27,X28,X30,X31)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X29))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X27),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X28))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f103,f87])).
fof(f1188,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',X11),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) = s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',sK5(X13,X11,X12,X14,X15)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X13))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X11),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X12)))) )),
  inference(resolution,[],[f103,f85])).
fof(f1184,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X16) != s(bool,X17)) ) | $spl70),
  inference(cnf_transformation,[],[f1184_D])).
fof(f1184_D,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X16) != s(bool,X17)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1183,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl10 | $spl54)),
  inference(subsumption_resolution,[],[f1178,f200])).
fof(f1178,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X15)) ) | $spl54),
  inference(superposition,[],[f1029,f115])).
fof(f1170,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,X64) != s(bool,X66) | s(bool,X67) != s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X67) | p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X65) != s(bool,X70)) ) | $spl54),
  inference(resolution,[],[f1029,f119])).
fof(f1169,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,X59) != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) != s(bool,X59) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X57) != s(bool,X61) | s(bool,X62) != s(bool,X63)) ) | $spl54),
  inference(resolution,[],[f1029,f91])).
fof(f1166,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X32) != s(bool,X34) | s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X33) != s(bool,X39)) ) | $spl54),
  inference(resolution,[],[f1029,f90])).
fof(f1165,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,X27) != s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) != s(bool,X27) | p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X29) != s(bool,X31) | s(bool,X25) != s(bool,X30)) ) | $spl54),
  inference(resolution,[],[f1029,f89])).
fof(f1164,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,X19) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,X18) != s(bool,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X17) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f1029,f87])).
fof(f1162,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X3) != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X2) != s(bool,X3) | p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X1) != s(bool,X6)) ) | $spl54),
  inference(resolution,[],[f1029,f85])).
fof(f1115,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X75,X76,X77,X78,X79)),s('type.closSem.v',X75)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X75))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X76),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X79),X78))))) | p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X81) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X79),X78),s('type.sptree.spt'('type.clos_known.val_approx'),X77)))) )),
  inference(resolution,[],[f104,f119])).
fof(f1114,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X67,X73] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X67,X68,X69,X70,X71)),s('type.closSem.v',X67)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X67))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X68),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X71),X70))))) | ~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X71),X70),s('type.sptree.spt'('type.clos_known.val_approx'),X69))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f104,f91])).
fof(f1112,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X51,X52,X53,X54,X55)),s('type.closSem.v',X51)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X51))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X52),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X55),X54))))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X55),X54),s('type.sptree.spt'('type.clos_known.val_approx'),X53)))) )),
  inference(resolution,[],[f104,f90])).
fof(f1110,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X35,X36,X37,X38,X39)),s('type.closSem.v',X35)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X35))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X36),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X39),X38))))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X39),X38),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) )),
  inference(resolution,[],[f104,f89])).
fof(f1109,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X34,X32] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X27,X28,X29,X30,X31)),s('type.closSem.v',X27)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X27))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X28),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X31),X30))))) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f104,f87])).
fof(f1107,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X18,X16] : (p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X11,X12,X13,X14,X15)),s('type.closSem.v',X11)))) | s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X11))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X12),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X15),X14))))) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X15),X14),s('type.sptree.spt'('type.clos_known.val_approx'),X13)))) )),
  inference(resolution,[],[f104,f85])).
fof(f1103,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X71))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X72),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X73),X74))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X73),X74),s('type.sptree.spt'('type.clos_known.val_approx'),X75)))) | p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X71,X72,X75,X74,X73)),s('type.closSem.v',X71)))) )),
  inference(resolution,[],[f104,f119])).
fof(f1102,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X63))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X64),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X65),X66))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X65),X66),s('type.sptree.spt'('type.clos_known.val_approx'),X67)))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X63,X64,X67,X66,X65)),s('type.closSem.v',X63))) | s(bool,X69) != s(bool,X70)) )),
  inference(resolution,[],[f104,f91])).
fof(f1099,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X39))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X40),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X41),X42))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X41),X42),s('type.sptree.spt'('type.clos_known.val_approx'),X43)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X39,X40,X43,X42,X41)),s('type.closSem.v',X39)))) )),
  inference(resolution,[],[f104,f90])).
fof(f1098,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X31))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X32),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X33),X34))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X33),X34),s('type.sptree.spt'('type.clos_known.val_approx'),X35)))) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X31,X32,X35,X34,X33)),s('type.closSem.v',X31)))) )),
  inference(resolution,[],[f104,f89])).
fof(f1097,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X23))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X24),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X25),X26))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X25),X26),s('type.sptree.spt'('type.clos_known.val_approx'),X27)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X23,X24,X27,X26,X25)),s('type.closSem.v',X23))) | s(bool,X29) != s(bool,X30)) )),
  inference(resolution,[],[f104,f87])).
fof(f1095,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',X7))))) != s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',X8),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))))) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X9),X10),s('type.sptree.spt'('type.clos_known.val_approx'),X11)))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',sK5(X7,X8,X11,X10,X9)),s('type.closSem.v',X7)))) )),
  inference(resolution,[],[f104,f85])).
fof(f1091,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) | $spl62),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1087,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl24 | $spl40)),
  inference(duplicate_literal_removal,[],[f1080])).
fof(f1080,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl40)),
  inference(superposition,[],[f1020,f115])).
fof(f1090,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) != s(bool,X21)) ) | $spl60),
  inference(cnf_transformation,[],[f1090_D])).
fof(f1090_D,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) != s(bool,X21)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1089,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl4 | $spl24 | $spl40)),
  inference(subsumption_resolution,[],[f1079,f115])).
fof(f1079,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X19) | s(bool,X20) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X19) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl4 | $spl24 | $spl40)),
  inference(superposition,[],[f1020,f128])).
fof(f1054,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X48,X49,X50)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X48,X49,X50)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X50),X49))))) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X50),X49),s('type.sptree.spt'('type.clos_known.val_approx'),X48)))) )),
  inference(resolution,[],[f105,f119])).
fof(f1053,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X42,X43,X44)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X42,X43,X44)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X44),X43))))) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X44),X43),s('type.sptree.spt'('type.clos_known.val_approx'),X42))) | s(bool,X46) != s(bool,X47)) )),
  inference(resolution,[],[f105,f91])).
fof(f1050,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X24,X25,X26)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X24,X25,X26)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X26),X25))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X26),X25),s('type.sptree.spt'('type.clos_known.val_approx'),X24)))) )),
  inference(resolution,[],[f105,f90])).
fof(f1049,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X18,X19,X20)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X18,X19,X20)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X20),X19))))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X20),X19),s('type.sptree.spt'('type.clos_known.val_approx'),X18)))) )),
  inference(resolution,[],[f105,f89])).
fof(f1048,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X12,X13,X14)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X12,X13,X14)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X14),X13))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X13),s('type.sptree.spt'('type.clos_known.val_approx'),X12))) | s(bool,X16) != s(bool,X17)) )),
  inference(resolution,[],[f105,f87])).
fof(f1046,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'('type.option.option'('type.closSem.v')),'const.option.SOME_1'(s('type.option.option'('type.closSem.v'),'const.option.SOME_1'(s('type.closSem.v',sK7(X0,X1,X2)))))) = s('type.option.option'('type.option.option'('type.closSem.v')),'const.closSem.get_global_2'(s('type.num.num',sK6(X0,X1,X2)),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X2),X1))))) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X2),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X0)))) )),
  inference(resolution,[],[f105,f85])).
fof(f1032,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s(bool,X13) != s(bool,X14)) ) | $spl56),
  inference(cnf_transformation,[],[f1032_D])).
fof(f1032_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) != s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s(bool,X13) != s(bool,X14)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1031,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | ($spl10 | $spl40)),
  inference(subsumption_resolution,[],[f1026,f200])).
fof(f1026,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,X12)) ) | $spl40),
  inference(superposition,[],[f729,f115])).
fof(f1019,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,X56) != s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,X58) != s(bool,X59) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X57) != s(bool,X61)) ) | $spl40),
  inference(resolution,[],[f729,f119])).
fof(f1018,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X49) != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X50) != s(bool,X53) | s(bool,X54) != s(bool,X55)) ) | $spl40),
  inference(resolution,[],[f729,f91])).
fof(f1015,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s(bool,X28) != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) != s(bool,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X29) != s(bool,X34)) ) | $spl40),
  inference(resolution,[],[f729,f90])).
fof(f1014,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,X21) != s(bool,X22) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) != s(bool,X24) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X22) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f729,f89])).
fof(f1013,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) != s(bool,X16) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) != s(bool,X17) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X15) != s(bool,X18) | s(bool,X19) != s(bool,X20)) ) | $spl40),
  inference(resolution,[],[f729,f87])).
fof(f1011,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X1) != s(bool,X5)) ) | $spl40),
  inference(resolution,[],[f729,f85])).
fof(f983,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X56),X57),s('type.sptree.spt'('type.clos_known.val_approx'),X58)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X59))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X58,X57,X56)),s('type.sptree.spt'('type.clos_known.val_approx'),X58))) | p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X59),s('type.closSem.v',sK7(X58,X57,X56))))) )),
  inference(resolution,[],[f106,f119])).
fof(f982,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X49),X50),s('type.sptree.spt'('type.clos_known.val_approx'),X51)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X52))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X51,X50,X49)),s('type.sptree.spt'('type.clos_known.val_approx'),X51))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X52),s('type.closSem.v',sK7(X51,X50,X49)))) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f106,f91])).
fof(f980,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X35),X36),s('type.sptree.spt'('type.clos_known.val_approx'),X37)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X38))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X37,X36,X35)),s('type.sptree.spt'('type.clos_known.val_approx'),X37))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X38),s('type.closSem.v',sK7(X37,X36,X35))))) )),
  inference(resolution,[],[f106,f90])).
fof(f978,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X21),X22),s('type.sptree.spt'('type.clos_known.val_approx'),X23)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X24))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X23,X22,X21)),s('type.sptree.spt'('type.clos_known.val_approx'),X23))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X24),s('type.closSem.v',sK7(X23,X22,X21))))) )),
  inference(resolution,[],[f106,f89])).
fof(f977,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X14),X15),s('type.sptree.spt'('type.clos_known.val_approx'),X16)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X17))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X16,X15,X14)),s('type.sptree.spt'('type.clos_known.val_approx'),X16))) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X17),s('type.closSem.v',sK7(X16,X15,X14)))) | s(bool,X18) != s(bool,X20)) )),
  inference(resolution,[],[f106,f87])).
fof(f975,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X0),X1),s('type.sptree.spt'('type.clos_known.val_approx'),X2)))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X3))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.sptree.spt'('type.clos_known.val_approx'),X2))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X3),s('type.closSem.v',sK7(X2,X1,X0))))) )),
  inference(resolution,[],[f106,f85])).
fof(f973,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X56),s('type.closSem.v',sK7(X57,X58,X59))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X56))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X57,X58,X59)),s('type.sptree.spt'('type.clos_known.val_approx'),X57))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X61) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X59),X58),s('type.sptree.spt'('type.clos_known.val_approx'),X57)))) )),
  inference(resolution,[],[f106,f119])).
fof(f972,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X49),s('type.closSem.v',sK7(X50,X51,X52))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X49))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X50,X51,X52)),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X52),X51),s('type.sptree.spt'('type.clos_known.val_approx'),X50))) | s(bool,X54) != s(bool,X55)) )),
  inference(resolution,[],[f106,f91])).
fof(f969,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X28),s('type.closSem.v',sK7(X29,X30,X31))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X28))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X29,X30,X31)),s('type.sptree.spt'('type.clos_known.val_approx'),X29))) | p(s(bool,X32)) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X31),X30),s('type.sptree.spt'('type.clos_known.val_approx'),X29)))) )),
  inference(resolution,[],[f106,f90])).
fof(f968,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X21),s('type.closSem.v',sK7(X22,X23,X24))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X21))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X22,X23,X24)),s('type.sptree.spt'('type.clos_known.val_approx'),X22))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X24),X23),s('type.sptree.spt'('type.clos_known.val_approx'),X22)))) )),
  inference(resolution,[],[f106,f89])).
fof(f967,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X14),s('type.closSem.v',sK7(X15,X16,X17))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X14))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X15,X16,X17)),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X17),X16),s('type.sptree.spt'('type.clos_known.val_approx'),X15))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f106,f87])).
fof(f965,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.clos_knownProof.val_approx_val_2'(s('type.clos_known.val_approx',X0),s('type.closSem.v',sK7(X1,X2,X3))))) | s('type.option.option'('type.clos_known.val_approx'),'const.option.SOME_1'(s('type.clos_known.val_approx',X0))) != s('type.option.option'('type.clos_known.val_approx'),'const.sptree.lookup_2'(s('type.num.num',sK6(X1,X2,X3)),s('type.sptree.spt'('type.clos_known.val_approx'),X1))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(X3),X2),s('type.sptree.spt'('type.clos_known.val_approx'),X1)))) )),
  inference(resolution,[],[f106,f85])).
fof(f952,plain,(
  ( ! [X161,X162] : (~p(s(bool,X161)) | s(bool,'T_0') = s(bool,X162)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f885,f115])).
fof(f885,plain,(
  ( ! [X161,X162] : (s(bool,'F_0') != s(bool,X162) | ~p(s(bool,X161)) | s(bool,'T_0') = s(bool,X162)) ) | ($spl24 | $spl44)),
  inference(superposition,[],[f462,f764])).
fof(f949,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X144) != s(bool,X146)) ) | $spl44),
  inference(subsumption_resolution,[],[f877,f116])).
fof(f877,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,'F_0')) | p(s(bool,X145)) | ~p(s(bool,X146)) | s(bool,X144) != s(bool,X146)) ) | $spl44),
  inference(superposition,[],[f119,f764])).
fof(f948,plain,(
  ( ! [X140,X138] : (p(s(bool,X140)) | ~p(s(bool,X138))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f947,f764])).
fof(f947,plain,(
  ( ! [X140,X138,X139] : (s(bool,'F_0') != s(bool,X139) | p(s(bool,X140)) | ~p(s(bool,X138))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f875,f304])).
fof(f875,plain,(
  ( ! [X140,X138,X139] : (s(bool,'F_0') != s(bool,X139) | p(s(bool,X140)) | ~p(s(bool,X138)) | p(s(bool,X139))) ) | $spl44),
  inference(superposition,[],[f119,f764])).
fof(f946,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl44),
  inference(subsumption_resolution,[],[f871,f116])).
fof(f871,plain,(
  ( ! [X127,X125,X126,X124] : (p(s(bool,'F_0')) | ~p(s(bool,X125)) | ~p(s(bool,X126)) | s(bool,X124) != s(bool,X126) | s(bool,X125) != s(bool,X127)) ) | $spl44),
  inference(superposition,[],[f91,f764])).
fof(f945,plain,(
  ( ! [X123,X122,X120] : (~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f944,f764])).
fof(f944,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,X121) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f870,f304])).
fof(f870,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,X121) | p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X120)) | s(bool,X122) != s(bool,X123)) ) | $spl44),
  inference(superposition,[],[f91,f764])).
fof(f942,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl44),
  inference(subsumption_resolution,[],[f863,f116])).
fof(f863,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'F_0')) | p(s(bool,X93)) | ~p(s(bool,X94)) | ~p(s(bool,X95)) | s(bool,X92) != s(bool,X94)) ) | $spl44),
  inference(superposition,[],[f90,f764])).
fof(f941,plain,(
  ( ! [X87,X86,X84] : (p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f940,f764])).
fof(f940,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87))) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f861,f304])).
fof(f861,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X85) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~p(s(bool,X87)) | p(s(bool,X85))) ) | $spl44),
  inference(superposition,[],[f90,f764])).
fof(f939,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl44),
  inference(subsumption_resolution,[],[f858,f116])).
fof(f858,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'F_0')) | p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,X75) | s(bool,X72) != s(bool,X74)) ) | $spl44),
  inference(superposition,[],[f89,f764])).
fof(f937,plain,(
  ( ! [X59,X58,X56] : (p(s(bool,X58)) | ~p(s(bool,X56)) | s(bool,X58) != s(bool,X59)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f936,f764])).
fof(f936,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | s(bool,X58) != s(bool,X59)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f854,f304])).
fof(f854,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | p(s(bool,X57)) | s(bool,X58) != s(bool,X59)) ) | $spl44),
  inference(superposition,[],[f89,f764])).
fof(f935,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl44),
  inference(subsumption_resolution,[],[f851,f116])).
fof(f851,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'F_0')) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) ) | $spl44),
  inference(superposition,[],[f87,f764])).
fof(f934,plain,(
  ( ! [X39,X38,X36] : (~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f933,f764])).
fof(f933,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f849,f304])).
fof(f849,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,X38) != s(bool,X39)) ) | $spl44),
  inference(superposition,[],[f87,f764])).
fof(f932,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl44),
  inference(subsumption_resolution,[],[f844,f116])).
fof(f844,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'F_0')) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl44),
  inference(superposition,[],[f85,f764])).
fof(f930,plain,(
  ( ! [X6,X4,X7] : (p(s(bool,X6)) | ~p(s(bool,X4)) | s(bool,X6) != s(bool,X7)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f929,f764])).
fof(f929,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X4)) | s(bool,X6) != s(bool,X7)) ) | ($spl24 | $spl44)),
  inference(subsumption_resolution,[],[f841,f304])).
fof(f841,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X6)) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,X7)) ) | $spl44),
  inference(superposition,[],[f85,f764])).
fof(f926,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,X159)) | p(s(bool,X160)) | p(s(bool,X161))) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f824,f512])).
fof(f824,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,X159)) | p(s(bool,X160)) | p(s(bool,X161)) | s(bool,'F_0') != s(bool,X161)) ) | $spl44),
  inference(superposition,[],[f119,f764])).
fof(f925,plain,(
  ( ! [X152,X151,X150] : (s(bool,X150) != s(bool,X151) | p(s(bool,X152)) | ~p(s(bool,X151))) ) | $spl44),
  inference(subsumption_resolution,[],[f821,f116])).
fof(f821,plain,(
  ( ! [X152,X151,X150] : (s(bool,X150) != s(bool,X151) | p(s(bool,X152)) | ~p(s(bool,X151)) | p(s(bool,'F_0'))) ) | $spl44),
  inference(superposition,[],[f119,f764])).
fof(f924,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,X143) != s(bool,X144)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f818,f512])).
fof(f818,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~p(s(bool,X143)) | s(bool,'F_0') != s(bool,X142) | s(bool,X143) != s(bool,X144)) ) | $spl44),
  inference(superposition,[],[f91,f764])).
fof(f922,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X125) != s(bool,X126) | ~p(s(bool,X127)) | ~p(s(bool,X126)) | s(bool,X127) != s(bool,X128)) ) | $spl44),
  inference(subsumption_resolution,[],[f814,f116])).
fof(f814,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X125) != s(bool,X126) | p(s(bool,'F_0')) | ~p(s(bool,X127)) | ~p(s(bool,X126)) | s(bool,X127) != s(bool,X128)) ) | $spl44),
  inference(superposition,[],[f91,f764])).
fof(f920,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | p(s(bool,X112))) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f810,f512])).
fof(f810,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | p(s(bool,X110)) | ~p(s(bool,X111)) | p(s(bool,X112)) | s(bool,'F_0') != s(bool,X112)) ) | $spl44),
  inference(superposition,[],[f90,f764])).
fof(f918,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | ~p(s(bool,X100))) ) | $spl44),
  inference(subsumption_resolution,[],[f807,f116])).
fof(f807,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X97) != s(bool,X98) | p(s(bool,X99)) | ~p(s(bool,X98)) | ~p(s(bool,X100)) | p(s(bool,'F_0'))) ) | $spl44),
  inference(superposition,[],[f90,f764])).
fof(f917,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X90) != s(bool,X92)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f805,f512])).
fof(f805,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X90) != s(bool,X92) | s(bool,'F_0') != s(bool,X91)) ) | $spl44),
  inference(superposition,[],[f89,f764])).
fof(f916,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X69) != s(bool,X70) | p(s(bool,X71)) | ~p(s(bool,X70)) | s(bool,X71) != s(bool,X72)) ) | $spl44),
  inference(subsumption_resolution,[],[f800,f116])).
fof(f800,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X69) != s(bool,X70) | p(s(bool,X71)) | ~p(s(bool,X70)) | p(s(bool,'F_0')) | s(bool,X71) != s(bool,X72)) ) | $spl44),
  inference(superposition,[],[f89,f764])).
fof(f915,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X63) != s(bool,X64)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f798,f512])).
fof(f798,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,'F_0') != s(bool,X62) | s(bool,X63) != s(bool,X64)) ) | $spl44),
  inference(superposition,[],[f87,f764])).
fof(f913,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X49) != s(bool,X50) | ~p(s(bool,X51)) | ~p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl44),
  inference(subsumption_resolution,[],[f795,f116])).
fof(f795,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,'F_0')) | ~p(s(bool,X51)) | ~p(s(bool,X50)) | s(bool,X51) != s(bool,X52)) ) | $spl44),
  inference(superposition,[],[f87,f764])).
fof(f911,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36)) ) | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f791,f512])).
fof(f791,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,'F_0') != s(bool,X35)) ) | $spl44),
  inference(superposition,[],[f85,f764])).
fof(f910,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | ~p(s(bool,X10)) | s(bool,X11) != s(bool,X12)) ) | $spl44),
  inference(subsumption_resolution,[],[f785,f116])).
fof(f785,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | ~p(s(bool,X10)) | p(s(bool,'F_0')) | s(bool,X11) != s(bool,X12)) ) | $spl44),
  inference(superposition,[],[f85,f764])).
fof(f766,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f760,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl38)),
  inference(duplicate_literal_removal,[],[f753])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | ($spl24 | $spl38)),
  inference(superposition,[],[f718,f115])).
fof(f765,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X15)) ) | $spl46),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f763,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl24 | $spl38)),
  inference(subsumption_resolution,[],[f752,f115])).
fof(f752,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | s(bool,X15) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl24 | $spl38)),
  inference(superposition,[],[f718,f128])).
fof(f732,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl42),
  inference(cnf_transformation,[],[f732_D])).
fof(f732_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f731,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | ($spl10 | $spl38)),
  inference(subsumption_resolution,[],[f726,f200])).
fof(f726,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X9)) ) | $spl38),
  inference(superposition,[],[f678,f115])).
fof(f717,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X48) != s(bool,X50) | s(bool,'F_0') = s(bool,X50) | p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X49) != s(bool,X52)) ) | $spl38),
  inference(resolution,[],[f678,f119])).
fof(f716,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(bool,X42) != s(bool,X43) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X43) != s(bool,X45) | s(bool,X46) != s(bool,X47)) ) | $spl38),
  inference(resolution,[],[f678,f91])).
fof(f713,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X25) != s(bool,X29)) ) | $spl38),
  inference(resolution,[],[f678,f90])).
fof(f712,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(bool,X18) != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X19) != s(bool,X22)) ) | $spl38),
  inference(resolution,[],[f678,f89])).
fof(f711,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X13) != s(bool,X15) | s(bool,X16) != s(bool,X17)) ) | $spl38),
  inference(resolution,[],[f678,f87])).
fof(f709,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X1) != s(bool,X4)) ) | $spl38),
  inference(resolution,[],[f678,f85])).
fof(f666,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40) | p(s(bool,X42)) | p(s(bool,X43)) | s(bool,X41) != s(bool,X43)) ) | $spl14),
  inference(resolution,[],[f516,f119])).
fof(f665,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) ) | $spl14),
  inference(resolution,[],[f516,f91])).
fof(f662,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) ) | $spl14),
  inference(resolution,[],[f516,f90])).
fof(f661,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) ) | $spl14),
  inference(resolution,[],[f516,f89])).
fof(f660,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f516,f87])).
fof(f658,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) ) | $spl14),
  inference(resolution,[],[f516,f85])).
fof(f625,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') = s(bool,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl34),
  inference(resolution,[],[f512,f119])).
fof(f624,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'F_0') = s(bool,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl34),
  inference(resolution,[],[f512,f91])).
fof(f621,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl34),
  inference(resolution,[],[f512,f90])).
fof(f620,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f512,f89])).
fof(f619,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl34),
  inference(resolution,[],[f512,f87])).
fof(f617,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl34),
  inference(resolution,[],[f512,f85])).
fof(f200,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8))) ) | $spl10),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f514,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f505,f304])).
fof(f505,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f205,f115])).
fof(f512,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f511,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f500,f472])).
fof(f500,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f205,f115])).
fof(f205,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl14),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X28) | p(s(bool,X27)) | ~p(s(bool,X28))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f483,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) ) | $spl24),
  inference(resolution,[],[f472,f119])).
fof(f496,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  ( ! [X2,X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f482,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34)) ) | $spl24),
  inference(resolution,[],[f472,f119])).
fof(f481,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f472,f91])).
fof(f498,plain,(
  ( ! [X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl24),
  inference(subsumption_resolution,[],[f480,f119])).
fof(f480,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl24),
  inference(resolution,[],[f472,f90])).
fof(f479,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f472,f90])).
fof(f478,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl24),
  inference(resolution,[],[f472,f89])).
fof(f477,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f472,f89])).
fof(f476,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f472,f87])).
fof(f475,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl24),
  inference(resolution,[],[f472,f85])).
fof(f474,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f472,f85])).
fof(f461,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f124,f251])).
fof(f251,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl6),
  inference(backward_demodulation,[],[f129,f83])).
fof(f458,plain,(
  $false | ($spl6 | $spl26)),
  inference(trivial_inequality_removal,[],[f456])).
fof(f456,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl26)),
  inference(superposition,[],[f306,f129])).
fof(f457,plain,(
  $false | $spl26),
  inference(equality_resolution,[],[f306])).
fof(f447,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f206])).
fof(f429,plain,(
  $false | $spl18),
  inference(resolution,[],[f393,f118])).
fof(f393,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(equality_resolution,[],[f215])).
fof(f430,plain,(
  $false | $spl18),
  inference(resolution,[],[f393,f117])).
fof(f117,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.state_globals_approx_ffifupd','thm.bool.TRUTH')).
fof(f431,plain,(
  $false | $spl18),
  inference(resolution,[],[f393,f120])).
fof(f408,plain,(
  $false | $spl20),
  inference(resolution,[],[f228,f118])).
fof(f409,plain,(
  $false | $spl20),
  inference(resolution,[],[f228,f117])).
fof(f410,plain,(
  $false | $spl20),
  inference(resolution,[],[f228,f120])).
fof(f395,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2))) ) | $spl18),
  inference(subsumption_resolution,[],[f384,f215])).
fof(f384,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(superposition,[],[f215,f115])).
fof(f394,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0))) ) | $spl18),
  inference(subsumption_resolution,[],[f383,f215])).
fof(f383,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f215,f115])).
fof(f377,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2))) ) | $spl12),
  inference(subsumption_resolution,[],[f366,f203])).
fof(f366,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(superposition,[],[f203,f115])).
fof(f376,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f365,f203])).
fof(f365,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f203,f115])).
fof(f364,plain,(
  $false | $spl8),
  inference(equality_resolution,[],[f199])).
fof(f345,plain,(
  $false | $spl22),
  inference(resolution,[],[f229,f116])).
fof(f335,plain,(
  ( ! [X111,X112] : (s(bool,'F_0') != s(bool,X112) | p(s(bool,X111)) | ~p(s(bool,X112))) ) | $spl6),
  inference(forward_demodulation,[],[f334,f129])).
fof(f334,plain,(
  ( ! [X111,X112] : (p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f295,f116])).
fof(f295,plain,(
  ( ! [X111,X112] : (p(s(bool,'F_0')) | p(s(bool,X111)) | ~p(s(bool,X112)) | s(bool,X112) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(superposition,[],[f119,f129])).
fof(f331,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X99) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(forward_demodulation,[],[f330,f129])).
fof(f330,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(subsumption_resolution,[],[f289,f116])).
fof(f289,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,'F_0')) | ~p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X99) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X98) != s(bool,X100)) ) | $spl6),
  inference(superposition,[],[f91,f129])).
fof(f199,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10)) ) | $spl8),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f328,plain,(
  ( ! [X94,X92,X93] : (s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(subsumption_resolution,[],[f327,f116])).
fof(f327,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(subsumption_resolution,[],[f326,f119])).
fof(f326,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X92) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(forward_demodulation,[],[f287,f129])).
fof(f287,plain,(
  ( ! [X94,X92,X93] : (s(bool,'F_0') != s(bool,X92) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X93)) | ~p(s(bool,X92)) | s(bool,X93) != s(bool,X94)) ) | $spl6),
  inference(superposition,[],[f91,f129])).
fof(f325,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl6),
  inference(subsumption_resolution,[],[f286,f89])).
fof(f286,plain,(
  ( ! [X90,X91,X89] : (s(bool,'F_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl6),
  inference(superposition,[],[f91,f129])).
fof(f324,plain,(
  ( ! [X88,X87,X86] : (s(bool,'F_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl6),
  inference(subsumption_resolution,[],[f285,f119])).
fof(f285,plain,(
  ( ! [X88,X87,X86] : (s(bool,'F_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl6),
  inference(superposition,[],[f91,f129])).
fof(f322,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl6),
  inference(forward_demodulation,[],[f321,f129])).
fof(f321,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f281,f116])).
fof(f281,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'F_0')) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~p(s(bool,X76)) | s(bool,X75) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(superposition,[],[f90,f129])).
fof(f317,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,X60) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61)) ) | $spl6),
  inference(forward_demodulation,[],[f316,f129])).
fof(f316,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f276,f116])).
fof(f276,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,'F_0')) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(superposition,[],[f89,f129])).
fof(f313,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(forward_demodulation,[],[f312,f129])).
fof(f312,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f269,f116])).
fof(f269,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'F_0')) | ~p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X39) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X38) != s(bool,X40)) ) | $spl6),
  inference(superposition,[],[f87,f129])).
fof(f311,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(subsumption_resolution,[],[f310,f116])).
fof(f310,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(subsumption_resolution,[],[f309,f119])).
fof(f309,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X36)) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(forward_demodulation,[],[f268,f129])).
fof(f268,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X36)) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f87,f129])).
fof(f307,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl6),
  inference(subsumption_resolution,[],[f266,f119])).
fof(f266,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl6),
  inference(superposition,[],[f87,f129])).
fof(f305,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl6),
  inference(subsumption_resolution,[],[f265,f89])).
fof(f265,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f87,f129])).
fof(f303,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19)) ) | $spl6),
  inference(forward_demodulation,[],[f302,f129])).
fof(f302,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f262,f116])).
fof(f262,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'F_0')) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl6),
  inference(superposition,[],[f85,f129])).
fof(f247,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f246,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f242,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') != s(bool,X116) | p(s(bool,X115)) | p(s(bool,X116))) ) | $spl2),
  inference(forward_demodulation,[],[f241,f125])).
fof(f241,plain,(
  ( ! [X116,X115] : (p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f192,f120])).
fof(f192,plain,(
  ( ! [X116,X115] : (~p(s(bool,'T_0')) | p(s(bool,X115)) | p(s(bool,X116)) | s(bool,X116) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f119,f125])).
fof(f238,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(forward_demodulation,[],[f237,f125])).
fof(f237,plain,(
  ( ! [X105,X106,X104] : (p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(subsumption_resolution,[],[f186,f120])).
fof(f186,plain,(
  ( ! [X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | ~p(s(bool,X105)) | s(bool,X104) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X105) != s(bool,X106)) ) | $spl2),
  inference(superposition,[],[f91,f125])).
fof(f236,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') != s(bool,X103) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102)) ) | $spl2),
  inference(forward_demodulation,[],[f235,f125])).
fof(f235,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102) | s(bool,X103) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f185,f119])).
fof(f185,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'T_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X102) | s(bool,X103) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f91,f125])).
fof(f231,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl2),
  inference(subsumption_resolution,[],[f181,f89])).
fof(f181,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X89) | p(s(bool,X90)) | ~p(s(bool,X89)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,X91)) ) | $spl2),
  inference(superposition,[],[f91,f125])).
fof(f230,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl2),
  inference(subsumption_resolution,[],[f180,f119])).
fof(f180,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') != s(bool,X86) | p(s(bool,X87)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,X88)) ) | $spl2),
  inference(superposition,[],[f91,f125])).
fof(f227,plain,(
  ( ! [X80,X81,X82] : (s(bool,'T_0') != s(bool,X82) | p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82))) ) | $spl2),
  inference(forward_demodulation,[],[f226,f125])).
fof(f226,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f178,f120])).
fof(f178,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f90,f125])).
fof(f221,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,X66) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67)) ) | $spl2),
  inference(forward_demodulation,[],[f220,f125])).
fof(f220,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,X66) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f173,f120])).
fof(f173,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | s(bool,X65) != s(bool,X67) | s(bool,X66) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f89,f125])).
fof(f218,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,X47) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(subsumption_resolution,[],[f217,f120])).
fof(f217,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X47) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(subsumption_resolution,[],[f216,f91])).
fof(f216,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X47) | p(s(bool,X48)) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(forward_demodulation,[],[f167,f125])).
fof(f167,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,X47) | p(s(bool,X48)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | p(s(bool,X47)) | s(bool,X48) != s(bool,X49)) ) | $spl2),
  inference(superposition,[],[f89,f125])).
fof(f214,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(forward_demodulation,[],[f213,f125])).
fof(f213,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(subsumption_resolution,[],[f166,f120])).
fof(f166,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,X45) != s(bool,X46)) ) | $spl2),
  inference(superposition,[],[f87,f125])).
fof(f212,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(forward_demodulation,[],[f211,f125])).
fof(f211,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f165,f119])).
fof(f165,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | s(bool,X43) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f87,f125])).
fof(f207,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(subsumption_resolution,[],[f161,f119])).
fof(f161,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(superposition,[],[f87,f125])).
fof(f204,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl2),
  inference(subsumption_resolution,[],[f160,f89])).
fof(f160,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl2),
  inference(superposition,[],[f87,f125])).
fof(f202,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25)) ) | $spl2),
  inference(forward_demodulation,[],[f201,f125])).
fof(f201,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(subsumption_resolution,[],[f159,f120])).
fof(f159,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) ) | $spl2),
  inference(superposition,[],[f85,f125])).
fof(f198,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f197,f120])).
fof(f197,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f196,f91])).
fof(f196,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X9)) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(forward_demodulation,[],[f154,f125])).
fof(f154,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X9)) | ~p(s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))) | p(s(bool,X8)) | s(bool,X9) != s(bool,X10)) ) | $spl2),
  inference(superposition,[],[f85,f125])).
fof(f141,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f140,f126])).
fof(f140,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl7),
  inference(trivial_inequality_removal,[],[f139])).
fof(f139,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | $spl7),
  inference(superposition,[],[f130,f115])).
fof(f122,plain,(
  s(bool,'F_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
fof(f121,plain,(
  s(bool,'T_0') != s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK1),sK2),s('type.sptree.spt'('type.clos_known.val_approx'),sK3))) | s(bool,'F_0') = s(bool,'const.clos_knownProof.state_globals_approx_2'(s('type.closSem.state'(sK0),'const.closSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(sK1),'type.ffi.ffi_state'(sK0)),sK4),s('type.closSem.state'(sK1),sK2))),s('type.sptree.spt'('type.clos_known.val_approx'),sK3)))),
  inference(superposition,[],[f83,f115])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 213
% Passive clauses: 533
% Generated clauses: 3730
% Final active clauses: 139
% Final passive clauses: 194
% Input formulas: 36
% Initial clauses: 36
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 51
% Trivial inequalities: 119
% Fw subsumption resolutions: 630
% Fw demodulations: 190
% Bw demodulations: 34
% 
% Simple tautologies: 96
% Equational tautologies: 6
% Forward subsumptions: 1662
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 596
% Forward superposition: 877
% Backward superposition: 994
% Self superposition: 4
% Equality resolution: 71
% 
% Split clauses: 470
% Split components: 1011
% SAT solver clauses: 498
% SAT solver unit clauses: 12
% SAT solver binary clauses: 104
% 
% Sat splits: 470
% Sat splitting refutations: 19
% 
% Memory used [KB]: 3837
% Time elapsed: 0.204 s
% ------------------------------
% ------------------------------
% Success in time 11.976 s
