Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% (109463)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 3891
% Generated clauses: 13356
% Final active clauses: 99
% Final passive clauses: 3554
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 559
% Trivial inequalities: 134
% Fw subsumption resolutions: 1833
% Fw demodulations: 1117
% Bw demodulations: 3
% 
% Simple tautologies: 246
% Equational tautologies: 73
% Forward subsumptions: 4459
% 
% Binary resolution: 544
% Factoring: 2
% Forward superposition: 5044
% Backward superposition: 3582
% Self superposition: 230
% Equality factoring: 170
% Equality resolution: 21
% 
% Split clauses: 496
% Split components: 1037
% SAT solver clauses: 1239
% SAT solver unit clauses: 26
% SAT solver binary clauses: 216
% 
% Sat splits: 496
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1232
% TWLsolver calls for satisfiability: 199
% 
% Memory used [KB]: 9083
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (109524)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 83
% Passive clauses: 1061
% Generated clauses: 9547
% Final active clauses: 66
% Final passive clauses: 539
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 318
% Trivial inequalities: 96
% Fw subsumption resolutions: 503
% Fw demodulations: 2045
% Bw demodulations: 416
% Condensations: 568
% 
% Simple tautologies: 146
% Equational tautologies: 34
% Forward subsumptions: 3467
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 37
% 
% Binary resolution: 106
% Factoring: 1
% Forward superposition: 2226
% Backward superposition: 1972
% Self superposition: 97
% Equality factoring: 62
% Equality resolution: 11
% 
% Split clauses: 41
% Split components: 91
% SAT solver clauses: 1297
% SAT solver unit clauses: 38
% SAT solver binary clauses: 24
% 
% Sat splits: 41
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1297
% TWLsolver calls for satisfiability: 239
% 
% Memory used [KB]: 11001
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f18851,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f77,f132,f131,f76,f134,f75,f136,f130,f161,f164,f162,f177,f181,f179,f187,f438,f447,f439,f570,f573,f571,f821,f823,f822,f825,f836,f838,f842,f844,f846,f848,f864,f865,f866,f1284,f1268,f1257,f1252,f1243,f1292,f1242,f1311,f1313,f1314,f1317,f1318,f1220,f1328,f1330,f1215,f1329,f1212,f1338,f1445,f1453,f2734,f2735,f3991,f3992,f4517,f4502,f4598,f4478,f1301,f4669,f4670,f4671,f5102,f5155,f6270,f6271,f6273,f6274,f6277,f6278,f8379,f8377,f178,f8402,f8448,f8460,f8449,f8450,f8451,f8464,f8457,f8452,f8463,f8458,f8461,f8503,f8504,f8505,f8516,f8517,f8518,f8521,f8524,f9164,f9166,f9168,f9169,f9174,f9175,f13341,f13342,f13345,f13343,f572,f17061,f17064,f17086,f17084,f18764,f18843])).
fof(f18843,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f18842,f78])).
fof(f78,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.evalProps_43_prove_ltb',aHL_TRUTH)).
fof(f18842,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f18762,f7959])).
fof(f7959,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(equality_resolution,[],[f574])).
fof(f574,plain,(
  ( ! [X17,X18,X16] : (s(X16,X17) != s(X16,X18) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X16,bool),'const.bool.LET_2'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X17))),s(X16,X18)))) )),
  inference(subsumption_resolution,[],[f545,f80])).
fof(f80,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.evalProps.evalProps_43_prove_ltb',aHL_FALSITY)).
fof(f545,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(X16,X17) != s(X16,X18) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X16,bool),'const.bool.LET_2'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f139,f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.evalProps_43_prove_ltb',aHL_BOOL_CASES)).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f138,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7,X6] : ! [X9,X22] : s(X7,'const.bool.LET_2'(s(fun(X6,X7),X9),s(X6,X22))) = s(X7,happ(s(fun(X6,X7),X9),s(X6,X22)))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.bool.LET_THM')).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f95,f91])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X19,X20,X21] : (p(s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) <=> s(X19,X20) = s(X19,X21))),
  file('thm.evalProps.evalProps_43_prove_ltb',arity_0_HL_EQ)).
fof(f18762,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f18761,f5312])).
fof(f5312,plain,(
  s(sK2,sK4) = s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5))) | $spl0),
  inference(unit_resulting_resolution,[],[f5177,f141])).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f140,f91])).
fof(f140,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f94,f91])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f5177,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5)))))) | $spl0),
  inference(forward_demodulation,[],[f5157,f1624])).
fof(f1624,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X5)),X8),s('type.pair.prod'(X6,X7),X9))) = s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X6,X7),X5),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,X5)),X8))),s('type.pair.prod'(X6,X7),X9)))) )),
  inference(superposition,[],[f1389,f91])).
fof(f1389,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(backward_demodulation,[],[f1375,f142])).
fof(f142,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.bool.LET_2'(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(forward_demodulation,[],[f103,f91])).
fof(f103,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X6,X7] : ! [X17,X18] : s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X17),s('type.pair.prod'(X6,X7),X18))) = s(X8,happ(s(fun('type.pair.prod'(X6,X7),X8),happ(s(fun(fun(X6,fun(X7,X8)),fun('type.pair.prod'(X6,X7),X8)),'const.pair.UNCURRY_0'),s(fun(X6,fun(X7,X8)),X17))),s('type.pair.prod'(X6,X7),X18)))),
  file('thm.evalProps.evalProps_43_prove_ltb','ARITY_0_2_const.pair.UNCURRY')).
fof(f1375,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun('type.pair.prod'(X4,X5),X6),'const.pair.UNCURRY_1'(s(fun(X4,fun(X5,X6)),X7))) = s(fun('type.pair.prod'(X4,X5),X6),'const.bool.LET_2'(s(fun(fun(X4,fun(X5,X6)),fun('type.pair.prod'(X4,X5),X6)),'const.pair.UNCURRY_0'),s(fun(X4,fun(X5,X6)),X7)))) )),
  inference(superposition,[],[f102,f91])).
fof(f102,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : ! [X3] : s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))) = s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7,X8] : ! [X17] : s(fun('type.pair.prod'(X6,X7),X8),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,X8)),X17))) = s(fun('type.pair.prod'(X6,X7),X8),happ(s(fun(fun(X6,fun(X7,X8)),fun('type.pair.prod'(X6,X7),X8)),'const.pair.UNCURRY_0'),s(fun(X6,fun(X7,X8)),X17)))),
  file('thm.evalProps.evalProps_43_prove_ltb','ARITY_0_1_const.pair.UNCURRY')).
fof(f5157,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f130,f139])).
fof(f18761,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5)))))) | $spl2),
  inference(backward_demodulation,[],[f18550,f5311])).
fof(f5311,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0))))))) | $spl2),
  inference(forward_demodulation,[],[f5305,f91])).
fof(f5305,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f5264,f141])).
fof(f5264,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))) | $spl2),
  inference(unit_resulting_resolution,[],[f90,f5156])).
fof(f5156,plain,(
  ( ! [X6,X7] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7)))) ) | $spl2),
  inference(forward_demodulation,[],[f132,f91])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(X2,X1,X0)),s(X1,sK9(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(X2,X1,X0)),s(X1,sK9(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : ! [X22] : ? [X32,X33] : s('type.pair.prod'(X6,X7),X22) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X32),s(X7,X33)))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f18550,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X3)),X4),s('type.pair.prod'(X0,X1),X2))) = s(X3,'const.bool.LET_2'(s(fun(X1,X3),'const.bool.LET_2'(s(fun(X0,fun(X1,X3)),X4),s(X0,sK8(X2,X1,X0)))),s(X1,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f1777,f90])).
fof(f1777,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X4,'const.bool.LET_2'(s(fun(X5,X4),'const.bool.LET_2'(s(fun(X6,fun(X5,X4)),X7),s(X6,X8))),s(X5,X9))) = s(X4,'const.pair.UNCURRY_2'(s(fun(X6,fun(X5,X4)),X7),s('type.pair.prod'(X6,X5),'const.pair.,_2'(s(X6,X8),s(X5,X9)))))) )),
  inference(superposition,[],[f143,f91])).
fof(f143,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'const.bool.LET_2'(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(forward_demodulation,[],[f104,f91])).
fof(f104,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X6,X7] : ! [X9,X22,X23] : s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X9),s(X6,X22))),s(X7,X23))) = s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X9),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X22),s(X7,X23)))))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.pair.UNCURRY_DEF')).
fof(f18764,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f18763,f78])).
fof(f18763,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f18762,f4725])).
fof(f4725,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4))) | $spl4),
  inference(unit_resulting_resolution,[],[f4703,f158])).
fof(f158,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f80,f81])).
fof(f4703,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4)))) | $spl4),
  inference(backward_demodulation,[],[f4683,f4702])).
fof(f4702,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7)))))) | $spl4),
  inference(forward_demodulation,[],[f4681,f91])).
fof(f4681,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7)))))) | $spl4),
  inference(unit_resulting_resolution,[],[f4643,f139])).
fof(f4643,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(forward_demodulation,[],[f134,f91])).
fof(f4683,plain,(
  s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(superposition,[],[f4643,f91])).
fof(f17084,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK10(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK10(X2,X3,bool,bool))))) ) | $spl66),
  inference(cnf_transformation,[],[f17084_D])).
fof(f17084_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK10(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK10(X2,X3,bool,bool))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f17086,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl69),
  inference(cnf_transformation,[],[f17086_D])).
fof(f17086_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f17064,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK10(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK10(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f787,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f81,f81])).
fof(f787,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK10(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK10(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f93,f81])).
fof(f93,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.evalProps_43_prove_ltb',aHL_EXT)).
fof(f17061,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK10(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK10(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f787,f156])).
fof(f572,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl20),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f13343,plain,(
  ( ! [X19] : (~p(s(bool,f(s(bool,X19)))) | s(bool,'T_0') = s(bool,X19)) ) | $spl64),
  inference(cnf_transformation,[],[f13343_D])).
fof(f13343_D,plain,(
  ( ! [X19] : (~p(s(bool,f(s(bool,X19)))) | s(bool,'T_0') = s(bool,X19)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13345,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,f(s(bool,X0)))) | p(s(bool,X1))) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(condensation,[],[f13344])).
fof(f13344,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,f(s(bool,X22)))) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(subsumption_resolution,[],[f13123,f12711])).
fof(f12711,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(subsumption_resolution,[],[f12639,f80])).
fof(f12639,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X5) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(superposition,[],[f86,f12495])).
fof(f12495,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(global_subsumption,[],[f12439,f77,f76,f75,f78,f81,f89,f86,f90,f91,f137,f93,f102,f142,f143,f128,f145,f154,f156,f158,f161,f177,f180,f181,f149,f153,f239,f438,f440,f441,f447,f439,f546,f550,f570,f573,f574,f688,f778,f780,f920,f921,f889,f1385,f1386,f1375,f1387,f1389,f1655,f1656,f1624,f1659,f1876,f1877,f1777,f1823,f2092,f2094,f1938,f1940,f1941,f2010,f2011,f2012,f2396,f2397,f2351,f2398,f2886,f7967,f7970,f7979,f7984,f7987,f7990,f8020,f8031,f8037,f8048,f8049,f8050,f8054,f7887,f7888,f7895,f8066,f7910,f8078,f8079,f8080,f8084,f7928,f7929,f7936,f8096,f7951,f7959,f8448,f8449,f8450,f8451,f8452,f8712,f8711,f8708,f8713,f8731,f308,f272,f307,f271,f1627,f1626,f2632,f2852,f2617,f2816,f2628,f2821,f2621,f2828,f2625,f2823,f2623,f2825,f2616,f2817,f2600,f2832,f2662,f2881,f2661,f2882,f2794,f2579,f2798,f2795,f2576,f2575,f2634,f2853,f2645,f2858,f2640,f2862,f2642,f2860,f2638,f2865,f2633,f2854,f2653,f2871,f2651,f2873,f2602,f2833,f2610,f2840,f2608,f2842,f2601,f2834,f2656,f2869,f2649,f2876,f2613,f2838,f2606,f2845,f1661,f1662,f2404,f2406,f2400,f2402,f2403,f2405,f2399,f2401,f655,f530,f8674,f646,f643,f535,f532,f522,f519,f641,f517,f1884,f1886,f2132,f2136,f2102,f2106,f1888,f1890,f2130,f2134,f2100,f2104,f2152,f2150,f2146,f2144,f2142,f2138,f2122,f2120,f2116,f2114,f2112,f2108,f1887,f1889,f1883,f1885,f1879,f1880,f2772,f2771,f2594,f2593,f2544,f2543,f1985,f1984,f892,f890,f463,f461,f141,f139,f783,f785,f8701,f8685,f8683,f8673,f8463,f8088,f8058,f7954,f7953,f7938,f7927,f7913,f7912,f7897,f7886,f6682,f6681,f6680,f6679,f6678,f6677,f6667,f6665,f6651,f6650,f6647,f6646,f6642,f6633,f6629,f6627,f6625,f6620,f6616,f6614,f6612,f6607,f6606,f6605,f6545,f6544,f6543,f6520,f6518,f6511,f6510,f6509,f6364,f6363,f6356,f6355,f6354,f6353,f6349,f6334,f6333,f5696,f5695,f5688,f5687,f5686,f5685,f5681,f5678,f5677,f5676,f2800,f2793,f2770,f2592,f2581,f2574,f2542,f2148,f2118,f2013,f1983,f1939,f1878,f1625,f894,f787,f782,f662,f659,f648,f645,f534,f529,f521,f516,f465,f214,f213,f179,f80,f9340,f9821,f12000,f12029,f12030,f12004,f12032,f12018,f12398,f12486,f12487,f12402,f12489,f12490,f12491,f12492,f12493])).
fof(f12493,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'T_0') != s(bool,f(s(bool,X23)))) ) | ($spl16 | $spl59)),
  inference(subsumption_resolution,[],[f12430,f439])).
fof(f12430,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'T_0') != s(bool,f(s(bool,X23))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl16 | $spl59)),
  inference(superposition,[],[f439,f9821])).
fof(f12492,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl16 | $spl59)),
  inference(subsumption_resolution,[],[f12429,f439])).
fof(f12429,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) ) | $spl59),
  inference(superposition,[],[f180,f9821])).
fof(f12491,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,'F_0') = s(bool,f(s(bool,X20)))) ) | ($spl14 | $spl16 | $spl59)),
  inference(subsumption_resolution,[],[f12428,f439])).
fof(f12428,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,'F_0') = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl14 | $spl59)),
  inference(superposition,[],[f179,f9821])).
fof(f12490,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X4) = s(bool,f(s(bool,X3)))) ) | ($spl16 | $spl59)),
  inference(subsumption_resolution,[],[f12419,f439])).
fof(f12419,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X4) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl59),
  inference(superposition,[],[f86,f9821])).
fof(f12489,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'const.bool.LET_2'(s(fun(X77,bool),'const.bool.LET_2'(s(fun(X76,fun(X77,bool)),X78),s(X76,X79))),s(X77,X80))) = s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X77,bool),'const.bool.LET_2'(s(fun(X76,fun(X77,bool)),X78),s(X76,X79))),s(X77,X80))))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X76,fun(X77,bool)),X78),s('type.pair.prod'(X76,X77),'const.pair.,_2'(s(X76,X79),s(X77,X80)))))) ) | $spl59),
  inference(forward_demodulation,[],[f12404,f91])).
fof(f12404,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,happ(s(fun(X77,bool),'const.bool.LET_2'(s(fun(X76,fun(X77,bool)),X78),s(X76,X79))),s(X77,X80))) = s(bool,f(s(bool,happ(s(fun(X77,bool),'const.bool.LET_2'(s(fun(X76,fun(X77,bool)),X78),s(X76,X79))),s(X77,X80))))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X76,fun(X77,bool)),X78),s('type.pair.prod'(X76,X77),'const.pair.,_2'(s(X76,X79),s(X77,X80)))))) ) | $spl59),
  inference(superposition,[],[f9821,f143])).
fof(f12402,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),X71),s('type.pair.prod'(X69,X70),X72))) = s(bool,f(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),X71),s('type.pair.prod'(X69,X70),X72))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X69,X70),bool),'const.pair.UNCURRY_1'(s(fun(X69,fun(X70,bool)),X71))),s('type.pair.prod'(X69,X70),X72)))) ) | $spl59),
  inference(superposition,[],[f9821,f1389])).
fof(f12487,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,'const.bool.LET_2'(s(fun(X65,bool),'_dst_x0x1x0_3'(s(fun(X63,fun(X62,bool)),X64),s(fun(X65,X63),X66),s(X62,X68))),s(X65,X67))) = s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X65,bool),'_dst_x0x1x0_3'(s(fun(X63,fun(X62,bool)),X64),s(fun(X65,X63),X66),s(X62,X68))),s(X65,X67))))) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),'const.bool.LET_2'(s(fun(X63,fun(X62,bool)),X64),s(X63,'const.bool.LET_2'(s(fun(X65,X63),X66),s(X65,X67))))),s(X62,X68)))) ) | $spl59),
  inference(forward_demodulation,[],[f12401,f91])).
fof(f12401,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,happ(s(fun(X65,bool),'_dst_x0x1x0_3'(s(fun(X63,fun(X62,bool)),X64),s(fun(X65,X63),X66),s(X62,X68))),s(X65,X67))) = s(bool,f(s(bool,happ(s(fun(X65,bool),'_dst_x0x1x0_3'(s(fun(X63,fun(X62,bool)),X64),s(fun(X65,X63),X66),s(X62,X68))),s(X65,X67))))) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),'const.bool.LET_2'(s(fun(X63,fun(X62,bool)),X64),s(X63,'const.bool.LET_2'(s(fun(X65,X63),X66),s(X65,X67))))),s(X62,X68)))) ) | $spl59),
  inference(superposition,[],[f9821,f145])).
fof(f12486,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,'const.bool.LET_2'(s(fun(X57,bool),'const.bool.LET_2'(s(fun(X56,fun(X57,bool)),X58),s(X56,'const.bool.LET_2'(s(fun(X55,X56),X59),s(X55,X61))))),s(X57,X60))) = s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X57,bool),'const.bool.LET_2'(s(fun(X56,fun(X57,bool)),X58),s(X56,'const.bool.LET_2'(s(fun(X55,X56),X59),s(X55,X61))))),s(X57,X60))))) | s(bool,'T_0') = s(bool,happ(s(fun(X55,bool),'_dst_x0x1x0_3'(s(fun(X56,fun(X57,bool)),X58),s(fun(X55,X56),X59),s(X57,X60))),s(X55,X61)))) ) | $spl59),
  inference(forward_demodulation,[],[f12400,f91])).
fof(f12400,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(bool,happ(s(fun(X57,bool),'const.bool.LET_2'(s(fun(X56,fun(X57,bool)),X58),s(X56,'const.bool.LET_2'(s(fun(X55,X56),X59),s(X55,X61))))),s(X57,X60))) = s(bool,f(s(bool,happ(s(fun(X57,bool),'const.bool.LET_2'(s(fun(X56,fun(X57,bool)),X58),s(X56,'const.bool.LET_2'(s(fun(X55,X56),X59),s(X55,X61))))),s(X57,X60))))) | s(bool,'T_0') = s(bool,happ(s(fun(X55,bool),'_dst_x0x1x0_3'(s(fun(X56,fun(X57,bool)),X58),s(fun(X55,X56),X59),s(X57,X60))),s(X55,X61)))) ) | $spl59),
  inference(superposition,[],[f9821,f145])).
fof(f12398,plain,(
  ( ! [X45,X46,X44] : (s(bool,'const.bool.LET_2'(s(fun(X44,bool),X45),s(X44,X46))) = s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X44,bool),X45),s(X44,X46))))) | s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),'const.bool.LET_2'(s(fun(fun(X44,bool),fun(X44,bool)),'const.bool.LET_0'),s(fun(X44,bool),X45))),s(X44,X46)))) ) | $spl59),
  inference(superposition,[],[f9821,f137])).
fof(f12018,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl59),
  inference(trivial_inequality_removal,[],[f12009])).
fof(f12009,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl59),
  inference(superposition,[],[f8713,f156])).
fof(f12032,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'T_0') != s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X74,bool),'const.bool.LET_2'(s(fun(X73,fun(X74,bool)),X75),s(X73,X76))),s(X74,X77))))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X73,fun(X74,bool)),X75),s('type.pair.prod'(X73,X74),'const.pair.,_2'(s(X73,X76),s(X74,X77)))))) ) | $spl59),
  inference(forward_demodulation,[],[f12006,f91])).
fof(f12006,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(X74,bool),'const.bool.LET_2'(s(fun(X73,fun(X74,bool)),X75),s(X73,X76))),s(X74,X77))))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X73,fun(X74,bool)),X75),s('type.pair.prod'(X73,X74),'const.pair.,_2'(s(X73,X76),s(X74,X77)))))) ) | $spl59),
  inference(superposition,[],[f8713,f143])).
fof(f12004,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,f(s(bool,'const.pair.UNCURRY_2'(s(fun(X66,fun(X67,bool)),X68),s('type.pair.prod'(X66,X67),X69))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X66,X67),bool),'const.pair.UNCURRY_1'(s(fun(X66,fun(X67,bool)),X68))),s('type.pair.prod'(X66,X67),X69)))) ) | $spl59),
  inference(superposition,[],[f8713,f1389])).
fof(f12030,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X62,bool),'_dst_x0x1x0_3'(s(fun(X60,fun(X59,bool)),X61),s(fun(X62,X60),X63),s(X59,X65))),s(X62,X64))))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),'const.bool.LET_2'(s(fun(X60,fun(X59,bool)),X61),s(X60,'const.bool.LET_2'(s(fun(X62,X60),X63),s(X62,X64))))),s(X59,X65)))) ) | $spl59),
  inference(forward_demodulation,[],[f12003,f91])).
fof(f12003,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(X62,bool),'_dst_x0x1x0_3'(s(fun(X60,fun(X59,bool)),X61),s(fun(X62,X60),X63),s(X59,X65))),s(X62,X64))))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),'const.bool.LET_2'(s(fun(X60,fun(X59,bool)),X61),s(X60,'const.bool.LET_2'(s(fun(X62,X60),X63),s(X62,X64))))),s(X59,X65)))) ) | $spl59),
  inference(superposition,[],[f8713,f145])).
fof(f12029,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') != s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X54,bool),'const.bool.LET_2'(s(fun(X53,fun(X54,bool)),X55),s(X53,'const.bool.LET_2'(s(fun(X52,X53),X56),s(X52,X58))))),s(X54,X57))))) | s(bool,'T_0') = s(bool,happ(s(fun(X52,bool),'_dst_x0x1x0_3'(s(fun(X53,fun(X54,bool)),X55),s(fun(X52,X53),X56),s(X54,X57))),s(X52,X58)))) ) | $spl59),
  inference(forward_demodulation,[],[f12002,f91])).
fof(f12002,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(X54,bool),'const.bool.LET_2'(s(fun(X53,fun(X54,bool)),X55),s(X53,'const.bool.LET_2'(s(fun(X52,X53),X56),s(X52,X58))))),s(X54,X57))))) | s(bool,'T_0') = s(bool,happ(s(fun(X52,bool),'_dst_x0x1x0_3'(s(fun(X53,fun(X54,bool)),X55),s(fun(X52,X53),X56),s(X54,X57))),s(X52,X58)))) ) | $spl59),
  inference(superposition,[],[f8713,f145])).
fof(f12000,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,f(s(bool,'const.bool.LET_2'(s(fun(X41,bool),X42),s(X41,X43))))) | s(bool,'T_0') = s(bool,happ(s(fun(X41,bool),'const.bool.LET_2'(s(fun(fun(X41,bool),fun(X41,bool)),'const.bool.LET_0'),s(fun(X41,bool),X42))),s(X41,X43)))) ) | $spl59),
  inference(superposition,[],[f8713,f137])).
fof(f9821,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl59),
  inference(superposition,[],[f8701,f81])).
fof(f9340,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),bool),'_dst_x0x1x0_3'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun(X0,fun(X1,X2)),X3))))) )),
  inference(forward_demodulation,[],[f9339,f91])).
fof(f9339,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),'_dst_x0x1x0_3'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun(X0,fun(X1,X2)),X3))))) )),
  inference(forward_demodulation,[],[f9338,f2011])).
fof(f9338,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(forward_demodulation,[],[f9224,f91])).
fof(f9224,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(unit_resulting_resolution,[],[f102,f139])).
fof(f213,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(bool,X1),X0) = s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK9(X0,X1,bool)))) | s(bool,'T_0') = s(bool,sK8(X0,X1,bool))) )),
  inference(superposition,[],[f90,f81])).
fof(f214,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(X1,bool),X0) = s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK8(X0,bool,X1)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X0,bool,X1))) )),
  inference(superposition,[],[f90,f81])).
fof(f465,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(superposition,[],[f91,f81])).
fof(f516,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X6) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X5))),s(bool,X6)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f139,f81])).
fof(f521,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X20)))) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f139,f81])).
fof(f529,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X6) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X6))),s(bool,X5)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f139,f81])).
fof(f534,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X20))),s(bool,'F_0')))) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f139,f81])).
fof(f645,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X18)))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f141,f81])).
fof(f648,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X26)))) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f141,f81])).
fof(f659,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X18))),s(bool,'F_0')))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f141,f81])).
fof(f662,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X26))),s(bool,X25)))) | s(bool,'F_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f141,f81])).
fof(f782,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK10(X16,X17,X18,bool))) )),
  inference(superposition,[],[f93,f81])).
fof(f894,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X24,bool),X25),s(X24,X26))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),'const.bool.LET_2'(s(fun(fun(X24,bool),fun(X24,bool)),'const.bool.LET_0'),s(fun(X24,bool),X25))),s(X24,X26)))) )),
  inference(superposition,[],[f137,f81])).
fof(f1625,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,bool)),X12),s('type.pair.prod'(X10,X11),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),'const.pair.UNCURRY_1'(s(fun(X10,fun(X11,bool)),X12))),s('type.pair.prod'(X10,X11),X13)))) )),
  inference(superposition,[],[f1389,f81])).
fof(f1878,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X7,bool),'const.bool.LET_2'(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,bool)),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(forward_demodulation,[],[f1824,f91])).
fof(f1824,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),'const.bool.LET_2'(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,bool)),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f143,f81])).
fof(f1939,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),'const.bool.LET_2'(s(fun(X17,fun(X18,bool)),X19),s(X17,'const.bool.LET_2'(s(fun(X16,X17),X20),s(X16,X22))))),s(X18,X21))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),'_dst_x0x1x0_3'(s(fun(X17,fun(X18,bool)),X19),s(fun(X16,X17),X20),s(X18,X21))),s(X16,X22)))) )),
  inference(superposition,[],[f145,f81])).
fof(f1983,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (s(X11,happ(s(fun(X12,X11),'const.bool.LET_2'(s(fun(bool,fun(X12,X11)),X13),s(bool,'F_0'))),s(X12,X14))) = s(X11,happ(s(fun(X8,X11),'_dst_x0x1x0_3'(s(fun(bool,fun(X12,X11)),X13),s(fun(X8,bool),X9),s(X12,X14))),s(X8,X10))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),X9),s(X8,X10)))) )),
  inference(superposition,[],[f145,f81])).
fof(f2013,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1x0_3'(s(fun(X31,fun(X30,bool)),X32),s(fun(X33,X31),X34),s(X30,X36))),s(X33,X35))) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),'const.bool.LET_2'(s(fun(X31,fun(X30,bool)),X32),s(X31,'const.bool.LET_2'(s(fun(X33,X31),X34),s(X33,X35))))),s(X30,X36)))) )),
  inference(superposition,[],[f145,f81])).
fof(f2118,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X166,bool),'_dst_x0x1x0_3'(s(fun(X167,fun(X168,bool)),X169),s(fun(X166,X167),X170),s(X168,X171))),s(X166,X172))) | p(s(bool,'const.bool.LET_2'(s(fun(X168,bool),'const.bool.LET_2'(s(fun(X167,fun(X168,bool)),X169),s(X167,'const.bool.LET_2'(s(fun(X166,X167),X170),s(X166,X172))))),s(X168,X171))))) ) | $spl14),
  inference(forward_demodulation,[],[f2117,f91])).
fof(f2117,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (p(s(bool,'const.bool.LET_2'(s(fun(X168,bool),'const.bool.LET_2'(s(fun(X167,fun(X168,bool)),X169),s(X167,'const.bool.LET_2'(s(fun(X166,X167),X170),s(X166,X172))))),s(X168,X171)))) | s(bool,'F_0') = s(bool,happ(s(fun(X166,bool),'_dst_x0x1x0_3'(s(fun(X167,fun(X168,bool)),X169),s(fun(X166,X167),X170),s(X168,X171))),s(X166,X172)))) ) | $spl14),
  inference(forward_demodulation,[],[f2037,f91])).
fof(f2037,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (p(s(bool,happ(s(fun(X168,bool),'const.bool.LET_2'(s(fun(X167,fun(X168,bool)),X169),s(X167,'const.bool.LET_2'(s(fun(X166,X167),X170),s(X166,X172))))),s(X168,X171)))) | s(bool,'F_0') = s(bool,happ(s(fun(X166,bool),'_dst_x0x1x0_3'(s(fun(X167,fun(X168,bool)),X169),s(fun(X166,X167),X170),s(X168,X171))),s(X166,X172)))) ) | $spl14),
  inference(superposition,[],[f179,f145])).
fof(f2148,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X180,bool),'const.bool.LET_2'(s(fun(X181,fun(X180,bool)),X182),s(X181,'const.bool.LET_2'(s(fun(X183,X181),X184),s(X183,X185))))),s(X180,X186))) | p(s(bool,'const.bool.LET_2'(s(fun(X183,bool),'_dst_x0x1x0_3'(s(fun(X181,fun(X180,bool)),X182),s(fun(X183,X181),X184),s(X180,X186))),s(X183,X185))))) ) | $spl14),
  inference(forward_demodulation,[],[f2147,f91])).
fof(f2147,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (p(s(bool,'const.bool.LET_2'(s(fun(X183,bool),'_dst_x0x1x0_3'(s(fun(X181,fun(X180,bool)),X182),s(fun(X183,X181),X184),s(X180,X186))),s(X183,X185)))) | s(bool,'F_0') = s(bool,happ(s(fun(X180,bool),'const.bool.LET_2'(s(fun(X181,fun(X180,bool)),X182),s(X181,'const.bool.LET_2'(s(fun(X183,X181),X184),s(X183,X185))))),s(X180,X186)))) ) | $spl14),
  inference(forward_demodulation,[],[f2070,f91])).
fof(f2070,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (p(s(bool,happ(s(fun(X183,bool),'_dst_x0x1x0_3'(s(fun(X181,fun(X180,bool)),X182),s(fun(X183,X181),X184),s(X180,X186))),s(X183,X185)))) | s(bool,'F_0') = s(bool,happ(s(fun(X180,bool),'const.bool.LET_2'(s(fun(X181,fun(X180,bool)),X182),s(X181,'const.bool.LET_2'(s(fun(X183,X181),X184),s(X183,X185))))),s(X180,X186)))) ) | $spl14),
  inference(superposition,[],[f179,f145])).
fof(f2542,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(X21,fun(X22,bool)),X23),s('type.pair.prod'(X21,X22),'const.bool.LET_2'(s(fun(X18,'type.pair.prod'(X21,X22)),X24),s(X18,X20))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X21,fun(X22,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X21,fun(X22,bool)),X23),s(bool,'const.bool.LET_2'(s(fun(X18,bool),X19),s(X18,X20))))),s('type.pair.prod'(X21,X22),'const.bool.LET_2'(s(fun(X18,'type.pair.prod'(X21,X22)),X24),s(X18,X20)))))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X18,bool),X19),s(X18,X20)))) )),
  inference(superposition,[],[f1661,f81])).
fof(f2574,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84] : (s(X84,'const.bool.LET_2'(s(fun(bool,X84),X85),s(bool,'F_0'))) != s(X84,'const.pair.UNCURRY_2'(s(fun(X86,fun(X87,X84)),X88),s('type.pair.prod'(X86,X87),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X86,X87)),X89),s(bool,'F_0'))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X86,fun(X87,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X86,fun(X87,X84)),X88),s(X84,'const.bool.LET_2'(s(fun(bool,X84),X85),s(bool,X83))))),s('type.pair.prod'(X86,X87),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X86,X87)),X89),s(bool,X83)))))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f1661,f81])).
fof(f2581,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139] : (s(X139,'const.bool.LET_2'(s(fun(bool,X139),X140),s(bool,X138))) != s(X139,'const.pair.UNCURRY_2'(s(fun(X141,fun(X142,X139)),X143),s('type.pair.prod'(X141,X142),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X141,X142)),X144),s(bool,X138))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X141,fun(X142,bool)),'_dst_x00x11x00_3'(s(fun(X139,fun(X139,bool)),'=_0'),s(fun(X141,fun(X142,X139)),X143),s(X139,'const.bool.LET_2'(s(fun(bool,X139),X140),s(bool,'F_0'))))),s('type.pair.prod'(X141,X142),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X141,X142)),X144),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f1661,f81])).
fof(f2592,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(X3,bool),X6),s(X3,X5))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X0,fun(X1,bool)),X2),s(bool,'const.bool.LET_2'(s(fun(X3,bool),X6),s(X3,X5))))),s('type.pair.prod'(X0,X1),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X0,X1)),X4),s(X3,X5)))))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X0,X1)),X4),s(X3,X5)))))) )),
  inference(superposition,[],[f1661,f81])).
fof(f2770,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X11,fun(X12,bool)),X13),s(bool,'F_0'))),s('type.pair.prod'(X11,X12),'const.bool.LET_2'(s(fun(X8,'type.pair.prod'(X11,X12)),X14),s(X8,X10)))))) | s(bool,'const.bool.LET_2'(s(fun(X8,bool),X9),s(X8,X10))) = s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),'const.bool.LET_2'(s(fun(X8,'type.pair.prod'(X11,X12)),X14),s(X8,X10))))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),X9),s(X8,X10)))) )),
  inference(superposition,[],[f1662,f81])).
fof(f2793,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X86,fun(X86,bool)),'=_0'),s(fun(X84,fun(X85,X86)),X87),s(X86,'const.bool.LET_2'(s(fun(bool,X86),X88),s(bool,'F_0'))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X84,X85)),X89),s(bool,'F_0')))))) | s(X86,'const.bool.LET_2'(s(fun(bool,X86),X88),s(bool,X83))) = s(X86,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X86)),X87),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X84,X85)),X89),s(bool,X83))))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f1662,f81])).
fof(f2800,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X139,fun(X140,bool)),'_dst_x00x11x00_3'(s(fun(X141,fun(X141,bool)),'=_0'),s(fun(X139,fun(X140,X141)),X142),s(X141,'const.bool.LET_2'(s(fun(bool,X141),X143),s(bool,X138))))),s('type.pair.prod'(X139,X140),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X139,X140)),X144),s(bool,X138)))))) | s(X141,'const.bool.LET_2'(s(fun(bool,X141),X143),s(bool,'F_0'))) = s(X141,'const.pair.UNCURRY_2'(s(fun(X139,fun(X140,X141)),X142),s('type.pair.prod'(X139,X140),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X139,X140)),X144),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f1662,f81])).
fof(f5676,plain,(
  ( ! [X0] : (s('type.pair.prod'(bool,bool),X0) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK8(X0,bool,bool)) | s(bool,'T_0') = s(bool,sK9(X0,bool,bool))) )),
  inference(superposition,[],[f213,f81])).
fof(f5677,plain,(
  ( ! [X2,X1] : (s('type.pair.prod'(bool,bool),X1) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK8(X1,bool,bool)) | s(bool,X2) = s(bool,sK9(X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f213,f156])).
fof(f5678,plain,(
  ( ! [X4,X3] : (s('type.pair.prod'(bool,bool),X3) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,sK8(X3,bool,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK9(X3,bool,bool))) )),
  inference(superposition,[],[f213,f156])).
fof(f5681,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(bool,fun(X0,X2)),X3),s('type.pair.prod'(bool,X0),X1))) = s(X2,happ(s(fun(X0,X2),'const.bool.LET_2'(s(fun(bool,fun(X0,X2)),X3),s(bool,'F_0'))),s(X0,sK9(X1,X0,bool)))) | s(bool,'T_0') = s(bool,sK8(X1,X0,bool))) )),
  inference(superposition,[],[f143,f213])).
fof(f5685,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X16),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X16),fun('type.pair.prod'(bool,X16),bool)),'=_0'),s('type.pair.prod'(bool,X16),X17))),s('type.pair.prod'(bool,X16),X18)))) | s('type.pair.prod'(bool,X16),X18) != s('type.pair.prod'(bool,X16),'const.pair.,_2'(s(bool,'F_0'),s(X16,sK9(X17,X16,bool)))) | s(bool,'T_0') = s(bool,sK8(X17,X16,bool))) )),
  inference(superposition,[],[f139,f213])).
fof(f5686,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X19),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X19),fun('type.pair.prod'(bool,X19),bool)),'=_0'),s('type.pair.prod'(bool,X19),X21))),s('type.pair.prod'(bool,X19),X20)))) | s('type.pair.prod'(bool,X19),X21) != s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,'F_0'),s(X19,sK9(X20,X19,bool)))) | s(bool,'T_0') = s(bool,sK8(X20,X19,bool))) )),
  inference(superposition,[],[f139,f213])).
fof(f5687,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X22),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X22),fun('type.pair.prod'(bool,X22),bool)),'=_0'),s('type.pair.prod'(bool,X22),X23))),s('type.pair.prod'(bool,X22),X24)))) | s('type.pair.prod'(bool,X22),X24) = s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,'F_0'),s(X22,sK9(X23,X22,bool)))) | s(bool,'T_0') = s(bool,sK8(X23,X22,bool))) )),
  inference(superposition,[],[f141,f213])).
fof(f5688,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X25),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X25),fun('type.pair.prod'(bool,X25),bool)),'=_0'),s('type.pair.prod'(bool,X25),X27))),s('type.pair.prod'(bool,X25),X26)))) | s('type.pair.prod'(bool,X25),X27) = s('type.pair.prod'(bool,X25),'const.pair.,_2'(s(bool,'F_0'),s(X25,sK9(X26,X25,bool)))) | s(bool,'T_0') = s(bool,sK8(X26,X25,bool))) )),
  inference(superposition,[],[f141,f213])).
fof(f5695,plain,(
  ( ! [X80,X78,X76,X74,X81,X79,X77,X75] : (s(X76,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X74),X76),X77),s('type.pair.prod'(bool,X74),X75))) != s(X76,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,X76)),X80),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X74),'type.pair.prod'(X78,X79)),X81),s('type.pair.prod'(bool,X74),X75))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,bool)),'_dst_x00x11x00_3'(s(fun(X76,fun(X76,bool)),'=_0'),s(fun(X78,fun(X79,X76)),X80),s(X76,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X74),X76),X77),s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'F_0'),s(X74,sK9(X75,X74,bool)))))))),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X74),'type.pair.prod'(X78,X79)),X81),s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'F_0'),s(X74,sK9(X75,X74,bool))))))))) | s(bool,'T_0') = s(bool,sK8(X75,X74,bool))) )),
  inference(superposition,[],[f1661,f213])).
fof(f5696,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X86,fun(X86,bool)),'=_0'),s(fun(X84,fun(X85,X86)),X87),s(X86,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X82),X86),X88),s('type.pair.prod'(bool,X82),X83))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X82),'type.pair.prod'(X84,X85)),X89),s('type.pair.prod'(bool,X82),X83)))))) | s(X86,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X82),X86),X88),s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,'F_0'),s(X82,sK9(X83,X82,bool)))))) = s(X86,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X86)),X87),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X82),'type.pair.prod'(X84,X85)),X89),s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,'F_0'),s(X82,sK9(X83,X82,bool)))))))) | s(bool,'T_0') = s(bool,sK8(X83,X82,bool))) )),
  inference(superposition,[],[f1662,f213])).
fof(f6333,plain,(
  ( ! [X2,X1] : (s('type.pair.prod'(bool,bool),X1) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X1,bool,bool)) | s(bool,X2) = s(bool,sK8(X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f214,f156])).
fof(f6334,plain,(
  ( ! [X4,X3] : (s('type.pair.prod'(bool,bool),X3) = s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X3,bool,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK8(X3,bool,bool))) )),
  inference(superposition,[],[f214,f156])).
fof(f6349,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(bool,X2)),X3),s('type.pair.prod'(X0,bool),X1))) = s(X2,happ(s(fun(bool,X2),'const.bool.LET_2'(s(fun(X0,fun(bool,X2)),X3),s(X0,sK8(X1,bool,X0)))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X1,bool,X0))) )),
  inference(superposition,[],[f143,f214])).
fof(f6353,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X16,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X16,bool),fun('type.pair.prod'(X16,bool),bool)),'=_0'),s('type.pair.prod'(X16,bool),X17))),s('type.pair.prod'(X16,bool),X18)))) | s('type.pair.prod'(X16,bool),X18) != s('type.pair.prod'(X16,bool),'const.pair.,_2'(s(X16,sK8(X17,bool,X16)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X17,bool,X16))) )),
  inference(superposition,[],[f139,f214])).
fof(f6354,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X19,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X19,bool),fun('type.pair.prod'(X19,bool),bool)),'=_0'),s('type.pair.prod'(X19,bool),X21))),s('type.pair.prod'(X19,bool),X20)))) | s('type.pair.prod'(X19,bool),X21) != s('type.pair.prod'(X19,bool),'const.pair.,_2'(s(X19,sK8(X20,bool,X19)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X20,bool,X19))) )),
  inference(superposition,[],[f139,f214])).
fof(f6355,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X22,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X22,bool),fun('type.pair.prod'(X22,bool),bool)),'=_0'),s('type.pair.prod'(X22,bool),X23))),s('type.pair.prod'(X22,bool),X24)))) | s('type.pair.prod'(X22,bool),X24) = s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,sK8(X23,bool,X22)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X23,bool,X22))) )),
  inference(superposition,[],[f141,f214])).
fof(f6356,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X25,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X25,bool),fun('type.pair.prod'(X25,bool),bool)),'=_0'),s('type.pair.prod'(X25,bool),X27))),s('type.pair.prod'(X25,bool),X26)))) | s('type.pair.prod'(X25,bool),X27) = s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,sK8(X26,bool,X25)),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X26,bool,X25))) )),
  inference(superposition,[],[f141,f214])).
fof(f6363,plain,(
  ( ! [X80,X78,X76,X74,X81,X79,X77,X75] : (s(X76,'const.bool.LET_2'(s(fun('type.pair.prod'(X74,bool),X76),X77),s('type.pair.prod'(X74,bool),X75))) != s(X76,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,X76)),X80),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun('type.pair.prod'(X74,bool),'type.pair.prod'(X78,X79)),X81),s('type.pair.prod'(X74,bool),X75))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,bool)),'_dst_x00x11x00_3'(s(fun(X76,fun(X76,bool)),'=_0'),s(fun(X78,fun(X79,X76)),X80),s(X76,'const.bool.LET_2'(s(fun('type.pair.prod'(X74,bool),X76),X77),s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,sK8(X75,bool,X74)),s(bool,'F_0'))))))),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun('type.pair.prod'(X74,bool),'type.pair.prod'(X78,X79)),X81),s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,sK8(X75,bool,X74)),s(bool,'F_0')))))))) | s(bool,'T_0') = s(bool,sK9(X75,bool,X74))) )),
  inference(superposition,[],[f1661,f214])).
fof(f6364,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X86,fun(X86,bool)),'=_0'),s(fun(X84,fun(X85,X86)),X87),s(X86,'const.bool.LET_2'(s(fun('type.pair.prod'(X82,bool),X86),X88),s('type.pair.prod'(X82,bool),X83))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun('type.pair.prod'(X82,bool),'type.pair.prod'(X84,X85)),X89),s('type.pair.prod'(X82,bool),X83)))))) | s(X86,'const.bool.LET_2'(s(fun('type.pair.prod'(X82,bool),X86),X88),s('type.pair.prod'(X82,bool),'const.pair.,_2'(s(X82,sK8(X83,bool,X82)),s(bool,'F_0'))))) = s(X86,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X86)),X87),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun('type.pair.prod'(X82,bool),'type.pair.prod'(X84,X85)),X89),s('type.pair.prod'(X82,bool),'const.pair.,_2'(s(X82,sK8(X83,bool,X82)),s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,sK9(X83,bool,X82))) )),
  inference(superposition,[],[f1662,f214])).
fof(f6509,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X49),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X49),s(bool,X48))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f465,f81])).
fof(f6510,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X51),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X51),s(bool,X50))) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f465,f156])).
fof(f6511,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X55),s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f465,f156])).
fof(f6518,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,happ(s(fun(bool,bool),X75),s(bool,X74))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X75),s(bool,'T_0'))) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f465,f156])).
fof(f6520,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X81),s(bool,X80))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X81),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f465,f81])).
fof(f6543,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X34,bool),X35),s(X34,X36))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X34,bool),'const.bool.LET_2'(s(fun(fun(X34,bool),fun(X34,bool)),'const.bool.LET_0'),s(fun(X34,bool),X35))),s(X34,X36)))) )),
  inference(superposition,[],[f465,f137])).
fof(f6544,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),'const.bool.LET_2'(s(fun(X43,fun(X44,bool)),X45),s(X43,'const.bool.LET_2'(s(fun(X42,X43),X46),s(X42,X48))))),s(X44,X47))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X42,bool),'_dst_x0x1x0_3'(s(fun(X43,fun(X44,bool)),X45),s(fun(X42,X43),X46),s(X44,X47))),s(X42,X48)))) )),
  inference(superposition,[],[f465,f145])).
fof(f6545,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X52,bool),'_dst_x0x1x0_3'(s(fun(X50,fun(X49,bool)),X51),s(fun(X52,X50),X53),s(X49,X55))),s(X52,X54))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X49,bool),'const.bool.LET_2'(s(fun(X50,fun(X49,bool)),X51),s(X50,'const.bool.LET_2'(s(fun(X52,X50),X53),s(X52,X54))))),s(X49,X55)))) )),
  inference(superposition,[],[f465,f145])).
fof(f6605,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3)))) )),
  inference(forward_demodulation,[],[f6604,f1375])).
fof(f6604,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,bool)),fun('type.pair.prod'(X0,X1),bool)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3)))) )),
  inference(forward_demodulation,[],[f6603,f91])).
fof(f6603,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun('type.pair.prod'(X0,X1),bool)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3)))) )),
  inference(forward_demodulation,[],[f6602,f1624])).
fof(f6602,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun('type.pair.prod'(X0,X1),bool)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3)))) )),
  inference(forward_demodulation,[],[f6492,f91])).
fof(f6492,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X0,fun(X1,bool)),fun('type.pair.prod'(X0,X1),bool)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,bool)),X2))),s('type.pair.prod'(X0,X1),X3)))) )),
  inference(superposition,[],[f465,f102])).
fof(f6606,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X11,X9] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X4,bool),'const.bool.LET_2'(s(fun(X5,fun(X4,bool)),'_dst_x00x11x00_3'(s(fun(X6,fun(X7,bool)),X8),s(fun(X5,fun(X4,X6)),X9),s(X7,X10))),s(X5,X11))),s(X4,X12))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0x1x0_3'(s(fun(X6,fun(X7,bool)),X8),s(fun(X4,X6),'const.bool.LET_2'(s(fun(X5,fun(X4,X6)),X9),s(X5,X11))),s(X7,X10))),s(X4,X12)))) )),
  inference(forward_demodulation,[],[f6493,f91])).
fof(f6493,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0x1x0_3'(s(fun(X6,fun(X7,bool)),X8),s(fun(X4,X6),'const.bool.LET_2'(s(fun(X5,fun(X4,X6)),X9),s(X5,X11))),s(X7,X10))),s(X4,X12))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X4,bool),happ(s(fun(X5,fun(X4,bool)),'_dst_x00x11x00_3'(s(fun(X6,fun(X7,bool)),X8),s(fun(X5,fun(X4,X6)),X9),s(X7,X10))),s(X5,X11))),s(X4,X12)))) )),
  inference(superposition,[],[f465,f154])).
fof(f6607,plain,(
  ( ! [X14,X21,X19,X17,X15,X13,X20,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),'const.bool.LET_2'(s(fun(X17,fun(X13,bool)),'_dst_x00x11x00_3'(s(fun(X14,fun(X15,bool)),X16),s(fun(X17,fun(X13,X14)),X18),s(X15,X20))),s(X17,X19))),s(X13,X21))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X13,bool),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,bool)),X16),s(fun(X13,X14),'const.bool.LET_2'(s(fun(X17,fun(X13,X14)),X18),s(X17,X19))),s(X15,X20))),s(X13,X21)))) )),
  inference(forward_demodulation,[],[f6494,f91])).
fof(f6494,plain,(
  ( ! [X14,X21,X19,X17,X15,X13,X20,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),happ(s(fun(X17,fun(X13,bool)),'_dst_x00x11x00_3'(s(fun(X14,fun(X15,bool)),X16),s(fun(X17,fun(X13,X14)),X18),s(X15,X20))),s(X17,X19))),s(X13,X21))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X13,bool),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,bool)),X16),s(fun(X13,X14),'const.bool.LET_2'(s(fun(X17,fun(X13,X14)),X18),s(X17,X19))),s(X15,X20))),s(X13,X21)))) )),
  inference(superposition,[],[f465,f154])).
fof(f6612,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X32,bool),'const.bool.LET_2'(s(fun(X33,fun(X32,bool)),'const.bool.LET_2'(s(fun(X34,fun(X33,fun(X32,bool))),X35),s(X34,X36))),s(X33,X37))),s(X32,X38))) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),'const.pair.UNCURRY_2'(s(fun(X34,fun(X33,fun(X32,bool))),X35),s('type.pair.prod'(X34,X33),'const.pair.,_2'(s(X34,X36),s(X33,X37))))),s(X32,X38)))) )),
  inference(forward_demodulation,[],[f6497,f91])).
fof(f6497,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),'const.pair.UNCURRY_2'(s(fun(X34,fun(X33,fun(X32,bool))),X35),s('type.pair.prod'(X34,X33),'const.pair.,_2'(s(X34,X36),s(X33,X37))))),s(X32,X38))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X32,bool),happ(s(fun(X33,fun(X32,bool)),'const.bool.LET_2'(s(fun(X34,fun(X33,fun(X32,bool))),X35),s(X34,X36))),s(X33,X37))),s(X32,X38)))) )),
  inference(superposition,[],[f465,f143])).
fof(f6614,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X39,bool),'const.bool.LET_2'(s(fun(X40,fun(X39,bool)),'_dst_x0x1x0_3'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(fun(X40,X41),X44),s(X42,X45))),s(X40,X46))),s(X39,X47))) | s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),'const.bool.LET_2'(s(fun(X42,fun(X39,bool)),'const.bool.LET_2'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(X41,'const.bool.LET_2'(s(fun(X40,X41),X44),s(X40,X46))))),s(X42,X45))),s(X39,X47)))) )),
  inference(forward_demodulation,[],[f6613,f91])).
fof(f6613,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),'const.bool.LET_2'(s(fun(X42,fun(X39,bool)),'const.bool.LET_2'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(X41,'const.bool.LET_2'(s(fun(X40,X41),X44),s(X40,X46))))),s(X42,X45))),s(X39,X47))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X39,bool),happ(s(fun(X40,fun(X39,bool)),'_dst_x0x1x0_3'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(fun(X40,X41),X44),s(X42,X45))),s(X40,X46))),s(X39,X47)))) )),
  inference(forward_demodulation,[],[f6498,f91])).
fof(f6498,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),happ(s(fun(X42,fun(X39,bool)),'const.bool.LET_2'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(X41,'const.bool.LET_2'(s(fun(X40,X41),X44),s(X40,X46))))),s(X42,X45))),s(X39,X47))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X39,bool),happ(s(fun(X40,fun(X39,bool)),'_dst_x0x1x0_3'(s(fun(X41,fun(X42,fun(X39,bool))),X43),s(fun(X40,X41),X44),s(X42,X45))),s(X40,X46))),s(X39,X47)))) )),
  inference(superposition,[],[f465,f145])).
fof(f6616,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X48,bool),'const.bool.LET_2'(s(fun(X49,fun(X48,bool)),'const.bool.LET_2'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(X50,'const.bool.LET_2'(s(fun(X52,X50),X53),s(X52,X54))))),s(X49,X55))),s(X48,X56))) | s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),'const.bool.LET_2'(s(fun(X52,fun(X48,bool)),'_dst_x0x1x0_3'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(fun(X52,X50),X53),s(X49,X55))),s(X52,X54))),s(X48,X56)))) )),
  inference(forward_demodulation,[],[f6615,f91])).
fof(f6615,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),'const.bool.LET_2'(s(fun(X52,fun(X48,bool)),'_dst_x0x1x0_3'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(fun(X52,X50),X53),s(X49,X55))),s(X52,X54))),s(X48,X56))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X48,bool),happ(s(fun(X49,fun(X48,bool)),'const.bool.LET_2'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(X50,'const.bool.LET_2'(s(fun(X52,X50),X53),s(X52,X54))))),s(X49,X55))),s(X48,X56)))) )),
  inference(forward_demodulation,[],[f6499,f91])).
fof(f6499,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X48,bool),happ(s(fun(X52,fun(X48,bool)),'_dst_x0x1x0_3'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(fun(X52,X50),X53),s(X49,X55))),s(X52,X54))),s(X48,X56))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X48,bool),happ(s(fun(X49,fun(X48,bool)),'const.bool.LET_2'(s(fun(X50,fun(X49,fun(X48,bool))),X51),s(X50,'const.bool.LET_2'(s(fun(X52,X50),X53),s(X52,X54))))),s(X49,X55))),s(X48,X56)))) )),
  inference(superposition,[],[f465,f145])).
fof(f6620,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'const.bool.LET_2'(s(fun(X65,fun(X63,bool)),'const.bool.LET_2'(s(fun(X64,fun(X65,fun(X63,bool))),X66),s(X64,X67))),s(X65,X68))),s(X63,X69))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X63,bool),'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,fun(X63,bool))),X66),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,X67),s(X65,X68))))),s(X63,X69)))) )),
  inference(forward_demodulation,[],[f6501,f91])).
fof(f6501,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),happ(s(fun(X65,fun(X63,bool)),'const.bool.LET_2'(s(fun(X64,fun(X65,fun(X63,bool))),X66),s(X64,X67))),s(X65,X68))),s(X63,X69))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X63,bool),'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,fun(X63,bool))),X66),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,X67),s(X65,X68))))),s(X63,X69)))) )),
  inference(superposition,[],[f465,f143])).
fof(f6625,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X10,bool),X16),s(X10,'const.bool.LET_2'(s(fun(X11,X10),'const.bool.LET_2'(s(fun(X12,fun(X11,X10)),X13),s(X12,X14))),s(X11,X15))))) | s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),X16),s(X10,'const.pair.UNCURRY_2'(s(fun(X12,fun(X11,X10)),X13),s('type.pair.prod'(X12,X11),'const.pair.,_2'(s(X12,X14),s(X11,X15)))))))) )),
  inference(forward_demodulation,[],[f6504,f91])).
fof(f6504,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),X16),s(X10,'const.pair.UNCURRY_2'(s(fun(X12,fun(X11,X10)),X13),s('type.pair.prod'(X12,X11),'const.pair.,_2'(s(X12,X14),s(X11,X15))))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X10,bool),X16),s(X10,happ(s(fun(X11,X10),'const.bool.LET_2'(s(fun(X12,fun(X11,X10)),X13),s(X12,X14))),s(X11,X15)))))) )),
  inference(superposition,[],[f465,f143])).
fof(f6627,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X17,bool),X25),s(X17,'const.bool.LET_2'(s(fun(X18,X17),'_dst_x0x1x0_3'(s(fun(X19,fun(X20,X17)),X21),s(fun(X18,X19),X22),s(X20,X23))),s(X18,X24))))) | s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),X25),s(X17,'const.bool.LET_2'(s(fun(X20,X17),'const.bool.LET_2'(s(fun(X19,fun(X20,X17)),X21),s(X19,'const.bool.LET_2'(s(fun(X18,X19),X22),s(X18,X24))))),s(X20,X23)))))) )),
  inference(forward_demodulation,[],[f6626,f91])).
fof(f6626,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),X25),s(X17,'const.bool.LET_2'(s(fun(X20,X17),'const.bool.LET_2'(s(fun(X19,fun(X20,X17)),X21),s(X19,'const.bool.LET_2'(s(fun(X18,X19),X22),s(X18,X24))))),s(X20,X23))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X17,bool),X25),s(X17,happ(s(fun(X18,X17),'_dst_x0x1x0_3'(s(fun(X19,fun(X20,X17)),X21),s(fun(X18,X19),X22),s(X20,X23))),s(X18,X24)))))) )),
  inference(forward_demodulation,[],[f6505,f91])).
fof(f6505,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),X25),s(X17,happ(s(fun(X20,X17),'const.bool.LET_2'(s(fun(X19,fun(X20,X17)),X21),s(X19,'const.bool.LET_2'(s(fun(X18,X19),X22),s(X18,X24))))),s(X20,X23))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X17,bool),X25),s(X17,happ(s(fun(X18,X17),'_dst_x0x1x0_3'(s(fun(X19,fun(X20,X17)),X21),s(fun(X18,X19),X22),s(X20,X23))),s(X18,X24)))))) )),
  inference(superposition,[],[f465,f145])).
fof(f6629,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X26,bool),X34),s(X26,'const.bool.LET_2'(s(fun(X27,X26),'const.bool.LET_2'(s(fun(X28,fun(X27,X26)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33))))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X34),s(X26,'const.bool.LET_2'(s(fun(X30,X26),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,X26)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32)))))) )),
  inference(forward_demodulation,[],[f6628,f91])).
fof(f6628,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X34),s(X26,'const.bool.LET_2'(s(fun(X30,X26),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,X26)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X26,bool),X34),s(X26,happ(s(fun(X27,X26),'const.bool.LET_2'(s(fun(X28,fun(X27,X26)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33)))))) )),
  inference(forward_demodulation,[],[f6506,f91])).
fof(f6506,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X34),s(X26,happ(s(fun(X30,X26),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,X26)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X26,bool),X34),s(X26,happ(s(fun(X27,X26),'const.bool.LET_2'(s(fun(X28,fun(X27,X26)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33)))))) )),
  inference(superposition,[],[f465,f145])).
fof(f6633,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X41,bool),X47),s(X41,'const.bool.LET_2'(s(fun(X43,X41),'const.bool.LET_2'(s(fun(X42,fun(X43,X41)),X44),s(X42,X45))),s(X43,X46))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X41,bool),X47),s(X41,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,X41)),X44),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,X45),s(X43,X46)))))))) )),
  inference(forward_demodulation,[],[f6508,f91])).
fof(f6508,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X41,bool),X47),s(X41,happ(s(fun(X43,X41),'const.bool.LET_2'(s(fun(X42,fun(X43,X41)),X44),s(X42,X45))),s(X43,X46))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X41,bool),X47),s(X41,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,X41)),X44),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,X45),s(X43,X46)))))))) )),
  inference(superposition,[],[f465,f143])).
fof(f6642,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X94),s(bool,'const.bool.LET_2'(s(fun(X91,bool),X92),s(X91,X93))))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X94),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X91,bool),X92),s(X91,X93)))) )),
  inference(forward_demodulation,[],[f6641,f91])).
fof(f6641,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X94),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X94),s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X91,bool),X92),s(X91,X93)))) )),
  inference(forward_demodulation,[],[f6526,f91])).
fof(f6526,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X94),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),X94),s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X91,bool),X92),s(X91,X93)))) )),
  inference(superposition,[],[f465,f465])).
fof(f6646,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X100,X101),bool),X109),s(fun(X100,X101),'const.bool.LET_2'(s(fun(X102,fun(X100,X101)),'_dst_x00x11x00_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X102,fun(X100,X103)),X106),s(X104,X107))),s(X102,X108))))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X100,X101),bool),X109),s(fun(X100,X101),'_dst_x0x1x0_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X100,X103),'const.bool.LET_2'(s(fun(X102,fun(X100,X103)),X106),s(X102,X108))),s(X104,X107)))))) )),
  inference(forward_demodulation,[],[f6528,f91])).
fof(f6528,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X100,X101),bool),X109),s(fun(X100,X101),'_dst_x0x1x0_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X100,X103),'const.bool.LET_2'(s(fun(X102,fun(X100,X103)),X106),s(X102,X108))),s(X104,X107))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X100,X101),bool),X109),s(fun(X100,X101),happ(s(fun(X102,fun(X100,X101)),'_dst_x00x11x00_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X102,fun(X100,X103)),X106),s(X104,X107))),s(X102,X108)))))) )),
  inference(superposition,[],[f465,f154])).
fof(f6647,plain,(
  ( ! [X111,X118,X116,X114,X112,X110,X119,X117,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X110,X111),bool),X119),s(fun(X110,X111),'const.bool.LET_2'(s(fun(X115,fun(X110,X111)),'_dst_x00x11x00_3'(s(fun(X112,fun(X113,X111)),X114),s(fun(X115,fun(X110,X112)),X116),s(X113,X118))),s(X115,X117))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X110,X111),bool),X119),s(fun(X110,X111),'_dst_x0x1x0_3'(s(fun(X112,fun(X113,X111)),X114),s(fun(X110,X112),'const.bool.LET_2'(s(fun(X115,fun(X110,X112)),X116),s(X115,X117))),s(X113,X118)))))) )),
  inference(forward_demodulation,[],[f6529,f91])).
fof(f6529,plain,(
  ( ! [X111,X118,X116,X114,X112,X110,X119,X117,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X110,X111),bool),X119),s(fun(X110,X111),happ(s(fun(X115,fun(X110,X111)),'_dst_x00x11x00_3'(s(fun(X112,fun(X113,X111)),X114),s(fun(X115,fun(X110,X112)),X116),s(X113,X118))),s(X115,X117))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X110,X111),bool),X119),s(fun(X110,X111),'_dst_x0x1x0_3'(s(fun(X112,fun(X113,X111)),X114),s(fun(X110,X112),'const.bool.LET_2'(s(fun(X115,fun(X110,X112)),X116),s(X115,X117))),s(X113,X118)))))) )),
  inference(superposition,[],[f465,f154])).
fof(f6650,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X124),bool),X126),s('type.pair.prod'(bool,X124),X125))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X124),bool),X126),s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'F_0'),s(X124,sK9(X125,X124,bool)))))) | s(bool,'T_0') = s(bool,sK8(X125,X124,bool))) )),
  inference(forward_demodulation,[],[f6531,f91])).
fof(f6531,plain,(
  ( ! [X125,X126,X124] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X124),bool),X126),s('type.pair.prod'(bool,X124),X125))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X124),bool),X126),s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'F_0'),s(X124,sK9(X125,X124,bool)))))) | s(bool,'T_0') = s(bool,sK8(X125,X124,bool))) )),
  inference(superposition,[],[f465,f213])).
fof(f6651,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X127,bool),bool),X129),s('type.pair.prod'(X127,bool),X128))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X127,bool),bool),X129),s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,sK8(X128,bool,X127)),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK9(X128,bool,X127))) )),
  inference(forward_demodulation,[],[f6532,f91])).
fof(f6532,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X127,bool),bool),X129),s('type.pair.prod'(X127,bool),X128))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X127,bool),bool),X129),s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,sK8(X128,bool,X127)),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK9(X128,bool,X127))) )),
  inference(superposition,[],[f465,f214])).
fof(f6665,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,'const.bool.LET_2'(s(fun(X11,bool),X12),s(X11,X13))) | ~p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X11,bool),X12),s(X11,X13)))) )),
  inference(forward_demodulation,[],[f6664,f91])).
fof(f6664,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X11,bool),X12),s(X11,X13)))) )),
  inference(subsumption_resolution,[],[f6550,f78])).
fof(f6550,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X11,bool),X12),s(X11,X13)))) )),
  inference(superposition,[],[f89,f465])).
fof(f6667,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17)))) )),
  inference(forward_demodulation,[],[f6666,f91])).
fof(f6666,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) != s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17)))) )),
  inference(subsumption_resolution,[],[f6551,f78])).
fof(f6551,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'T_0')) | s(bool,X18) != s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17)))) )),
  inference(superposition,[],[f128,f465])).
fof(f6677,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,'const.bool.LET_2'(s(fun(X129,bool),X130),s(X129,X131))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X132)))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X129,bool),X130),s(X129,X131)))) )),
  inference(forward_demodulation,[],[f6577,f91])).
fof(f6577,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X132)))) | s(bool,X132) != s(bool,happ(s(fun(X129,bool),X130),s(X129,X131))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X129,bool),X130),s(X129,X131)))) )),
  inference(superposition,[],[f139,f465])).
fof(f6678,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,X136) != s(bool,'const.bool.LET_2'(s(fun(X133,bool),X134),s(X133,X135))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X136))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X133,bool),X134),s(X133,X135)))) )),
  inference(forward_demodulation,[],[f6578,f91])).
fof(f6578,plain,(
  ( ! [X134,X136,X135,X133] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X136))),s(bool,'T_0')))) | s(bool,X136) != s(bool,happ(s(fun(X133,bool),X134),s(X133,X135))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X133,bool),X134),s(X133,X135)))) )),
  inference(superposition,[],[f139,f465])).
fof(f6679,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X140) = s(bool,'const.bool.LET_2'(s(fun(X137,bool),X138),s(X137,X139))) | ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X140)))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X137,bool),X138),s(X137,X139)))) )),
  inference(forward_demodulation,[],[f6579,f91])).
fof(f6579,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X140)))) | s(bool,X140) = s(bool,happ(s(fun(X137,bool),X138),s(X137,X139))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X137,bool),X138),s(X137,X139)))) )),
  inference(superposition,[],[f141,f465])).
fof(f6680,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X144) = s(bool,'const.bool.LET_2'(s(fun(X141,bool),X142),s(X141,X143))) | ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X144))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X141,bool),X142),s(X141,X143)))) )),
  inference(forward_demodulation,[],[f6580,f91])).
fof(f6580,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X144))),s(bool,'T_0')))) | s(bool,X144) = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X141,bool),X142),s(X141,X143)))) )),
  inference(superposition,[],[f141,f465])).
fof(f6681,plain,(
  ( ! [X198,X204,X202,X200,X199,X197,X205,X203,X201] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,bool)),'_dst_x00x11x00_3'(s(fun(X200,fun(X200,bool)),'=_0'),s(fun(X202,fun(X203,X200)),X204),s(X200,'const.bool.LET_2'(s(fun(bool,X200),X201),s(bool,'const.bool.LET_2'(s(fun(X197,bool),X198),s(X197,X199))))))),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X202,X203)),X205),s(bool,'const.bool.LET_2'(s(fun(X197,bool),X198),s(X197,X199)))))))) | s(X200,'const.bool.LET_2'(s(fun(bool,X200),X201),s(bool,'T_0'))) != s(X200,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,X200)),X204),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X202,X203)),X205),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X197,bool),X198),s(X197,X199)))) )),
  inference(forward_demodulation,[],[f6587,f91])).
fof(f6587,plain,(
  ( ! [X198,X204,X202,X200,X199,X197,X205,X203,X201] : (s(X200,'const.bool.LET_2'(s(fun(bool,X200),X201),s(bool,'T_0'))) != s(X200,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,X200)),X204),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X202,X203)),X205),s(bool,'T_0'))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,bool)),'_dst_x00x11x00_3'(s(fun(X200,fun(X200,bool)),'=_0'),s(fun(X202,fun(X203,X200)),X204),s(X200,'const.bool.LET_2'(s(fun(bool,X200),X201),s(bool,happ(s(fun(X197,bool),X198),s(X197,X199))))))),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X202,X203)),X205),s(bool,happ(s(fun(X197,bool),X198),s(X197,X199)))))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X197,bool),X198),s(X197,X199)))) )),
  inference(superposition,[],[f1661,f465])).
fof(f6682,plain,(
  ( ! [X212,X210,X208,X206,X213,X211,X209,X207,X214] : (s(X211,'const.bool.LET_2'(s(fun(bool,X211),X213),s(bool,'const.bool.LET_2'(s(fun(X206,bool),X207),s(X206,X208))))) = s(X211,'const.pair.UNCURRY_2'(s(fun(X209,fun(X210,X211)),X212),s('type.pair.prod'(X209,X210),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X209,X210)),X214),s(bool,'const.bool.LET_2'(s(fun(X206,bool),X207),s(X206,X208))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X209,fun(X210,bool)),'_dst_x00x11x00_3'(s(fun(X211,fun(X211,bool)),'=_0'),s(fun(X209,fun(X210,X211)),X212),s(X211,'const.bool.LET_2'(s(fun(bool,X211),X213),s(bool,'T_0'))))),s('type.pair.prod'(X209,X210),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X209,X210)),X214),s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X206,bool),X207),s(X206,X208)))) )),
  inference(forward_demodulation,[],[f6588,f91])).
fof(f6588,plain,(
  ( ! [X212,X210,X208,X206,X213,X211,X209,X207,X214] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X209,fun(X210,bool)),'_dst_x00x11x00_3'(s(fun(X211,fun(X211,bool)),'=_0'),s(fun(X209,fun(X210,X211)),X212),s(X211,'const.bool.LET_2'(s(fun(bool,X211),X213),s(bool,'T_0'))))),s('type.pair.prod'(X209,X210),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X209,X210)),X214),s(bool,'T_0')))))) | s(X211,'const.bool.LET_2'(s(fun(bool,X211),X213),s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))))) = s(X211,'const.pair.UNCURRY_2'(s(fun(X209,fun(X210,X211)),X212),s('type.pair.prod'(X209,X210),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X209,X210)),X214),s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X206,bool),X207),s(X206,X208)))) )),
  inference(superposition,[],[f1662,f465])).
fof(f7886,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,X59) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,X59))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f574,f81])).
fof(f7897,plain,(
  ( ! [X90,X91] : (s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X91))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f574,f81])).
fof(f7912,plain,(
  ( ! [X144,X142,X143] : (s('type.pair.prod'(bool,X142),X143) != s('type.pair.prod'(bool,X142),X144) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X142),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X142),fun('type.pair.prod'(bool,X142),bool)),'=_0'),s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,'F_0'),s(X142,sK9(X143,X142,bool)))))),s('type.pair.prod'(bool,X142),X144))) | s(bool,'T_0') = s(bool,sK8(X143,X142,bool))) )),
  inference(superposition,[],[f574,f213])).
fof(f7913,plain,(
  ( ! [X146,X147,X145] : (s('type.pair.prod'(X145,bool),X146) != s('type.pair.prod'(X145,bool),X147) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X145,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X145,bool),fun('type.pair.prod'(X145,bool),bool)),'=_0'),s('type.pair.prod'(X145,bool),'const.pair.,_2'(s(X145,sK8(X146,bool,X145)),s(bool,'F_0'))))),s('type.pair.prod'(X145,bool),X147))) | s(bool,'T_0') = s(bool,sK9(X146,bool,X145))) )),
  inference(superposition,[],[f574,f214])).
fof(f7927,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,X59) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X59))),s(bool,X58))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f574,f81])).
fof(f7938,plain,(
  ( ! [X90,X91] : (s(bool,X90) != s(bool,X91) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f574,f81])).
fof(f7953,plain,(
  ( ! [X144,X142,X143] : (s('type.pair.prod'(bool,X142),X143) != s('type.pair.prod'(bool,X142),X144) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X142),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(bool,X142),fun('type.pair.prod'(bool,X142),bool)),'=_0'),s('type.pair.prod'(bool,X142),X144))),s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,'F_0'),s(X142,sK9(X143,X142,bool)))))) | s(bool,'T_0') = s(bool,sK8(X143,X142,bool))) )),
  inference(superposition,[],[f574,f213])).
fof(f7954,plain,(
  ( ! [X146,X147,X145] : (s('type.pair.prod'(X145,bool),X146) != s('type.pair.prod'(X145,bool),X147) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X145,bool),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X145,bool),fun('type.pair.prod'(X145,bool),bool)),'=_0'),s('type.pair.prod'(X145,bool),X147))),s('type.pair.prod'(X145,bool),'const.pair.,_2'(s(X145,sK8(X146,bool,X145)),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK9(X146,bool,X145))) )),
  inference(superposition,[],[f574,f214])).
fof(f8058,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107))))),s(bool,X108))) | s(bool,'T_0') != s(bool,X108) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107)))) )),
  inference(forward_demodulation,[],[f7904,f91])).
fof(f7904,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,X108) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(X105,bool),X106),s(X105,X107))))),s(bool,X108))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107)))) )),
  inference(superposition,[],[f574,f465])).
fof(f8088,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X108))),s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107))))) | s(bool,'T_0') != s(bool,X108) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107)))) )),
  inference(forward_demodulation,[],[f7945,f91])).
fof(f7945,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,X108) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X108))),s(bool,happ(s(fun(X105,bool),X106),s(X105,X107))))) | s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(X105,bool),X106),s(X105,X107)))) )),
  inference(superposition,[],[f574,f465])).
fof(f8673,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(condensation,[],[f8672])).
fof(f8672,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) )),
  inference(trivial_inequality_removal,[],[f8650])).
fof(f8650,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f516,f81])).
fof(f8683,plain,(
  ( ! [X78,X76,X74,X72,X71,X77,X75,X73] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,'const.bool.LET_2'(s(fun(X71,bool),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,bool)),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77)))))) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(X73,bool),'const.bool.LET_2'(s(fun(X72,fun(X73,bool)),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(forward_demodulation,[],[f8682,f91])).
fof(f8682,plain,(
  ( ! [X78,X76,X74,X72,X71,X77,X75,X73] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(X73,bool),'const.bool.LET_2'(s(fun(X72,fun(X73,bool)),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,happ(s(fun(X71,bool),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,bool)),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77)))))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(forward_demodulation,[],[f8666,f91])).
fof(f8666,plain,(
  ( ! [X78,X76,X74,X72,X71,X77,X75,X73] : (s(bool,'F_0') != s(bool,happ(s(fun(X73,bool),'const.bool.LET_2'(s(fun(X72,fun(X73,bool)),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,happ(s(fun(X71,bool),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,bool)),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77)))))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f516,f145])).
fof(f8685,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X86))),s(bool,'const.bool.LET_2'(s(fun(X79,bool),'const.bool.LET_2'(s(fun(X80,fun(X79,bool)),X81),s(X80,'const.bool.LET_2'(s(fun(X82,X80),X83),s(X82,X84))))),s(X79,X85)))))) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(X82,bool),'_dst_x0x1x0_3'(s(fun(X80,fun(X79,bool)),X81),s(fun(X82,X80),X83),s(X79,X85))),s(X82,X84))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(forward_demodulation,[],[f8684,f91])).
fof(f8684,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(X82,bool),'_dst_x0x1x0_3'(s(fun(X80,fun(X79,bool)),X81),s(fun(X82,X80),X83),s(X79,X85))),s(X82,X84))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X86))),s(bool,happ(s(fun(X79,bool),'const.bool.LET_2'(s(fun(X80,fun(X79,bool)),X81),s(X80,'const.bool.LET_2'(s(fun(X82,X80),X83),s(X82,X84))))),s(X79,X85)))))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(forward_demodulation,[],[f8667,f91])).
fof(f8667,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (s(bool,'F_0') != s(bool,happ(s(fun(X82,bool),'_dst_x0x1x0_3'(s(fun(X80,fun(X79,bool)),X81),s(fun(X82,X80),X83),s(X79,X85))),s(X82,X84))) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X86))),s(bool,happ(s(fun(X79,bool),'const.bool.LET_2'(s(fun(X80,fun(X79,bool)),X81),s(X80,'const.bool.LET_2'(s(fun(X82,X80),X83),s(X82,X84))))),s(X79,X85)))))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f516,f145])).
fof(f8701,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl59),
  inference(unit_resulting_resolution,[],[f80,f8463,f180])).
fof(f785,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK10(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK10(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f93,f156])).
fof(f783,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK10(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK10(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f93,f156])).
fof(f461,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f91,f156])).
fof(f463,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) = s(bool,'const.bool.LET_2'(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f91,f156])).
fof(f890,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),X9),s(X8,X10))) | s(bool,X11) = s(bool,happ(s(fun(X8,bool),'const.bool.LET_2'(s(fun(fun(X8,bool),fun(X8,bool)),'const.bool.LET_0'),s(fun(X8,bool),X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f137,f156])).
fof(f892,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,'const.bool.LET_2'(s(fun(X16,bool),X17),s(X16,X18))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),'const.bool.LET_2'(s(fun(fun(X16,bool),fun(X16,bool)),'const.bool.LET_0'),s(fun(X16,bool),X17))),s(X16,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f137,f156])).
fof(f1984,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s(X18,happ(s(fun(X19,X18),'const.bool.LET_2'(s(fun(bool,fun(X19,X18)),X20),s(bool,'T_0'))),s(X19,X21))) = s(X18,happ(s(fun(X15,X18),'_dst_x0x1x0_3'(s(fun(bool,fun(X19,X18)),X20),s(fun(X15,bool),X16),s(X19,X21))),s(X15,X17))) | s(bool,X22) = s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f145,f156])).
fof(f1985,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s(X27,happ(s(fun(X28,X27),'const.bool.LET_2'(s(fun(bool,fun(X28,X27)),X29),s(bool,X26))),s(X28,X30))) = s(X27,happ(s(fun(X23,X27),'_dst_x0x1x0_3'(s(fun(bool,fun(X28,X27)),X29),s(fun(X23,bool),X24),s(X28,X30))),s(X23,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X23,bool),X24),s(X23,X25)))) )),
  inference(superposition,[],[f145,f156])).
fof(f2543,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25,X32] : (s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,bool)),X30),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X25,'type.pair.prod'(X28,X29)),X31),s(X25,X27))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X28,fun(X29,bool)),X30),s(bool,'const.bool.LET_2'(s(fun(X25,bool),X26),s(X25,X27))))),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X25,'type.pair.prod'(X28,X29)),X31),s(X25,X27)))))) | s(bool,X32) = s(bool,'const.bool.LET_2'(s(fun(X25,bool),X26),s(X25,X27))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f1661,f156])).
fof(f2544,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (s(bool,X36) != s(bool,'const.pair.UNCURRY_2'(s(fun(X37,fun(X38,bool)),X39),s('type.pair.prod'(X37,X38),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X37,X38)),X40),s(X33,X35))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X37,fun(X38,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X37,fun(X38,bool)),X39),s(bool,'const.bool.LET_2'(s(fun(X33,bool),X34),s(X33,X35))))),s('type.pair.prod'(X37,X38),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X37,X38)),X40),s(X33,X35)))))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X33,bool),X34),s(X33,X35)))) )),
  inference(superposition,[],[f1661,f156])).
fof(f2593,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X10,bool),X13),s(X10,X12))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X7,fun(X8,bool)),X9),s(bool,'const.bool.LET_2'(s(fun(X10,bool),X13),s(X10,X12))))),s('type.pair.prod'(X7,X8),'const.bool.LET_2'(s(fun(X10,'type.pair.prod'(X7,X8)),X11),s(X10,X12)))))) | s(bool,X14) = s(bool,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,bool)),X9),s('type.pair.prod'(X7,X8),'const.bool.LET_2'(s(fun(X10,'type.pair.prod'(X7,X8)),X11),s(X10,X12))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f1661,f156])).
fof(f2594,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (s(bool,X21) != s(bool,'const.bool.LET_2'(s(fun(X18,bool),X22),s(X18,X20))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X15,fun(X16,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X15,fun(X16,bool)),X17),s(bool,'const.bool.LET_2'(s(fun(X18,bool),X22),s(X18,X20))))),s('type.pair.prod'(X15,X16),'const.bool.LET_2'(s(fun(X18,'type.pair.prod'(X15,X16)),X19),s(X18,X20)))))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X15,fun(X16,bool)),X17),s('type.pair.prod'(X15,X16),'const.bool.LET_2'(s(fun(X18,'type.pair.prod'(X15,X16)),X19),s(X18,X20)))))) )),
  inference(superposition,[],[f1661,f156])).
fof(f2771,plain,(
  ( ! [X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X18,fun(X19,bool)),X20),s(bool,'T_0'))),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X15,'type.pair.prod'(X18,X19)),X21),s(X15,X17)))))) | s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17))) = s(bool,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,bool)),X20),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X15,'type.pair.prod'(X18,X19)),X21),s(X15,X17))))) | s(bool,X22) = s(bool,'const.bool.LET_2'(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1662,f156])).
fof(f2772,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X27,fun(X28,bool)),'_dst_x00x11x00_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X27,fun(X28,bool)),X29),s(bool,X26))),s('type.pair.prod'(X27,X28),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X27,X28)),X30),s(X23,X25)))))) | s(bool,'const.bool.LET_2'(s(fun(X23,bool),X24),s(X23,X25))) = s(bool,'const.pair.UNCURRY_2'(s(fun(X27,fun(X28,bool)),X29),s('type.pair.prod'(X27,X28),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X27,X28)),X30),s(X23,X25))))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X23,bool),X24),s(X23,X25)))) )),
  inference(superposition,[],[f1662,f156])).
fof(f1880,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,'const.bool.LET_2'(s(fun(X18,bool),'const.bool.LET_2'(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X17,fun(X18,bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(forward_demodulation,[],[f1826,f91])).
fof(f1826,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,happ(s(fun(X18,bool),'const.bool.LET_2'(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X17,fun(X18,bool)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f143,f156])).
fof(f1879,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X12,bool),'const.bool.LET_2'(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f1825,f91])).
fof(f1825,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),'const.bool.LET_2'(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,'const.pair.UNCURRY_2'(s(fun(X11,fun(X12,bool)),X13),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f143,f156])).
fof(f1885,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (s(X40,X46) = s(X40,'const.bool.LET_2'(s(fun(X41,X40),'const.bool.LET_2'(s(fun(X42,fun(X41,X40)),X43),s(X42,X44))),s(X41,X45))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X40,bool),'const.bool.LET_2'(s(fun(X40,fun(X40,bool)),'=_0'),s(X40,'const.pair.UNCURRY_2'(s(fun(X42,fun(X41,X40)),X43),s('type.pair.prod'(X42,X41),'const.pair.,_2'(s(X42,X44),s(X41,X45))))))),s(X40,X46))))) )),
  inference(forward_demodulation,[],[f1834,f91])).
fof(f1834,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,'const.bool.LET_2'(s(fun(X40,bool),'const.bool.LET_2'(s(fun(X40,fun(X40,bool)),'=_0'),s(X40,'const.pair.UNCURRY_2'(s(fun(X42,fun(X41,X40)),X43),s('type.pair.prod'(X42,X41),'const.pair.,_2'(s(X42,X44),s(X41,X45))))))),s(X40,X46)))) | s(X40,X46) = s(X40,happ(s(fun(X41,X40),'const.bool.LET_2'(s(fun(X42,fun(X41,X40)),X43),s(X42,X44))),s(X41,X45)))) )),
  inference(superposition,[],[f141,f143])).
fof(f1883,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s(X26,X32) != s(X26,'const.bool.LET_2'(s(fun(X27,X26),'const.bool.LET_2'(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))),s(X27,X31))) | p(s(bool,'const.bool.LET_2'(s(fun(X26,bool),'const.bool.LET_2'(s(fun(X26,fun(X26,bool)),'=_0'),s(X26,'const.pair.UNCURRY_2'(s(fun(X28,fun(X27,X26)),X29),s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))))))),s(X26,X32))))) )),
  inference(forward_demodulation,[],[f1832,f91])).
fof(f1832,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (p(s(bool,'const.bool.LET_2'(s(fun(X26,bool),'const.bool.LET_2'(s(fun(X26,fun(X26,bool)),'=_0'),s(X26,'const.pair.UNCURRY_2'(s(fun(X28,fun(X27,X26)),X29),s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))))))),s(X26,X32)))) | s(X26,X32) != s(X26,happ(s(fun(X27,X26),'const.bool.LET_2'(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))),s(X27,X31)))) )),
  inference(superposition,[],[f139,f143])).
fof(f1889,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'const.bool.LET_2'(s(fun(X30,bool),'const.bool.LET_2'(s(fun(X30,fun(X30,bool)),'=_0'),s(X30,'const.bool.LET_2'(s(fun(X32,X30),'const.bool.LET_2'(s(fun(X31,fun(X32,X30)),X33),s(X31,X34))),s(X32,X35))))),s(X30,X36)))) | s(X30,X36) = s(X30,'const.pair.UNCURRY_2'(s(fun(X31,fun(X32,X30)),X33),s('type.pair.prod'(X31,X32),'const.pair.,_2'(s(X31,X34),s(X32,X35)))))) )),
  inference(forward_demodulation,[],[f1845,f91])).
fof(f1845,plain,(
  ( ! [X30,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'const.bool.LET_2'(s(fun(X30,bool),'const.bool.LET_2'(s(fun(X30,fun(X30,bool)),'=_0'),s(X30,happ(s(fun(X32,X30),'const.bool.LET_2'(s(fun(X31,fun(X32,X30)),X33),s(X31,X34))),s(X32,X35))))),s(X30,X36)))) | s(X30,X36) = s(X30,'const.pair.UNCURRY_2'(s(fun(X31,fun(X32,X30)),X33),s('type.pair.prod'(X31,X32),'const.pair.,_2'(s(X31,X34),s(X32,X35)))))) )),
  inference(superposition,[],[f141,f143])).
fof(f1887,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.bool.LET_2'(s(fun(X16,bool),'const.bool.LET_2'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,'const.bool.LET_2'(s(fun(X18,X16),'const.bool.LET_2'(s(fun(X17,fun(X18,X16)),X19),s(X17,X20))),s(X18,X21))))),s(X16,X22)))) | s(X16,X22) != s(X16,'const.pair.UNCURRY_2'(s(fun(X17,fun(X18,X16)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(forward_demodulation,[],[f1843,f91])).
fof(f1843,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,'const.bool.LET_2'(s(fun(X16,bool),'const.bool.LET_2'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun(X18,X16),'const.bool.LET_2'(s(fun(X17,fun(X18,X16)),X19),s(X17,X20))),s(X18,X21))))),s(X16,X22)))) | s(X16,X22) != s(X16,'const.pair.UNCURRY_2'(s(fun(X17,fun(X18,X16)),X19),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f139,f143])).
fof(f2108,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (s(bool,X102) = s(bool,'const.bool.LET_2'(s(fun(X95,bool),'_dst_x0x1x0_3'(s(fun(X96,fun(X97,bool)),X98),s(fun(X95,X96),X99),s(X97,X100))),s(X95,X101))) | p(s(bool,'const.bool.LET_2'(s(fun(X97,bool),'const.bool.LET_2'(s(fun(X96,fun(X97,bool)),X98),s(X96,'const.bool.LET_2'(s(fun(X95,X96),X99),s(X95,X101))))),s(X97,X100)))) | p(s(bool,X102))) )),
  inference(forward_demodulation,[],[f2107,f91])).
fof(f2107,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (p(s(bool,'const.bool.LET_2'(s(fun(X97,bool),'const.bool.LET_2'(s(fun(X96,fun(X97,bool)),X98),s(X96,'const.bool.LET_2'(s(fun(X95,X96),X99),s(X95,X101))))),s(X97,X100)))) | p(s(bool,X102)) | s(bool,X102) = s(bool,happ(s(fun(X95,bool),'_dst_x0x1x0_3'(s(fun(X96,fun(X97,bool)),X98),s(fun(X95,X96),X99),s(X97,X100))),s(X95,X101)))) )),
  inference(forward_demodulation,[],[f2028,f91])).
fof(f2028,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),'const.bool.LET_2'(s(fun(X96,fun(X97,bool)),X98),s(X96,'const.bool.LET_2'(s(fun(X95,X96),X99),s(X95,X101))))),s(X97,X100)))) | p(s(bool,X102)) | s(bool,X102) = s(bool,happ(s(fun(X95,bool),'_dst_x0x1x0_3'(s(fun(X96,fun(X97,bool)),X98),s(fun(X95,X96),X99),s(X97,X100))),s(X95,X101)))) )),
  inference(superposition,[],[f86,f145])).
fof(f2112,plain,(
  ( ! [X111,X118,X116,X114,X112,X117,X115,X113] : (s(bool,X118) = s(bool,'const.bool.LET_2'(s(fun(X111,bool),'_dst_x0x1x0_3'(s(fun(X112,fun(X113,bool)),X114),s(fun(X111,X112),X115),s(X113,X116))),s(X111,X117))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X113,bool),'const.bool.LET_2'(s(fun(X112,fun(X113,bool)),X114),s(X112,'const.bool.LET_2'(s(fun(X111,X112),X115),s(X111,X117))))),s(X113,X116)))) | ~p(s(bool,X118))) )),
  inference(forward_demodulation,[],[f2111,f91])).
fof(f2111,plain,(
  ( ! [X111,X118,X116,X114,X112,X117,X115,X113] : (~p(s(bool,'const.bool.LET_2'(s(fun(X113,bool),'const.bool.LET_2'(s(fun(X112,fun(X113,bool)),X114),s(X112,'const.bool.LET_2'(s(fun(X111,X112),X115),s(X111,X117))))),s(X113,X116)))) | ~p(s(bool,X118)) | s(bool,X118) = s(bool,happ(s(fun(X111,bool),'_dst_x0x1x0_3'(s(fun(X112,fun(X113,bool)),X114),s(fun(X111,X112),X115),s(X113,X116))),s(X111,X117)))) )),
  inference(forward_demodulation,[],[f2030,f91])).
fof(f2030,plain,(
  ( ! [X111,X118,X116,X114,X112,X117,X115,X113] : (~p(s(bool,happ(s(fun(X113,bool),'const.bool.LET_2'(s(fun(X112,fun(X113,bool)),X114),s(X112,'const.bool.LET_2'(s(fun(X111,X112),X115),s(X111,X117))))),s(X113,X116)))) | ~p(s(bool,X118)) | s(bool,X118) = s(bool,happ(s(fun(X111,bool),'_dst_x0x1x0_3'(s(fun(X112,fun(X113,bool)),X114),s(fun(X111,X112),X115),s(X113,X116))),s(X111,X117)))) )),
  inference(superposition,[],[f89,f145])).
fof(f2114,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120,X119] : (s(bool,X126) != s(bool,'const.bool.LET_2'(s(fun(X119,bool),'_dst_x0x1x0_3'(s(fun(X120,fun(X121,bool)),X122),s(fun(X119,X120),X123),s(X121,X124))),s(X119,X125))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X121,bool),'const.bool.LET_2'(s(fun(X120,fun(X121,bool)),X122),s(X120,'const.bool.LET_2'(s(fun(X119,X120),X123),s(X119,X125))))),s(X121,X124)))) | p(s(bool,X126))) )),
  inference(forward_demodulation,[],[f2113,f91])).
fof(f2113,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120,X119] : (~p(s(bool,'const.bool.LET_2'(s(fun(X121,bool),'const.bool.LET_2'(s(fun(X120,fun(X121,bool)),X122),s(X120,'const.bool.LET_2'(s(fun(X119,X120),X123),s(X119,X125))))),s(X121,X124)))) | s(bool,X126) != s(bool,happ(s(fun(X119,bool),'_dst_x0x1x0_3'(s(fun(X120,fun(X121,bool)),X122),s(fun(X119,X120),X123),s(X121,X124))),s(X119,X125))) | p(s(bool,X126))) )),
  inference(forward_demodulation,[],[f2031,f91])).
fof(f2031,plain,(
  ( ! [X125,X123,X121,X126,X124,X122,X120,X119] : (~p(s(bool,happ(s(fun(X121,bool),'const.bool.LET_2'(s(fun(X120,fun(X121,bool)),X122),s(X120,'const.bool.LET_2'(s(fun(X119,X120),X123),s(X119,X125))))),s(X121,X124)))) | s(bool,X126) != s(bool,happ(s(fun(X119,bool),'_dst_x0x1x0_3'(s(fun(X120,fun(X121,bool)),X122),s(fun(X119,X120),X123),s(X121,X124))),s(X119,X125))) | p(s(bool,X126))) )),
  inference(superposition,[],[f128,f145])).
fof(f2116,plain,(
  ( ! [X165,X163,X161,X159,X164,X162,X160] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X159,bool),'_dst_x0x1x0_3'(s(fun(X160,fun(X161,bool)),X162),s(fun(X159,X160),X163),s(X161,X164))),s(X159,X165))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X161,bool),'const.bool.LET_2'(s(fun(X160,fun(X161,bool)),X162),s(X160,'const.bool.LET_2'(s(fun(X159,X160),X163),s(X159,X165))))),s(X161,X164))))) )),
  inference(forward_demodulation,[],[f2115,f91])).
fof(f2115,plain,(
  ( ! [X165,X163,X161,X159,X164,X162,X160] : (~p(s(bool,'const.bool.LET_2'(s(fun(X161,bool),'const.bool.LET_2'(s(fun(X160,fun(X161,bool)),X162),s(X160,'const.bool.LET_2'(s(fun(X159,X160),X163),s(X159,X165))))),s(X161,X164)))) | s(bool,'T_0') = s(bool,happ(s(fun(X159,bool),'_dst_x0x1x0_3'(s(fun(X160,fun(X161,bool)),X162),s(fun(X159,X160),X163),s(X161,X164))),s(X159,X165)))) )),
  inference(forward_demodulation,[],[f2036,f91])).
fof(f2036,plain,(
  ( ! [X165,X163,X161,X159,X164,X162,X160] : (~p(s(bool,happ(s(fun(X161,bool),'const.bool.LET_2'(s(fun(X160,fun(X161,bool)),X162),s(X160,'const.bool.LET_2'(s(fun(X159,X160),X163),s(X159,X165))))),s(X161,X164)))) | s(bool,'T_0') = s(bool,happ(s(fun(X159,bool),'_dst_x0x1x0_3'(s(fun(X160,fun(X161,bool)),X162),s(fun(X159,X160),X163),s(X161,X164))),s(X159,X165)))) )),
  inference(superposition,[],[f158,f145])).
fof(f2120,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X179] : (s(bool,X180) = s(bool,'const.bool.LET_2'(s(fun(X173,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X175,bool)),X176),s(fun(X173,X174),X177),s(X175,X178))),s(X173,X179))) | p(s(bool,'const.bool.LET_2'(s(fun(X175,bool),'const.bool.LET_2'(s(fun(X174,fun(X175,bool)),X176),s(X174,'const.bool.LET_2'(s(fun(X173,X174),X177),s(X173,X179))))),s(X175,X178)))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(forward_demodulation,[],[f2119,f91])).
fof(f2119,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X179] : (p(s(bool,'const.bool.LET_2'(s(fun(X175,bool),'const.bool.LET_2'(s(fun(X174,fun(X175,bool)),X176),s(X174,'const.bool.LET_2'(s(fun(X173,X174),X177),s(X173,X179))))),s(X175,X178)))) | s(bool,X180) = s(bool,happ(s(fun(X173,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X175,bool)),X176),s(fun(X173,X174),X177),s(X175,X178))),s(X173,X179))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(forward_demodulation,[],[f2038,f91])).
fof(f2038,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X179] : (p(s(bool,happ(s(fun(X175,bool),'const.bool.LET_2'(s(fun(X174,fun(X175,bool)),X176),s(X174,'const.bool.LET_2'(s(fun(X173,X174),X177),s(X173,X179))))),s(X175,X178)))) | s(bool,X180) = s(bool,happ(s(fun(X173,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X175,bool)),X176),s(fun(X173,X174),X177),s(X175,X178))),s(X173,X179))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f180,f145])).
fof(f2122,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X181,bool),'_dst_x0x1x0_3'(s(fun(X182,fun(X183,bool)),X184),s(fun(X181,X182),X185),s(X183,X186))),s(X181,X187))) | p(s(bool,'const.bool.LET_2'(s(fun(X183,bool),'const.bool.LET_2'(s(fun(X182,fun(X183,bool)),X184),s(X182,'const.bool.LET_2'(s(fun(X181,X182),X185),s(X181,X187))))),s(X183,X186))))) ) | $spl16),
  inference(forward_demodulation,[],[f2121,f91])).
fof(f2121,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (p(s(bool,'const.bool.LET_2'(s(fun(X183,bool),'const.bool.LET_2'(s(fun(X182,fun(X183,bool)),X184),s(X182,'const.bool.LET_2'(s(fun(X181,X182),X185),s(X181,X187))))),s(X183,X186)))) | s(bool,'T_0') != s(bool,happ(s(fun(X181,bool),'_dst_x0x1x0_3'(s(fun(X182,fun(X183,bool)),X184),s(fun(X181,X182),X185),s(X183,X186))),s(X181,X187)))) ) | $spl16),
  inference(forward_demodulation,[],[f2039,f91])).
fof(f2039,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (p(s(bool,happ(s(fun(X183,bool),'const.bool.LET_2'(s(fun(X182,fun(X183,bool)),X184),s(X182,'const.bool.LET_2'(s(fun(X181,X182),X185),s(X181,X187))))),s(X183,X186)))) | s(bool,'T_0') != s(bool,happ(s(fun(X181,bool),'_dst_x0x1x0_3'(s(fun(X182,fun(X183,bool)),X184),s(fun(X181,X182),X185),s(X183,X186))),s(X181,X187)))) ) | $spl16),
  inference(superposition,[],[f439,f145])).
fof(f2138,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(bool,X116) = s(bool,'const.bool.LET_2'(s(fun(X109,bool),'const.bool.LET_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,'const.bool.LET_2'(s(fun(X112,X110),X113),s(X112,X114))))),s(X109,X115))) | p(s(bool,'const.bool.LET_2'(s(fun(X112,bool),'_dst_x0x1x0_3'(s(fun(X110,fun(X109,bool)),X111),s(fun(X112,X110),X113),s(X109,X115))),s(X112,X114)))) | p(s(bool,X116))) )),
  inference(forward_demodulation,[],[f2137,f91])).
fof(f2137,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (p(s(bool,'const.bool.LET_2'(s(fun(X112,bool),'_dst_x0x1x0_3'(s(fun(X110,fun(X109,bool)),X111),s(fun(X112,X110),X113),s(X109,X115))),s(X112,X114)))) | p(s(bool,X116)) | s(bool,X116) = s(bool,happ(s(fun(X109,bool),'const.bool.LET_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,'const.bool.LET_2'(s(fun(X112,X110),X113),s(X112,X114))))),s(X109,X115)))) )),
  inference(forward_demodulation,[],[f2061,f91])).
fof(f2061,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (p(s(bool,happ(s(fun(X112,bool),'_dst_x0x1x0_3'(s(fun(X110,fun(X109,bool)),X111),s(fun(X112,X110),X113),s(X109,X115))),s(X112,X114)))) | p(s(bool,X116)) | s(bool,X116) = s(bool,happ(s(fun(X109,bool),'const.bool.LET_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,'const.bool.LET_2'(s(fun(X112,X110),X113),s(X112,X114))))),s(X109,X115)))) )),
  inference(superposition,[],[f86,f145])).
fof(f2142,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (s(bool,X132) = s(bool,'const.bool.LET_2'(s(fun(X125,bool),'const.bool.LET_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,'const.bool.LET_2'(s(fun(X128,X126),X129),s(X128,X130))))),s(X125,X131))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X128,bool),'_dst_x0x1x0_3'(s(fun(X126,fun(X125,bool)),X127),s(fun(X128,X126),X129),s(X125,X131))),s(X128,X130)))) | ~p(s(bool,X132))) )),
  inference(forward_demodulation,[],[f2141,f91])).
fof(f2141,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,'const.bool.LET_2'(s(fun(X128,bool),'_dst_x0x1x0_3'(s(fun(X126,fun(X125,bool)),X127),s(fun(X128,X126),X129),s(X125,X131))),s(X128,X130)))) | ~p(s(bool,X132)) | s(bool,X132) = s(bool,happ(s(fun(X125,bool),'const.bool.LET_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,'const.bool.LET_2'(s(fun(X128,X126),X129),s(X128,X130))))),s(X125,X131)))) )),
  inference(forward_demodulation,[],[f2063,f91])).
fof(f2063,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,happ(s(fun(X128,bool),'_dst_x0x1x0_3'(s(fun(X126,fun(X125,bool)),X127),s(fun(X128,X126),X129),s(X125,X131))),s(X128,X130)))) | ~p(s(bool,X132)) | s(bool,X132) = s(bool,happ(s(fun(X125,bool),'const.bool.LET_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,'const.bool.LET_2'(s(fun(X128,X126),X129),s(X128,X130))))),s(X125,X131)))) )),
  inference(superposition,[],[f89,f145])).
fof(f2144,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (s(bool,X140) != s(bool,'const.bool.LET_2'(s(fun(X133,bool),'const.bool.LET_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,'const.bool.LET_2'(s(fun(X136,X134),X137),s(X136,X138))))),s(X133,X139))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X136,bool),'_dst_x0x1x0_3'(s(fun(X134,fun(X133,bool)),X135),s(fun(X136,X134),X137),s(X133,X139))),s(X136,X138)))) | p(s(bool,X140))) )),
  inference(forward_demodulation,[],[f2143,f91])).
fof(f2143,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,'const.bool.LET_2'(s(fun(X136,bool),'_dst_x0x1x0_3'(s(fun(X134,fun(X133,bool)),X135),s(fun(X136,X134),X137),s(X133,X139))),s(X136,X138)))) | s(bool,X140) != s(bool,happ(s(fun(X133,bool),'const.bool.LET_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,'const.bool.LET_2'(s(fun(X136,X134),X137),s(X136,X138))))),s(X133,X139))) | p(s(bool,X140))) )),
  inference(forward_demodulation,[],[f2064,f91])).
fof(f2064,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,happ(s(fun(X136,bool),'_dst_x0x1x0_3'(s(fun(X134,fun(X133,bool)),X135),s(fun(X136,X134),X137),s(X133,X139))),s(X136,X138)))) | s(bool,X140) != s(bool,happ(s(fun(X133,bool),'const.bool.LET_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,'const.bool.LET_2'(s(fun(X136,X134),X137),s(X136,X138))))),s(X133,X139))) | p(s(bool,X140))) )),
  inference(superposition,[],[f128,f145])).
fof(f2146,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X179] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X173,bool),'const.bool.LET_2'(s(fun(X174,fun(X173,bool)),X175),s(X174,'const.bool.LET_2'(s(fun(X176,X174),X177),s(X176,X178))))),s(X173,X179))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X176,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X173,bool)),X175),s(fun(X176,X174),X177),s(X173,X179))),s(X176,X178))))) )),
  inference(forward_demodulation,[],[f2145,f91])).
fof(f2145,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X179] : (~p(s(bool,'const.bool.LET_2'(s(fun(X176,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X173,bool)),X175),s(fun(X176,X174),X177),s(X173,X179))),s(X176,X178)))) | s(bool,'T_0') = s(bool,happ(s(fun(X173,bool),'const.bool.LET_2'(s(fun(X174,fun(X173,bool)),X175),s(X174,'const.bool.LET_2'(s(fun(X176,X174),X177),s(X176,X178))))),s(X173,X179)))) )),
  inference(forward_demodulation,[],[f2069,f91])).
fof(f2069,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X179] : (~p(s(bool,happ(s(fun(X176,bool),'_dst_x0x1x0_3'(s(fun(X174,fun(X173,bool)),X175),s(fun(X176,X174),X177),s(X173,X179))),s(X176,X178)))) | s(bool,'T_0') = s(bool,happ(s(fun(X173,bool),'const.bool.LET_2'(s(fun(X174,fun(X173,bool)),X175),s(X174,'const.bool.LET_2'(s(fun(X176,X174),X177),s(X176,X178))))),s(X173,X179)))) )),
  inference(superposition,[],[f158,f145])).
fof(f2150,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X193] : (s(bool,X194) = s(bool,'const.bool.LET_2'(s(fun(X187,bool),'const.bool.LET_2'(s(fun(X188,fun(X187,bool)),X189),s(X188,'const.bool.LET_2'(s(fun(X190,X188),X191),s(X190,X192))))),s(X187,X193))) | p(s(bool,'const.bool.LET_2'(s(fun(X190,bool),'_dst_x0x1x0_3'(s(fun(X188,fun(X187,bool)),X189),s(fun(X190,X188),X191),s(X187,X193))),s(X190,X192)))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(forward_demodulation,[],[f2149,f91])).
fof(f2149,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X193] : (p(s(bool,'const.bool.LET_2'(s(fun(X190,bool),'_dst_x0x1x0_3'(s(fun(X188,fun(X187,bool)),X189),s(fun(X190,X188),X191),s(X187,X193))),s(X190,X192)))) | s(bool,X194) = s(bool,happ(s(fun(X187,bool),'const.bool.LET_2'(s(fun(X188,fun(X187,bool)),X189),s(X188,'const.bool.LET_2'(s(fun(X190,X188),X191),s(X190,X192))))),s(X187,X193))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(forward_demodulation,[],[f2071,f91])).
fof(f2071,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X193] : (p(s(bool,happ(s(fun(X190,bool),'_dst_x0x1x0_3'(s(fun(X188,fun(X187,bool)),X189),s(fun(X190,X188),X191),s(X187,X193))),s(X190,X192)))) | s(bool,X194) = s(bool,happ(s(fun(X187,bool),'const.bool.LET_2'(s(fun(X188,fun(X187,bool)),X189),s(X188,'const.bool.LET_2'(s(fun(X190,X188),X191),s(X190,X192))))),s(X187,X193))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f180,f145])).
fof(f2152,plain,(
  ( ! [X198,X196,X200,X199,X197,X195,X201] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X195,bool),'const.bool.LET_2'(s(fun(X196,fun(X195,bool)),X197),s(X196,'const.bool.LET_2'(s(fun(X198,X196),X199),s(X198,X200))))),s(X195,X201))) | p(s(bool,'const.bool.LET_2'(s(fun(X198,bool),'_dst_x0x1x0_3'(s(fun(X196,fun(X195,bool)),X197),s(fun(X198,X196),X199),s(X195,X201))),s(X198,X200))))) ) | $spl16),
  inference(forward_demodulation,[],[f2151,f91])).
fof(f2151,plain,(
  ( ! [X198,X196,X200,X199,X197,X195,X201] : (p(s(bool,'const.bool.LET_2'(s(fun(X198,bool),'_dst_x0x1x0_3'(s(fun(X196,fun(X195,bool)),X197),s(fun(X198,X196),X199),s(X195,X201))),s(X198,X200)))) | s(bool,'T_0') != s(bool,happ(s(fun(X195,bool),'const.bool.LET_2'(s(fun(X196,fun(X195,bool)),X197),s(X196,'const.bool.LET_2'(s(fun(X198,X196),X199),s(X198,X200))))),s(X195,X201)))) ) | $spl16),
  inference(forward_demodulation,[],[f2072,f91])).
fof(f2072,plain,(
  ( ! [X198,X196,X200,X199,X197,X195,X201] : (p(s(bool,happ(s(fun(X198,bool),'_dst_x0x1x0_3'(s(fun(X196,fun(X195,bool)),X197),s(fun(X198,X196),X199),s(X195,X201))),s(X198,X200)))) | s(bool,'T_0') != s(bool,happ(s(fun(X195,bool),'const.bool.LET_2'(s(fun(X196,fun(X195,bool)),X197),s(X196,'const.bool.LET_2'(s(fun(X198,X196),X199),s(X198,X200))))),s(X195,X201)))) ) | $spl16),
  inference(superposition,[],[f439,f145])).
fof(f2104,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (s(X46,X54) = s(X46,'const.bool.LET_2'(s(fun(X47,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X49,X46)),X50),s(fun(X47,X48),X51),s(X49,X52))),s(X47,X53))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X46,bool),'const.bool.LET_2'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,'const.bool.LET_2'(s(fun(X49,X46),'const.bool.LET_2'(s(fun(X48,fun(X49,X46)),X50),s(X48,'const.bool.LET_2'(s(fun(X47,X48),X51),s(X47,X53))))),s(X49,X52))))),s(X46,X54))))) )),
  inference(forward_demodulation,[],[f2103,f91])).
fof(f2103,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (~p(s(bool,'const.bool.LET_2'(s(fun(X46,bool),'const.bool.LET_2'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,'const.bool.LET_2'(s(fun(X49,X46),'const.bool.LET_2'(s(fun(X48,fun(X49,X46)),X50),s(X48,'const.bool.LET_2'(s(fun(X47,X48),X51),s(X47,X53))))),s(X49,X52))))),s(X46,X54)))) | s(X46,X54) = s(X46,happ(s(fun(X47,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X49,X46)),X50),s(fun(X47,X48),X51),s(X49,X52))),s(X47,X53)))) )),
  inference(forward_demodulation,[],[f2023,f91])).
fof(f2023,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (~p(s(bool,'const.bool.LET_2'(s(fun(X46,bool),'const.bool.LET_2'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,happ(s(fun(X49,X46),'const.bool.LET_2'(s(fun(X48,fun(X49,X46)),X50),s(X48,'const.bool.LET_2'(s(fun(X47,X48),X51),s(X47,X53))))),s(X49,X52))))),s(X46,X54)))) | s(X46,X54) = s(X46,happ(s(fun(X47,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X49,X46)),X50),s(fun(X47,X48),X51),s(X49,X52))),s(X47,X53)))) )),
  inference(superposition,[],[f141,f145])).
fof(f2100,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (s(X28,X36) != s(X28,'const.bool.LET_2'(s(fun(X29,X28),'_dst_x0x1x0_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X29,X30),X33),s(X31,X34))),s(X29,X35))) | p(s(bool,'const.bool.LET_2'(s(fun(X28,bool),'const.bool.LET_2'(s(fun(X28,fun(X28,bool)),'=_0'),s(X28,'const.bool.LET_2'(s(fun(X31,X28),'const.bool.LET_2'(s(fun(X30,fun(X31,X28)),X32),s(X30,'const.bool.LET_2'(s(fun(X29,X30),X33),s(X29,X35))))),s(X31,X34))))),s(X28,X36))))) )),
  inference(forward_demodulation,[],[f2099,f91])).
fof(f2099,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,'const.bool.LET_2'(s(fun(X28,bool),'const.bool.LET_2'(s(fun(X28,fun(X28,bool)),'=_0'),s(X28,'const.bool.LET_2'(s(fun(X31,X28),'const.bool.LET_2'(s(fun(X30,fun(X31,X28)),X32),s(X30,'const.bool.LET_2'(s(fun(X29,X30),X33),s(X29,X35))))),s(X31,X34))))),s(X28,X36)))) | s(X28,X36) != s(X28,happ(s(fun(X29,X28),'_dst_x0x1x0_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X29,X30),X33),s(X31,X34))),s(X29,X35)))) )),
  inference(forward_demodulation,[],[f2021,f91])).
fof(f2021,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (p(s(bool,'const.bool.LET_2'(s(fun(X28,bool),'const.bool.LET_2'(s(fun(X28,fun(X28,bool)),'=_0'),s(X28,happ(s(fun(X31,X28),'const.bool.LET_2'(s(fun(X30,fun(X31,X28)),X32),s(X30,'const.bool.LET_2'(s(fun(X29,X30),X33),s(X29,X35))))),s(X31,X34))))),s(X28,X36)))) | s(X28,X36) != s(X28,happ(s(fun(X29,X28),'_dst_x0x1x0_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X29,X30),X33),s(X31,X34))),s(X29,X35)))) )),
  inference(superposition,[],[f139,f145])).
fof(f2134,plain,(
  ( ! [X61,X68,X66,X64,X62,X60,X67,X65,X63] : (s(X60,X68) = s(X60,'const.bool.LET_2'(s(fun(X61,X60),'const.bool.LET_2'(s(fun(X62,fun(X61,X60)),X63),s(X62,'const.bool.LET_2'(s(fun(X64,X62),X65),s(X64,X66))))),s(X61,X67))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X60,bool),'const.bool.LET_2'(s(fun(X60,fun(X60,bool)),'=_0'),s(X60,'const.bool.LET_2'(s(fun(X64,X60),'_dst_x0x1x0_3'(s(fun(X62,fun(X61,X60)),X63),s(fun(X64,X62),X65),s(X61,X67))),s(X64,X66))))),s(X60,X68))))) )),
  inference(forward_demodulation,[],[f2133,f91])).
fof(f2133,plain,(
  ( ! [X61,X68,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,'const.bool.LET_2'(s(fun(X60,bool),'const.bool.LET_2'(s(fun(X60,fun(X60,bool)),'=_0'),s(X60,'const.bool.LET_2'(s(fun(X64,X60),'_dst_x0x1x0_3'(s(fun(X62,fun(X61,X60)),X63),s(fun(X64,X62),X65),s(X61,X67))),s(X64,X66))))),s(X60,X68)))) | s(X60,X68) = s(X60,happ(s(fun(X61,X60),'const.bool.LET_2'(s(fun(X62,fun(X61,X60)),X63),s(X62,'const.bool.LET_2'(s(fun(X64,X62),X65),s(X64,X66))))),s(X61,X67)))) )),
  inference(forward_demodulation,[],[f2056,f91])).
fof(f2056,plain,(
  ( ! [X61,X68,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,'const.bool.LET_2'(s(fun(X60,bool),'const.bool.LET_2'(s(fun(X60,fun(X60,bool)),'=_0'),s(X60,happ(s(fun(X64,X60),'_dst_x0x1x0_3'(s(fun(X62,fun(X61,X60)),X63),s(fun(X64,X62),X65),s(X61,X67))),s(X64,X66))))),s(X60,X68)))) | s(X60,X68) = s(X60,happ(s(fun(X61,X60),'const.bool.LET_2'(s(fun(X62,fun(X61,X60)),X63),s(X62,'const.bool.LET_2'(s(fun(X64,X62),X65),s(X64,X66))))),s(X61,X67)))) )),
  inference(superposition,[],[f141,f145])).
fof(f2130,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (s(X42,X50) != s(X42,'const.bool.LET_2'(s(fun(X43,X42),'const.bool.LET_2'(s(fun(X44,fun(X43,X42)),X45),s(X44,'const.bool.LET_2'(s(fun(X46,X44),X47),s(X46,X48))))),s(X43,X49))) | p(s(bool,'const.bool.LET_2'(s(fun(X42,bool),'const.bool.LET_2'(s(fun(X42,fun(X42,bool)),'=_0'),s(X42,'const.bool.LET_2'(s(fun(X46,X42),'_dst_x0x1x0_3'(s(fun(X44,fun(X43,X42)),X45),s(fun(X46,X44),X47),s(X43,X49))),s(X46,X48))))),s(X42,X50))))) )),
  inference(forward_demodulation,[],[f2129,f91])).
fof(f2129,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (p(s(bool,'const.bool.LET_2'(s(fun(X42,bool),'const.bool.LET_2'(s(fun(X42,fun(X42,bool)),'=_0'),s(X42,'const.bool.LET_2'(s(fun(X46,X42),'_dst_x0x1x0_3'(s(fun(X44,fun(X43,X42)),X45),s(fun(X46,X44),X47),s(X43,X49))),s(X46,X48))))),s(X42,X50)))) | s(X42,X50) != s(X42,happ(s(fun(X43,X42),'const.bool.LET_2'(s(fun(X44,fun(X43,X42)),X45),s(X44,'const.bool.LET_2'(s(fun(X46,X44),X47),s(X46,X48))))),s(X43,X49)))) )),
  inference(forward_demodulation,[],[f2054,f91])).
fof(f2054,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (p(s(bool,'const.bool.LET_2'(s(fun(X42,bool),'const.bool.LET_2'(s(fun(X42,fun(X42,bool)),'=_0'),s(X42,happ(s(fun(X46,X42),'_dst_x0x1x0_3'(s(fun(X44,fun(X43,X42)),X45),s(fun(X46,X44),X47),s(X43,X49))),s(X46,X48))))),s(X42,X50)))) | s(X42,X50) != s(X42,happ(s(fun(X43,X42),'const.bool.LET_2'(s(fun(X44,fun(X43,X42)),X45),s(X44,'const.bool.LET_2'(s(fun(X46,X44),X47),s(X46,X48))))),s(X43,X49)))) )),
  inference(superposition,[],[f139,f145])).
fof(f1890,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.bool.LET_2'(s(fun(X37,bool),'const.bool.LET_2'(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X43))),s(X37,'const.bool.LET_2'(s(fun(X39,X37),'const.bool.LET_2'(s(fun(X38,fun(X39,X37)),X40),s(X38,X41))),s(X39,X42)))))) | s(X37,X43) = s(X37,'const.pair.UNCURRY_2'(s(fun(X38,fun(X39,X37)),X40),s('type.pair.prod'(X38,X39),'const.pair.,_2'(s(X38,X41),s(X39,X42)))))) )),
  inference(forward_demodulation,[],[f1846,f91])).
fof(f1846,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (~p(s(bool,'const.bool.LET_2'(s(fun(X37,bool),'const.bool.LET_2'(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X43))),s(X37,happ(s(fun(X39,X37),'const.bool.LET_2'(s(fun(X38,fun(X39,X37)),X40),s(X38,X41))),s(X39,X42)))))) | s(X37,X43) = s(X37,'const.pair.UNCURRY_2'(s(fun(X38,fun(X39,X37)),X40),s('type.pair.prod'(X38,X39),'const.pair.,_2'(s(X38,X41),s(X39,X42)))))) )),
  inference(superposition,[],[f141,f143])).
fof(f1888,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.bool.LET_2'(s(fun(X23,bool),'const.bool.LET_2'(s(fun(X23,fun(X23,bool)),'=_0'),s(X23,X29))),s(X23,'const.bool.LET_2'(s(fun(X25,X23),'const.bool.LET_2'(s(fun(X24,fun(X25,X23)),X26),s(X24,X27))),s(X25,X28)))))) | s(X23,X29) != s(X23,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X23)),X26),s('type.pair.prod'(X24,X25),'const.pair.,_2'(s(X24,X27),s(X25,X28)))))) )),
  inference(forward_demodulation,[],[f1844,f91])).
fof(f1844,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (p(s(bool,'const.bool.LET_2'(s(fun(X23,bool),'const.bool.LET_2'(s(fun(X23,fun(X23,bool)),'=_0'),s(X23,X29))),s(X23,happ(s(fun(X25,X23),'const.bool.LET_2'(s(fun(X24,fun(X25,X23)),X26),s(X24,X27))),s(X25,X28)))))) | s(X23,X29) != s(X23,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X23)),X26),s('type.pair.prod'(X24,X25),'const.pair.,_2'(s(X24,X27),s(X25,X28)))))) )),
  inference(superposition,[],[f139,f143])).
fof(f2106,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (s(X55,X63) = s(X55,'const.bool.LET_2'(s(fun(X56,X55),'_dst_x0x1x0_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X56,X57),X60),s(X58,X61))),s(X56,X62))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X63))),s(X55,'const.bool.LET_2'(s(fun(X58,X55),'const.bool.LET_2'(s(fun(X57,fun(X58,X55)),X59),s(X57,'const.bool.LET_2'(s(fun(X56,X57),X60),s(X56,X62))))),s(X58,X61))))))) )),
  inference(forward_demodulation,[],[f2105,f91])).
fof(f2105,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (~p(s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X63))),s(X55,'const.bool.LET_2'(s(fun(X58,X55),'const.bool.LET_2'(s(fun(X57,fun(X58,X55)),X59),s(X57,'const.bool.LET_2'(s(fun(X56,X57),X60),s(X56,X62))))),s(X58,X61)))))) | s(X55,X63) = s(X55,happ(s(fun(X56,X55),'_dst_x0x1x0_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X56,X57),X60),s(X58,X61))),s(X56,X62)))) )),
  inference(forward_demodulation,[],[f2024,f91])).
fof(f2024,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (~p(s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X63))),s(X55,happ(s(fun(X58,X55),'const.bool.LET_2'(s(fun(X57,fun(X58,X55)),X59),s(X57,'const.bool.LET_2'(s(fun(X56,X57),X60),s(X56,X62))))),s(X58,X61)))))) | s(X55,X63) = s(X55,happ(s(fun(X56,X55),'_dst_x0x1x0_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X56,X57),X60),s(X58,X61))),s(X56,X62)))) )),
  inference(superposition,[],[f141,f145])).
fof(f2102,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (s(X37,X45) != s(X37,'const.bool.LET_2'(s(fun(X38,X37),'_dst_x0x1x0_3'(s(fun(X39,fun(X40,X37)),X41),s(fun(X38,X39),X42),s(X40,X43))),s(X38,X44))) | p(s(bool,'const.bool.LET_2'(s(fun(X37,bool),'const.bool.LET_2'(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X45))),s(X37,'const.bool.LET_2'(s(fun(X40,X37),'const.bool.LET_2'(s(fun(X39,fun(X40,X37)),X41),s(X39,'const.bool.LET_2'(s(fun(X38,X39),X42),s(X38,X44))))),s(X40,X43))))))) )),
  inference(forward_demodulation,[],[f2101,f91])).
fof(f2101,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(X37,bool),'const.bool.LET_2'(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X45))),s(X37,'const.bool.LET_2'(s(fun(X40,X37),'const.bool.LET_2'(s(fun(X39,fun(X40,X37)),X41),s(X39,'const.bool.LET_2'(s(fun(X38,X39),X42),s(X38,X44))))),s(X40,X43)))))) | s(X37,X45) != s(X37,happ(s(fun(X38,X37),'_dst_x0x1x0_3'(s(fun(X39,fun(X40,X37)),X41),s(fun(X38,X39),X42),s(X40,X43))),s(X38,X44)))) )),
  inference(forward_demodulation,[],[f2022,f91])).
fof(f2022,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(X37,bool),'const.bool.LET_2'(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X45))),s(X37,happ(s(fun(X40,X37),'const.bool.LET_2'(s(fun(X39,fun(X40,X37)),X41),s(X39,'const.bool.LET_2'(s(fun(X38,X39),X42),s(X38,X44))))),s(X40,X43)))))) | s(X37,X45) != s(X37,happ(s(fun(X38,X37),'_dst_x0x1x0_3'(s(fun(X39,fun(X40,X37)),X41),s(fun(X38,X39),X42),s(X40,X43))),s(X38,X44)))) )),
  inference(superposition,[],[f139,f145])).
fof(f2136,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (s(X69,X77) = s(X69,'const.bool.LET_2'(s(fun(X70,X69),'const.bool.LET_2'(s(fun(X71,fun(X70,X69)),X72),s(X71,'const.bool.LET_2'(s(fun(X73,X71),X74),s(X73,X75))))),s(X70,X76))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X69,bool),'const.bool.LET_2'(s(fun(X69,fun(X69,bool)),'=_0'),s(X69,X77))),s(X69,'const.bool.LET_2'(s(fun(X73,X69),'_dst_x0x1x0_3'(s(fun(X71,fun(X70,X69)),X72),s(fun(X73,X71),X74),s(X70,X76))),s(X73,X75))))))) )),
  inference(forward_demodulation,[],[f2135,f91])).
fof(f2135,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (~p(s(bool,'const.bool.LET_2'(s(fun(X69,bool),'const.bool.LET_2'(s(fun(X69,fun(X69,bool)),'=_0'),s(X69,X77))),s(X69,'const.bool.LET_2'(s(fun(X73,X69),'_dst_x0x1x0_3'(s(fun(X71,fun(X70,X69)),X72),s(fun(X73,X71),X74),s(X70,X76))),s(X73,X75)))))) | s(X69,X77) = s(X69,happ(s(fun(X70,X69),'const.bool.LET_2'(s(fun(X71,fun(X70,X69)),X72),s(X71,'const.bool.LET_2'(s(fun(X73,X71),X74),s(X73,X75))))),s(X70,X76)))) )),
  inference(forward_demodulation,[],[f2057,f91])).
fof(f2057,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (~p(s(bool,'const.bool.LET_2'(s(fun(X69,bool),'const.bool.LET_2'(s(fun(X69,fun(X69,bool)),'=_0'),s(X69,X77))),s(X69,happ(s(fun(X73,X69),'_dst_x0x1x0_3'(s(fun(X71,fun(X70,X69)),X72),s(fun(X73,X71),X74),s(X70,X76))),s(X73,X75)))))) | s(X69,X77) = s(X69,happ(s(fun(X70,X69),'const.bool.LET_2'(s(fun(X71,fun(X70,X69)),X72),s(X71,'const.bool.LET_2'(s(fun(X73,X71),X74),s(X73,X75))))),s(X70,X76)))) )),
  inference(superposition,[],[f141,f145])).
fof(f2132,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (s(X51,X59) != s(X51,'const.bool.LET_2'(s(fun(X52,X51),'const.bool.LET_2'(s(fun(X53,fun(X52,X51)),X54),s(X53,'const.bool.LET_2'(s(fun(X55,X53),X56),s(X55,X57))))),s(X52,X58))) | p(s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,X59))),s(X51,'const.bool.LET_2'(s(fun(X55,X51),'_dst_x0x1x0_3'(s(fun(X53,fun(X52,X51)),X54),s(fun(X55,X53),X56),s(X52,X58))),s(X55,X57))))))) )),
  inference(forward_demodulation,[],[f2131,f91])).
fof(f2131,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,X59))),s(X51,'const.bool.LET_2'(s(fun(X55,X51),'_dst_x0x1x0_3'(s(fun(X53,fun(X52,X51)),X54),s(fun(X55,X53),X56),s(X52,X58))),s(X55,X57)))))) | s(X51,X59) != s(X51,happ(s(fun(X52,X51),'const.bool.LET_2'(s(fun(X53,fun(X52,X51)),X54),s(X53,'const.bool.LET_2'(s(fun(X55,X53),X56),s(X55,X57))))),s(X52,X58)))) )),
  inference(forward_demodulation,[],[f2055,f91])).
fof(f2055,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53,X51] : (p(s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,X59))),s(X51,happ(s(fun(X55,X51),'_dst_x0x1x0_3'(s(fun(X53,fun(X52,X51)),X54),s(fun(X55,X53),X56),s(X52,X58))),s(X55,X57)))))) | s(X51,X59) != s(X51,happ(s(fun(X52,X51),'const.bool.LET_2'(s(fun(X53,fun(X52,X51)),X54),s(X53,'const.bool.LET_2'(s(fun(X55,X53),X56),s(X55,X57))))),s(X52,X58)))) )),
  inference(superposition,[],[f139,f145])).
fof(f1886,plain,(
  ( ! [X47,X52,X50,X48,X53,X51,X49] : (s(X47,X53) = s(X47,'const.bool.LET_2'(s(fun(X48,X47),'const.bool.LET_2'(s(fun(X49,fun(X48,X47)),X50),s(X49,X51))),s(X48,X52))) | ~p(s(bool,'const.bool.LET_2'(s(fun(X47,bool),'const.bool.LET_2'(s(fun(X47,fun(X47,bool)),'=_0'),s(X47,X53))),s(X47,'const.pair.UNCURRY_2'(s(fun(X49,fun(X48,X47)),X50),s('type.pair.prod'(X49,X48),'const.pair.,_2'(s(X49,X51),s(X48,X52))))))))) )),
  inference(forward_demodulation,[],[f1835,f91])).
fof(f1835,plain,(
  ( ! [X47,X52,X50,X48,X53,X51,X49] : (~p(s(bool,'const.bool.LET_2'(s(fun(X47,bool),'const.bool.LET_2'(s(fun(X47,fun(X47,bool)),'=_0'),s(X47,X53))),s(X47,'const.pair.UNCURRY_2'(s(fun(X49,fun(X48,X47)),X50),s('type.pair.prod'(X49,X48),'const.pair.,_2'(s(X49,X51),s(X48,X52)))))))) | s(X47,X53) = s(X47,happ(s(fun(X48,X47),'const.bool.LET_2'(s(fun(X49,fun(X48,X47)),X50),s(X49,X51))),s(X48,X52)))) )),
  inference(superposition,[],[f141,f143])).
fof(f1884,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (s(X33,X39) != s(X33,'const.bool.LET_2'(s(fun(X34,X33),'const.bool.LET_2'(s(fun(X35,fun(X34,X33)),X36),s(X35,X37))),s(X34,X38))) | p(s(bool,'const.bool.LET_2'(s(fun(X33,bool),'const.bool.LET_2'(s(fun(X33,fun(X33,bool)),'=_0'),s(X33,X39))),s(X33,'const.pair.UNCURRY_2'(s(fun(X35,fun(X34,X33)),X36),s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X37),s(X34,X38))))))))) )),
  inference(forward_demodulation,[],[f1833,f91])).
fof(f1833,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (p(s(bool,'const.bool.LET_2'(s(fun(X33,bool),'const.bool.LET_2'(s(fun(X33,fun(X33,bool)),'=_0'),s(X33,X39))),s(X33,'const.pair.UNCURRY_2'(s(fun(X35,fun(X34,X33)),X36),s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X37),s(X34,X38)))))))) | s(X33,X39) != s(X33,happ(s(fun(X34,X33),'const.bool.LET_2'(s(fun(X35,fun(X34,X33)),X36),s(X35,X37))),s(X34,X38)))) )),
  inference(superposition,[],[f139,f143])).
fof(f517,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X8)))) | s(bool,X7) != s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f139,f156])).
fof(f641,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X6)))) | s(bool,X5) = s(bool,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f141,f156])).
fof(f519,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X14))),s(bool,X15)))) | s(bool,X13) != s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f139,f156])).
fof(f522,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X22)))) | s(bool,'T_0') != s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f139,f156])).
fof(f532,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,X14)))) | s(bool,X13) != s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f139,f156])).
fof(f535,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X22))),s(bool,X21)))) | s(bool,'T_0') != s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f139,f156])).
fof(f643,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X13)))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f141,f156])).
fof(f646,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X19))),s(bool,X20)))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f141,f156])).
fof(f8674,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f8642])).
fof(f8642,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f516,f81])).
fof(f530,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X8))),s(bool,'T_0')))) | s(bool,X7) != s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f139,f156])).
fof(f655,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X6))),s(bool,'T_0')))) | s(bool,X5) = s(bool,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f141,f156])).
fof(f2401,plain,(
  ( ! [X134,X142,X140,X138,X136,X135,X133,X141,X139,X137] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X133,X134),bool),'const.bool.LET_2'(s(fun(fun(X133,X134),fun(fun(X133,X134),bool)),'=_0'),s(fun(X133,X134),'const.bool.LET_2'(s(fun(X138,fun(X133,X134)),'_dst_x00x11x00_3'(s(fun(X135,fun(X136,X134)),X137),s(fun(X138,fun(X133,X135)),X139),s(X136,X141))),s(X138,X140))))),s(fun(X133,X134),X142)))) | s(fun(X133,X134),X142) = s(fun(X133,X134),'_dst_x0x1x0_3'(s(fun(X135,fun(X136,X134)),X137),s(fun(X133,X135),'const.bool.LET_2'(s(fun(X138,fun(X133,X135)),X139),s(X138,X140))),s(X136,X141)))) )),
  inference(forward_demodulation,[],[f2364,f91])).
fof(f2364,plain,(
  ( ! [X134,X142,X140,X138,X136,X135,X133,X141,X139,X137] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X133,X134),bool),'const.bool.LET_2'(s(fun(fun(X133,X134),fun(fun(X133,X134),bool)),'=_0'),s(fun(X133,X134),happ(s(fun(X138,fun(X133,X134)),'_dst_x00x11x00_3'(s(fun(X135,fun(X136,X134)),X137),s(fun(X138,fun(X133,X135)),X139),s(X136,X141))),s(X138,X140))))),s(fun(X133,X134),X142)))) | s(fun(X133,X134),X142) = s(fun(X133,X134),'_dst_x0x1x0_3'(s(fun(X135,fun(X136,X134)),X137),s(fun(X133,X135),'const.bool.LET_2'(s(fun(X138,fun(X133,X135)),X139),s(X138,X140))),s(X136,X141)))) )),
  inference(superposition,[],[f141,f154])).
fof(f2399,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115,X113] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X113,X114),bool),'const.bool.LET_2'(s(fun(fun(X113,X114),fun(fun(X113,X114),bool)),'=_0'),s(fun(X113,X114),'const.bool.LET_2'(s(fun(X118,fun(X113,X114)),'_dst_x00x11x00_3'(s(fun(X115,fun(X116,X114)),X117),s(fun(X118,fun(X113,X115)),X119),s(X116,X121))),s(X118,X120))))),s(fun(X113,X114),X122)))) | s(fun(X113,X114),X122) != s(fun(X113,X114),'_dst_x0x1x0_3'(s(fun(X115,fun(X116,X114)),X117),s(fun(X113,X115),'const.bool.LET_2'(s(fun(X118,fun(X113,X115)),X119),s(X118,X120))),s(X116,X121)))) )),
  inference(forward_demodulation,[],[f2362,f91])).
fof(f2362,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115,X113] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X113,X114),bool),'const.bool.LET_2'(s(fun(fun(X113,X114),fun(fun(X113,X114),bool)),'=_0'),s(fun(X113,X114),happ(s(fun(X118,fun(X113,X114)),'_dst_x00x11x00_3'(s(fun(X115,fun(X116,X114)),X117),s(fun(X118,fun(X113,X115)),X119),s(X116,X121))),s(X118,X120))))),s(fun(X113,X114),X122)))) | s(fun(X113,X114),X122) != s(fun(X113,X114),'_dst_x0x1x0_3'(s(fun(X115,fun(X116,X114)),X117),s(fun(X113,X115),'const.bool.LET_2'(s(fun(X118,fun(X113,X115)),X119),s(X118,X120))),s(X116,X121)))) )),
  inference(superposition,[],[f139,f154])).
fof(f2405,plain,(
  ( ! [X134,X132,X140,X138,X136,X135,X133,X141,X139,X137] : (s(fun(X132,X133),X141) = s(fun(X132,X133),'const.bool.LET_2'(s(fun(X134,fun(X132,X133)),'_dst_x00x11x00_3'(s(fun(X135,fun(X136,X133)),X137),s(fun(X134,fun(X132,X135)),X138),s(X136,X139))),s(X134,X140))) | ~p(s(bool,'const.bool.LET_2'(s(fun(fun(X132,X133),bool),'const.bool.LET_2'(s(fun(fun(X132,X133),fun(fun(X132,X133),bool)),'=_0'),s(fun(X132,X133),'_dst_x0x1x0_3'(s(fun(X135,fun(X136,X133)),X137),s(fun(X132,X135),'const.bool.LET_2'(s(fun(X134,fun(X132,X135)),X138),s(X134,X140))),s(X136,X139))))),s(fun(X132,X133),X141))))) )),
  inference(forward_demodulation,[],[f2382,f91])).
fof(f2382,plain,(
  ( ! [X134,X132,X140,X138,X136,X135,X133,X141,X139,X137] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X132,X133),bool),'const.bool.LET_2'(s(fun(fun(X132,X133),fun(fun(X132,X133),bool)),'=_0'),s(fun(X132,X133),'_dst_x0x1x0_3'(s(fun(X135,fun(X136,X133)),X137),s(fun(X132,X135),'const.bool.LET_2'(s(fun(X134,fun(X132,X135)),X138),s(X134,X140))),s(X136,X139))))),s(fun(X132,X133),X141)))) | s(fun(X132,X133),X141) = s(fun(X132,X133),happ(s(fun(X134,fun(X132,X133)),'_dst_x00x11x00_3'(s(fun(X135,fun(X136,X133)),X137),s(fun(X134,fun(X132,X135)),X138),s(X136,X139))),s(X134,X140)))) )),
  inference(superposition,[],[f141,f154])).
fof(f2403,plain,(
  ( ! [X121,X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s(fun(X112,X113),X121) != s(fun(X112,X113),'const.bool.LET_2'(s(fun(X114,fun(X112,X113)),'_dst_x00x11x00_3'(s(fun(X115,fun(X116,X113)),X117),s(fun(X114,fun(X112,X115)),X118),s(X116,X119))),s(X114,X120))) | p(s(bool,'const.bool.LET_2'(s(fun(fun(X112,X113),bool),'const.bool.LET_2'(s(fun(fun(X112,X113),fun(fun(X112,X113),bool)),'=_0'),s(fun(X112,X113),'_dst_x0x1x0_3'(s(fun(X115,fun(X116,X113)),X117),s(fun(X112,X115),'const.bool.LET_2'(s(fun(X114,fun(X112,X115)),X118),s(X114,X120))),s(X116,X119))))),s(fun(X112,X113),X121))))) )),
  inference(forward_demodulation,[],[f2380,f91])).
fof(f2380,plain,(
  ( ! [X121,X118,X116,X114,X112,X120,X119,X117,X115,X113] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X112,X113),bool),'const.bool.LET_2'(s(fun(fun(X112,X113),fun(fun(X112,X113),bool)),'=_0'),s(fun(X112,X113),'_dst_x0x1x0_3'(s(fun(X115,fun(X116,X113)),X117),s(fun(X112,X115),'const.bool.LET_2'(s(fun(X114,fun(X112,X115)),X118),s(X114,X120))),s(X116,X119))))),s(fun(X112,X113),X121)))) | s(fun(X112,X113),X121) != s(fun(X112,X113),happ(s(fun(X114,fun(X112,X113)),'_dst_x00x11x00_3'(s(fun(X115,fun(X116,X113)),X117),s(fun(X114,fun(X112,X115)),X118),s(X116,X119))),s(X114,X120)))) )),
  inference(superposition,[],[f139,f154])).
fof(f2402,plain,(
  ( ! [X146,X144,X152,X151,X149,X147,X145,X143,X150,X148] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X143,X144),bool),'const.bool.LET_2'(s(fun(fun(X143,X144),fun(fun(X143,X144),bool)),'=_0'),s(fun(X143,X144),X152))),s(fun(X143,X144),'const.bool.LET_2'(s(fun(X148,fun(X143,X144)),'_dst_x00x11x00_3'(s(fun(X145,fun(X146,X144)),X147),s(fun(X148,fun(X143,X145)),X149),s(X146,X151))),s(X148,X150)))))) | s(fun(X143,X144),X152) = s(fun(X143,X144),'_dst_x0x1x0_3'(s(fun(X145,fun(X146,X144)),X147),s(fun(X143,X145),'const.bool.LET_2'(s(fun(X148,fun(X143,X145)),X149),s(X148,X150))),s(X146,X151)))) )),
  inference(forward_demodulation,[],[f2365,f91])).
fof(f2365,plain,(
  ( ! [X146,X144,X152,X151,X149,X147,X145,X143,X150,X148] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X143,X144),bool),'const.bool.LET_2'(s(fun(fun(X143,X144),fun(fun(X143,X144),bool)),'=_0'),s(fun(X143,X144),X152))),s(fun(X143,X144),happ(s(fun(X148,fun(X143,X144)),'_dst_x00x11x00_3'(s(fun(X145,fun(X146,X144)),X147),s(fun(X148,fun(X143,X145)),X149),s(X146,X151))),s(X148,X150)))))) | s(fun(X143,X144),X152) = s(fun(X143,X144),'_dst_x0x1x0_3'(s(fun(X145,fun(X146,X144)),X147),s(fun(X143,X145),'const.bool.LET_2'(s(fun(X148,fun(X143,X145)),X149),s(X148,X150))),s(X146,X151)))) )),
  inference(superposition,[],[f141,f154])).
fof(f2400,plain,(
  ( ! [X127,X125,X123,X132,X130,X128,X126,X124,X131,X129] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X123,X124),bool),'const.bool.LET_2'(s(fun(fun(X123,X124),fun(fun(X123,X124),bool)),'=_0'),s(fun(X123,X124),X132))),s(fun(X123,X124),'const.bool.LET_2'(s(fun(X128,fun(X123,X124)),'_dst_x00x11x00_3'(s(fun(X125,fun(X126,X124)),X127),s(fun(X128,fun(X123,X125)),X129),s(X126,X131))),s(X128,X130)))))) | s(fun(X123,X124),X132) != s(fun(X123,X124),'_dst_x0x1x0_3'(s(fun(X125,fun(X126,X124)),X127),s(fun(X123,X125),'const.bool.LET_2'(s(fun(X128,fun(X123,X125)),X129),s(X128,X130))),s(X126,X131)))) )),
  inference(forward_demodulation,[],[f2363,f91])).
fof(f2363,plain,(
  ( ! [X127,X125,X123,X132,X130,X128,X126,X124,X131,X129] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X123,X124),bool),'const.bool.LET_2'(s(fun(fun(X123,X124),fun(fun(X123,X124),bool)),'=_0'),s(fun(X123,X124),X132))),s(fun(X123,X124),happ(s(fun(X128,fun(X123,X124)),'_dst_x00x11x00_3'(s(fun(X125,fun(X126,X124)),X127),s(fun(X128,fun(X123,X125)),X129),s(X126,X131))),s(X128,X130)))))) | s(fun(X123,X124),X132) != s(fun(X123,X124),'_dst_x0x1x0_3'(s(fun(X125,fun(X126,X124)),X127),s(fun(X123,X125),'const.bool.LET_2'(s(fun(X128,fun(X123,X125)),X129),s(X128,X130))),s(X126,X131)))) )),
  inference(superposition,[],[f139,f154])).
fof(f2406,plain,(
  ( ! [X146,X144,X142,X151,X149,X147,X145,X143,X150,X148] : (s(fun(X142,X143),X151) = s(fun(X142,X143),'const.bool.LET_2'(s(fun(X144,fun(X142,X143)),'_dst_x00x11x00_3'(s(fun(X145,fun(X146,X143)),X147),s(fun(X144,fun(X142,X145)),X148),s(X146,X149))),s(X144,X150))) | ~p(s(bool,'const.bool.LET_2'(s(fun(fun(X142,X143),bool),'const.bool.LET_2'(s(fun(fun(X142,X143),fun(fun(X142,X143),bool)),'=_0'),s(fun(X142,X143),X151))),s(fun(X142,X143),'_dst_x0x1x0_3'(s(fun(X145,fun(X146,X143)),X147),s(fun(X142,X145),'const.bool.LET_2'(s(fun(X144,fun(X142,X145)),X148),s(X144,X150))),s(X146,X149))))))) )),
  inference(forward_demodulation,[],[f2383,f91])).
fof(f2383,plain,(
  ( ! [X146,X144,X142,X151,X149,X147,X145,X143,X150,X148] : (~p(s(bool,'const.bool.LET_2'(s(fun(fun(X142,X143),bool),'const.bool.LET_2'(s(fun(fun(X142,X143),fun(fun(X142,X143),bool)),'=_0'),s(fun(X142,X143),X151))),s(fun(X142,X143),'_dst_x0x1x0_3'(s(fun(X145,fun(X146,X143)),X147),s(fun(X142,X145),'const.bool.LET_2'(s(fun(X144,fun(X142,X145)),X148),s(X144,X150))),s(X146,X149)))))) | s(fun(X142,X143),X151) = s(fun(X142,X143),happ(s(fun(X144,fun(X142,X143)),'_dst_x00x11x00_3'(s(fun(X145,fun(X146,X143)),X147),s(fun(X144,fun(X142,X145)),X148),s(X146,X149))),s(X144,X150)))) )),
  inference(superposition,[],[f141,f154])).
fof(f2404,plain,(
  ( ! [X127,X125,X123,X130,X128,X126,X124,X122,X131,X129] : (s(fun(X122,X123),X131) != s(fun(X122,X123),'const.bool.LET_2'(s(fun(X124,fun(X122,X123)),'_dst_x00x11x00_3'(s(fun(X125,fun(X126,X123)),X127),s(fun(X124,fun(X122,X125)),X128),s(X126,X129))),s(X124,X130))) | p(s(bool,'const.bool.LET_2'(s(fun(fun(X122,X123),bool),'const.bool.LET_2'(s(fun(fun(X122,X123),fun(fun(X122,X123),bool)),'=_0'),s(fun(X122,X123),X131))),s(fun(X122,X123),'_dst_x0x1x0_3'(s(fun(X125,fun(X126,X123)),X127),s(fun(X122,X125),'const.bool.LET_2'(s(fun(X124,fun(X122,X125)),X128),s(X124,X130))),s(X126,X129))))))) )),
  inference(forward_demodulation,[],[f2381,f91])).
fof(f2381,plain,(
  ( ! [X127,X125,X123,X130,X128,X126,X124,X122,X131,X129] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X122,X123),bool),'const.bool.LET_2'(s(fun(fun(X122,X123),fun(fun(X122,X123),bool)),'=_0'),s(fun(X122,X123),X131))),s(fun(X122,X123),'_dst_x0x1x0_3'(s(fun(X125,fun(X126,X123)),X127),s(fun(X122,X125),'const.bool.LET_2'(s(fun(X124,fun(X122,X125)),X128),s(X124,X130))),s(X126,X129)))))) | s(fun(X122,X123),X131) != s(fun(X122,X123),happ(s(fun(X124,fun(X122,X123)),'_dst_x00x11x00_3'(s(fun(X125,fun(X126,X123)),X127),s(fun(X124,fun(X122,X125)),X128),s(X126,X129))),s(X124,X130)))) )),
  inference(superposition,[],[f139,f154])).
fof(f1662,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X5),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f1658,f1624])).
fof(f1658,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X5),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | ~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(backward_demodulation,[],[f1624,f149])).
fof(f1661,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X5),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f1657,f1624])).
fof(f1657,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X5),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(backward_demodulation,[],[f1624,f153])).
fof(f2845,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X115,fun(X116,bool)),'_dst_x00x11x00_3'(s(fun(X109,fun(X109,bool)),'=_0'),s(fun(X115,fun(X116,X109)),X117),s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.bool.LET_2'(s(fun(X111,fun(X108,X109)),'const.bool.LET_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s(X110,X113))),s(X111,X114))),s(X108,X118))))),s('type.pair.prod'(X115,X116),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X115,X116)),X119),s(X108,X118)))))) | s(X109,'const.pair.UNCURRY_2'(s(fun(X115,fun(X116,X109)),X117),s('type.pair.prod'(X115,X116),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X115,X116)),X119),s(X108,X118))))) = s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.pair.UNCURRY_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s('type.pair.prod'(X110,X111),'const.pair.,_2'(s(X110,X113),s(X111,X114))))),s(X108,X118)))) )),
  inference(forward_demodulation,[],[f2767,f91])).
fof(f2767,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X115,fun(X116,bool)),'_dst_x00x11x00_3'(s(fun(X109,fun(X109,bool)),'=_0'),s(fun(X115,fun(X116,X109)),X117),s(X109,'const.bool.LET_2'(s(fun(X108,X109),happ(s(fun(X111,fun(X108,X109)),'const.bool.LET_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s(X110,X113))),s(X111,X114))),s(X108,X118))))),s('type.pair.prod'(X115,X116),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X115,X116)),X119),s(X108,X118)))))) | s(X109,'const.pair.UNCURRY_2'(s(fun(X115,fun(X116,X109)),X117),s('type.pair.prod'(X115,X116),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X115,X116)),X119),s(X108,X118))))) = s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.pair.UNCURRY_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s('type.pair.prod'(X110,X111),'const.pair.,_2'(s(X110,X113),s(X111,X114))))),s(X108,X118)))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2606,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X65,fun(X66,bool)),'_dst_x00x11x00_3'(s(fun(X58,fun(X58,bool)),'=_0'),s(fun(X65,fun(X66,X58)),X67),s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.bool.LET_2'(s(fun(X59,fun(X57,X58)),'const.bool.LET_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s(X60,X62))),s(X59,X63))),s(X57,X64))))),s('type.pair.prod'(X65,X66),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X65,X66)),X68),s(X57,X64)))))) | s(X58,'const.pair.UNCURRY_2'(s(fun(X65,fun(X66,X58)),X67),s('type.pair.prod'(X65,X66),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X65,X66)),X68),s(X57,X64))))) != s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.pair.UNCURRY_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s('type.pair.prod'(X60,X59),'const.pair.,_2'(s(X60,X62),s(X59,X63))))),s(X57,X64)))) )),
  inference(forward_demodulation,[],[f2535,f91])).
fof(f2535,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X58,'const.pair.UNCURRY_2'(s(fun(X65,fun(X66,X58)),X67),s('type.pair.prod'(X65,X66),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X65,X66)),X68),s(X57,X64))))) != s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.pair.UNCURRY_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s('type.pair.prod'(X60,X59),'const.pair.,_2'(s(X60,X62),s(X59,X63))))),s(X57,X64))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X65,fun(X66,bool)),'_dst_x00x11x00_3'(s(fun(X58,fun(X58,bool)),'=_0'),s(fun(X65,fun(X66,X58)),X67),s(X58,'const.bool.LET_2'(s(fun(X57,X58),happ(s(fun(X59,fun(X57,X58)),'const.bool.LET_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s(X60,X62))),s(X59,X63))),s(X57,X64))))),s('type.pair.prod'(X65,X66),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X65,X66)),X68),s(X57,X64))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2838,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X58,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,X58)),X66),s('type.pair.prod'(X64,X65),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X64,X65)),X68),s(X57,X67))))) = s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.bool.LET_2'(s(fun(X59,fun(X57,X58)),'const.bool.LET_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s(X60,X62))),s(X59,X63))),s(X57,X67))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),'_dst_x00x11x00_3'(s(fun(X58,fun(X58,bool)),'=_0'),s(fun(X64,fun(X65,X58)),X66),s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.pair.UNCURRY_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s('type.pair.prod'(X60,X59),'const.pair.,_2'(s(X60,X62),s(X59,X63))))),s(X57,X67))))),s('type.pair.prod'(X64,X65),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X64,X65)),X68),s(X57,X67))))))) )),
  inference(forward_demodulation,[],[f2763,f91])).
fof(f2763,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,bool)),'_dst_x00x11x00_3'(s(fun(X58,fun(X58,bool)),'=_0'),s(fun(X64,fun(X65,X58)),X66),s(X58,'const.bool.LET_2'(s(fun(X57,X58),'const.pair.UNCURRY_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s('type.pair.prod'(X60,X59),'const.pair.,_2'(s(X60,X62),s(X59,X63))))),s(X57,X67))))),s('type.pair.prod'(X64,X65),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X64,X65)),X68),s(X57,X67)))))) | s(X58,'const.pair.UNCURRY_2'(s(fun(X64,fun(X65,X58)),X66),s('type.pair.prod'(X64,X65),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X64,X65)),X68),s(X57,X67))))) = s(X58,'const.bool.LET_2'(s(fun(X57,X58),happ(s(fun(X59,fun(X57,X58)),'const.bool.LET_2'(s(fun(X60,fun(X59,fun(X57,X58))),X61),s(X60,X62))),s(X59,X63))),s(X57,X67)))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2613,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X109,'const.pair.UNCURRY_2'(s(fun(X116,fun(X117,X109)),X118),s('type.pair.prod'(X116,X117),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X116,X117)),X119),s(X108,X115))))) != s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.bool.LET_2'(s(fun(X111,fun(X108,X109)),'const.bool.LET_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s(X110,X113))),s(X111,X114))),s(X108,X115))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X116,fun(X117,bool)),'_dst_x00x11x00_3'(s(fun(X109,fun(X109,bool)),'=_0'),s(fun(X116,fun(X117,X109)),X118),s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.pair.UNCURRY_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s('type.pair.prod'(X110,X111),'const.pair.,_2'(s(X110,X113),s(X111,X114))))),s(X108,X115))))),s('type.pair.prod'(X116,X117),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X116,X117)),X119),s(X108,X115))))))) )),
  inference(forward_demodulation,[],[f2539,f91])).
fof(f2539,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X109,'const.pair.UNCURRY_2'(s(fun(X116,fun(X117,X109)),X118),s('type.pair.prod'(X116,X117),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X116,X117)),X119),s(X108,X115))))) != s(X109,'const.bool.LET_2'(s(fun(X108,X109),happ(s(fun(X111,fun(X108,X109)),'const.bool.LET_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s(X110,X113))),s(X111,X114))),s(X108,X115))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X116,fun(X117,bool)),'_dst_x00x11x00_3'(s(fun(X109,fun(X109,bool)),'=_0'),s(fun(X116,fun(X117,X109)),X118),s(X109,'const.bool.LET_2'(s(fun(X108,X109),'const.pair.UNCURRY_2'(s(fun(X110,fun(X111,fun(X108,X109))),X112),s('type.pair.prod'(X110,X111),'const.pair.,_2'(s(X110,X113),s(X111,X114))))),s(X108,X115))))),s('type.pair.prod'(X116,X117),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X116,X117)),X119),s(X108,X115))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2876,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X77,fun(X78,X79)),X80),s(X79,'const.bool.LET_2'(s(fun(X71,X79),X81),s(X71,'const.bool.LET_2'(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))))),s('type.pair.prod'(X77,X78),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X77,X78)),X82),s(X71,'const.bool.LET_2'(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76)))))))) | s(X79,'const.bool.LET_2'(s(fun(X71,X79),X81),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))) = s(X79,'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,X79)),X80),s('type.pair.prod'(X77,X78),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X77,X78)),X82),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76)))))))))) )),
  inference(forward_demodulation,[],[f2792,f91])).
fof(f2792,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X77,fun(X78,X79)),X80),s(X79,'const.bool.LET_2'(s(fun(X71,X79),X81),s(X71,happ(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))))),s('type.pair.prod'(X77,X78),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X77,X78)),X82),s(X71,happ(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76)))))))) | s(X79,'const.bool.LET_2'(s(fun(X71,X79),X81),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))) = s(X79,'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,X79)),X80),s('type.pair.prod'(X77,X78),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X77,X78)),X82),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76)))))))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2649,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,bool)),'_dst_x00x11x00_3'(s(fun(X26,fun(X26,bool)),'=_0'),s(fun(X28,fun(X29,X26)),X30),s(X26,'const.bool.LET_2'(s(fun(X20,X26),X27),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))))),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X28,X29)),X31),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25)))))))) | s(X26,'const.bool.LET_2'(s(fun(X20,X26),X27),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))) != s(X26,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,X26)),X30),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X28,X29)),X31),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25)))))))))) )),
  inference(forward_demodulation,[],[f2569,f91])).
fof(f2569,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (s(X26,'const.bool.LET_2'(s(fun(X20,X26),X27),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))) != s(X26,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,X26)),X30),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X28,X29)),X31),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X28,fun(X29,bool)),'_dst_x00x11x00_3'(s(fun(X26,fun(X26,bool)),'=_0'),s(fun(X28,fun(X29,X26)),X30),s(X26,'const.bool.LET_2'(s(fun(X20,X26),X27),s(X20,happ(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))))),s('type.pair.prod'(X28,X29),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X28,X29)),X31),s(X20,happ(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2869,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (s(X28,'const.bool.LET_2'(s(fun(X20,X28),X30),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))) = s(X28,'const.pair.UNCURRY_2'(s(fun(X26,fun(X27,X28)),X29),s('type.pair.prod'(X26,X27),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X26,X27)),X31),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X26,fun(X27,bool)),'_dst_x00x11x00_3'(s(fun(X28,fun(X28,bool)),'=_0'),s(fun(X26,fun(X27,X28)),X29),s(X28,'const.bool.LET_2'(s(fun(X20,X28),X30),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))))),s('type.pair.prod'(X26,X27),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X26,X27)),X31),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))))))) )),
  inference(forward_demodulation,[],[f2788,f91])).
fof(f2788,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X26,fun(X27,bool)),'_dst_x00x11x00_3'(s(fun(X28,fun(X28,bool)),'=_0'),s(fun(X26,fun(X27,X28)),X29),s(X28,'const.bool.LET_2'(s(fun(X20,X28),X30),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25))))))))),s('type.pair.prod'(X26,X27),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X26,X27)),X31),s(X20,'const.pair.UNCURRY_2'(s(fun(X22,fun(X21,X20)),X23),s('type.pair.prod'(X22,X21),'const.pair.,_2'(s(X22,X24),s(X21,X25)))))))))) | s(X28,'const.bool.LET_2'(s(fun(X20,X28),X30),s(X20,happ(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25))))) = s(X28,'const.pair.UNCURRY_2'(s(fun(X26,fun(X27,X28)),X29),s('type.pair.prod'(X26,X27),'const.bool.LET_2'(s(fun(X20,'type.pair.prod'(X26,X27)),X31),s(X20,happ(s(fun(X21,X20),'const.bool.LET_2'(s(fun(X22,fun(X21,X20)),X23),s(X22,X24))),s(X21,X25)))))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2656,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s(X77,'const.bool.LET_2'(s(fun(X71,X77),X78),s(X71,'const.bool.LET_2'(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))) != s(X77,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,X77)),X81),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X79,X80)),X82),s(X71,'const.bool.LET_2'(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,bool)),'_dst_x00x11x00_3'(s(fun(X77,fun(X77,bool)),'=_0'),s(fun(X79,fun(X80,X77)),X81),s(X77,'const.bool.LET_2'(s(fun(X71,X77),X78),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))))),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X79,X80)),X82),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))))))) )),
  inference(forward_demodulation,[],[f2573,f91])).
fof(f2573,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s(X77,'const.bool.LET_2'(s(fun(X71,X77),X78),s(X71,happ(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))) != s(X77,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,X77)),X81),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X79,X80)),X82),s(X71,happ(s(fun(X73,X71),'const.bool.LET_2'(s(fun(X72,fun(X73,X71)),X74),s(X72,X75))),s(X73,X76))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,bool)),'_dst_x00x11x00_3'(s(fun(X77,fun(X77,bool)),'=_0'),s(fun(X79,fun(X80,X77)),X81),s(X77,'const.bool.LET_2'(s(fun(X71,X77),X78),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))))),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X71,'type.pair.prod'(X79,X80)),X82),s(X71,'const.pair.UNCURRY_2'(s(fun(X72,fun(X73,X71)),X74),s('type.pair.prod'(X72,X73),'const.pair.,_2'(s(X72,X75),s(X73,X76))))))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2834,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X32,fun(X33,bool)),'_dst_x00x11x00_3'(s(fun(X24,fun(X24,bool)),'=_0'),s(fun(X32,fun(X33,X24)),X34),s(X24,'const.bool.LET_2'(s(fun(X23,X24),'const.bool.LET_2'(s(fun(X28,fun(X23,X24)),'_dst_x00x11x00_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X28,fun(X23,X25)),X29),s(X26,X31))),s(X28,X30))),s(X23,X35))))),s('type.pair.prod'(X32,X33),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X32,X33)),X36),s(X23,X35)))))) | s(X24,'const.pair.UNCURRY_2'(s(fun(X32,fun(X33,X24)),X34),s('type.pair.prod'(X32,X33),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X32,X33)),X36),s(X23,X35))))) = s(X24,'const.bool.LET_2'(s(fun(X23,X24),'_dst_x0x1x0_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X23,X25),'const.bool.LET_2'(s(fun(X28,fun(X23,X25)),X29),s(X28,X30))),s(X26,X31))),s(X23,X35)))) )),
  inference(forward_demodulation,[],[f2760,f91])).
fof(f2760,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X32,fun(X33,bool)),'_dst_x00x11x00_3'(s(fun(X24,fun(X24,bool)),'=_0'),s(fun(X32,fun(X33,X24)),X34),s(X24,'const.bool.LET_2'(s(fun(X23,X24),happ(s(fun(X28,fun(X23,X24)),'_dst_x00x11x00_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X28,fun(X23,X25)),X29),s(X26,X31))),s(X28,X30))),s(X23,X35))))),s('type.pair.prod'(X32,X33),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X32,X33)),X36),s(X23,X35)))))) | s(X24,'const.pair.UNCURRY_2'(s(fun(X32,fun(X33,X24)),X34),s('type.pair.prod'(X32,X33),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X32,X33)),X36),s(X23,X35))))) = s(X24,'const.bool.LET_2'(s(fun(X23,X24),'_dst_x0x1x0_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X23,X25),'const.bool.LET_2'(s(fun(X28,fun(X23,X25)),X29),s(X28,X30))),s(X26,X31))),s(X23,X35)))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2601,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X19,fun(X20,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X19,fun(X20,X10)),X21),s(X10,'const.bool.LET_2'(s(fun(X9,X10),'const.bool.LET_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x11x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(X13,X16))),s(X11,X17))),s(X9,X18))))),s('type.pair.prod'(X19,X20),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X19,X20)),X22),s(X9,X18)))))) | s(X10,'const.pair.UNCURRY_2'(s(fun(X19,fun(X20,X10)),X21),s('type.pair.prod'(X19,X20),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X19,X20)),X22),s(X9,X18))))) != s(X10,'const.bool.LET_2'(s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),'const.bool.LET_2'(s(fun(X11,fun(X9,X12)),X15),s(X11,X17))),s(X13,X16))),s(X9,X18)))) )),
  inference(forward_demodulation,[],[f2531,f91])).
fof(f2531,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X10,'const.pair.UNCURRY_2'(s(fun(X19,fun(X20,X10)),X21),s('type.pair.prod'(X19,X20),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X19,X20)),X22),s(X9,X18))))) != s(X10,'const.bool.LET_2'(s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),'const.bool.LET_2'(s(fun(X11,fun(X9,X12)),X15),s(X11,X17))),s(X13,X16))),s(X9,X18))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X19,fun(X20,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X19,fun(X20,X10)),X21),s(X10,'const.bool.LET_2'(s(fun(X9,X10),happ(s(fun(X11,fun(X9,X10)),'_dst_x00x11x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(X13,X16))),s(X11,X17))),s(X9,X18))))),s('type.pair.prod'(X19,X20),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X19,X20)),X22),s(X9,X18))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2842,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X84,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,X84)),X94),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95))))) = s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X95))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X92,fun(X93,X84)),X94),s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X95))))),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95))))))) )),
  inference(forward_demodulation,[],[f2841,f91])).
fof(f2841,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X92,fun(X93,X84)),X94),s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X95))))),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95)))))) | s(X84,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,X84)),X94),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95))))) = s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X95)))) )),
  inference(forward_demodulation,[],[f2765,f91])).
fof(f2765,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X92,fun(X93,X84)),X94),s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X95))))),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95)))))) | s(X84,'const.pair.UNCURRY_2'(s(fun(X92,fun(X93,X84)),X94),s('type.pair.prod'(X92,X93),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X92,X93)),X96),s(X83,X95))))) = s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X95)))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2608,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X79,fun(X80,X70)),X81),s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X78))))),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78)))))) | s(X70,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,X70)),X81),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78))))) != s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X78)))) )),
  inference(forward_demodulation,[],[f2607,f91])).
fof(f2607,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X70,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,X70)),X81),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78))))) != s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X78))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X79,fun(X80,X70)),X81),s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X78))))),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78))))))) )),
  inference(forward_demodulation,[],[f2536,f91])).
fof(f2536,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X70,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,X70)),X81),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78))))) != s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X78))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X79,fun(X80,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X79,fun(X80,X70)),X81),s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X78))))),s('type.pair.prod'(X79,X80),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X79,X80)),X82),s(X69,X78))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2840,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X70,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,X70)),X80),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81))))) = s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X81))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X78,fun(X79,X70)),X80),s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X81))))),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81))))))) )),
  inference(forward_demodulation,[],[f2839,f91])).
fof(f2839,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X78,fun(X79,X70)),X80),s(X70,'const.bool.LET_2'(s(fun(X69,X70),'const.bool.LET_2'(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X81))))),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81)))))) | s(X70,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,X70)),X80),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81))))) = s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X81)))) )),
  inference(forward_demodulation,[],[f2764,f91])).
fof(f2764,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,bool)),'_dst_x00x11x00_3'(s(fun(X70,fun(X70,bool)),'=_0'),s(fun(X78,fun(X79,X70)),X80),s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X73,fun(X69,X70)),'const.bool.LET_2'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(X72,'const.bool.LET_2'(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))),s(X69,X81))))),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81)))))) | s(X70,'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,X70)),X80),s('type.pair.prod'(X78,X79),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X78,X79)),X82),s(X69,X81))))) = s(X70,'const.bool.LET_2'(s(fun(X69,X70),happ(s(fun(X71,fun(X69,X70)),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,fun(X69,X70))),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X71,X77))),s(X69,X81)))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2610,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X93,fun(X94,X84)),X95),s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X92))))),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92)))))) | s(X84,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,X84)),X95),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92))))) != s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X92)))) )),
  inference(forward_demodulation,[],[f2609,f91])).
fof(f2609,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X84,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,X84)),X95),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92))))) != s(X84,'const.bool.LET_2'(s(fun(X83,X84),'const.bool.LET_2'(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X92))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X93,fun(X94,X84)),X95),s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X92))))),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92))))))) )),
  inference(forward_demodulation,[],[f2537,f91])).
fof(f2537,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X84,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,X84)),X95),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92))))) != s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X88,fun(X83,X84)),'_dst_x0x1x0_3'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(fun(X88,X86),X89),s(X85,X91))),s(X88,X90))),s(X83,X92))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,bool)),'_dst_x00x11x00_3'(s(fun(X84,fun(X84,bool)),'=_0'),s(fun(X93,fun(X94,X84)),X95),s(X84,'const.bool.LET_2'(s(fun(X83,X84),happ(s(fun(X85,fun(X83,X84)),'const.bool.LET_2'(s(fun(X86,fun(X85,fun(X83,X84))),X87),s(X86,'const.bool.LET_2'(s(fun(X88,X86),X89),s(X88,X90))))),s(X85,X91))),s(X83,X92))))),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X93,X94)),X96),s(X83,X92))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2833,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X10,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,X10)),X20),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X18,X19)),X22),s(X9,X21))))) = s(X10,'const.bool.LET_2'(s(fun(X9,X10),'const.bool.LET_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x11x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(X13,X16))),s(X11,X17))),s(X9,X21))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X18,fun(X19,X10)),X20),s(X10,'const.bool.LET_2'(s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),'const.bool.LET_2'(s(fun(X11,fun(X9,X12)),X15),s(X11,X17))),s(X13,X16))),s(X9,X21))))),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X18,X19)),X22),s(X9,X21))))))) )),
  inference(forward_demodulation,[],[f2759,f91])).
fof(f2759,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X18,fun(X19,X10)),X20),s(X10,'const.bool.LET_2'(s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),'const.bool.LET_2'(s(fun(X11,fun(X9,X12)),X15),s(X11,X17))),s(X13,X16))),s(X9,X21))))),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X18,X19)),X22),s(X9,X21)))))) | s(X10,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,X10)),X20),s('type.pair.prod'(X18,X19),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X18,X19)),X22),s(X9,X21))))) = s(X10,'const.bool.LET_2'(s(fun(X9,X10),happ(s(fun(X11,fun(X9,X10)),'_dst_x00x11x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(X13,X16))),s(X11,X17))),s(X9,X21)))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2602,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X24,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,X24)),X35),s('type.pair.prod'(X33,X34),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X33,X34)),X36),s(X23,X32))))) != s(X24,'const.bool.LET_2'(s(fun(X23,X24),'const.bool.LET_2'(s(fun(X28,fun(X23,X24)),'_dst_x00x11x00_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X28,fun(X23,X25)),X29),s(X26,X31))),s(X28,X30))),s(X23,X32))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,bool)),'_dst_x00x11x00_3'(s(fun(X24,fun(X24,bool)),'=_0'),s(fun(X33,fun(X34,X24)),X35),s(X24,'const.bool.LET_2'(s(fun(X23,X24),'_dst_x0x1x0_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X23,X25),'const.bool.LET_2'(s(fun(X28,fun(X23,X25)),X29),s(X28,X30))),s(X26,X31))),s(X23,X32))))),s('type.pair.prod'(X33,X34),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X33,X34)),X36),s(X23,X32))))))) )),
  inference(forward_demodulation,[],[f2532,f91])).
fof(f2532,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X24,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,X24)),X35),s('type.pair.prod'(X33,X34),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X33,X34)),X36),s(X23,X32))))) != s(X24,'const.bool.LET_2'(s(fun(X23,X24),happ(s(fun(X28,fun(X23,X24)),'_dst_x00x11x00_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X28,fun(X23,X25)),X29),s(X26,X31))),s(X28,X30))),s(X23,X32))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X33,fun(X34,bool)),'_dst_x00x11x00_3'(s(fun(X24,fun(X24,bool)),'=_0'),s(fun(X33,fun(X34,X24)),X35),s(X24,'const.bool.LET_2'(s(fun(X23,X24),'_dst_x0x1x0_3'(s(fun(X25,fun(X26,X24)),X27),s(fun(X23,X25),'const.bool.LET_2'(s(fun(X28,fun(X23,X25)),X29),s(X28,X30))),s(X26,X31))),s(X23,X32))))),s('type.pair.prod'(X33,X34),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X33,X34)),X36),s(X23,X32))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2873,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,'const.bool.LET_2'(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))) = s(X56,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,X56)),X57),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,'const.bool.LET_2'(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,bool)),'_dst_x00x11x00_3'(s(fun(X56,fun(X56,bool)),'=_0'),s(fun(X54,fun(X55,X56)),X57),s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))))) )),
  inference(forward_demodulation,[],[f2872,f91])).
fof(f2872,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,bool)),'_dst_x00x11x00_3'(s(fun(X56,fun(X56,bool)),'=_0'),s(fun(X54,fun(X55,X56)),X57),s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52)))))))) | s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))) = s(X56,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,X56)),X57),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53)))))))) )),
  inference(forward_demodulation,[],[f2790,f91])).
fof(f2790,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,bool)),'_dst_x00x11x00_3'(s(fun(X56,fun(X56,bool)),'=_0'),s(fun(X54,fun(X55,X56)),X57),s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,happ(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,happ(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52)))))))) | s(X56,'const.bool.LET_2'(s(fun(X46,X56),X58),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))) = s(X56,'const.pair.UNCURRY_2'(s(fun(X54,fun(X55,X56)),X57),s('type.pair.prod'(X54,X55),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X54,X55)),X59),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53)))))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2651,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),'_dst_x00x11x00_3'(s(fun(X40,fun(X40,bool)),'=_0'),s(fun(X42,fun(X43,X40)),X44),s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,'const.bool.LET_2'(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,'const.bool.LET_2'(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39)))))))) | s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))) != s(X40,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,X40)),X44),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38)))))))) )),
  inference(forward_demodulation,[],[f2650,f91])).
fof(f2650,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))) != s(X40,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,X40)),X44),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),'_dst_x00x11x00_3'(s(fun(X40,fun(X40,bool)),'=_0'),s(fun(X42,fun(X43,X40)),X44),s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))))) )),
  inference(forward_demodulation,[],[f2570,f91])).
fof(f2570,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,happ(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))) != s(X40,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,X40)),X44),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,happ(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),'_dst_x00x11x00_3'(s(fun(X40,fun(X40,bool)),'=_0'),s(fun(X42,fun(X43,X40)),X44),s(X40,'const.bool.LET_2'(s(fun(X32,X40),X41),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))),s('type.pair.prod'(X42,X43),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X42,X43)),X45),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2871,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,'const.bool.LET_2'(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))) = s(X42,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,X42)),X43),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,'const.bool.LET_2'(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,bool)),'_dst_x00x11x00_3'(s(fun(X42,fun(X42,bool)),'=_0'),s(fun(X40,fun(X41,X42)),X43),s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))))) )),
  inference(forward_demodulation,[],[f2870,f91])).
fof(f2870,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,bool)),'_dst_x00x11x00_3'(s(fun(X42,fun(X42,bool)),'=_0'),s(fun(X40,fun(X41,X42)),X43),s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,'const.bool.LET_2'(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38)))))))) | s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))) = s(X42,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,X42)),X43),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39)))))))) )),
  inference(forward_demodulation,[],[f2789,f91])).
fof(f2789,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X32,X44,X42,X40] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,bool)),'_dst_x00x11x00_3'(s(fun(X42,fun(X42,bool)),'=_0'),s(fun(X40,fun(X41,X42)),X43),s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,happ(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38))))))),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,happ(s(fun(X35,X32),'const.bool.LET_2'(s(fun(X34,fun(X35,X32)),X36),s(X34,'const.bool.LET_2'(s(fun(X33,X34),X37),s(X33,X39))))),s(X35,X38)))))))) | s(X42,'const.bool.LET_2'(s(fun(X32,X42),X44),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39))))) = s(X42,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,X42)),X43),s('type.pair.prod'(X40,X41),'const.bool.LET_2'(s(fun(X32,'type.pair.prod'(X40,X41)),X45),s(X32,happ(s(fun(X33,X32),'_dst_x0x1x0_3'(s(fun(X34,fun(X35,X32)),X36),s(fun(X33,X34),X37),s(X35,X38))),s(X33,X39)))))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2653,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,bool)),'_dst_x00x11x00_3'(s(fun(X54,fun(X54,bool)),'=_0'),s(fun(X56,fun(X57,X54)),X58),s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,'const.bool.LET_2'(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,'const.bool.LET_2'(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53)))))))) | s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))) != s(X54,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,X54)),X58),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52)))))))) )),
  inference(forward_demodulation,[],[f2652,f91])).
fof(f2652,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))) != s(X54,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,X54)),X58),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,'const.bool.LET_2'(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,bool)),'_dst_x00x11x00_3'(s(fun(X54,fun(X54,bool)),'=_0'),s(fun(X56,fun(X57,X54)),X58),s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))))) )),
  inference(forward_demodulation,[],[f2571,f91])).
fof(f2571,plain,(
  ( ! [X59,X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,happ(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))) != s(X54,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,X54)),X58),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,happ(s(fun(X50,X46),'_dst_x0x1x0_3'(s(fun(X48,fun(X47,X46)),X49),s(fun(X50,X48),X51),s(X47,X53))),s(X50,X52))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,bool)),'_dst_x00x11x00_3'(s(fun(X54,fun(X54,bool)),'=_0'),s(fun(X56,fun(X57,X54)),X58),s(X54,'const.bool.LET_2'(s(fun(X46,X54),X55),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))),s('type.pair.prod'(X56,X57),'const.bool.LET_2'(s(fun(X46,'type.pair.prod'(X56,X57)),X59),s(X46,happ(s(fun(X47,X46),'const.bool.LET_2'(s(fun(X48,fun(X47,X46)),X49),s(X48,'const.bool.LET_2'(s(fun(X50,X48),X51),s(X50,X52))))),s(X47,X53))))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2854,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'_dst_x00x11x00_3'(s(fun(X33,fun(X33,bool)),'=_0'),s(fun(X24,fun(X25,X33)),X34),s(X33,'const.bool.LET_2'(s(fun(X23,X33),X35),s(X23,X36))))),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'const.bool.LET_2'(s(fun(X29,fun(X23,'type.pair.prod'(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X23,X36)))))) | s(X33,'const.bool.LET_2'(s(fun(X23,X33),X35),s(X23,X36))) = s(X33,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X33)),X34),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X23,X36)))))) )),
  inference(forward_demodulation,[],[f2778,f91])).
fof(f2778,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'_dst_x00x11x00_3'(s(fun(X33,fun(X33,bool)),'=_0'),s(fun(X24,fun(X25,X33)),X34),s(X33,'const.bool.LET_2'(s(fun(X23,X33),X35),s(X23,X36))))),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),happ(s(fun(X29,fun(X23,'type.pair.prod'(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X23,X36)))))) | s(X33,'const.bool.LET_2'(s(fun(X23,X33),X35),s(X23,X36))) = s(X33,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X33)),X34),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X23,X36)))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2633,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,bool)),'_dst_x00x11x00_3'(s(fun(X19,fun(X19,bool)),'=_0'),s(fun(X10,fun(X11,X19)),X22),s(X19,'const.bool.LET_2'(s(fun(X9,X19),X20),s(X9,X21))))),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'const.bool.LET_2'(s(fun(X12,fun(X9,'type.pair.prod'(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X9,X21)))))) | s(X19,'const.bool.LET_2'(s(fun(X9,X19),X20),s(X9,X21))) != s(X19,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,X19)),X22),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X9,X21)))))) )),
  inference(forward_demodulation,[],[f2558,f91])).
fof(f2558,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X19,'const.bool.LET_2'(s(fun(X9,X19),X20),s(X9,X21))) != s(X19,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,X19)),X22),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X9,X21))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,bool)),'_dst_x00x11x00_3'(s(fun(X19,fun(X19,bool)),'=_0'),s(fun(X10,fun(X11,X19)),X22),s(X19,'const.bool.LET_2'(s(fun(X9,X19),X20),s(X9,X21))))),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),happ(s(fun(X12,fun(X9,'type.pair.prod'(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X9,X21))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2865,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,bool)),'_dst_x00x11x00_3'(s(fun(X116,fun(X116,bool)),'=_0'),s(fun(X109,fun(X110,X116)),X117),s(X116,'const.bool.LET_2'(s(fun(X108,X116),X118),s(X108,X119))))),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.bool.LET_2'(s(fun(X112,fun(X108,'type.pair.prod'(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X108,X119)))))) | s(X116,'const.bool.LET_2'(s(fun(X108,X116),X118),s(X108,X119))) = s(X116,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,X116)),X117),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X108,X119)))))) )),
  inference(forward_demodulation,[],[f2785,f91])).
fof(f2785,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,bool)),'_dst_x00x11x00_3'(s(fun(X116,fun(X116,bool)),'=_0'),s(fun(X109,fun(X110,X116)),X117),s(X116,'const.bool.LET_2'(s(fun(X108,X116),X118),s(X108,X119))))),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),happ(s(fun(X112,fun(X108,'type.pair.prod'(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X108,X119)))))) | s(X116,'const.bool.LET_2'(s(fun(X108,X116),X118),s(X108,X119))) = s(X116,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,X116)),X117),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X108,X119)))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2638,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),'_dst_x00x11x00_3'(s(fun(X65,fun(X65,bool)),'=_0'),s(fun(X58,fun(X59,X65)),X68),s(X65,'const.bool.LET_2'(s(fun(X57,X65),X66),s(X57,X67))))),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.bool.LET_2'(s(fun(X60,fun(X57,'type.pair.prod'(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X57,X67)))))) | s(X65,'const.bool.LET_2'(s(fun(X57,X65),X66),s(X57,X67))) != s(X65,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,X65)),X68),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X57,X67)))))) )),
  inference(forward_demodulation,[],[f2562,f91])).
fof(f2562,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X65,'const.bool.LET_2'(s(fun(X57,X65),X66),s(X57,X67))) != s(X65,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,X65)),X68),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X57,X67))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),'_dst_x00x11x00_3'(s(fun(X65,fun(X65,bool)),'=_0'),s(fun(X58,fun(X59,X65)),X68),s(X65,'const.bool.LET_2'(s(fun(X57,X65),X66),s(X57,X67))))),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),happ(s(fun(X60,fun(X57,'type.pair.prod'(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X57,X67))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2860,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))) = s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X80),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X82))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X80),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X82))))))) )),
  inference(forward_demodulation,[],[f2859,f91])).
fof(f2859,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X80),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X82)))))) | s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))) = s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X80),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X82)))))) )),
  inference(forward_demodulation,[],[f2782,f91])).
fof(f2782,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X80),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X82)))))) | s(X79,'const.bool.LET_2'(s(fun(X69,X79),X81),s(X69,X82))) = s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X80),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X82)))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2642,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X96),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X95)))))) | s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))) != s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X96),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X95)))))) )),
  inference(forward_demodulation,[],[f2641,f91])).
fof(f2641,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))) != s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X96),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X95))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X96),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X95))))))) )),
  inference(forward_demodulation,[],[f2564,f91])).
fof(f2564,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))) != s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X96),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X95))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X96),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X94),s(X83,X95))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X95))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2862,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))) = s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X94),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X96))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X94),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X96))))))) )),
  inference(forward_demodulation,[],[f2861,f91])).
fof(f2861,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X94),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X96)))))) | s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))) = s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X94),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X96)))))) )),
  inference(forward_demodulation,[],[f2783,f91])).
fof(f2783,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X84,fun(X85,X93)),X94),s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))))),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X89,fun(X83,'type.pair.prod'(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X83,X96)))))) | s(X93,'const.bool.LET_2'(s(fun(X83,X93),X95),s(X83,X96))) = s(X93,'const.pair.UNCURRY_2'(s(fun(X84,fun(X85,X93)),X94),s('type.pair.prod'(X84,X85),'const.bool.LET_2'(s(fun(X83,'type.pair.prod'(X84,X85)),happ(s(fun(X86,fun(X83,'type.pair.prod'(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,'type.pair.prod'(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X83,X96)))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2640,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X82),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X81)))))) | s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))) != s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X82),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X81)))))) )),
  inference(forward_demodulation,[],[f2639,f91])).
fof(f2639,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))) != s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X82),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X81))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X82),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X81))))))) )),
  inference(forward_demodulation,[],[f2563,f91])).
fof(f2563,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))) != s(X79,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,X79)),X82),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X74,fun(X69,'type.pair.prod'(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X69,X81))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X70,fun(X71,bool)),'_dst_x00x11x00_3'(s(fun(X79,fun(X79,bool)),'=_0'),s(fun(X70,fun(X71,X79)),X82),s(X79,'const.bool.LET_2'(s(fun(X69,X79),X80),s(X69,X81))))),s('type.pair.prod'(X70,X71),'const.bool.LET_2'(s(fun(X69,'type.pair.prod'(X70,X71)),happ(s(fun(X72,fun(X69,'type.pair.prod'(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,'type.pair.prod'(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X69,X81))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2858,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X65,'const.bool.LET_2'(s(fun(X57,X65),X67),s(X57,X68))) = s(X65,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,X65)),X66),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.bool.LET_2'(s(fun(X60,fun(X57,'type.pair.prod'(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X57,X68))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),'_dst_x00x11x00_3'(s(fun(X65,fun(X65,bool)),'=_0'),s(fun(X58,fun(X59,X65)),X66),s(X65,'const.bool.LET_2'(s(fun(X57,X65),X67),s(X57,X68))))),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X57,X68))))))) )),
  inference(forward_demodulation,[],[f2781,f91])).
fof(f2781,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,bool)),'_dst_x00x11x00_3'(s(fun(X65,fun(X65,bool)),'=_0'),s(fun(X58,fun(X59,X65)),X66),s(X65,'const.bool.LET_2'(s(fun(X57,X65),X67),s(X57,X68))))),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X57,X68)))))) | s(X65,'const.bool.LET_2'(s(fun(X57,X65),X67),s(X57,X68))) = s(X65,'const.pair.UNCURRY_2'(s(fun(X58,fun(X59,X65)),X66),s('type.pair.prod'(X58,X59),'const.bool.LET_2'(s(fun(X57,'type.pair.prod'(X58,X59)),happ(s(fun(X60,fun(X57,'type.pair.prod'(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,'type.pair.prod'(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X57,X68)))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2645,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X116,'const.bool.LET_2'(s(fun(X108,X116),X117),s(X108,X118))) != s(X116,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,X116)),X119),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.bool.LET_2'(s(fun(X112,fun(X108,'type.pair.prod'(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X108,X118))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,bool)),'_dst_x00x11x00_3'(s(fun(X116,fun(X116,bool)),'=_0'),s(fun(X109,fun(X110,X116)),X119),s(X116,'const.bool.LET_2'(s(fun(X108,X116),X117),s(X108,X118))))),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X108,X118))))))) )),
  inference(forward_demodulation,[],[f2566,f91])).
fof(f2566,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X116,'const.bool.LET_2'(s(fun(X108,X116),X117),s(X108,X118))) != s(X116,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,X116)),X119),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),happ(s(fun(X112,fun(X108,'type.pair.prod'(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X108,X118))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X109,fun(X110,bool)),'_dst_x00x11x00_3'(s(fun(X116,fun(X116,bool)),'=_0'),s(fun(X109,fun(X110,X116)),X119),s(X116,'const.bool.LET_2'(s(fun(X108,X116),X117),s(X108,X118))))),s('type.pair.prod'(X109,X110),'const.bool.LET_2'(s(fun(X108,'type.pair.prod'(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,'type.pair.prod'(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X108,X118))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2853,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X19,'const.bool.LET_2'(s(fun(X9,X19),X21),s(X9,X22))) = s(X19,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,X19)),X20),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'const.bool.LET_2'(s(fun(X12,fun(X9,'type.pair.prod'(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X9,X22))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,bool)),'_dst_x00x11x00_3'(s(fun(X19,fun(X19,bool)),'=_0'),s(fun(X10,fun(X11,X19)),X20),s(X19,'const.bool.LET_2'(s(fun(X9,X19),X21),s(X9,X22))))),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X9,X22))))))) )),
  inference(forward_demodulation,[],[f2777,f91])).
fof(f2777,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,bool)),'_dst_x00x11x00_3'(s(fun(X19,fun(X19,bool)),'=_0'),s(fun(X10,fun(X11,X19)),X20),s(X19,'const.bool.LET_2'(s(fun(X9,X19),X21),s(X9,X22))))),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X9,X22)))))) | s(X19,'const.bool.LET_2'(s(fun(X9,X19),X21),s(X9,X22))) = s(X19,'const.pair.UNCURRY_2'(s(fun(X10,fun(X11,X19)),X20),s('type.pair.prod'(X10,X11),'const.bool.LET_2'(s(fun(X9,'type.pair.prod'(X10,X11)),happ(s(fun(X12,fun(X9,'type.pair.prod'(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,'type.pair.prod'(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X9,X22)))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2634,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X33,'const.bool.LET_2'(s(fun(X23,X33),X34),s(X23,X35))) != s(X33,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X33)),X36),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'const.bool.LET_2'(s(fun(X29,fun(X23,'type.pair.prod'(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X23,X35))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'_dst_x00x11x00_3'(s(fun(X33,fun(X33,bool)),'=_0'),s(fun(X24,fun(X25,X33)),X36),s(X33,'const.bool.LET_2'(s(fun(X23,X33),X34),s(X23,X35))))),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X23,X35))))))) )),
  inference(forward_demodulation,[],[f2559,f91])).
fof(f2559,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X33,'const.bool.LET_2'(s(fun(X23,X33),X34),s(X23,X35))) != s(X33,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,X33)),X36),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),happ(s(fun(X29,fun(X23,'type.pair.prod'(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X23,X35))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X24,fun(X25,bool)),'_dst_x00x11x00_3'(s(fun(X33,fun(X33,bool)),'=_0'),s(fun(X24,fun(X25,X33)),X36),s(X33,'const.bool.LET_2'(s(fun(X23,X33),X34),s(X23,X35))))),s('type.pair.prod'(X24,X25),'const.bool.LET_2'(s(fun(X23,'type.pair.prod'(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,'type.pair.prod'(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X23,X35))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2575,plain,(
  ( ! [X94,X92,X90,X97,X95,X93,X91,X96] : (s(X91,'const.bool.LET_2'(s(fun(bool,X91),X92),s(bool,'T_0'))) != s(X91,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,X91)),X95),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X93,X94)),X96),s(bool,'T_0'))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X93,fun(X94,bool)),'_dst_x00x11x00_3'(s(fun(X91,fun(X91,bool)),'=_0'),s(fun(X93,fun(X94,X91)),X95),s(X91,'const.bool.LET_2'(s(fun(bool,X91),X92),s(bool,X90))))),s('type.pair.prod'(X93,X94),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X93,X94)),X96),s(bool,X90)))))) | s(bool,X90) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f1661,f156])).
fof(f2576,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X98,X104] : (s(X100,'const.bool.LET_2'(s(fun(bool,X100),X101),s(bool,X99))) != s(X100,'const.pair.UNCURRY_2'(s(fun(X102,fun(X103,X100)),X104),s('type.pair.prod'(X102,X103),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X102,X103)),X105),s(bool,X99))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X102,fun(X103,bool)),'_dst_x00x11x00_3'(s(fun(X100,fun(X100,bool)),'=_0'),s(fun(X102,fun(X103,X100)),X104),s(X100,'const.bool.LET_2'(s(fun(bool,X100),X101),s(bool,X98))))),s('type.pair.prod'(X102,X103),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X102,X103)),X105),s(bool,X98)))))) | s(bool,'T_0') = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f1661,f156])).
fof(f2795,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X98,X104] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X100,fun(X101,bool)),'_dst_x00x11x00_3'(s(fun(X102,fun(X102,bool)),'=_0'),s(fun(X100,fun(X101,X102)),X103),s(X102,'const.bool.LET_2'(s(fun(bool,X102),X104),s(bool,X99))))),s('type.pair.prod'(X100,X101),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X100,X101)),X105),s(bool,X99)))))) | s(X102,'const.bool.LET_2'(s(fun(bool,X102),X104),s(bool,X98))) = s(X102,'const.pair.UNCURRY_2'(s(fun(X100,fun(X101,X102)),X103),s('type.pair.prod'(X100,X101),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X100,X101)),X105),s(bool,X98))))) | s(bool,'T_0') = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f1662,f156])).
fof(f2798,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122,X129] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X123,fun(X124,bool)),'_dst_x00x11x00_3'(s(fun(X125,fun(X125,bool)),'=_0'),s(fun(X123,fun(X124,X125)),X126),s(X125,'const.bool.LET_2'(s(fun(bool,X125),X127),s(bool,X122))))),s('type.pair.prod'(X123,X124),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X123,X124)),X128),s(bool,X122)))))) | s(X125,'const.bool.LET_2'(s(fun(bool,X125),X127),s(bool,'T_0'))) = s(X125,'const.pair.UNCURRY_2'(s(fun(X123,fun(X124,X125)),X126),s('type.pair.prod'(X123,X124),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X123,X124)),X128),s(bool,'T_0'))))) | s(bool,X122) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f1662,f156])).
fof(f2579,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122,X129] : (s(X123,'const.bool.LET_2'(s(fun(bool,X123),X124),s(bool,X122))) != s(X123,'const.pair.UNCURRY_2'(s(fun(X125,fun(X126,X123)),X127),s('type.pair.prod'(X125,X126),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X125,X126)),X128),s(bool,X122))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X125,fun(X126,bool)),'_dst_x00x11x00_3'(s(fun(X123,fun(X123,bool)),'=_0'),s(fun(X125,fun(X126,X123)),X127),s(X123,'const.bool.LET_2'(s(fun(bool,X123),X124),s(bool,'T_0'))))),s('type.pair.prod'(X125,X126),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X125,X126)),X128),s(bool,'T_0')))))) | s(bool,X122) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f1661,f156])).
fof(f2794,plain,(
  ( ! [X94,X92,X90,X97,X95,X93,X91,X96] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X91,fun(X92,bool)),'_dst_x00x11x00_3'(s(fun(X93,fun(X93,bool)),'=_0'),s(fun(X91,fun(X92,X93)),X94),s(X93,'const.bool.LET_2'(s(fun(bool,X93),X95),s(bool,'T_0'))))),s('type.pair.prod'(X91,X92),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X91,X92)),X96),s(bool,'T_0')))))) | s(X93,'const.bool.LET_2'(s(fun(bool,X93),X95),s(bool,X90))) = s(X93,'const.pair.UNCURRY_2'(s(fun(X91,fun(X92,X93)),X94),s('type.pair.prod'(X91,X92),'const.bool.LET_2'(s(fun(bool,'type.pair.prod'(X91,X92)),X96),s(bool,X90))))) | s(bool,X90) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f1662,f156])).
fof(f2882,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X196,X182,X190,X188,X186,X195,X193,X184,X183] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X191,fun(X192,bool)),'_dst_x00x11x00_3'(s(fun(X193,fun(X193,bool)),'=_0'),s(fun(X191,fun(X192,X193)),X194),s(X193,'const.bool.LET_2'(s(fun(fun(X182,X183),X193),X195),s(fun(X182,X183),'const.bool.LET_2'(s(fun(X187,fun(X182,X183)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X187,fun(X182,X184)),X188),s(X185,X190))),s(X187,X189))))))),s('type.pair.prod'(X191,X192),'const.bool.LET_2'(s(fun(fun(X182,X183),'type.pair.prod'(X191,X192)),X196),s(fun(X182,X183),'const.bool.LET_2'(s(fun(X187,fun(X182,X183)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X187,fun(X182,X184)),X188),s(X185,X190))),s(X187,X189)))))))) | s(X193,'const.bool.LET_2'(s(fun(fun(X182,X183),X193),X195),s(fun(X182,X183),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X182,X184),'const.bool.LET_2'(s(fun(X187,fun(X182,X184)),X188),s(X187,X189))),s(X185,X190))))) = s(X193,'const.pair.UNCURRY_2'(s(fun(X191,fun(X192,X193)),X194),s('type.pair.prod'(X191,X192),'const.bool.LET_2'(s(fun(fun(X182,X183),'type.pair.prod'(X191,X192)),X196),s(fun(X182,X183),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X182,X184),'const.bool.LET_2'(s(fun(X187,fun(X182,X184)),X188),s(X187,X189))),s(X185,X190)))))))) )),
  inference(forward_demodulation,[],[f2805,f91])).
fof(f2805,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X196,X182,X190,X188,X186,X195,X193,X184,X183] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X191,fun(X192,bool)),'_dst_x00x11x00_3'(s(fun(X193,fun(X193,bool)),'=_0'),s(fun(X191,fun(X192,X193)),X194),s(X193,'const.bool.LET_2'(s(fun(fun(X182,X183),X193),X195),s(fun(X182,X183),happ(s(fun(X187,fun(X182,X183)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X187,fun(X182,X184)),X188),s(X185,X190))),s(X187,X189))))))),s('type.pair.prod'(X191,X192),'const.bool.LET_2'(s(fun(fun(X182,X183),'type.pair.prod'(X191,X192)),X196),s(fun(X182,X183),happ(s(fun(X187,fun(X182,X183)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X187,fun(X182,X184)),X188),s(X185,X190))),s(X187,X189)))))))) | s(X193,'const.bool.LET_2'(s(fun(fun(X182,X183),X193),X195),s(fun(X182,X183),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X182,X184),'const.bool.LET_2'(s(fun(X187,fun(X182,X184)),X188),s(X187,X189))),s(X185,X190))))) = s(X193,'const.pair.UNCURRY_2'(s(fun(X191,fun(X192,X193)),X194),s('type.pair.prod'(X191,X192),'const.bool.LET_2'(s(fun(fun(X182,X183),'type.pair.prod'(X191,X192)),X196),s(fun(X182,X183),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X183)),X186),s(fun(X182,X184),'const.bool.LET_2'(s(fun(X187,fun(X182,X184)),X188),s(X187,X189))),s(X185,X190)))))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2661,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X182,X190,X188,X186,X184,X193,X195,X183,X181] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X192,fun(X193,bool)),'_dst_x00x11x00_3'(s(fun(X190,fun(X190,bool)),'=_0'),s(fun(X192,fun(X193,X190)),X194),s(X190,'const.bool.LET_2'(s(fun(fun(X181,X182),X190),X191),s(fun(X181,X182),'const.bool.LET_2'(s(fun(X183,fun(X181,X182)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X183,fun(X181,X184)),X187),s(X185,X188))),s(X183,X189))))))),s('type.pair.prod'(X192,X193),'const.bool.LET_2'(s(fun(fun(X181,X182),'type.pair.prod'(X192,X193)),X195),s(fun(X181,X182),'const.bool.LET_2'(s(fun(X183,fun(X181,X182)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X183,fun(X181,X184)),X187),s(X185,X188))),s(X183,X189)))))))) | s(X190,'const.bool.LET_2'(s(fun(fun(X181,X182),X190),X191),s(fun(X181,X182),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X181,X184),'const.bool.LET_2'(s(fun(X183,fun(X181,X184)),X187),s(X183,X189))),s(X185,X188))))) != s(X190,'const.pair.UNCURRY_2'(s(fun(X192,fun(X193,X190)),X194),s('type.pair.prod'(X192,X193),'const.bool.LET_2'(s(fun(fun(X181,X182),'type.pair.prod'(X192,X193)),X195),s(fun(X181,X182),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X181,X184),'const.bool.LET_2'(s(fun(X183,fun(X181,X184)),X187),s(X183,X189))),s(X185,X188)))))))) )),
  inference(forward_demodulation,[],[f2587,f91])).
fof(f2587,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X182,X190,X188,X186,X184,X193,X195,X183,X181] : (s(X190,'const.bool.LET_2'(s(fun(fun(X181,X182),X190),X191),s(fun(X181,X182),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X181,X184),'const.bool.LET_2'(s(fun(X183,fun(X181,X184)),X187),s(X183,X189))),s(X185,X188))))) != s(X190,'const.pair.UNCURRY_2'(s(fun(X192,fun(X193,X190)),X194),s('type.pair.prod'(X192,X193),'const.bool.LET_2'(s(fun(fun(X181,X182),'type.pair.prod'(X192,X193)),X195),s(fun(X181,X182),'_dst_x0x1x0_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X181,X184),'const.bool.LET_2'(s(fun(X183,fun(X181,X184)),X187),s(X183,X189))),s(X185,X188))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X192,fun(X193,bool)),'_dst_x00x11x00_3'(s(fun(X190,fun(X190,bool)),'=_0'),s(fun(X192,fun(X193,X190)),X194),s(X190,'const.bool.LET_2'(s(fun(fun(X181,X182),X190),X191),s(fun(X181,X182),happ(s(fun(X183,fun(X181,X182)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X183,fun(X181,X184)),X187),s(X185,X188))),s(X183,X189))))))),s('type.pair.prod'(X192,X193),'const.bool.LET_2'(s(fun(fun(X181,X182),'type.pair.prod'(X192,X193)),X195),s(fun(X181,X182),happ(s(fun(X183,fun(X181,X182)),'_dst_x00x11x00_3'(s(fun(X184,fun(X185,X182)),X186),s(fun(X183,fun(X181,X184)),X187),s(X185,X188))),s(X183,X189))))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2881,plain,(
  ( ! [X177,X167,X175,X173,X171,X169,X180,X178,X176,X174,X172,X170,X168,X181,X179] : (s(X178,'const.bool.LET_2'(s(fun(fun(X167,X168),X178),X180),s(fun(X167,X168),'const.bool.LET_2'(s(fun(X169,fun(X167,X168)),'_dst_x00x11x00_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X169,fun(X167,X170)),X173),s(X171,X174))),s(X169,X175))))) = s(X178,'const.pair.UNCURRY_2'(s(fun(X176,fun(X177,X178)),X179),s('type.pair.prod'(X176,X177),'const.bool.LET_2'(s(fun(fun(X167,X168),'type.pair.prod'(X176,X177)),X181),s(fun(X167,X168),'const.bool.LET_2'(s(fun(X169,fun(X167,X168)),'_dst_x00x11x00_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X169,fun(X167,X170)),X173),s(X171,X174))),s(X169,X175))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X176,fun(X177,bool)),'_dst_x00x11x00_3'(s(fun(X178,fun(X178,bool)),'=_0'),s(fun(X176,fun(X177,X178)),X179),s(X178,'const.bool.LET_2'(s(fun(fun(X167,X168),X178),X180),s(fun(X167,X168),'_dst_x0x1x0_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X167,X170),'const.bool.LET_2'(s(fun(X169,fun(X167,X170)),X173),s(X169,X175))),s(X171,X174))))))),s('type.pair.prod'(X176,X177),'const.bool.LET_2'(s(fun(fun(X167,X168),'type.pair.prod'(X176,X177)),X181),s(fun(X167,X168),'_dst_x0x1x0_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X167,X170),'const.bool.LET_2'(s(fun(X169,fun(X167,X170)),X173),s(X169,X175))),s(X171,X174))))))))) )),
  inference(forward_demodulation,[],[f2804,f91])).
fof(f2804,plain,(
  ( ! [X177,X167,X175,X173,X171,X169,X180,X178,X176,X174,X172,X170,X168,X181,X179] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X176,fun(X177,bool)),'_dst_x00x11x00_3'(s(fun(X178,fun(X178,bool)),'=_0'),s(fun(X176,fun(X177,X178)),X179),s(X178,'const.bool.LET_2'(s(fun(fun(X167,X168),X178),X180),s(fun(X167,X168),'_dst_x0x1x0_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X167,X170),'const.bool.LET_2'(s(fun(X169,fun(X167,X170)),X173),s(X169,X175))),s(X171,X174))))))),s('type.pair.prod'(X176,X177),'const.bool.LET_2'(s(fun(fun(X167,X168),'type.pair.prod'(X176,X177)),X181),s(fun(X167,X168),'_dst_x0x1x0_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X167,X170),'const.bool.LET_2'(s(fun(X169,fun(X167,X170)),X173),s(X169,X175))),s(X171,X174)))))))) | s(X178,'const.bool.LET_2'(s(fun(fun(X167,X168),X178),X180),s(fun(X167,X168),happ(s(fun(X169,fun(X167,X168)),'_dst_x00x11x00_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X169,fun(X167,X170)),X173),s(X171,X174))),s(X169,X175))))) = s(X178,'const.pair.UNCURRY_2'(s(fun(X176,fun(X177,X178)),X179),s('type.pair.prod'(X176,X177),'const.bool.LET_2'(s(fun(fun(X167,X168),'type.pair.prod'(X176,X177)),X181),s(fun(X167,X168),happ(s(fun(X169,fun(X167,X168)),'_dst_x00x11x00_3'(s(fun(X170,fun(X171,X168)),X172),s(fun(X169,fun(X167,X170)),X173),s(X171,X174))),s(X169,X175)))))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2662,plain,(
  ( ! [X210,X208,X198,X196,X206,X204,X202,X200,X209,X199,X197,X207,X205,X203,X201] : (s(X205,'const.bool.LET_2'(s(fun(fun(X196,X197),X205),X206),s(fun(X196,X197),'const.bool.LET_2'(s(fun(X201,fun(X196,X197)),'_dst_x00x11x00_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X201,fun(X196,X198)),X202),s(X199,X204))),s(X201,X203))))) != s(X205,'const.pair.UNCURRY_2'(s(fun(X207,fun(X208,X205)),X209),s('type.pair.prod'(X207,X208),'const.bool.LET_2'(s(fun(fun(X196,X197),'type.pair.prod'(X207,X208)),X210),s(fun(X196,X197),'const.bool.LET_2'(s(fun(X201,fun(X196,X197)),'_dst_x00x11x00_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X201,fun(X196,X198)),X202),s(X199,X204))),s(X201,X203))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X207,fun(X208,bool)),'_dst_x00x11x00_3'(s(fun(X205,fun(X205,bool)),'=_0'),s(fun(X207,fun(X208,X205)),X209),s(X205,'const.bool.LET_2'(s(fun(fun(X196,X197),X205),X206),s(fun(X196,X197),'_dst_x0x1x0_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X196,X198),'const.bool.LET_2'(s(fun(X201,fun(X196,X198)),X202),s(X201,X203))),s(X199,X204))))))),s('type.pair.prod'(X207,X208),'const.bool.LET_2'(s(fun(fun(X196,X197),'type.pair.prod'(X207,X208)),X210),s(fun(X196,X197),'_dst_x0x1x0_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X196,X198),'const.bool.LET_2'(s(fun(X201,fun(X196,X198)),X202),s(X201,X203))),s(X199,X204))))))))) )),
  inference(forward_demodulation,[],[f2588,f91])).
fof(f2588,plain,(
  ( ! [X210,X208,X198,X196,X206,X204,X202,X200,X209,X199,X197,X207,X205,X203,X201] : (s(X205,'const.bool.LET_2'(s(fun(fun(X196,X197),X205),X206),s(fun(X196,X197),happ(s(fun(X201,fun(X196,X197)),'_dst_x00x11x00_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X201,fun(X196,X198)),X202),s(X199,X204))),s(X201,X203))))) != s(X205,'const.pair.UNCURRY_2'(s(fun(X207,fun(X208,X205)),X209),s('type.pair.prod'(X207,X208),'const.bool.LET_2'(s(fun(fun(X196,X197),'type.pair.prod'(X207,X208)),X210),s(fun(X196,X197),happ(s(fun(X201,fun(X196,X197)),'_dst_x00x11x00_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X201,fun(X196,X198)),X202),s(X199,X204))),s(X201,X203))))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X207,fun(X208,bool)),'_dst_x00x11x00_3'(s(fun(X205,fun(X205,bool)),'=_0'),s(fun(X207,fun(X208,X205)),X209),s(X205,'const.bool.LET_2'(s(fun(fun(X196,X197),X205),X206),s(fun(X196,X197),'_dst_x0x1x0_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X196,X198),'const.bool.LET_2'(s(fun(X201,fun(X196,X198)),X202),s(X201,X203))),s(X199,X204))))))),s('type.pair.prod'(X207,X208),'const.bool.LET_2'(s(fun(fun(X196,X197),'type.pair.prod'(X207,X208)),X210),s(fun(X196,X197),'_dst_x0x1x0_3'(s(fun(X198,fun(X199,X197)),X200),s(fun(X196,X198),'const.bool.LET_2'(s(fun(X201,fun(X196,X198)),X202),s(X201,X203))),s(X199,X204))))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2832,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X7))) = s(X2,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,X2)),X6),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X4,fun(X5,X2)),X6),s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))))) )),
  inference(forward_demodulation,[],[f2831,f1624])).
fof(f2831,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X7))) = s(X2,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,X2)),X6),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X4,fun(X5,X2)),X6),s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))))) )),
  inference(forward_demodulation,[],[f2830,f1375])).
fof(f2830,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,X2)),X6),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X7))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X4,fun(X5,X2)),X6),s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))))) )),
  inference(forward_demodulation,[],[f2829,f91])).
fof(f2829,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X4,fun(X5,X2)),X6),s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7)))))) | s(X2,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,X2)),X6),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X7)))) )),
  inference(forward_demodulation,[],[f2758,f1624])).
fof(f2758,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X4,fun(X5,X2)),X6),s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7)))))) | s(X2,'const.pair.UNCURRY_2'(s(fun(X4,fun(X5,X2)),X6),s('type.pair.prod'(X4,X5),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X4,X5)),X8),s('type.pair.prod'(X0,X1),X7))))) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X7)))) )),
  inference(superposition,[],[f1662,f102])).
fof(f2600,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X5,fun(X6,X2)),X7),s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))))),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) | s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X2)),X7),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) )),
  inference(forward_demodulation,[],[f2599,f1624])).
fof(f2599,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X5,fun(X6,X2)),X7),s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X4))))),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) | s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X2)),X7),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) )),
  inference(forward_demodulation,[],[f2598,f1375])).
fof(f2598,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X5,fun(X6,X2)),X7),s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X4))))),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) | s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X2)),X7),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4)))))) )),
  inference(forward_demodulation,[],[f2597,f91])).
fof(f2597,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,X2)),X3),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X2)),X7),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X5,fun(X6,X2)),X7),s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X4))))),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4))))))) )),
  inference(forward_demodulation,[],[f2530,f1624])).
fof(f2530,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X4))) != s(X2,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,X2)),X7),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X5,fun(X6,bool)),'_dst_x00x11x00_3'(s(fun(X2,fun(X2,bool)),'=_0'),s(fun(X5,fun(X6,X2)),X7),s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X4))))),s('type.pair.prod'(X5,X6),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X5,X6)),X8),s('type.pair.prod'(X0,X1),X4))))))) )),
  inference(superposition,[],[f1661,f102])).
fof(f2817,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,bool)),'_dst_x00x11x00_3'(s(fun(X25,fun(X25,bool)),'=_0'),s(fun(X23,fun(X24,X25)),'const.bool.LET_2'(s(fun(X29,fun(X23,fun(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35)))))) | s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))) = s(X25,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35)))))) )),
  inference(forward_demodulation,[],[f2750,f91])).
fof(f2750,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,bool)),'_dst_x00x11x00_3'(s(fun(X25,fun(X25,bool)),'=_0'),s(fun(X23,fun(X24,X25)),happ(s(fun(X29,fun(X23,fun(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35)))))) | s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))) = s(X25,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35)))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2616,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,bool)),'_dst_x00x11x00_3'(s(fun(X11,fun(X11,bool)),'=_0'),s(fun(X9,fun(X10,X11)),'const.bool.LET_2'(s(fun(X12,fun(X9,fun(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21)))))) | s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))) != s(X11,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21)))))) )),
  inference(forward_demodulation,[],[f2548,f91])).
fof(f2548,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))) != s(X11,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,bool)),'_dst_x00x11x00_3'(s(fun(X11,fun(X11,bool)),'=_0'),s(fun(X9,fun(X10,X11)),happ(s(fun(X12,fun(X9,fun(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2825,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) = s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))))) )),
  inference(forward_demodulation,[],[f2824,f91])).
fof(f2824,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) | s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) = s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),happ(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) )),
  inference(forward_demodulation,[],[f2755,f91])).
fof(f2755,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),happ(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) | s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) = s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),happ(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2623,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) | s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) != s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) )),
  inference(forward_demodulation,[],[f2622,f91])).
fof(f2622,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) != s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),happ(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))))) )),
  inference(forward_demodulation,[],[f2553,f91])).
fof(f2553,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) != s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),happ(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),happ(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2823,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) = s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81))))))) )),
  inference(forward_demodulation,[],[f2822,f91])).
fof(f2822,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),'const.bool.LET_2'(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) | s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) = s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),happ(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) )),
  inference(forward_demodulation,[],[f2754,f91])).
fof(f2754,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X69,X79,X77,X75,X73,X82] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,bool)),'_dst_x00x11x00_3'(s(fun(X71,fun(X71,bool)),'=_0'),s(fun(X69,fun(X70,X71)),happ(s(fun(X74,fun(X69,fun(X70,X71))),'const.bool.LET_2'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(X73,'const.bool.LET_2'(s(fun(X72,X73),X76),s(X72,X78))))),s(X74,X77))),s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) | s(X71,'const.bool.LET_2'(s(fun(X79,X71),X80),s(X79,X81))) = s(X71,'const.pair.UNCURRY_2'(s(fun(X69,fun(X70,X71)),happ(s(fun(X72,fun(X69,fun(X70,X71))),'_dst_x0x1x0_3'(s(fun(X73,fun(X74,fun(X69,fun(X70,X71)))),X75),s(fun(X72,X73),X76),s(X74,X77))),s(X72,X78))),s('type.pair.prod'(X69,X70),'const.bool.LET_2'(s(fun(X79,'type.pair.prod'(X69,X70)),X82),s(X79,X81)))))) )),
  inference(superposition,[],[f1662,f145])).
fof(f2625,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) | s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) != s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95)))))) )),
  inference(forward_demodulation,[],[f2624,f91])).
fof(f2624,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) != s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),'const.bool.LET_2'(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),happ(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))))) )),
  inference(forward_demodulation,[],[f2554,f91])).
fof(f2554,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X96,X86,X84] : (s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))) != s(X85,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,X85)),happ(s(fun(X89,fun(X83,fun(X84,X85))),'_dst_x0x1x0_3'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(fun(X89,X87),X90),s(X86,X92))),s(X89,X91))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X83,fun(X84,bool)),'_dst_x00x11x00_3'(s(fun(X85,fun(X85,bool)),'=_0'),s(fun(X83,fun(X84,X85)),happ(s(fun(X86,fun(X83,fun(X84,X85))),'const.bool.LET_2'(s(fun(X87,fun(X86,fun(X83,fun(X84,X85)))),X88),s(X87,'const.bool.LET_2'(s(fun(X89,X87),X90),s(X89,X91))))),s(X86,X92))),s(X85,'const.bool.LET_2'(s(fun(X93,X85),X94),s(X93,X95))))),s('type.pair.prod'(X83,X84),'const.bool.LET_2'(s(fun(X93,'type.pair.prod'(X83,X84)),X96),s(X93,X95))))))) )),
  inference(superposition,[],[f1661,f145])).
fof(f2828,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,bool)),'_dst_x00x11x00_3'(s(fun(X110,fun(X110,bool)),'=_0'),s(fun(X108,fun(X109,X110)),'const.bool.LET_2'(s(fun(X112,fun(X108,fun(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118)))))) | s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))) = s(X110,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118)))))) )),
  inference(forward_demodulation,[],[f2757,f91])).
fof(f2757,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,bool)),'_dst_x00x11x00_3'(s(fun(X110,fun(X110,bool)),'=_0'),s(fun(X108,fun(X109,X110)),happ(s(fun(X112,fun(X108,fun(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118)))))) | s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))) = s(X110,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118)))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2621,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,bool)),'_dst_x00x11x00_3'(s(fun(X59,fun(X59,bool)),'=_0'),s(fun(X57,fun(X58,X59)),'const.bool.LET_2'(s(fun(X60,fun(X57,fun(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67)))))) | s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))) != s(X59,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67)))))) )),
  inference(forward_demodulation,[],[f2552,f91])).
fof(f2552,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))) != s(X59,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,bool)),'_dst_x00x11x00_3'(s(fun(X59,fun(X59,bool)),'=_0'),s(fun(X57,fun(X58,X59)),happ(s(fun(X60,fun(X57,fun(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2821,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))) = s(X59,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,X59)),'const.bool.LET_2'(s(fun(X60,fun(X57,fun(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,bool)),'_dst_x00x11x00_3'(s(fun(X59,fun(X59,bool)),'=_0'),s(fun(X57,fun(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67))))))) )),
  inference(forward_demodulation,[],[f2753,f91])).
fof(f2753,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X67,X65,X63] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,bool)),'_dst_x00x11x00_3'(s(fun(X59,fun(X59,bool)),'=_0'),s(fun(X57,fun(X58,X59)),'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))))),s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67)))))) | s(X59,'const.bool.LET_2'(s(fun(X65,X59),X66),s(X65,X67))) = s(X59,'const.pair.UNCURRY_2'(s(fun(X57,fun(X58,X59)),happ(s(fun(X60,fun(X57,fun(X58,X59))),'const.bool.LET_2'(s(fun(X61,fun(X60,fun(X57,fun(X58,X59)))),X62),s(X61,X63))),s(X60,X64))),s('type.pair.prod'(X57,X58),'const.bool.LET_2'(s(fun(X65,'type.pair.prod'(X57,X58)),X68),s(X65,X67)))))) )),
  inference(superposition,[],[f1662,f143])).
fof(f2628,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))) != s(X110,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,X110)),'const.bool.LET_2'(s(fun(X112,fun(X108,fun(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,bool)),'_dst_x00x11x00_3'(s(fun(X110,fun(X110,bool)),'=_0'),s(fun(X108,fun(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118))))))) )),
  inference(forward_demodulation,[],[f2556,f91])).
fof(f2556,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))) != s(X110,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,X110)),happ(s(fun(X112,fun(X108,fun(X109,X110))),'const.bool.LET_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s(X111,X114))),s(X112,X115))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X108,fun(X109,bool)),'_dst_x00x11x00_3'(s(fun(X110,fun(X110,bool)),'=_0'),s(fun(X108,fun(X109,X110)),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X108,fun(X109,X110)))),X113),s('type.pair.prod'(X111,X112),'const.pair.,_2'(s(X111,X114),s(X112,X115))))),s(X110,'const.bool.LET_2'(s(fun(X116,X110),X117),s(X116,X118))))),s('type.pair.prod'(X108,X109),'const.bool.LET_2'(s(fun(X116,'type.pair.prod'(X108,X109)),X119),s(X116,X118))))))) )),
  inference(superposition,[],[f1661,f143])).
fof(f2816,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,X11)),'const.bool.LET_2'(s(fun(X12,fun(X9,fun(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,bool)),'_dst_x00x11x00_3'(s(fun(X11,fun(X11,bool)),'=_0'),s(fun(X9,fun(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21))))))) )),
  inference(forward_demodulation,[],[f2749,f91])).
fof(f2749,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X9,X22,X20,X18,X16] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,bool)),'_dst_x00x11x00_3'(s(fun(X11,fun(X11,bool)),'=_0'),s(fun(X9,fun(X10,X11)),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X9,X13),'const.bool.LET_2'(s(fun(X12,fun(X9,X13)),X16),s(X12,X18))),s(X14,X17))),s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21)))))) | s(X11,'const.bool.LET_2'(s(fun(X19,X11),X20),s(X19,X21))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X9,fun(X10,X11)),happ(s(fun(X12,fun(X9,fun(X10,X11))),'_dst_x00x11x00_3'(s(fun(X13,fun(X14,fun(X10,X11))),X15),s(fun(X12,fun(X9,X13)),X16),s(X14,X17))),s(X12,X18))),s('type.pair.prod'(X9,X10),'const.bool.LET_2'(s(fun(X19,'type.pair.prod'(X9,X10)),X22),s(X19,X21)))))) )),
  inference(superposition,[],[f1662,f154])).
fof(f2617,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))) != s(X25,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,X25)),'const.bool.LET_2'(s(fun(X29,fun(X23,fun(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,bool)),'_dst_x00x11x00_3'(s(fun(X25,fun(X25,bool)),'=_0'),s(fun(X23,fun(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35))))))) )),
  inference(forward_demodulation,[],[f2549,f91])).
fof(f2549,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X23,X31,X29,X27,X25,X36,X34,X32] : (s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))) != s(X25,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,X25)),happ(s(fun(X29,fun(X23,fun(X24,X25))),'_dst_x00x11x00_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X29,fun(X23,X26)),X30),s(X27,X32))),s(X29,X31))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,bool)),'_dst_x00x11x00_3'(s(fun(X25,fun(X25,bool)),'=_0'),s(fun(X23,fun(X24,X25)),'_dst_x0x1x0_3'(s(fun(X26,fun(X27,fun(X24,X25))),X28),s(fun(X23,X26),'const.bool.LET_2'(s(fun(X29,fun(X23,X26)),X30),s(X29,X31))),s(X27,X32))),s(X25,'const.bool.LET_2'(s(fun(X33,X25),X34),s(X33,X35))))),s('type.pair.prod'(X23,X24),'const.bool.LET_2'(s(fun(X33,'type.pair.prod'(X23,X24)),X36),s(X33,X35))))))) )),
  inference(superposition,[],[f1661,f154])).
fof(f2852,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))) = s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X8))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X6),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))))),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X8))))))) )),
  inference(forward_demodulation,[],[f2851,f1624])).
fof(f2851,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))) = s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X8))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X6),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))))),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X8))))))) )),
  inference(forward_demodulation,[],[f2850,f1375])).
fof(f2850,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))) = s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X8))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X6),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))))),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X8))))))) )),
  inference(forward_demodulation,[],[f2849,f91])).
fof(f2849,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X6),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))))),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X8)))))) | s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))) = s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),happ(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X8)))))) )),
  inference(forward_demodulation,[],[f2776,f1624])).
fof(f2776,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X6),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X8)))))) | s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X7),s('type.pair.prod'(X0,X1),X8))) = s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),happ(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X8)))))) )),
  inference(superposition,[],[f1662,f102])).
fof(f2632,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X8),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X7)))))) | s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))) != s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X8),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X7)))))) )),
  inference(forward_demodulation,[],[f2631,f1624])).
fof(f2631,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X8),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X7)))))) | s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))) != s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X8),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X7)))))) )),
  inference(forward_demodulation,[],[f2630,f1375])).
fof(f2630,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X8),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X7)))))) | s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))) != s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X8),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X7)))))) )),
  inference(forward_demodulation,[],[f2629,f91])).
fof(f2629,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))) != s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X8),s('type.pair.prod'(X2,X3),'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4),s('type.pair.prod'(X0,X1),X7))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X8),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),happ(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X7))))))) )),
  inference(forward_demodulation,[],[f2557,f1624])).
fof(f2557,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))) != s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X8),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X7))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X5,fun(X5,bool)),'=_0'),s(fun(X2,fun(X3,X5)),X8),s(X5,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X5),X6),s('type.pair.prod'(X0,X1),X7))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3)),happ(s(fun(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),fun('type.pair.prod'(X0,X1),'type.pair.prod'(X2,X3))),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,'type.pair.prod'(X2,X3))),X4))),s('type.pair.prod'(X0,X1),X7))))))) )),
  inference(superposition,[],[f1661,f102])).
fof(f1626,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X14,fun(X15,bool)),X16),s('type.pair.prod'(X14,X15),X17))) | s(bool,X18) = s(bool,happ(s(fun('type.pair.prod'(X14,X15),bool),'const.pair.UNCURRY_1'(s(fun(X14,fun(X15,bool)),X16))),s('type.pair.prod'(X14,X15),X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1389,f156])).
fof(f1627,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X23) = s(bool,'const.pair.UNCURRY_2'(s(fun(X19,fun(X20,bool)),X21),s('type.pair.prod'(X19,X20),X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X19,X20),bool),'const.pair.UNCURRY_1'(s(fun(X19,fun(X20,bool)),X21))),s('type.pair.prod'(X19,X20),X22)))) )),
  inference(superposition,[],[f1389,f156])).
fof(f271,plain,(
  ( ! [X28,X29,X27] : (s('type.pair.prod'(bool,X28),X27) = s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,sK9(X27,X28,bool)))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,sK8(X27,X28,bool))) )),
  inference(superposition,[],[f90,f156])).
fof(f307,plain,(
  ( ! [X28,X26,X27] : (s('type.pair.prod'(bool,X27),X26) = s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,'T_0'),s(X27,sK9(X26,X27,bool)))) | s(bool,X28) = s(bool,sK8(X26,X27,bool)) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f90,f156])).
fof(f272,plain,(
  ( ! [X30,X31,X32] : (s('type.pair.prod'(X31,bool),X30) = s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,sK8(X30,bool,X31)),s(bool,X32))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,sK9(X30,bool,X31))) )),
  inference(superposition,[],[f90,f156])).
fof(f308,plain,(
  ( ! [X30,X31,X29] : (s('type.pair.prod'(X30,bool),X29) = s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,sK8(X29,bool,X30)),s(bool,'T_0'))) | s(bool,X31) = s(bool,sK9(X29,bool,X30)) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f90,f156])).
fof(f8731,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl59),
  inference(trivial_inequality_removal,[],[f8726])).
fof(f8726,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl59),
  inference(superposition,[],[f8463,f156])).
fof(f8713,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl59),
  inference(superposition,[],[f8463,f81])).
fof(f8708,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,f(s(bool,'F_0')))))) | $spl59),
  inference(unit_resulting_resolution,[],[f8463,f141])).
fof(f8711,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl59),
  inference(unit_resulting_resolution,[],[f78,f8463,f89])).
fof(f8712,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,f(s(bool,'F_0'))))),s(bool,'T_0')))) | $spl59),
  inference(unit_resulting_resolution,[],[f8463,f141])).
fof(f7951,plain,(
  ( ! [X134,X132,X130,X128,X136,X135,X133,X131,X129,X137] : (s(fun(X128,X129),X137) != s(fun(X128,X129),happ(s(fun(X133,fun(X128,X129)),'_dst_x00x11x00_3'(s(fun(X130,fun(X131,X129)),X132),s(fun(X133,fun(X128,X130)),X134),s(X131,X136))),s(X133,X135))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X128,X129),bool),'const.bool.LET_2'(s(fun(fun(X128,X129),fun(fun(X128,X129),bool)),'=_0'),s(fun(X128,X129),X137))),s(fun(X128,X129),'_dst_x0x1x0_3'(s(fun(X130,fun(X131,X129)),X132),s(fun(X128,X130),'const.bool.LET_2'(s(fun(X133,fun(X128,X130)),X134),s(X133,X135))),s(X131,X136)))))) )),
  inference(superposition,[],[f574,f154])).
fof(f8096,plain,(
  ( ! [X127,X125,X123,X121,X118,X126,X124,X122,X120,X119] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X118,X119),bool),'const.bool.LET_2'(s(fun(fun(X118,X119),fun(fun(X118,X119),bool)),'=_0'),s(fun(X118,X119),X127))),s(fun(X118,X119),'const.bool.LET_2'(s(fun(X120,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X120,fun(X118,X121)),X124),s(X122,X125))),s(X120,X126))))) | s(fun(X118,X119),X127) != s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X118,X121),'const.bool.LET_2'(s(fun(X120,fun(X118,X121)),X124),s(X120,X126))),s(X122,X125)))) )),
  inference(forward_demodulation,[],[f7950,f91])).
fof(f7950,plain,(
  ( ! [X127,X125,X123,X121,X118,X126,X124,X122,X120,X119] : (s(fun(X118,X119),X127) != s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X118,X121),'const.bool.LET_2'(s(fun(X120,fun(X118,X121)),X124),s(X120,X126))),s(X122,X125))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X118,X119),bool),'const.bool.LET_2'(s(fun(fun(X118,X119),fun(fun(X118,X119),bool)),'=_0'),s(fun(X118,X119),X127))),s(fun(X118,X119),happ(s(fun(X120,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X120,fun(X118,X121)),X124),s(X122,X125))),s(X120,X126)))))) )),
  inference(superposition,[],[f574,f154])).
fof(f7936,plain,(
  ( ! [X85,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X85))),s(bool,'T_0'))) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f574,f156])).
fof(f7929,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X65))),s(bool,X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f574,f156])).
fof(f7928,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X61))),s(bool,X60))) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f574,f156])).
fof(f8084,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(X51,X57) != s(X51,'const.bool.LET_2'(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,X57))),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))) )),
  inference(forward_demodulation,[],[f7926,f91])).
fof(f7926,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(X51,X57) != s(X51,happ(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,X57))),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))) )),
  inference(superposition,[],[f574,f143])).
fof(f8080,plain,(
  ( ! [X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X31,bool),'const.bool.LET_2'(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X39))),s(X31,'const.bool.LET_2'(s(fun(X32,X31),'const.bool.LET_2'(s(fun(X33,fun(X32,X31)),X34),s(X33,'const.bool.LET_2'(s(fun(X35,X33),X36),s(X35,X37))))),s(X32,X38))))) | s(X31,X39) != s(X31,happ(s(fun(X35,X31),'_dst_x0x1x0_3'(s(fun(X33,fun(X32,X31)),X34),s(fun(X35,X33),X36),s(X32,X38))),s(X35,X37)))) )),
  inference(forward_demodulation,[],[f7923,f91])).
fof(f7923,plain,(
  ( ! [X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(X31,X39) != s(X31,happ(s(fun(X35,X31),'_dst_x0x1x0_3'(s(fun(X33,fun(X32,X31)),X34),s(fun(X35,X33),X36),s(X32,X38))),s(X35,X37))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X31,bool),'const.bool.LET_2'(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X39))),s(X31,happ(s(fun(X32,X31),'const.bool.LET_2'(s(fun(X33,fun(X32,X31)),X34),s(X33,'const.bool.LET_2'(s(fun(X35,X33),X36),s(X35,X37))))),s(X32,X38)))))) )),
  inference(superposition,[],[f574,f145])).
fof(f8079,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X22,bool),'const.bool.LET_2'(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,X30))),s(X22,'const.bool.LET_2'(s(fun(X23,X22),'_dst_x0x1x0_3'(s(fun(X24,fun(X25,X22)),X26),s(fun(X23,X24),X27),s(X25,X28))),s(X23,X29))))) | s(X22,X30) != s(X22,happ(s(fun(X25,X22),'const.bool.LET_2'(s(fun(X24,fun(X25,X22)),X26),s(X24,'const.bool.LET_2'(s(fun(X23,X24),X27),s(X23,X29))))),s(X25,X28)))) )),
  inference(forward_demodulation,[],[f7922,f91])).
fof(f7922,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25,X22] : (s(X22,X30) != s(X22,happ(s(fun(X25,X22),'const.bool.LET_2'(s(fun(X24,fun(X25,X22)),X26),s(X24,'const.bool.LET_2'(s(fun(X23,X24),X27),s(X23,X29))))),s(X25,X28))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X22,bool),'const.bool.LET_2'(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,X30))),s(X22,happ(s(fun(X23,X22),'_dst_x0x1x0_3'(s(fun(X24,fun(X25,X22)),X26),s(fun(X23,X24),X27),s(X25,X28))),s(X23,X29)))))) )),
  inference(superposition,[],[f574,f145])).
fof(f8078,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),'const.bool.LET_2'(s(fun(X15,fun(X15,bool)),'=_0'),s(X15,X21))),s(X15,'const.bool.LET_2'(s(fun(X16,X15),'const.bool.LET_2'(s(fun(X17,fun(X16,X15)),X18),s(X17,X19))),s(X16,X20))))) | s(X15,X21) != s(X15,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,X15)),X18),s('type.pair.prod'(X17,X16),'const.pair.,_2'(s(X17,X19),s(X16,X20)))))) )),
  inference(forward_demodulation,[],[f7921,f91])).
fof(f7921,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X15,X21) != s(X15,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,X15)),X18),s('type.pair.prod'(X17,X16),'const.pair.,_2'(s(X17,X19),s(X16,X20))))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),'const.bool.LET_2'(s(fun(X15,fun(X15,bool)),'=_0'),s(X15,X21))),s(X15,happ(s(fun(X16,X15),'const.bool.LET_2'(s(fun(X17,fun(X16,X15)),X18),s(X17,X19))),s(X16,X20)))))) )),
  inference(superposition,[],[f574,f143])).
fof(f7910,plain,(
  ( ! [X134,X132,X130,X128,X136,X135,X133,X131,X129,X137] : (s(fun(X128,X129),X137) != s(fun(X128,X129),happ(s(fun(X133,fun(X128,X129)),'_dst_x00x11x00_3'(s(fun(X130,fun(X131,X129)),X132),s(fun(X133,fun(X128,X130)),X134),s(X131,X136))),s(X133,X135))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X128,X129),bool),'const.bool.LET_2'(s(fun(fun(X128,X129),fun(fun(X128,X129),bool)),'=_0'),s(fun(X128,X129),'_dst_x0x1x0_3'(s(fun(X130,fun(X131,X129)),X132),s(fun(X128,X130),'const.bool.LET_2'(s(fun(X133,fun(X128,X130)),X134),s(X133,X135))),s(X131,X136))))),s(fun(X128,X129),X137)))) )),
  inference(superposition,[],[f574,f154])).
fof(f8066,plain,(
  ( ! [X127,X125,X123,X121,X118,X126,X124,X122,X120,X119] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X118,X119),bool),'const.bool.LET_2'(s(fun(fun(X118,X119),fun(fun(X118,X119),bool)),'=_0'),s(fun(X118,X119),'const.bool.LET_2'(s(fun(X120,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X120,fun(X118,X121)),X124),s(X122,X125))),s(X120,X126))))),s(fun(X118,X119),X127))) | s(fun(X118,X119),X127) != s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X118,X121),'const.bool.LET_2'(s(fun(X120,fun(X118,X121)),X124),s(X120,X126))),s(X122,X125)))) )),
  inference(forward_demodulation,[],[f7909,f91])).
fof(f7909,plain,(
  ( ! [X127,X125,X123,X121,X118,X126,X124,X122,X120,X119] : (s(fun(X118,X119),X127) != s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X118,X121),'const.bool.LET_2'(s(fun(X120,fun(X118,X121)),X124),s(X120,X126))),s(X122,X125))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X118,X119),bool),'const.bool.LET_2'(s(fun(fun(X118,X119),fun(fun(X118,X119),bool)),'=_0'),s(fun(X118,X119),happ(s(fun(X120,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X121,fun(X122,X119)),X123),s(fun(X120,fun(X118,X121)),X124),s(X122,X125))),s(X120,X126))))),s(fun(X118,X119),X127)))) )),
  inference(superposition,[],[f574,f154])).
fof(f7895,plain,(
  ( ! [X85,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X85))) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f574,f156])).
fof(f7888,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X63))),s(bool,X65))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f574,f156])).
fof(f7887,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X60))),s(bool,X61))) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f574,f156])).
fof(f8054,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(X51,X57) != s(X51,'const.bool.LET_2'(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56))))))),s(X51,X57)))) )),
  inference(forward_demodulation,[],[f7885,f91])).
fof(f7885,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(X51,X57) != s(X51,happ(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X51,bool),'const.bool.LET_2'(s(fun(X51,fun(X51,bool)),'=_0'),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56))))))),s(X51,X57)))) )),
  inference(superposition,[],[f574,f143])).
fof(f8050,plain,(
  ( ! [X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X31,bool),'const.bool.LET_2'(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,'const.bool.LET_2'(s(fun(X32,X31),'const.bool.LET_2'(s(fun(X33,fun(X32,X31)),X34),s(X33,'const.bool.LET_2'(s(fun(X35,X33),X36),s(X35,X37))))),s(X32,X38))))),s(X31,X39))) | s(X31,X39) != s(X31,happ(s(fun(X35,X31),'_dst_x0x1x0_3'(s(fun(X33,fun(X32,X31)),X34),s(fun(X35,X33),X36),s(X32,X38))),s(X35,X37)))) )),
  inference(forward_demodulation,[],[f7882,f91])).
fof(f7882,plain,(
  ( ! [X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(X31,X39) != s(X31,happ(s(fun(X35,X31),'_dst_x0x1x0_3'(s(fun(X33,fun(X32,X31)),X34),s(fun(X35,X33),X36),s(X32,X38))),s(X35,X37))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X31,bool),'const.bool.LET_2'(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,happ(s(fun(X32,X31),'const.bool.LET_2'(s(fun(X33,fun(X32,X31)),X34),s(X33,'const.bool.LET_2'(s(fun(X35,X33),X36),s(X35,X37))))),s(X32,X38))))),s(X31,X39)))) )),
  inference(superposition,[],[f574,f145])).
fof(f8049,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X22,bool),'const.bool.LET_2'(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,'const.bool.LET_2'(s(fun(X23,X22),'_dst_x0x1x0_3'(s(fun(X24,fun(X25,X22)),X26),s(fun(X23,X24),X27),s(X25,X28))),s(X23,X29))))),s(X22,X30))) | s(X22,X30) != s(X22,happ(s(fun(X25,X22),'const.bool.LET_2'(s(fun(X24,fun(X25,X22)),X26),s(X24,'const.bool.LET_2'(s(fun(X23,X24),X27),s(X23,X29))))),s(X25,X28)))) )),
  inference(forward_demodulation,[],[f7881,f91])).
fof(f7881,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25,X22] : (s(X22,X30) != s(X22,happ(s(fun(X25,X22),'const.bool.LET_2'(s(fun(X24,fun(X25,X22)),X26),s(X24,'const.bool.LET_2'(s(fun(X23,X24),X27),s(X23,X29))))),s(X25,X28))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X22,bool),'const.bool.LET_2'(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,happ(s(fun(X23,X22),'_dst_x0x1x0_3'(s(fun(X24,fun(X25,X22)),X26),s(fun(X23,X24),X27),s(X25,X28))),s(X23,X29))))),s(X22,X30)))) )),
  inference(superposition,[],[f574,f145])).
fof(f8048,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),'const.bool.LET_2'(s(fun(X15,fun(X15,bool)),'=_0'),s(X15,'const.bool.LET_2'(s(fun(X16,X15),'const.bool.LET_2'(s(fun(X17,fun(X16,X15)),X18),s(X17,X19))),s(X16,X20))))),s(X15,X21))) | s(X15,X21) != s(X15,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,X15)),X18),s('type.pair.prod'(X17,X16),'const.pair.,_2'(s(X17,X19),s(X16,X20)))))) )),
  inference(forward_demodulation,[],[f7880,f91])).
fof(f7880,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X15,X21) != s(X15,'const.pair.UNCURRY_2'(s(fun(X17,fun(X16,X15)),X18),s('type.pair.prod'(X17,X16),'const.pair.,_2'(s(X17,X19),s(X16,X20))))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X15,bool),'const.bool.LET_2'(s(fun(X15,fun(X15,bool)),'=_0'),s(X15,happ(s(fun(X16,X15),'const.bool.LET_2'(s(fun(X17,fun(X16,X15)),X18),s(X17,X19))),s(X16,X20))))),s(X15,X21)))) )),
  inference(superposition,[],[f574,f143])).
fof(f8037,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f8036,f91])).
fof(f8036,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f7844,f91])).
fof(f7844,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,happ(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f145,f574])).
fof(f8031,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5)))))))) )),
  inference(forward_demodulation,[],[f7846,f91])).
fof(f7846,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5)))))))) )),
  inference(unit_resulting_resolution,[],[f143,f574])).
fof(f8020,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X3,X8))))),s(fun(X0,X1),'const.bool.LET_2'(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7)))))) )),
  inference(forward_demodulation,[],[f7853,f91])).
fof(f7853,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X3,X8))))),s(fun(X0,X1),happ(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f154,f574])).
fof(f7990,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))))) )),
  inference(forward_demodulation,[],[f7989,f889])).
fof(f7989,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))))) )),
  inference(forward_demodulation,[],[f7860,f91])).
fof(f7860,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3)))))) )),
  inference(unit_resulting_resolution,[],[f137,f574])).
fof(f7987,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(forward_demodulation,[],[f7986,f1624])).
fof(f7986,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(forward_demodulation,[],[f7861,f91])).
fof(f7861,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1389,f574])).
fof(f7984,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f7983,f91])).
fof(f7983,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f7862,f91])).
fof(f7862,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f145,f574])).
fof(f7979,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,'const.bool.LET_2'(s(fun(X2,X0),'const.bool.LET_2'(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))))) )),
  inference(forward_demodulation,[],[f7864,f91])).
fof(f7864,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,happ(s(fun(X2,X0),'const.bool.LET_2'(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))))) )),
  inference(unit_resulting_resolution,[],[f143,f574])).
fof(f7970,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(forward_demodulation,[],[f7969,f1375])).
fof(f7969,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(forward_demodulation,[],[f7870,f91])).
fof(f7870,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3)))))) )),
  inference(unit_resulting_resolution,[],[f102,f574])).
fof(f7967,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'const.bool.LET_2'(s(fun(X2,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(X4,X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),'const.bool.LET_2'(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(X4,X7)))))) )),
  inference(forward_demodulation,[],[f7871,f91])).
fof(f7871,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(X4,X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),'const.bool.LET_2'(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(X4,X7)))))) )),
  inference(unit_resulting_resolution,[],[f154,f574])).
fof(f2886,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (s(X10,'const.bool.LET_2'(s(fun(X12,X10),X13),s(X12,X14))) = s(X10,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,X10)),X11),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(X12,'type.pair.prod'(X8,X9)),X15),s(X12,X14))))) | s(bool,X16) = s(bool,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X8,fun(X9,X10)),X11),s(X10,'const.bool.LET_2'(s(fun(X12,X10),X13),s(X12,X14))))),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(X12,'type.pair.prod'(X8,X9)),X15),s(X12,X14))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f2810,f78])).
fof(f2810,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9,X16] : (~p(s(bool,'T_0')) | s(X10,'const.bool.LET_2'(s(fun(X12,X10),X13),s(X12,X14))) = s(X10,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,X10)),X11),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(X12,'type.pair.prod'(X8,X9)),X15),s(X12,X14))))) | s(bool,X16) = s(bool,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,bool)),'_dst_x00x11x00_3'(s(fun(X10,fun(X10,bool)),'=_0'),s(fun(X8,fun(X9,X10)),X11),s(X10,'const.bool.LET_2'(s(fun(X12,X10),X13),s(X12,X14))))),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(X12,'type.pair.prod'(X8,X9)),X15),s(X12,X14))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1662,f156])).
fof(f2398,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),'const.bool.LET_2'(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7))),s(X0,X9))) = s(X1,happ(s(fun(X3,X1),'const.bool.LET_2'(s(fun(X2,fun(X3,X1)),X4),s(X2,'const.bool.LET_2'(s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X0,X9))))),s(X3,X8)))) )),
  inference(forward_demodulation,[],[f2352,f91])).
fof(f2352,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),happ(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7))),s(X0,X9))) = s(X1,happ(s(fun(X3,X1),'const.bool.LET_2'(s(fun(X2,fun(X3,X1)),X4),s(X2,'const.bool.LET_2'(s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X0,X9))))),s(X3,X8)))) )),
  inference(superposition,[],[f145,f154])).
fof(f2351,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),'const.bool.LET_2'(s(fun(X11,fun(X9,X12)),X15),s(X11,X17))),s(X13,X16))) = s(fun(X9,X10),'const.bool.LET_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x11x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(X13,X16))),s(X11,X17)))) )),
  inference(superposition,[],[f154,f91])).
fof(f2397,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X3,X8))))),s(fun(X0,X1),'const.bool.LET_2'(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7))))))) )),
  inference(forward_demodulation,[],[f2281,f91])).
fof(f2281,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X0,X2),'const.bool.LET_2'(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(X3,X8))))),s(fun(X0,X1),happ(s(fun(X5,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(X3,X8))),s(X5,X7))))))) )),
  inference(unit_resulting_resolution,[],[f154,f139])).
fof(f2396,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'const.bool.LET_2'(s(fun(X2,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(X4,X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),'const.bool.LET_2'(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(X4,X7))))))) )),
  inference(forward_demodulation,[],[f2282,f91])).
fof(f2282,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun(X0,X1),bool),'const.bool.LET_2'(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x00x11x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(X4,X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),'const.bool.LET_2'(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(X4,X7))))))) )),
  inference(unit_resulting_resolution,[],[f154,f139])).
fof(f2012,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s(X14,happ(s(fun(X18,X14),'_dst_x0x1x0_3'(s(fun(X16,fun(X15,X14)),X17),s(fun(X18,X16),X19),s(X15,X21))),s(X18,X20))) = s(X14,'const.pair.UNCURRY_2'(s(fun(X16,fun(X15,X14)),X17),s('type.pair.prod'(X16,X15),'const.pair.,_2'(s(X16,'const.bool.LET_2'(s(fun(X18,X16),X19),s(X18,X20))),s(X15,X21)))))) )),
  inference(superposition,[],[f145,f143])).
fof(f2011,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X6,happ(s(fun(X10,X6),'_dst_x0x1x0_3'(s(fun(X8,fun(X7,X6)),X9),s(fun(X10,X8),X11),s(X7,X13))),s(X10,X12))) = s(X6,'const.bool.LET_2'(s(fun(X7,X6),'const.bool.LET_2'(s(fun(X8,fun(X7,X6)),X9),s(X8,'const.bool.LET_2'(s(fun(X10,X8),X11),s(X10,X12))))),s(X7,X13)))) )),
  inference(superposition,[],[f145,f91])).
fof(f2010,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X2,fun(X1,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(superposition,[],[f145,f137])).
fof(f1941,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (s(bool,X38) = s(bool,happ(s(fun(X33,bool),'const.bool.LET_2'(s(fun(X32,fun(X33,bool)),X34),s(X32,'const.bool.LET_2'(s(fun(X31,X32),X35),s(X31,X37))))),s(X33,X36))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,happ(s(fun(X31,bool),'_dst_x0x1x0_3'(s(fun(X32,fun(X33,bool)),X34),s(fun(X31,X32),X35),s(X33,X36))),s(X31,X37)))) )),
  inference(superposition,[],[f145,f156])).
fof(f1940,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'const.bool.LET_2'(s(fun(X24,fun(X25,bool)),X26),s(X24,'const.bool.LET_2'(s(fun(X23,X24),X27),s(X23,X29))))),s(X25,X28))) | s(bool,X30) = s(bool,happ(s(fun(X23,bool),'_dst_x0x1x0_3'(s(fun(X24,fun(X25,bool)),X26),s(fun(X23,X24),X27),s(X25,X28))),s(X23,X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f145,f156])).
fof(f1938,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))) = s(X0,happ(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6)))) )),
  inference(superposition,[],[f145,f91])).
fof(f2094,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f2093,f91])).
fof(f2093,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f1936,f91])).
fof(f1936,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,happ(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X2,fun(X3,X0)),X4),s(X2,'const.bool.LET_2'(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6))))))) )),
  inference(unit_resulting_resolution,[],[f145,f139])).
fof(f2092,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6))))))) )),
  inference(forward_demodulation,[],[f2091,f91])).
fof(f2091,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,'const.bool.LET_2'(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6))))))) )),
  inference(forward_demodulation,[],[f1937,f91])).
fof(f1937,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,'const.bool.LET_2'(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6))))))) )),
  inference(unit_resulting_resolution,[],[f145,f139])).
fof(f1823,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(X5,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X5)),X6),s('type.pair.prod'(X2,X3),X4))) = s(X5,happ(s(fun(X3,X5),'const.bool.LET_2'(s(fun(X2,fun(X3,X5)),X6),s(X2,sK8(X4,X3,X2)))),s(X3,sK9(X4,X3,X2))))) )),
  inference(superposition,[],[f143,f90])).
fof(f1877,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))))))))) )),
  inference(forward_demodulation,[],[f1775,f91])).
fof(f1775,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f143,f139])).
fof(f1876,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,'const.bool.LET_2'(s(fun(X2,X0),'const.bool.LET_2'(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))))) )),
  inference(forward_demodulation,[],[f1776,f91])).
fof(f1776,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,happ(s(fun(X2,X0),'const.bool.LET_2'(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))))) )),
  inference(unit_resulting_resolution,[],[f143,f139])).
fof(f1659,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))))) )),
  inference(backward_demodulation,[],[f1624,f1656])).
fof(f1656,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))))) )),
  inference(forward_demodulation,[],[f1613,f91])).
fof(f1613,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))),s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))))) )),
  inference(unit_resulting_resolution,[],[f1389,f139])).
fof(f1655,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))))) )),
  inference(forward_demodulation,[],[f1614,f91])).
fof(f1614,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))))) )),
  inference(unit_resulting_resolution,[],[f1389,f139])).
fof(f1387,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(backward_demodulation,[],[f1375,f1386])).
fof(f1386,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(forward_demodulation,[],[f1370,f91])).
fof(f1370,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(unit_resulting_resolution,[],[f102,f139])).
fof(f1385,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(forward_demodulation,[],[f1371,f91])).
fof(f1371,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),bool),'const.bool.LET_2'(s(fun(fun('type.pair.prod'(X0,X1),X2),fun(fun('type.pair.prod'(X0,X1),X2),bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X2),happ(s(fun(fun(X0,fun(X1,X2)),fun('type.pair.prod'(X0,X1),X2)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,X2)),X3))))),s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))))))) )),
  inference(unit_resulting_resolution,[],[f102,f139])).
fof(f889,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.bool.LET_2'(s(fun(X5,X4),X6),s(X5,X7))) = s(X4,'const.bool.LET_2'(s(fun(X5,X4),'const.bool.LET_2'(s(fun(fun(X5,X4),fun(X5,X4)),'const.bool.LET_0'),s(fun(X5,X4),X6))),s(X5,X7)))) )),
  inference(superposition,[],[f137,f91])).
fof(f921,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f869,f91])).
fof(f869,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))))) )),
  inference(unit_resulting_resolution,[],[f137,f139])).
fof(f920,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f870,f91])).
fof(f870,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))))) )),
  inference(unit_resulting_resolution,[],[f137,f139])).
fof(f780,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK10(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f93,f156])).
fof(f778,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK10(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f93,f156])).
fof(f688,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,X2) | s(bool,X3) = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f669,f78])).
fof(f669,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2) | s(bool,X3) = s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f141,f156])).
fof(f550,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),X2))),s('type.pair.prod'(X0,X1),X2))))) )),
  inference(forward_demodulation,[],[f511,f90])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),bool),'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),fun('type.pair.prod'(X0,X1),bool)),'=_0'),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(X2,X1,X0)),s(X1,sK9(X2,X1,X0)))))),s('type.pair.prod'(X0,X1),X2))))) )),
  inference(unit_resulting_resolution,[],[f90,f139])).
fof(f546,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f514,f91])).
fof(f514,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))))))) )),
  inference(unit_resulting_resolution,[],[f91,f139])).
fof(f441,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f303,f78])).
fof(f303,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f128,f156])).
fof(f440,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f302,f78])).
fof(f302,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f89,f156])).
fof(f239,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f156,f156])).
fof(f153,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f152,f91])).
fof(f152,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f151,f91])).
fof(f151,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f150,f91])).
fof(f150,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f114,f91])).
fof(f114,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) != s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) & (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) <=> p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : (s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) <=> p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X8,X14] : ! [X11,X9,X15,X16] : (s(X6,happ(s(fun(X14,X6),X16),s(X14,X15))) = s(X6,'const.bool.LET_2'(s(fun('type.pair.prod'(X7,X8),X6),'const.pair.UNCURRY_1'(s(fun(X7,fun(X8,X6)),X9))),s('type.pair.prod'(X7,X8),happ(s(fun(X14,'type.pair.prod'(X7,X8)),X11),s(X14,X15))))) <=> p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X7,X8),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X8,bool)),'_dst_x00x11x00_3'(s(fun(X6,fun(X6,bool)),'=_0'),s(fun(X7,fun(X8,X6)),X9),s(X6,happ(s(fun(X14,X6),X16),s(X14,X15))))))),s('type.pair.prod'(X7,X8),happ(s(fun(X14,'type.pair.prod'(X7,X8)),X11),s(X14,X15)))))))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.lexer_fun.lexer_fun_5|49|_prove_ltb')).
fof(f149,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f148,f91])).
fof(f148,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) | s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6)))))) )),
  inference(forward_demodulation,[],[f147,f91])).
fof(f147,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | ~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f146,f91])).
fof(f146,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),'const.bool.LET_2'(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | ~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(forward_demodulation,[],[f115,f91])).
fof(f115,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))) = s(X0,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X5))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))) | ~p(s(bool,'const.bool.LET_2'(s(fun('type.pair.prod'(X1,X2),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,fun(X2,X0)),X5),s(X0,happ(s(fun(X3,X0),X7),s(X3,X6))))))),s('type.pair.prod'(X1,X2),happ(s(fun(X3,'type.pair.prod'(X1,X2)),X4),s(X3,X6))))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f180,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f174,f80])).
fof(f174,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f86,f81])).
fof(f154,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),'const.bool.LET_2'(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(forward_demodulation,[],[f116,f91])).
fof(f116,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X24,X25,X27,X26,X34] : ! [X28,X29,X30,X31] : s(fun(X25,X24),'_dst_x0x1x0_3'(s(fun(X27,fun(X34,X24)),X28),s(fun(X25,X27),happ(s(fun(X26,fun(X25,X27)),X29),s(X26,X31))),s(X34,X30))) = s(fun(X25,X24),happ(s(fun(X26,fun(X25,X24)),'_dst_x00x11x00_3'(s(fun(X27,fun(X34,X24)),X28),s(fun(X26,fun(X25,X27)),X29),s(X34,X30))),s(X26,X31)))),
  file('thm.evalProps.evalProps_43_prove_ltb','_dst_x00x11x00_ax')).
fof(f145,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X1,fun(X3,X0)),X4),s(X1,'const.bool.LET_2'(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(forward_demodulation,[],[f144,f91])).
fof(f144,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),'const.bool.LET_2'(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(forward_demodulation,[],[f113,f91])).
fof(f113,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X24,X25,X26,X27] : ! [X28,X29,X30,X31] : s(X24,happ(s(fun(X26,X24),'_dst_x0x1x0_3'(s(fun(X25,fun(X27,X24)),X28),s(fun(X26,X25),X29),s(X27,X30))),s(X26,X31))) = s(X24,happ(s(fun(X27,X24),happ(s(fun(X25,fun(X27,X24)),X28),s(X25,happ(s(fun(X26,X25),X29),s(X26,X31))))),s(X27,X30)))),
  file('thm.evalProps.evalProps_43_prove_ltb','_dst_x0x1x0_ax')).
fof(f128,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f127])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f112])).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X22,X38,X23,X39] : ((s(bool,X22) = s(bool,X38) & (p(s(bool,X38)) => s(bool,X23) = s(bool,X39))) => ((p(s(bool,X22)) => p(s(bool,X23))) <=> (p(s(bool,X38)) => p(s(bool,X39)))))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.bool.IMP_CONG')).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.bool.LET_2'(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f92,f91])).
fof(f92,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.bool.LET_2'(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),'const.bool.LET_0'),s(fun(X1,X0),X2))),s(X1,X3)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X17,X18] : s(X7,'const.bool.LET_2'(s(fun(X6,X7),X17),s(X6,X18))) = s(X7,happ(s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'const.bool.LET_0'),s(fun(X6,X7),X17))),s(X6,X18)))),
  file('thm.evalProps.evalProps_43_prove_ltb','ARITY_0_2_const.bool.LET')).
fof(f89,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X35,X36] : ((p(s(bool,X35)) => p(s(bool,X36))) => ((p(s(bool,X36)) => p(s(bool,X35))) => s(bool,X35) = s(bool,X36)))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f12439,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X36))),s(bool,X36))) | s(bool,'F_0') = s(bool,f(s(bool,X36))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl22 | $spl26 | $spl28 | $spl59)),
  inference(superposition,[],[f6873,f9821])).
fof(f6873,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X2))),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | ($spl22 | $spl26 | $spl28)),
  inference(trivial_inequality_removal,[],[f6860])).
fof(f6860,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X2))),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | ($spl22 | $spl26 | $spl28)),
  inference(superposition,[],[f6470,f5012])).
fof(f5012,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f4971])).
fof(f4971,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) | $spl26),
  inference(superposition,[],[f1301,f81])).
fof(f6470,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X23,bool),X24),s(X23,X25))) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) ) | ($spl22 | $spl28)),
  inference(forward_demodulation,[],[f6426,f2252])).
fof(f2252,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl28),
  inference(equality_resolution,[],[f1314])).
fof(f6426,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.bool.LET_2'(s(fun(X23,bool),X24),s(X23,X25))) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25)))) ) | $spl22),
  inference(superposition,[],[f6111,f91])).
fof(f6111,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f6089])).
fof(f6089,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(equality_factoring,[],[f822])).
fof(f86,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f13123,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,f(s(bool,X22)))) | s(bool,X23) != s(bool,f(s(bool,X21))) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(superposition,[],[f128,f12779])).
fof(f12779,plain,(
  ( ! [X39,X40] : (s(bool,f(s(bool,X39))) = s(bool,f(s(bool,X40))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(subsumption_resolution,[],[f12739,f80])).
fof(f12739,plain,(
  ( ! [X39,X40] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X39))) = s(bool,f(s(bool,X40))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(superposition,[],[f12711,f12495])).
fof(f13342,plain,(
  ( ! [X20,X18] : (s(bool,X20) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) | $spl62),
  inference(cnf_transformation,[],[f13342_D])).
fof(f13342_D,plain,(
  ( ! [X20,X18] : (s(bool,X20) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13341,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,f(s(bool,X19)))) | s(bool,X20) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(subsumption_resolution,[],[f13122,f12711])).
fof(f13122,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,f(s(bool,X19)))) | ~p(s(bool,X20)) | s(bool,X20) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | ($spl14 | $spl16 | $spl22 | $spl26 | $spl28 | $spl59)),
  inference(superposition,[],[f89,f12779])).
fof(f9175,plain,(
  $false | ($spl2 | $spl59 | $spl60)),
  inference(subsumption_resolution,[],[f9020,f78])).
fof(f9020,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f5311])).
fof(f8738,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl59 | $spl60)),
  inference(subsumption_resolution,[],[f8732,f8464])).
fof(f8732,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl59),
  inference(trivial_inequality_removal,[],[f8723])).
fof(f8723,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl59),
  inference(superposition,[],[f8463,f156])).
fof(f9174,plain,(
  $false | ($spl2 | $spl59 | $spl60)),
  inference(subsumption_resolution,[],[f9019,f78])).
fof(f9019,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f5310])).
fof(f5310,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK2,sK4))),s(sK1,sK9(sK5,sK1,sK0))))) | $spl2),
  inference(forward_demodulation,[],[f5309,f91])).
fof(f5309,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK2,sK4))),s(sK1,sK9(sK5,sK1,sK0))))) | $spl2),
  inference(forward_demodulation,[],[f5308,f2011])).
fof(f5308,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))))),s(sK2,sK4)))) | $spl2),
  inference(forward_demodulation,[],[f5306,f91])).
fof(f5306,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))))),s(sK2,sK4)))) | $spl2),
  inference(unit_resulting_resolution,[],[f5264,f141])).
fof(f9169,plain,(
  $false | ($spl59 | $spl60)),
  inference(subsumption_resolution,[],[f8905,f8738])).
fof(f8905,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f8463])).
fof(f9168,plain,(
  $false | ($spl11 | $spl28 | $spl59 | $spl60)),
  inference(subsumption_resolution,[],[f9167,f78])).
fof(f9167,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl28 | $spl59 | $spl60)),
  inference(forward_demodulation,[],[f8749,f2252])).
fof(f8749,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl11 | $spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f639])).
fof(f639,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl11),
  inference(unit_resulting_resolution,[],[f164,f141])).
fof(f9166,plain,(
  $false | ($spl11 | $spl28 | $spl59 | $spl60)),
  inference(subsumption_resolution,[],[f9165,f78])).
fof(f9165,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl28 | $spl59 | $spl60)),
  inference(forward_demodulation,[],[f8748,f2252])).
fof(f8748,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | ($spl11 | $spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f637])).
fof(f637,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl11),
  inference(unit_resulting_resolution,[],[f164,f141])).
fof(f9164,plain,(
  $false | ($spl59 | $spl60)),
  inference(subsumption_resolution,[],[f8739,f78])).
fof(f8739,plain,(
  ~p(s(bool,'T_0')) | ($spl59 | $spl60)),
  inference(backward_demodulation,[],[f8738,f80])).
fof(f8524,plain,(
  $false | ($spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8523,f8463])).
fof(f8523,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(condensation,[],[f8522])).
fof(f8522,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl55),
  inference(subsumption_resolution,[],[f8497,f81])).
fof(f8497,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl55),
  inference(superposition,[],[f8460,f156])).
fof(f8521,plain,(
  $false | ($spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8520,f8463])).
fof(f8520,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(condensation,[],[f8519])).
fof(f8519,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl55),
  inference(subsumption_resolution,[],[f8496,f81])).
fof(f8496,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl55),
  inference(superposition,[],[f8460,f156])).
fof(f8518,plain,(
  $false | ($spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8502,f8463])).
fof(f8502,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(trivial_inequality_removal,[],[f8494])).
fof(f8494,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(superposition,[],[f8460,f81])).
fof(f8517,plain,(
  $false | ($spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8465,f8463])).
fof(f8465,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(unit_resulting_resolution,[],[f8460,f81])).
fof(f8516,plain,(
  $false | ($spl26 | $spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8515,f8463])).
fof(f8515,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl26 | $spl55)),
  inference(forward_demodulation,[],[f8467,f7959])).
fof(f8467,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,f(s(bool,'F_0'))))),s(bool,f(s(bool,'F_0'))))) | ($spl26 | $spl55)),
  inference(unit_resulting_resolution,[],[f8460,f1301])).
fof(f8505,plain,(
  $false | ($spl11 | $spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8476,f8463])).
fof(f8476,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl11 | $spl55)),
  inference(unit_resulting_resolution,[],[f164,f8460,f156])).
fof(f8504,plain,(
  $false | ($spl11 | $spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8481,f8463])).
fof(f8481,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl11 | $spl55)),
  inference(unit_resulting_resolution,[],[f164,f8460,f156])).
fof(f8503,plain,(
  $false | ($spl55 | $spl59)),
  inference(subsumption_resolution,[],[f8482,f8463])).
fof(f8482,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(unit_resulting_resolution,[],[f80,f8460,f180])).
fof(f8461,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl56),
  inference(cnf_transformation,[],[f8461_D])).
fof(f8461_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8458,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK10(X0,X1,bool,bool))) ) | $spl52),
  inference(cnf_transformation,[],[f8458_D])).
fof(f8458_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK10(X0,X1,bool,bool))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8463,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl59),
  inference(cnf_transformation,[],[f8463_D])).
fof(f8463_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f8452,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK10(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f782,f156])).
fof(f8457,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl50),
  inference(cnf_transformation,[],[f8457_D])).
fof(f8457_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8464,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl60),
  inference(cnf_transformation,[],[f8464_D])).
fof(f8464_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8451,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK10(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f782,f156])).
fof(f8450,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK10(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f782,f156])).
fof(f8449,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK10(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f782,f156])).
fof(f8460,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(cnf_transformation,[],[f8460_D])).
fof(f8460_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f8448,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK10(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f782,f81])).
fof(f8402,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X1)) ) | $spl48),
  inference(cnf_transformation,[],[f8402_D])).
fof(f8402_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f178,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f8377,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X1)) ) | $spl26),
  inference(condensation,[],[f8376])).
fof(f8376,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X30))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f8344])).
fof(f8344,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,X29) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X30))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl26),
  inference(superposition,[],[f8109,f156])).
fof(f8109,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') != s(bool,X48) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,'F_0') = s(bool,X47)) ) | $spl26),
  inference(backward_demodulation,[],[f7959,f4998])).
fof(f4998,plain,(
  ( ! [X47,X48] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,X48) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X47))) | s(bool,'F_0') = s(bool,X47)) ) | $spl26),
  inference(superposition,[],[f139,f1301])).
fof(f8379,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X1)) ) | $spl26),
  inference(condensation,[],[f8378])).
fof(f8378,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f8343])).
fof(f8343,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl26),
  inference(superposition,[],[f8109,f156])).
fof(f6278,plain,(
  ( ! [X76,X77,X75] : (s(X75,X76) != s(X75,X77) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77))))),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77)))))) ) | $spl46),
  inference(cnf_transformation,[],[f6278_D])).
fof(f6278_D,plain,(
  ( ! [X76,X77,X75] : (s(X75,X76) != s(X75,X77) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77))))),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6277,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X78)) | s(X75,X76) != s(X75,X77) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77))))),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77)))))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5851,f439])).
fof(f5851,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X78)) | s(X75,X76) != s(X75,X77) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77))))),s(bool,'const.bool.LET_2'(s(fun(X75,bool),'const.bool.LET_2'(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X76))),s(X75,X77))))) | s(bool,'T_0') = s(bool,X78)) ) | $spl22),
  inference(superposition,[],[f139,f822])).
fof(f6274,plain,(
  ( ! [X50,X48] : (s(bool,X48) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X50))),s(bool,X50))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X48))) | s(bool,'F_0') = s(bool,X50)) ) | $spl44),
  inference(cnf_transformation,[],[f6274_D])).
fof(f6274_D,plain,(
  ( ! [X50,X48] : (s(bool,X48) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X50))),s(bool,X50))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X48))) | s(bool,'F_0') = s(bool,X50)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6273,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X49)) | s(bool,X48) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X50))),s(bool,X50))) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X48)))) ) | ($spl14 | $spl16 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f5840,f439])).
fof(f5840,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X49)) | s(bool,X48) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X50))),s(bool,X50))) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X48))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl14 | $spl22 | $spl26)),
  inference(superposition,[],[f5026,f822])).
fof(f5026,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl14 | $spl26)),
  inference(subsumption_resolution,[],[f4982,f179])).
fof(f4982,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,X10) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,X9)) ) | $spl26),
  inference(superposition,[],[f128,f1301])).
fof(f6271,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) | $spl42),
  inference(cnf_transformation,[],[f6271_D])).
fof(f6271_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f6270,plain,(
  ( ! [X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f5833,f439])).
fof(f5833,plain,(
  ( ! [X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f439,f822])).
fof(f5155,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5096,f4927])).
fof(f4927,plain,(
  s(sK2,sK4) != s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5))) | $spl1),
  inference(unit_resulting_resolution,[],[f4776,f139])).
fof(f4776,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5)))))) | $spl1),
  inference(forward_demodulation,[],[f4772,f1624])).
fof(f4772,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f131,f141])).
fof(f5096,plain,(
  s(sK2,sK4) = s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5))) | ($spl4 | $spl6)),
  inference(superposition,[],[f4643,f4648])).
fof(f4648,plain,(
  ( ! [X0,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))) = s(X0,happ(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f143,f136])).
fof(f5102,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5047,f4927])).
fof(f5047,plain,(
  s(sK2,sK4) = s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5))) | ($spl4 | $spl6)),
  inference(superposition,[],[f4648,f4643])).
fof(f4671,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f77,f76,f75,f78,f81,f89,f86,f90,f91,f137,f93,f102,f142,f143,f128,f145,f154,f155,f156,f158,f161,f177,f180,f181,f149,f153,f139,f141,f80,f213,f214,f239,f271,f272,f438,f440,f441,f307,f308,f447,f461,f463,f465,f466,f546,f547,f550,f516,f517,f519,f521,f522,f529,f530,f532,f534,f535,f574,f581,f626,f641,f643,f645,f646,f648,f655,f659,f662,f688,f696,f697,f724,f778,f780,f782,f783,f785,f787,f824,f920,f921,f889,f890,f892,f894,f1284,f965,f1275,f1268,f1257,f1252,f1006,f1007,f1008,f1009,f1243,f1242,f1025,f1026,f1302,f1303,f1045,f1046,f1047,f1048,f1077,f1078,f1079,f1080,f1082,f1085,f1086,f1087,f1220,f1328,f1330,f1215,f1115,f1116,f1117,f1118,f1120,f1123,f1124,f1125,f1126,f1212,f1144,f1145,f1339,f1340,f1342,f1343,f1170,f1171,f1172,f1173,f1175,f1178,f1179,f1180,f1203,f1183,f1197,f1385,f1386,f1375,f1387,f1389,f1401,f1400,f1445,f1454,f1455,f1418,f1419,f1469,f1517,f1526,f1530,f1522,f1655,f1656,f1624,f1661,f1662,f1659,f1625,f1626,f1627,f1876,f1877,f1777,f1822,f1823,f1878,f1879,f1880,f1883,f1884,f1885,f1886,f1887,f1888,f1889,f1890,f1932,f1933,f1931,f2092,f2094,f1938,f1939,f1940,f1941,f1983,f1984,f1985,f2010,f2011,f2098,f2012,f2013,f2100,f2102,f2104,f2106,f2108,f2112,f2114,f2116,f2120,f2126,f2128,f2130,f2132,f2134,f2136,f2138,f2142,f2144,f2146,f2150,f2156,f2158,f2082,f2396,f2397,f2351,f2398,f2399,f2400,f2401,f2402,f2403,f2404,f2405,f2406,f2600,f2601,f2602,f2606,f2608,f2610,f2613,f2542,f2543,f2544,f2616,f2617,f2621,f2623,f2625,f2628,f2632,f2633,f2634,f2638,f2640,f2642,f2645,f2649,f2651,f2653,f2656,f2574,f2575,f2576,f2579,f2581,f2661,f2662,f2592,f2593,f2594,f2816,f2817,f2821,f2823,f2825,f2828,f2832,f2833,f2834,f2838,f2840,f2842,f2845,f2770,f2771,f2772,f2852,f2853,f2854,f2858,f2860,f2862,f2865,f2869,f2871,f2873,f2876,f2793,f2794,f2795,f2798,f2800,f2881,f2882,f2886,f2960,f2961,f2899,f2901,f2902,f2903,f3011,f3013,f3108,f3110,f3122,f3124,f3145,f3146,f3147,f3213,f3214,f3152,f3183,f3233,f3249,f3251,f3360,f3361,f3362,f3366,f3368,f3370,f3376,f3382,f3384,f3386,f3391,f3332,f3333,f3334,f3337,f3339,f3398,f3399,f3349,f3407,f3409,f3584,f3585,f3589,f3591,f3593,f3599,f3605,f3607,f3609,f3614,f3553,f3554,f3555,f3558,f3560,f3630,f3631,f3864,f3865,f3869,f3871,f3873,f3879,f3885,f3887,f3889,f3894,f3825,f3826,f3827,f3830,f3832,f3902,f3903,f136,f4643,f4644,f4648])).
fof(f4644,plain,(
  s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f136,f1522])).
fof(f3903,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X131,X129] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X130))),s(sK1,'const.bool.LET_2'(s(fun(fun(X121,X122),sK1),X131),s(fun(X121,X122),'const.bool.LET_2'(s(fun(X126,fun(X121,X122)),'_dst_x00x11x00_3'(s(fun(X123,fun(X124,X122)),X125),s(fun(X126,fun(X121,X123)),X127),s(X124,X129))),s(X126,X128))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X130),s(sK1,'const.bool.LET_2'(s(fun(fun(X121,X122),sK1),X131),s(fun(X121,X122),'_dst_x0x1x0_3'(s(fun(X123,fun(X124,X122)),X125),s(fun(X121,X123),'const.bool.LET_2'(s(fun(X126,fun(X121,X123)),X127),s(X126,X128))),s(X124,X129)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3840,f91])).
fof(f3840,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X131,X129] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X130))),s(sK1,'const.bool.LET_2'(s(fun(fun(X121,X122),sK1),X131),s(fun(X121,X122),happ(s(fun(X126,fun(X121,X122)),'_dst_x00x11x00_3'(s(fun(X123,fun(X124,X122)),X125),s(fun(X126,fun(X121,X123)),X127),s(X124,X129))),s(X126,X128))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X130),s(sK1,'const.bool.LET_2'(s(fun(fun(X121,X122),sK1),X131),s(fun(X121,X122),'_dst_x0x1x0_3'(s(fun(X123,fun(X124,X122)),X125),s(fun(X121,X123),'const.bool.LET_2'(s(fun(X126,fun(X121,X123)),X127),s(X126,X128))),s(X124,X129)))))))) ) | $spl2),
  inference(superposition,[],[f1530,f154])).
fof(f3902,plain,(
  ( ! [X111,X118,X116,X114,X112,X120,X110,X119,X117,X115,X113] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X119),s(sK1,'const.bool.LET_2'(s(fun(fun(X110,X111),sK1),X120),s(fun(X110,X111),'const.bool.LET_2'(s(fun(X112,fun(X110,X111)),'_dst_x00x11x00_3'(s(fun(X113,fun(X114,X111)),X115),s(fun(X112,fun(X110,X113)),X116),s(X114,X117))),s(X112,X118))))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X119))),s(sK1,'const.bool.LET_2'(s(fun(fun(X110,X111),sK1),X120),s(fun(X110,X111),'_dst_x0x1x0_3'(s(fun(X113,fun(X114,X111)),X115),s(fun(X110,X113),'const.bool.LET_2'(s(fun(X112,fun(X110,X113)),X116),s(X112,X118))),s(X114,X117)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3839,f91])).
fof(f3839,plain,(
  ( ! [X111,X118,X116,X114,X112,X120,X110,X119,X117,X115,X113] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X119))),s(sK1,'const.bool.LET_2'(s(fun(fun(X110,X111),sK1),X120),s(fun(X110,X111),'_dst_x0x1x0_3'(s(fun(X113,fun(X114,X111)),X115),s(fun(X110,X113),'const.bool.LET_2'(s(fun(X112,fun(X110,X113)),X116),s(X112,X118))),s(X114,X117))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X119),s(sK1,'const.bool.LET_2'(s(fun(fun(X110,X111),sK1),X120),s(fun(X110,X111),happ(s(fun(X112,fun(X110,X111)),'_dst_x00x11x00_3'(s(fun(X113,fun(X114,X111)),X115),s(fun(X112,fun(X110,X113)),X116),s(X114,X117))),s(X112,X118)))))))) ) | $spl2),
  inference(superposition,[],[f1530,f154])).
fof(f3832,plain,(
  ( ! [X92,X90,X91] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X91))),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X92),s(bool,X90))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X91),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X92),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X90)) ) | $spl2),
  inference(superposition,[],[f1530,f81])).
fof(f3830,plain,(
  ( ! [X85,X83,X84,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X83))),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X84),s(bool,X82))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X83),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X84),s(bool,'T_0'))))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl2),
  inference(superposition,[],[f1530,f156])).
fof(f3827,plain,(
  ( ! [X70,X72,X71,X73] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X72))),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X73),s(bool,X71))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X72),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X73),s(bool,X70))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) ) | $spl2),
  inference(superposition,[],[f1530,f156])).
fof(f3826,plain,(
  ( ! [X68,X66,X69,X67] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X67))),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X68),s(bool,'T_0'))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X67),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X68),s(bool,X66))))) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) ) | $spl2),
  inference(superposition,[],[f1530,f156])).
fof(f3825,plain,(
  ( ! [X64,X65,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X64))),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X65),s(bool,'F_0'))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X64),s(sK1,'const.bool.LET_2'(s(fun(bool,sK1),X65),s(bool,X63))))) | s(bool,'T_0') = s(bool,X63)) ) | $spl2),
  inference(superposition,[],[f1530,f81])).
fof(f3894,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X57))),s(sK1,'const.bool.LET_2'(s(fun(X51,sK1),X58),s(X51,'const.bool.LET_2'(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X57),s(sK1,'const.bool.LET_2'(s(fun(X51,sK1),X58),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3823,f91])).
fof(f3823,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X57))),s(sK1,'const.bool.LET_2'(s(fun(X51,sK1),X58),s(X51,happ(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X57),s(sK1,'const.bool.LET_2'(s(fun(X51,sK1),X58),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))))) ) | $spl2),
  inference(superposition,[],[f1530,f143])).
fof(f3889,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X38),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,'const.bool.LET_2'(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37))))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X38))),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3888,f91])).
fof(f3888,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X38))),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X38),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3820,f91])).
fof(f3820,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X38))),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,happ(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X38),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),X39),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37)))))))) ) | $spl2),
  inference(superposition,[],[f1530,f145])).
fof(f3887,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27))))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3886,f91])).
fof(f3886,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3819,f91])).
fof(f3819,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,happ(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,'const.bool.LET_2'(s(fun(X20,sK1),X29),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27)))))))) ) | $spl2),
  inference(superposition,[],[f1530,f145])).
fof(f3885,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X18),s(sK1,'const.bool.LET_2'(s(fun(X12,sK1),X19),s(X12,'const.bool.LET_2'(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17))))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X18))),s(sK1,'const.bool.LET_2'(s(fun(X12,sK1),X19),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17)))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3818,f91])).
fof(f3818,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X18))),s(sK1,'const.bool.LET_2'(s(fun(X12,sK1),X19),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17))))))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X18),s(sK1,'const.bool.LET_2'(s(fun(X12,sK1),X19),s(X12,happ(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17)))))))) ) | $spl2),
  inference(superposition,[],[f1530,f143])).
fof(f3879,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X82))),s(sK1,'const.bool.LET_2'(s(fun(X76,sK1),'const.bool.LET_2'(s(fun(X78,fun(X76,sK1)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK1))),X79),s(X77,X80))),s(X78,X81))),s(X76,X83))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X82),s(sK1,'const.bool.LET_2'(s(fun(X76,sK1),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK1))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(X76,X83)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3814,f91])).
fof(f3814,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X82))),s(sK1,'const.bool.LET_2'(s(fun(X76,sK1),happ(s(fun(X78,fun(X76,sK1)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK1))),X79),s(X77,X80))),s(X78,X81))),s(X76,X83))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X82),s(sK1,'const.bool.LET_2'(s(fun(X76,sK1),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK1))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(X76,X83)))))) ) | $spl2),
  inference(superposition,[],[f1530,f143])).
fof(f3873,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X63),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),'const.bool.LET_2'(s(fun(X56,fun(X55,sK1)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X64))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X63))),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),'const.bool.LET_2'(s(fun(X59,fun(X55,sK1)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X64)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3872,f91])).
fof(f3872,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X63))),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),'const.bool.LET_2'(s(fun(X59,fun(X55,sK1)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X64))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X63),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),happ(s(fun(X56,fun(X55,sK1)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X64)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3811,f91])).
fof(f3811,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X63))),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),happ(s(fun(X59,fun(X55,sK1)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X64))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X63),s(sK1,'const.bool.LET_2'(s(fun(X55,sK1),happ(s(fun(X56,fun(X55,sK1)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK1))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X64)))))) ) | $spl2),
  inference(superposition,[],[f1530,f145])).
fof(f3871,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X53),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),'const.bool.LET_2'(s(fun(X46,fun(X45,sK1)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X54))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X53))),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),'const.bool.LET_2'(s(fun(X48,fun(X45,sK1)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X54)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3870,f91])).
fof(f3870,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X53))),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),'const.bool.LET_2'(s(fun(X48,fun(X45,sK1)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X54))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X53),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),happ(s(fun(X46,fun(X45,sK1)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X54)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3810,f91])).
fof(f3810,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X53))),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),happ(s(fun(X48,fun(X45,sK1)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X54))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X53),s(sK1,'const.bool.LET_2'(s(fun(X45,sK1),happ(s(fun(X46,fun(X45,sK1)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK1))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X54)))))) ) | $spl2),
  inference(superposition,[],[f1530,f145])).
fof(f3869,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X43),s(sK1,'const.bool.LET_2'(s(fun(X37,sK1),'const.bool.LET_2'(s(fun(X38,fun(X37,sK1)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK1))),X40),s(X39,X41))),s(X38,X42))),s(X37,X44))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X43))),s(sK1,'const.bool.LET_2'(s(fun(X37,sK1),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK1))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X44)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3809,f91])).
fof(f3809,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X43))),s(sK1,'const.bool.LET_2'(s(fun(X37,sK1),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK1))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X44))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X43),s(sK1,'const.bool.LET_2'(s(fun(X37,sK1),happ(s(fun(X38,fun(X37,sK1)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK1))),X40),s(X39,X41))),s(X38,X42))),s(X37,X44)))))) ) | $spl2),
  inference(superposition,[],[f1530,f143])).
fof(f3865,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X23))),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),'const.bool.LET_2'(s(fun(X19,fun(X15,sK1)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK1)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X24))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X23),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK1)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(X15,X24)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3806,f91])).
fof(f3806,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X23))),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),happ(s(fun(X19,fun(X15,sK1)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK1)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X24))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X23),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK1)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(X15,X24)))))) ) | $spl2),
  inference(superposition,[],[f1530,f154])).
fof(f3864,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X13),s(sK1,'const.bool.LET_2'(s(fun(X5,sK1),'const.bool.LET_2'(s(fun(X6,fun(X5,sK1)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK1)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(X5,X14))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X13))),s(sK1,'const.bool.LET_2'(s(fun(X5,sK1),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK1)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X14)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3805,f91])).
fof(f3805,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X13))),s(sK1,'const.bool.LET_2'(s(fun(X5,sK1),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK1)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X14))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X13),s(sK1,'const.bool.LET_2'(s(fun(X5,sK1),happ(s(fun(X6,fun(X5,sK1)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK1)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(X5,X14)))))) ) | $spl2),
  inference(superposition,[],[f1530,f154])).
fof(f3631,plain,(
  ( ! [X127,X125,X123,X121,X128,X118,X126,X124,X122,X120,X119] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(fun(X118,X119),sK0),X127),s(fun(X118,X119),'const.bool.LET_2'(s(fun(X123,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X120,fun(X121,X119)),X122),s(fun(X123,fun(X118,X120)),X124),s(X121,X126))),s(X123,X125))))))),s(sK1,X128))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X118,X119),sK0),X127),s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X120,fun(X121,X119)),X122),s(fun(X118,X120),'const.bool.LET_2'(s(fun(X123,fun(X118,X120)),X124),s(X123,X125))),s(X121,X126))))),s(sK1,X128)))) ) | $spl2),
  inference(forward_demodulation,[],[f3567,f91])).
fof(f3567,plain,(
  ( ! [X127,X125,X123,X121,X128,X118,X126,X124,X122,X120,X119] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(fun(X118,X119),sK0),X127),s(fun(X118,X119),happ(s(fun(X123,fun(X118,X119)),'_dst_x00x11x00_3'(s(fun(X120,fun(X121,X119)),X122),s(fun(X123,fun(X118,X120)),X124),s(X121,X126))),s(X123,X125))))))),s(sK1,X128))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X118,X119),sK0),X127),s(fun(X118,X119),'_dst_x0x1x0_3'(s(fun(X120,fun(X121,X119)),X122),s(fun(X118,X120),'const.bool.LET_2'(s(fun(X123,fun(X118,X120)),X124),s(X123,X125))),s(X121,X126))))),s(sK1,X128)))) ) | $spl2),
  inference(superposition,[],[f1526,f154])).
fof(f3630,plain,(
  ( ! [X111,X109,X107,X116,X114,X112,X110,X108,X117,X115,X113] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X107,X108),sK0),X116),s(fun(X107,X108),'const.bool.LET_2'(s(fun(X109,fun(X107,X108)),'_dst_x00x11x00_3'(s(fun(X110,fun(X111,X108)),X112),s(fun(X109,fun(X107,X110)),X113),s(X111,X114))),s(X109,X115))))),s(sK1,X117))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(fun(X107,X108),sK0),X116),s(fun(X107,X108),'_dst_x0x1x0_3'(s(fun(X110,fun(X111,X108)),X112),s(fun(X107,X110),'const.bool.LET_2'(s(fun(X109,fun(X107,X110)),X113),s(X109,X115))),s(X111,X114))))))),s(sK1,X117)))) ) | $spl2),
  inference(forward_demodulation,[],[f3566,f91])).
fof(f3566,plain,(
  ( ! [X111,X109,X107,X116,X114,X112,X110,X108,X117,X115,X113] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(fun(X107,X108),sK0),X116),s(fun(X107,X108),'_dst_x0x1x0_3'(s(fun(X110,fun(X111,X108)),X112),s(fun(X107,X110),'const.bool.LET_2'(s(fun(X109,fun(X107,X110)),X113),s(X109,X115))),s(X111,X114))))))),s(sK1,X117))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X107,X108),sK0),X116),s(fun(X107,X108),happ(s(fun(X109,fun(X107,X108)),'_dst_x00x11x00_3'(s(fun(X110,fun(X111,X108)),X112),s(fun(X109,fun(X107,X110)),X113),s(X111,X114))),s(X109,X115))))),s(sK1,X117)))) ) | $spl2),
  inference(superposition,[],[f1526,f154])).
fof(f3560,plain,(
  ( ! [X92,X90,X91] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X91),s(bool,X90))))),s(sK1,X92))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X91),s(bool,'F_0'))),s(sK1,X92))) | s(bool,'T_0') = s(bool,X90)) ) | $spl2),
  inference(superposition,[],[f1526,f81])).
fof(f3558,plain,(
  ( ! [X85,X83,X84,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X83),s(bool,X82))))),s(sK1,X84))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X83),s(bool,'T_0'))),s(sK1,X84))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl2),
  inference(superposition,[],[f1526,f156])).
fof(f3555,plain,(
  ( ! [X70,X72,X71,X73] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X72),s(bool,X71))))),s(sK1,X73))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X72),s(bool,X70))),s(sK1,X73))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) ) | $spl2),
  inference(superposition,[],[f1526,f156])).
fof(f3554,plain,(
  ( ! [X68,X66,X69,X67] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X67),s(bool,'T_0'))))),s(sK1,X68))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X67),s(bool,X66))),s(sK1,X68))) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) ) | $spl2),
  inference(superposition,[],[f1526,f156])).
fof(f3553,plain,(
  ( ! [X64,X65,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X64),s(bool,'F_0'))))),s(sK1,X65))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X64),s(bool,X63))),s(sK1,X65))) | s(bool,'T_0') = s(bool,X63)) ) | $spl2),
  inference(superposition,[],[f1526,f81])).
fof(f3614,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,'const.bool.LET_2'(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))))),s(sK1,X58))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56))))))),s(sK1,X58)))) ) | $spl2),
  inference(forward_demodulation,[],[f3551,f91])).
fof(f3551,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,happ(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))))),s(sK1,X58))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56))))))),s(sK1,X58)))) ) | $spl2),
  inference(superposition,[],[f1526,f143])).
fof(f3609,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,'const.bool.LET_2'(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37))))),s(sK1,X39))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))))),s(sK1,X39)))) ) | $spl2),
  inference(forward_demodulation,[],[f3608,f91])).
fof(f3608,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))))),s(sK1,X39))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37))))),s(sK1,X39)))) ) | $spl2),
  inference(forward_demodulation,[],[f3548,f91])).
fof(f3548,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,happ(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))))),s(sK1,X39))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37))))),s(sK1,X39)))) ) | $spl2),
  inference(superposition,[],[f1526,f145])).
fof(f3607,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27))))),s(sK1,X29))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))))),s(sK1,X29)))) ) | $spl2),
  inference(forward_demodulation,[],[f3606,f91])).
fof(f3606,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))))),s(sK1,X29))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27))))),s(sK1,X29)))) ) | $spl2),
  inference(forward_demodulation,[],[f3547,f91])).
fof(f3547,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,happ(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))))),s(sK1,X29))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27))))),s(sK1,X29)))) ) | $spl2),
  inference(superposition,[],[f1526,f145])).
fof(f3605,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,'const.bool.LET_2'(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17))))),s(sK1,X19))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17))))))))),s(sK1,X19)))) ) | $spl2),
  inference(forward_demodulation,[],[f3546,f91])).
fof(f3546,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17))))))))),s(sK1,X19))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,happ(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17))))),s(sK1,X19)))) ) | $spl2),
  inference(superposition,[],[f1526,f143])).
fof(f3599,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),'const.bool.LET_2'(s(fun(X78,fun(X76,sK0)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s(X77,X80))),s(X78,X81))),s(X76,X82))))),s(sK1,X83))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(X76,X82))),s(sK1,X83)))) ) | $spl2),
  inference(forward_demodulation,[],[f3542,f91])).
fof(f3542,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),happ(s(fun(X78,fun(X76,sK0)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s(X77,X80))),s(X78,X81))),s(X76,X82))))),s(sK1,X83))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(X76,X82))),s(sK1,X83)))) ) | $spl2),
  inference(superposition,[],[f1526,f143])).
fof(f3593,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X63))),s(sK1,X64))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))))),s(sK1,X64)))) ) | $spl2),
  inference(forward_demodulation,[],[f3592,f91])).
fof(f3592,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))))),s(sK1,X64))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),happ(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X63))),s(sK1,X64)))) ) | $spl2),
  inference(forward_demodulation,[],[f3539,f91])).
fof(f3539,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),happ(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))))),s(sK1,X64))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),happ(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(X55,X63))),s(sK1,X64)))) ) | $spl2),
  inference(superposition,[],[f1526,f145])).
fof(f3591,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X53))),s(sK1,X54))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))))),s(sK1,X54)))) ) | $spl2),
  inference(forward_demodulation,[],[f3590,f91])).
fof(f3590,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))))),s(sK1,X54))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),happ(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X53))),s(sK1,X54)))) ) | $spl2),
  inference(forward_demodulation,[],[f3538,f91])).
fof(f3538,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),happ(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))))),s(sK1,X54))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),happ(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(X45,X53))),s(sK1,X54)))) ) | $spl2),
  inference(superposition,[],[f1526,f145])).
fof(f3589,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),'const.bool.LET_2'(s(fun(X38,fun(X37,sK0)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s(X39,X41))),s(X38,X42))),s(X37,X43))),s(sK1,X44))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X43))))),s(sK1,X44)))) ) | $spl2),
  inference(forward_demodulation,[],[f3537,f91])).
fof(f3537,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X43))))),s(sK1,X44))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),happ(s(fun(X38,fun(X37,sK0)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s(X39,X41))),s(X38,X42))),s(X37,X43))),s(sK1,X44)))) ) | $spl2),
  inference(superposition,[],[f1526,f143])).
fof(f3585,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'const.bool.LET_2'(s(fun(X19,fun(X15,sK0)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X23))))),s(sK1,X24))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(X15,X23))),s(sK1,X24)))) ) | $spl2),
  inference(forward_demodulation,[],[f3534,f91])).
fof(f3534,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),happ(s(fun(X19,fun(X15,sK0)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X23))))),s(sK1,X24))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(X15,X23))),s(sK1,X24)))) ) | $spl2),
  inference(superposition,[],[f1526,f154])).
fof(f3584,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),'const.bool.LET_2'(s(fun(X6,fun(X5,sK0)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(X5,X13))),s(sK1,X14))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X13))))),s(sK1,X14)))) ) | $spl2),
  inference(forward_demodulation,[],[f3533,f91])).
fof(f3533,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X13))))),s(sK1,X14))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),happ(s(fun(X6,fun(X5,sK0)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(X5,X13))),s(sK1,X14)))) ) | $spl2),
  inference(superposition,[],[f1526,f154])).
fof(f3409,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X36,X34,X32] : (s(sK2,sK4) != s(sK2,happ(s(fun(X34,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X34,sK0),X35),s(sK1,'const.bool.LET_2'(s(fun(X27,sK1),'const.bool.LET_2'(s(fun(X28,fun(X27,sK1)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33))))),s(X34,X36))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X34,sK0),X35),s(X34,X36))),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,sK1)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3408,f91])).
fof(f3408,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X34,sK0),X35),s(X34,X36))),s(sK1,'const.bool.LET_2'(s(fun(X30,sK1),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,sK1)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32))))) | s(sK2,sK4) != s(sK2,happ(s(fun(X34,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X34,sK0),X35),s(sK1,happ(s(fun(X27,sK1),'const.bool.LET_2'(s(fun(X28,fun(X27,sK1)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33))))),s(X34,X36)))) ) | $spl2),
  inference(forward_demodulation,[],[f3353,f91])).
fof(f3353,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X34,sK0),X35),s(X34,X36))),s(sK1,happ(s(fun(X30,sK1),'_dst_x0x1x0_3'(s(fun(X28,fun(X27,sK1)),X29),s(fun(X30,X28),X31),s(X27,X33))),s(X30,X32))))) | s(sK2,sK4) != s(sK2,happ(s(fun(X34,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X34,sK0),X35),s(sK1,happ(s(fun(X27,sK1),'const.bool.LET_2'(s(fun(X28,fun(X27,sK1)),X29),s(X28,'const.bool.LET_2'(s(fun(X30,X28),X31),s(X30,X32))))),s(X27,X33))))),s(X34,X36)))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3407,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s(sK2,sK4) != s(sK2,happ(s(fun(X24,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X24,sK0),X25),s(sK1,'const.bool.LET_2'(s(fun(X17,sK1),'_dst_x0x1x0_3'(s(fun(X18,fun(X19,sK1)),X20),s(fun(X17,X18),X21),s(X19,X22))),s(X17,X23))))),s(X24,X26))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X24,sK0),X25),s(X24,X26))),s(sK1,'const.bool.LET_2'(s(fun(X19,sK1),'const.bool.LET_2'(s(fun(X18,fun(X19,sK1)),X20),s(X18,'const.bool.LET_2'(s(fun(X17,X18),X21),s(X17,X23))))),s(X19,X22)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3406,f91])).
fof(f3406,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X24,sK0),X25),s(X24,X26))),s(sK1,'const.bool.LET_2'(s(fun(X19,sK1),'const.bool.LET_2'(s(fun(X18,fun(X19,sK1)),X20),s(X18,'const.bool.LET_2'(s(fun(X17,X18),X21),s(X17,X23))))),s(X19,X22))))) | s(sK2,sK4) != s(sK2,happ(s(fun(X24,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X24,sK0),X25),s(sK1,happ(s(fun(X17,sK1),'_dst_x0x1x0_3'(s(fun(X18,fun(X19,sK1)),X20),s(fun(X17,X18),X21),s(X19,X22))),s(X17,X23))))),s(X24,X26)))) ) | $spl2),
  inference(forward_demodulation,[],[f3352,f91])).
fof(f3352,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X24,sK0),X25),s(X24,X26))),s(sK1,happ(s(fun(X19,sK1),'const.bool.LET_2'(s(fun(X18,fun(X19,sK1)),X20),s(X18,'const.bool.LET_2'(s(fun(X17,X18),X21),s(X17,X23))))),s(X19,X22))))) | s(sK2,sK4) != s(sK2,happ(s(fun(X24,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X24,sK0),X25),s(sK1,happ(s(fun(X17,sK1),'_dst_x0x1x0_3'(s(fun(X18,fun(X19,sK1)),X20),s(fun(X17,X18),X21),s(X19,X22))),s(X17,X23))))),s(X24,X26)))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3349,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK0)),X0),s('type.pair.prod'(sK0,sK1),sK5))),s(sK1,X1))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(sK1,sK0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK0)),X0),s(sK0,sK6))),s(sK1,X1))),s(sK1,sK7)))) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f2082,f2899])).
fof(f3399,plain,(
  ( ! [X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X116,X117),sK0),X125),s(fun(X116,X117),'const.bool.LET_2'(s(fun(X121,fun(X116,X117)),'_dst_x00x11x00_3'(s(fun(X118,fun(X119,X117)),X120),s(fun(X121,fun(X116,X118)),X122),s(X119,X124))),s(X121,X123))))),s(sK1,X126))) | s(sK2,sK4) != s(sK2,happ(s(fun(fun(X116,X117),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(fun(X116,X117),sK0),X125),s(sK1,X126))),s(fun(X116,X117),'_dst_x0x1x0_3'(s(fun(X118,fun(X119,X117)),X120),s(fun(X116,X118),'const.bool.LET_2'(s(fun(X121,fun(X116,X118)),X122),s(X121,X123))),s(X119,X124)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3345,f91])).
fof(f3345,plain,(
  ( ! [X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X116,X117),sK0),X125),s(fun(X116,X117),happ(s(fun(X121,fun(X116,X117)),'_dst_x00x11x00_3'(s(fun(X118,fun(X119,X117)),X120),s(fun(X121,fun(X116,X118)),X122),s(X119,X124))),s(X121,X123))))),s(sK1,X126))) | s(sK2,sK4) != s(sK2,happ(s(fun(fun(X116,X117),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(fun(X116,X117),sK0),X125),s(sK1,X126))),s(fun(X116,X117),'_dst_x0x1x0_3'(s(fun(X118,fun(X119,X117)),X120),s(fun(X116,X118),'const.bool.LET_2'(s(fun(X121,fun(X116,X118)),X122),s(X121,X123))),s(X119,X124)))))) ) | $spl2),
  inference(superposition,[],[f2082,f154])).
fof(f3398,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X115,X113] : (s(sK2,sK4) != s(sK2,happ(s(fun(fun(X105,X106),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(fun(X105,X106),sK0),X114),s(sK1,X115))),s(fun(X105,X106),'const.bool.LET_2'(s(fun(X107,fun(X105,X106)),'_dst_x00x11x00_3'(s(fun(X108,fun(X109,X106)),X110),s(fun(X107,fun(X105,X108)),X111),s(X109,X112))),s(X107,X113))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X105,X106),sK0),X114),s(fun(X105,X106),'_dst_x0x1x0_3'(s(fun(X108,fun(X109,X106)),X110),s(fun(X105,X108),'const.bool.LET_2'(s(fun(X107,fun(X105,X108)),X111),s(X107,X113))),s(X109,X112))))),s(sK1,X115)))) ) | $spl2),
  inference(forward_demodulation,[],[f3344,f91])).
fof(f3344,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X115,X113] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(fun(X105,X106),sK0),X114),s(fun(X105,X106),'_dst_x0x1x0_3'(s(fun(X108,fun(X109,X106)),X110),s(fun(X105,X108),'const.bool.LET_2'(s(fun(X107,fun(X105,X108)),X111),s(X107,X113))),s(X109,X112))))),s(sK1,X115))) | s(sK2,sK4) != s(sK2,happ(s(fun(fun(X105,X106),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(fun(X105,X106),sK0),X114),s(sK1,X115))),s(fun(X105,X106),happ(s(fun(X107,fun(X105,X106)),'_dst_x00x11x00_3'(s(fun(X108,fun(X109,X106)),X110),s(fun(X107,fun(X105,X108)),X111),s(X109,X112))),s(X107,X113)))))) ) | $spl2),
  inference(superposition,[],[f2082,f154])).
fof(f3339,plain,(
  ( ! [X92,X90,X91] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X91),s(bool,X90))),s(sK1,X92))) | s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(bool,sK0),X91),s(sK1,X92))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X90)) ) | $spl2),
  inference(superposition,[],[f2082,f81])).
fof(f3337,plain,(
  ( ! [X85,X83,X84,X82] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X83),s(bool,X82))),s(sK1,X84))) | s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(bool,sK0),X83),s(sK1,X84))),s(bool,'T_0'))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl2),
  inference(superposition,[],[f2082,f156])).
fof(f3334,plain,(
  ( ! [X70,X72,X71,X73] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X72),s(bool,X71))),s(sK1,X73))) | s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(bool,sK0),X72),s(sK1,X73))),s(bool,X70))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) ) | $spl2),
  inference(superposition,[],[f2082,f156])).
fof(f3333,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X67),s(bool,'T_0'))),s(sK1,X68))) | s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(bool,sK0),X67),s(sK1,X68))),s(bool,X66))) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) ) | $spl2),
  inference(superposition,[],[f2082,f156])).
fof(f3332,plain,(
  ( ! [X64,X65,X63] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(bool,sK0),X64),s(bool,'F_0'))),s(sK1,X65))) | s(sK2,sK4) != s(sK2,happ(s(fun(bool,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(bool,sK0),X64),s(sK1,X65))),s(bool,X63))) | s(bool,'T_0') = s(bool,X63)) ) | $spl2),
  inference(superposition,[],[f2082,f81])).
fof(f3391,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,'const.bool.LET_2'(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))),s(sK1,X58))) | s(sK2,sK4) != s(sK2,happ(s(fun(X51,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X51,sK0),X57),s(sK1,X58))),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3330,f91])).
fof(f3330,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X51,sK0),X57),s(X51,happ(s(fun(X53,X51),'const.bool.LET_2'(s(fun(X52,fun(X53,X51)),X54),s(X52,X55))),s(X53,X56))))),s(sK1,X58))) | s(sK2,sK4) != s(sK2,happ(s(fun(X51,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X51,sK0),X57),s(sK1,X58))),s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),X54),s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X55),s(X53,X56)))))))) ) | $spl2),
  inference(superposition,[],[f2082,f143])).
fof(f3386,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s(sK2,sK4) != s(sK2,happ(s(fun(X30,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X30,sK0),X38),s(sK1,X39))),s(X30,'const.bool.LET_2'(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))),s(sK1,X39)))) ) | $spl2),
  inference(forward_demodulation,[],[f3385,f91])).
fof(f3385,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,'const.bool.LET_2'(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))),s(sK1,X39))) | s(sK2,sK4) != s(sK2,happ(s(fun(X30,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X30,sK0),X38),s(sK1,X39))),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3327,f91])).
fof(f3327,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X30,sK0),X38),s(X30,happ(s(fun(X34,X30),'_dst_x0x1x0_3'(s(fun(X32,fun(X31,X30)),X33),s(fun(X34,X32),X35),s(X31,X37))),s(X34,X36))))),s(sK1,X39))) | s(sK2,sK4) != s(sK2,happ(s(fun(X30,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X30,sK0),X38),s(sK1,X39))),s(X30,happ(s(fun(X31,X30),'const.bool.LET_2'(s(fun(X32,fun(X31,X30)),X33),s(X32,'const.bool.LET_2'(s(fun(X34,X32),X35),s(X34,X36))))),s(X31,X37)))))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3384,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s(sK2,sK4) != s(sK2,happ(s(fun(X20,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X20,sK0),X28),s(sK1,X29))),s(X20,'const.bool.LET_2'(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))),s(sK1,X29)))) ) | $spl2),
  inference(forward_demodulation,[],[f3383,f91])).
fof(f3383,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,'const.bool.LET_2'(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))),s(sK1,X29))) | s(sK2,sK4) != s(sK2,happ(s(fun(X20,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X20,sK0),X28),s(sK1,X29))),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3326,f91])).
fof(f3326,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22,X20] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X20,sK0),X28),s(X20,happ(s(fun(X23,X20),'const.bool.LET_2'(s(fun(X22,fun(X23,X20)),X24),s(X22,'const.bool.LET_2'(s(fun(X21,X22),X25),s(X21,X27))))),s(X23,X26))))),s(sK1,X29))) | s(sK2,sK4) != s(sK2,happ(s(fun(X20,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X20,sK0),X28),s(sK1,X29))),s(X20,happ(s(fun(X21,X20),'_dst_x0x1x0_3'(s(fun(X22,fun(X23,X20)),X24),s(fun(X21,X22),X25),s(X23,X26))),s(X21,X27)))))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3382,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(sK2,sK4) != s(sK2,happ(s(fun(X12,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X12,sK0),X18),s(sK1,X19))),s(X12,'const.bool.LET_2'(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17))))))),s(sK1,X19)))) ) | $spl2),
  inference(forward_demodulation,[],[f3325,f91])).
fof(f3325,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X12,sK0),X18),s(X12,'const.pair.UNCURRY_2'(s(fun(X14,fun(X13,X12)),X15),s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,X17))))))),s(sK1,X19))) | s(sK2,sK4) != s(sK2,happ(s(fun(X12,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X12,sK0),X18),s(sK1,X19))),s(X12,happ(s(fun(X13,X12),'const.bool.LET_2'(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(X13,X17)))))) ) | $spl2),
  inference(superposition,[],[f2082,f143])).
fof(f3376,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),'const.bool.LET_2'(s(fun(X78,fun(X76,sK0)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s(X77,X80))),s(X78,X81))),s(X76,X82))),s(sK1,X83))) | s(sK2,sK4) != s(sK2,happ(s(fun(X76,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X76,sK0),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(sK1,X83))),s(X76,X82)))) ) | $spl2),
  inference(forward_demodulation,[],[f3321,f91])).
fof(f3321,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X76,sK0),happ(s(fun(X78,fun(X76,sK0)),'const.bool.LET_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s(X77,X80))),s(X78,X81))),s(X76,X82))),s(sK1,X83))) | s(sK2,sK4) != s(sK2,happ(s(fun(X76,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X76,sK0),'const.pair.UNCURRY_2'(s(fun(X77,fun(X78,fun(X76,sK0))),X79),s('type.pair.prod'(X77,X78),'const.pair.,_2'(s(X77,X80),s(X78,X81))))),s(sK1,X83))),s(X76,X82)))) ) | $spl2),
  inference(superposition,[],[f2082,f143])).
fof(f3370,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s(sK2,sK4) != s(sK2,happ(s(fun(X55,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(sK1,X64))),s(X55,X63))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))),s(sK1,X64)))) ) | $spl2),
  inference(forward_demodulation,[],[f3369,f91])).
fof(f3369,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),'const.bool.LET_2'(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))),s(sK1,X64))) | s(sK2,sK4) != s(sK2,happ(s(fun(X55,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X55,sK0),happ(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(sK1,X64))),s(X55,X63)))) ) | $spl2),
  inference(forward_demodulation,[],[f3318,f91])).
fof(f3318,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X55,X63] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X55,sK0),happ(s(fun(X59,fun(X55,sK0)),'_dst_x0x1x0_3'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(fun(X59,X57),X60),s(X56,X62))),s(X59,X61))),s(X55,X63))),s(sK1,X64))) | s(sK2,sK4) != s(sK2,happ(s(fun(X55,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X55,sK0),happ(s(fun(X56,fun(X55,sK0)),'const.bool.LET_2'(s(fun(X57,fun(X56,fun(X55,sK0))),X58),s(X57,'const.bool.LET_2'(s(fun(X59,X57),X60),s(X59,X61))))),s(X56,X62))),s(sK1,X64))),s(X55,X63)))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3368,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s(sK2,sK4) != s(sK2,happ(s(fun(X45,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(sK1,X54))),s(X45,X53))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))),s(sK1,X54)))) ) | $spl2),
  inference(forward_demodulation,[],[f3367,f91])).
fof(f3367,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),'const.bool.LET_2'(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))),s(sK1,X54))) | s(sK2,sK4) != s(sK2,happ(s(fun(X45,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X45,sK0),happ(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(sK1,X54))),s(X45,X53)))) ) | $spl2),
  inference(forward_demodulation,[],[f3317,f91])).
fof(f3317,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X53,X51,X49] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X45,sK0),happ(s(fun(X48,fun(X45,sK0)),'const.bool.LET_2'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(X47,'const.bool.LET_2'(s(fun(X46,X47),X50),s(X46,X52))))),s(X48,X51))),s(X45,X53))),s(sK1,X54))) | s(sK2,sK4) != s(sK2,happ(s(fun(X45,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X45,sK0),happ(s(fun(X46,fun(X45,sK0)),'_dst_x0x1x0_3'(s(fun(X47,fun(X48,fun(X45,sK0))),X49),s(fun(X46,X47),X50),s(X48,X51))),s(X46,X52))),s(sK1,X54))),s(X45,X53)))) ) | $spl2),
  inference(superposition,[],[f2082,f145])).
fof(f3366,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s(sK2,sK4) != s(sK2,happ(s(fun(X37,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X37,sK0),'const.bool.LET_2'(s(fun(X38,fun(X37,sK0)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s(X39,X41))),s(X38,X42))),s(sK1,X44))),s(X37,X43))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X43))),s(sK1,X44)))) ) | $spl2),
  inference(forward_demodulation,[],[f3316,f91])).
fof(f3316,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X37,sK0),'const.pair.UNCURRY_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s('type.pair.prod'(X39,X38),'const.pair.,_2'(s(X39,X41),s(X38,X42))))),s(X37,X43))),s(sK1,X44))) | s(sK2,sK4) != s(sK2,happ(s(fun(X37,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X37,sK0),happ(s(fun(X38,fun(X37,sK0)),'const.bool.LET_2'(s(fun(X39,fun(X38,fun(X37,sK0))),X40),s(X39,X41))),s(X38,X42))),s(sK1,X44))),s(X37,X43)))) ) | $spl2),
  inference(superposition,[],[f2082,f143])).
fof(f3362,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'const.bool.LET_2'(s(fun(X19,fun(X15,sK0)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X23))),s(sK1,X24))) | s(sK2,sK4) != s(sK2,happ(s(fun(X15,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X15,sK0),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(sK1,X24))),s(X15,X23)))) ) | $spl2),
  inference(forward_demodulation,[],[f3313,f91])).
fof(f3313,plain,(
  ( ! [X24,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),happ(s(fun(X19,fun(X15,sK0)),'_dst_x00x11x00_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X19,fun(X15,X16)),X20),s(X17,X22))),s(X19,X21))),s(X15,X23))),s(sK1,X24))) | s(sK2,sK4) != s(sK2,happ(s(fun(X15,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X15,sK0),'_dst_x0x1x0_3'(s(fun(X16,fun(X17,sK0)),X18),s(fun(X15,X16),'const.bool.LET_2'(s(fun(X19,fun(X15,X16)),X20),s(X19,X21))),s(X17,X22))),s(sK1,X24))),s(X15,X23)))) ) | $spl2),
  inference(superposition,[],[f2082,f154])).
fof(f3361,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s(sK2,sK4) != s(sK2,happ(s(fun(X5,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X5,sK0),'const.bool.LET_2'(s(fun(X6,fun(X5,sK0)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(sK1,X14))),s(X5,X13))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X13))),s(sK1,X14)))) ) | $spl2),
  inference(forward_demodulation,[],[f3312,f91])).
fof(f3312,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X5,sK0),'_dst_x0x1x0_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X5,X7),'const.bool.LET_2'(s(fun(X6,fun(X5,X7)),X10),s(X6,X12))),s(X8,X11))),s(X5,X13))),s(sK1,X14))) | s(sK2,sK4) != s(sK2,happ(s(fun(X5,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X5,sK0),happ(s(fun(X6,fun(X5,sK0)),'_dst_x00x11x00_3'(s(fun(X7,fun(X8,sK0)),X9),s(fun(X6,fun(X5,X7)),X10),s(X8,X11))),s(X6,X12))),s(sK1,X14))),s(X5,X13)))) ) | $spl2),
  inference(superposition,[],[f2082,f154])).
fof(f3360,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.pair.prod'(X0,X1),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun('type.pair.prod'(X0,X1),sK0),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,sK0)),X2))),s(sK1,X4))),s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,sK0)),X2),s('type.pair.prod'(X0,X1),X3))),s(sK1,X4)))) ) | $spl2),
  inference(forward_demodulation,[],[f3359,f1375])).
fof(f3359,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(sK2,sK4) != s(sK2,happ(s(fun('type.pair.prod'(X0,X1),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun('type.pair.prod'(X0,X1),sK0),'const.bool.LET_2'(s(fun(fun(X0,fun(X1,sK0)),fun('type.pair.prod'(X0,X1),sK0)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,sK0)),X2))),s(sK1,X4))),s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,sK0)),X2),s('type.pair.prod'(X0,X1),X3))),s(sK1,X4)))) ) | $spl2),
  inference(forward_demodulation,[],[f3358,f91])).
fof(f3358,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,sK0)),X2),s('type.pair.prod'(X0,X1),X3))),s(sK1,X4))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.pair.prod'(X0,X1),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun('type.pair.prod'(X0,X1),sK0),happ(s(fun(fun(X0,fun(X1,sK0)),fun('type.pair.prod'(X0,X1),sK0)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,sK0)),X2))),s(sK1,X4))),s('type.pair.prod'(X0,X1),X3)))) ) | $spl2),
  inference(forward_demodulation,[],[f3311,f1624])).
fof(f3311,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),sK0),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,sK0)),X2))),s('type.pair.prod'(X0,X1),X3))),s(sK1,X4))) | s(sK2,sK4) != s(sK2,happ(s(fun('type.pair.prod'(X0,X1),sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun('type.pair.prod'(X0,X1),sK0),happ(s(fun(fun(X0,fun(X1,sK0)),fun('type.pair.prod'(X0,X1),sK0)),'const.pair.UNCURRY_0'),s(fun(X0,fun(X1,sK0)),X2))),s(sK1,X4))),s('type.pair.prod'(X0,X1),X3)))) ) | $spl2),
  inference(superposition,[],[f2082,f102])).
fof(f3251,plain,(
  ( ! [X255,X253,X251,X254,X252,X250] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,bool)),'_dst_x00x11x00_3'(s(fun(X253,fun(X253,bool)),'=_0'),s(fun(X250,fun(X251,X253)),X255),s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))))),s('type.pair.prod'(X250,X251),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s('type.pair.prod'(sK0,sK1),sK5)))))) | s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))) != s(X253,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,X253)),X255),s('type.pair.prod'(X250,X251),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s('type.pair.prod'(sK0,sK1),sK5)))))) ) | $spl6),
  inference(forward_demodulation,[],[f3250,f136])).
fof(f3250,plain,(
  ( ! [X255,X253,X251,X254,X252,X250] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,bool)),'_dst_x00x11x00_3'(s(fun(X253,fun(X253,bool)),'=_0'),s(fun(X250,fun(X251,X253)),X255),s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))))),s('type.pair.prod'(X250,X251),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7)))))))) | s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))) != s(X253,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,X253)),X255),s('type.pair.prod'(X250,X251),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s('type.pair.prod'(sK0,sK1),sK5)))))) ) | $spl6),
  inference(forward_demodulation,[],[f3202,f1777])).
fof(f3202,plain,(
  ( ! [X255,X253,X251,X254,X252,X250] : (s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))) != s(X253,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,X253)),X255),s('type.pair.prod'(X250,X251),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s('type.pair.prod'(sK0,sK1),sK5))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X250,fun(X251,bool)),'_dst_x00x11x00_3'(s(fun(X253,fun(X253,bool)),'=_0'),s(fun(X250,fun(X251,X253)),X255),s(X253,'const.bool.LET_2'(s(fun(sK1,X253),X254),s(sK1,sK7))))),s('type.pair.prod'(X250,X251),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X250,X251)),'const.bool.LET_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X250,X251))),X252),s(sK0,sK6))),s(sK1,sK7))))))) ) | $spl6),
  inference(superposition,[],[f1661,f2899])).
fof(f3249,plain,(
  ( ! [X249,X246,X244,X248,X247,X245] : (s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))) = s(X247,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,X247)),X248),s('type.pair.prod'(X244,X245),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s('type.pair.prod'(sK0,sK1),sK5))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,bool)),'_dst_x00x11x00_3'(s(fun(X247,fun(X247,bool)),'=_0'),s(fun(X244,fun(X245,X247)),X248),s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))))),s('type.pair.prod'(X244,X245),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s('type.pair.prod'(sK0,sK1),sK5))))))) ) | $spl6),
  inference(forward_demodulation,[],[f3248,f136])).
fof(f3248,plain,(
  ( ! [X249,X246,X244,X248,X247,X245] : (s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))) = s(X247,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,X247)),X248),s('type.pair.prod'(X244,X245),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,bool)),'_dst_x00x11x00_3'(s(fun(X247,fun(X247,bool)),'=_0'),s(fun(X244,fun(X245,X247)),X248),s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))))),s('type.pair.prod'(X244,X245),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s('type.pair.prod'(sK0,sK1),sK5))))))) ) | $spl6),
  inference(forward_demodulation,[],[f3201,f1777])).
fof(f3201,plain,(
  ( ! [X249,X246,X244,X248,X247,X245] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,bool)),'_dst_x00x11x00_3'(s(fun(X247,fun(X247,bool)),'=_0'),s(fun(X244,fun(X245,X247)),X248),s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))))),s('type.pair.prod'(X244,X245),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s('type.pair.prod'(sK0,sK1),sK5)))))) | s(X247,'const.bool.LET_2'(s(fun(sK1,X247),X249),s(sK1,sK7))) = s(X247,'const.pair.UNCURRY_2'(s(fun(X244,fun(X245,X247)),X248),s('type.pair.prod'(X244,X245),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X244,X245)),'const.bool.LET_2'(s(fun(sK0,fun(sK1,'type.pair.prod'(X244,X245))),X246),s(sK0,sK6))),s(sK1,sK7)))))) ) | $spl6),
  inference(superposition,[],[f1662,f2899])).
fof(f3233,plain,(
  ( ! [X134,X136,X135,X133] : (s(X134,'const.bool.LET_2'(s(fun(X133,X134),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X133,X134))),X135),s('type.pair.prod'(sK0,sK1),sK5))),s(X133,X136))) = s(X134,'const.pair.UNCURRY_2'(s(fun(sK1,fun(X133,X134)),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X133,X134))),X135),s(sK0,sK6))),s('type.pair.prod'(sK1,X133),'const.pair.,_2'(s(sK1,sK7),s(X133,X136)))))) ) | $spl6),
  inference(forward_demodulation,[],[f3184,f91])).
fof(f3184,plain,(
  ( ! [X134,X136,X135,X133] : (s(X134,happ(s(fun(X133,X134),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X133,X134))),X135),s('type.pair.prod'(sK0,sK1),sK5))),s(X133,X136))) = s(X134,'const.pair.UNCURRY_2'(s(fun(sK1,fun(X133,X134)),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X133,X134))),X135),s(sK0,sK6))),s('type.pair.prod'(sK1,X133),'const.pair.,_2'(s(sK1,sK7),s(X133,X136)))))) ) | $spl6),
  inference(superposition,[],[f143,f2899])).
fof(f3183,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (s(fun(X126,X129),'_dst_x0x1x0_3'(s(fun(X127,fun(X130,X129)),X131),s(fun(X126,X127),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X126,X127))),X128),s('type.pair.prod'(sK0,sK1),sK5))),s(X130,X132))) = s(fun(X126,X129),happ(s(fun(sK1,fun(X126,X129)),'_dst_x00x11x00_3'(s(fun(X127,fun(X130,X129)),X131),s(fun(sK1,fun(X126,X127)),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X126,X127))),X128),s(sK0,sK6))),s(X130,X132))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f154,f2899])).
fof(f3152,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(X14,happ(s(fun(X15,X14),'const.bool.LET_2'(s(fun(X12,fun(X15,X14)),X16),s(X12,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X12)),X13),s('type.pair.prod'(sK0,sK1),sK5))))),s(X15,X17))) = s(X14,happ(s(fun(sK1,X14),'_dst_x0x1x0_3'(s(fun(X12,fun(X15,X14)),X16),s(fun(sK1,X12),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X12)),X13),s(sK0,sK6))),s(X15,X17))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f145,f2899])).
fof(f3214,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,bool)),'_dst_x00x11x00_3'(s(fun(X6,fun(X6,bool)),'=_0'),s(fun(X8,fun(X9,X6)),X10),s(X6,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X6)),X7),s('type.pair.prod'(sK0,sK1),sK5))))),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X8,X9)),X11),s(sK1,sK7)))))) | s(X6,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X6)),X7),s('type.pair.prod'(sK0,sK1),sK5))) != s(X6,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,X6)),X10),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X8,X9)),X11),s(sK1,sK7)))))) ) | $spl6),
  inference(forward_demodulation,[],[f3151,f2899])).
fof(f3151,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X6)),X7),s('type.pair.prod'(sK0,sK1),sK5))) != s(X6,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,X6)),X10),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X8,X9)),X11),s(sK1,sK7))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X8,fun(X9,bool)),'_dst_x00x11x00_3'(s(fun(X6,fun(X6,bool)),'=_0'),s(fun(X8,fun(X9,X6)),X10),s(X6,'const.bool.LET_2'(s(fun(sK1,X6),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X6)),X7),s(sK0,sK6))),s(sK1,sK7))))),s('type.pair.prod'(X8,X9),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X8,X9)),X11),s(sK1,sK7))))))) ) | $spl6),
  inference(superposition,[],[f1661,f2899])).
fof(f3213,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X0)),X4),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X2,X3)),X5),s(sK1,sK7))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X2,fun(X3,X0)),X4),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X2,X3)),X5),s(sK1,sK7))))))) ) | $spl6),
  inference(forward_demodulation,[],[f3150,f2899])).
fof(f3150,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,bool)),'_dst_x00x11x00_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X2,fun(X3,X0)),X4),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X2,X3)),X5),s(sK1,sK7)))))) | s(X0,'const.bool.LET_2'(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X3,X0)),X4),s('type.pair.prod'(X2,X3),'const.bool.LET_2'(s(fun(sK1,'type.pair.prod'(X2,X3)),X5),s(sK1,sK7)))))) ) | $spl6),
  inference(superposition,[],[f1662,f2899])).
fof(f3147,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X5),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X5),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f2899,f156])).
fof(f3146,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X3),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,X4) = s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X3),s(sK0,sK6))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f2899,f156])).
fof(f3145,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X2),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X2),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f2899,f81])).
fof(f3124,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,'const.bool.LET_2'(s(fun(X21,sK1),'const.bool.LET_2'(s(fun(X22,fun(X21,sK1)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,'const.bool.LET_2'(s(fun(X24,sK1),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK1)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3123,f91])).
fof(f3123,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,'const.bool.LET_2'(s(fun(X24,sK1),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK1)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,happ(s(fun(X21,sK1),'const.bool.LET_2'(s(fun(X22,fun(X21,sK1)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3100,f91])).
fof(f3100,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X28))),s(sK1,happ(s(fun(X24,sK1),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK1)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X28),s(sK1,happ(s(fun(X21,sK1),'const.bool.LET_2'(s(fun(X22,fun(X21,sK1)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27)))))) ) | $spl2),
  inference(superposition,[],[f1522,f145])).
fof(f3122,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X20),s(sK1,'const.bool.LET_2'(s(fun(X13,sK1),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK1)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X20))),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),'const.bool.LET_2'(s(fun(X14,fun(X15,sK1)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3121,f91])).
fof(f3121,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X20))),s(sK1,'const.bool.LET_2'(s(fun(X15,sK1),'const.bool.LET_2'(s(fun(X14,fun(X15,sK1)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X20),s(sK1,happ(s(fun(X13,sK1),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK1)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3099,f91])).
fof(f3099,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X20))),s(sK1,happ(s(fun(X15,sK1),'const.bool.LET_2'(s(fun(X14,fun(X15,sK1)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X20),s(sK1,happ(s(fun(X13,sK1),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK1)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19)))))) ) | $spl2),
  inference(superposition,[],[f1522,f145])).
fof(f3110,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X21,sK0),'const.bool.LET_2'(s(fun(X22,fun(X21,sK0)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27))),s(sK1,X28))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X24,sK0),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK0)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26))))),s(sK1,X28)))) ) | $spl2),
  inference(forward_demodulation,[],[f3109,f91])).
fof(f3109,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X24,sK0),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK0)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26))))),s(sK1,X28))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X21,sK0),'const.bool.LET_2'(s(fun(X22,fun(X21,sK0)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27))),s(sK1,X28)))) ) | $spl2),
  inference(forward_demodulation,[],[f3093,f91])).
fof(f3093,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,happ(s(fun(X24,sK0),'_dst_x0x1x0_3'(s(fun(X22,fun(X21,sK0)),X23),s(fun(X24,X22),X25),s(X21,X27))),s(X24,X26))))),s(sK1,X28))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X21,sK0),'const.bool.LET_2'(s(fun(X22,fun(X21,sK0)),X23),s(X22,'const.bool.LET_2'(s(fun(X24,X22),X25),s(X24,X26))))),s(X21,X27))),s(sK1,X28)))) ) | $spl2),
  inference(superposition,[],[f1522,f145])).
fof(f3108,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X13,sK0),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK0)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19))),s(sK1,X20))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'const.bool.LET_2'(s(fun(X14,fun(X15,sK0)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18))))),s(sK1,X20)))) ) | $spl2),
  inference(forward_demodulation,[],[f3107,f91])).
fof(f3107,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X15,sK0),'const.bool.LET_2'(s(fun(X14,fun(X15,sK0)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18))))),s(sK1,X20))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X13,sK0),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK0)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19))),s(sK1,X20)))) ) | $spl2),
  inference(forward_demodulation,[],[f3092,f91])).
fof(f3092,plain,(
  ( ! [X14,X19,X17,X15,X13,X20,X18,X16] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,happ(s(fun(X15,sK0),'const.bool.LET_2'(s(fun(X14,fun(X15,sK0)),X16),s(X14,'const.bool.LET_2'(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18))))),s(sK1,X20))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X13,sK0),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,sK0)),X16),s(fun(X13,X14),X17),s(X15,X18))),s(X13,X19))),s(sK1,X20)))) ) | $spl2),
  inference(superposition,[],[f1522,f145])).
fof(f3013,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X209,X214] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,bool)),'_dst_x00x11x00_3'(s(fun(X213,fun(X213,bool)),'=_0'),s(fun(X210,fun(X211,X213)),X216),s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))))),s('type.pair.prod'(X210,X211),'const.pair.UNCURRY_2'(s(fun(sK1,fun(X209,'type.pair.prod'(X210,X211))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s(sK0,sK6))),s('type.pair.prod'(sK1,X209),'const.pair.,_2'(s(sK1,sK7),s(X209,X215)))))))) | s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))) != s(X213,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,X213)),X216),s('type.pair.prod'(X210,X211),'const.bool.LET_2'(s(fun(X209,'type.pair.prod'(X210,X211)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s('type.pair.prod'(sK0,sK1),sK5))),s(X209,X215)))))) ) | $spl6),
  inference(forward_demodulation,[],[f3012,f1777])).
fof(f3012,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X209,X214] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,bool)),'_dst_x00x11x00_3'(s(fun(X213,fun(X213,bool)),'=_0'),s(fun(X210,fun(X211,X213)),X216),s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))))),s('type.pair.prod'(X210,X211),'const.bool.LET_2'(s(fun(X209,'type.pair.prod'(X210,X211)),'const.bool.LET_2'(s(fun(sK1,fun(X209,'type.pair.prod'(X210,X211))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s(sK0,sK6))),s(sK1,sK7))),s(X209,X215)))))) | s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))) != s(X213,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,X213)),X216),s('type.pair.prod'(X210,X211),'const.bool.LET_2'(s(fun(X209,'type.pair.prod'(X210,X211)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s('type.pair.prod'(sK0,sK1),sK5))),s(X209,X215)))))) ) | $spl6),
  inference(forward_demodulation,[],[f2952,f91])).
fof(f2952,plain,(
  ( ! [X212,X210,X216,X215,X213,X211,X209,X214] : (s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))) != s(X213,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,X213)),X216),s('type.pair.prod'(X210,X211),'const.bool.LET_2'(s(fun(X209,'type.pair.prod'(X210,X211)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s('type.pair.prod'(sK0,sK1),sK5))),s(X209,X215))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X210,fun(X211,bool)),'_dst_x00x11x00_3'(s(fun(X213,fun(X213,bool)),'=_0'),s(fun(X210,fun(X211,X213)),X216),s(X213,'const.bool.LET_2'(s(fun(X209,X213),X214),s(X209,X215))))),s('type.pair.prod'(X210,X211),'const.bool.LET_2'(s(fun(X209,'type.pair.prod'(X210,X211)),happ(s(fun(sK1,fun(X209,'type.pair.prod'(X210,X211))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X209,'type.pair.prod'(X210,X211)))),X212),s(sK0,sK6))),s(sK1,sK7))),s(X209,X215))))))) ) | $spl6),
  inference(superposition,[],[f1661,f1822])).
fof(f3011,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203,X201] : (s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))) = s(X205,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,X205)),X206),s('type.pair.prod'(X202,X203),'const.pair.UNCURRY_2'(s(fun(sK1,fun(X201,'type.pair.prod'(X202,X203))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s(sK0,sK6))),s('type.pair.prod'(sK1,X201),'const.pair.,_2'(s(sK1,sK7),s(X201,X208))))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,bool)),'_dst_x00x11x00_3'(s(fun(X205,fun(X205,bool)),'=_0'),s(fun(X202,fun(X203,X205)),X206),s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))))),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(X201,'type.pair.prod'(X202,X203)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s('type.pair.prod'(sK0,sK1),sK5))),s(X201,X208))))))) ) | $spl6),
  inference(forward_demodulation,[],[f3010,f1777])).
fof(f3010,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203,X201] : (s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))) = s(X205,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,X205)),X206),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(X201,'type.pair.prod'(X202,X203)),'const.bool.LET_2'(s(fun(sK1,fun(X201,'type.pair.prod'(X202,X203))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s(sK0,sK6))),s(sK1,sK7))),s(X201,X208))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,bool)),'_dst_x00x11x00_3'(s(fun(X205,fun(X205,bool)),'=_0'),s(fun(X202,fun(X203,X205)),X206),s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))))),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(X201,'type.pair.prod'(X202,X203)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s('type.pair.prod'(sK0,sK1),sK5))),s(X201,X208))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2951,f91])).
fof(f2951,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203,X201] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,bool)),'_dst_x00x11x00_3'(s(fun(X205,fun(X205,bool)),'=_0'),s(fun(X202,fun(X203,X205)),X206),s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))))),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(X201,'type.pair.prod'(X202,X203)),'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s('type.pair.prod'(sK0,sK1),sK5))),s(X201,X208)))))) | s(X205,'const.bool.LET_2'(s(fun(X201,X205),X207),s(X201,X208))) = s(X205,'const.pair.UNCURRY_2'(s(fun(X202,fun(X203,X205)),X206),s('type.pair.prod'(X202,X203),'const.bool.LET_2'(s(fun(X201,'type.pair.prod'(X202,X203)),happ(s(fun(sK1,fun(X201,'type.pair.prod'(X202,X203))),'const.bool.LET_2'(s(fun(sK0,fun(sK1,fun(X201,'type.pair.prod'(X202,X203)))),X204),s(sK0,sK6))),s(sK1,sK7))),s(X201,X208)))))) ) | $spl6),
  inference(superposition,[],[f1662,f1822])).
fof(f2903,plain,(
  ( ! [X10,X9] : (s(bool,X10) = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X9),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X9),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f1822,f156])).
fof(f2902,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X7),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,X8) = s(bool,happ(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X7),s(sK0,sK6))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f1822,f156])).
fof(f2901,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,bool)),X6),s('type.pair.prod'(sK0,sK1),sK5))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),'const.bool.LET_2'(s(fun(sK0,fun(sK1,bool)),X6),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f1822,f81])).
fof(f2899,plain,(
  ( ! [X2,X3] : (s(X2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X2)),X3),s('type.pair.prod'(sK0,sK1),sK5))) = s(X2,'const.bool.LET_2'(s(fun(sK1,X2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X2)),X3),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f1822,f91])).
fof(f2961,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s(X0,'const.bool.LET_2'(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2887,f91])).
fof(f2887,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s(X0,happ(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1822,f139])).
fof(f2960,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s(sK1,sK7))))) ) | $spl6),
  inference(forward_demodulation,[],[f2959,f91])).
fof(f2959,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))),s(sK1,sK7))))) ) | $spl6),
  inference(forward_demodulation,[],[f2958,f2011])).
fof(f2958,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.LET_2'(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7))))),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2888,f91])).
fof(f2888,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(X0,bool),'const.bool.LET_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7))))),s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1822,f139])).
fof(f2082,plain,(
  ( ! [X292,X290,X289,X291] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X289,sK0),X290),s(X289,X291))),s(sK1,X292))) | s(sK2,sK4) != s(sK2,happ(s(fun(X289,sK2),'_dst_x0x1x0_3'(s(fun(sK0,fun(sK1,sK2)),sK3),s(fun(X289,sK0),X290),s(sK1,X292))),s(X289,X291)))) ) | $spl2),
  inference(superposition,[],[f1469,f145])).
fof(f2158,plain,(
  ( ! [X282,X288,X285,X287,X281,X283,X284,X286] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X288),s(sK1,'const.bool.LET_2'(s(fun(X281,sK1),'const.bool.LET_2'(s(fun(X282,fun(X281,sK1)),X283),s(X282,'const.bool.LET_2'(s(fun(X284,X282),X285),s(X284,X286))))),s(X281,X287))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X288))),s(sK1,'const.bool.LET_2'(s(fun(X284,sK1),'_dst_x0x1x0_3'(s(fun(X282,fun(X281,sK1)),X283),s(fun(X284,X282),X285),s(X281,X287))),s(X284,X286)))))) ) | $spl2),
  inference(forward_demodulation,[],[f2157,f91])).
fof(f2157,plain,(
  ( ! [X282,X288,X285,X287,X281,X283,X284,X286] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X288))),s(sK1,'const.bool.LET_2'(s(fun(X284,sK1),'_dst_x0x1x0_3'(s(fun(X282,fun(X281,sK1)),X283),s(fun(X284,X282),X285),s(X281,X287))),s(X284,X286))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X288),s(sK1,happ(s(fun(X281,sK1),'const.bool.LET_2'(s(fun(X282,fun(X281,sK1)),X283),s(X282,'const.bool.LET_2'(s(fun(X284,X282),X285),s(X284,X286))))),s(X281,X287)))))) ) | $spl2),
  inference(forward_demodulation,[],[f2081,f91])).
fof(f2081,plain,(
  ( ! [X282,X288,X285,X287,X281,X283,X284,X286] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X288))),s(sK1,happ(s(fun(X284,sK1),'_dst_x0x1x0_3'(s(fun(X282,fun(X281,sK1)),X283),s(fun(X284,X282),X285),s(X281,X287))),s(X284,X286))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X288),s(sK1,happ(s(fun(X281,sK1),'const.bool.LET_2'(s(fun(X282,fun(X281,sK1)),X283),s(X282,'const.bool.LET_2'(s(fun(X284,X282),X285),s(X284,X286))))),s(X281,X287)))))) ) | $spl2),
  inference(superposition,[],[f1469,f145])).
fof(f2156,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X274] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X273,sK0),'const.bool.LET_2'(s(fun(X274,fun(X273,sK0)),X275),s(X274,'const.bool.LET_2'(s(fun(X276,X274),X277),s(X276,X278))))),s(X273,X279))),s(sK1,X280))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X276,sK0),'_dst_x0x1x0_3'(s(fun(X274,fun(X273,sK0)),X275),s(fun(X276,X274),X277),s(X273,X279))),s(X276,X278))))),s(sK1,X280)))) ) | $spl2),
  inference(forward_demodulation,[],[f2155,f91])).
fof(f2155,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X274] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X276,sK0),'_dst_x0x1x0_3'(s(fun(X274,fun(X273,sK0)),X275),s(fun(X276,X274),X277),s(X273,X279))),s(X276,X278))))),s(sK1,X280))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X273,sK0),'const.bool.LET_2'(s(fun(X274,fun(X273,sK0)),X275),s(X274,'const.bool.LET_2'(s(fun(X276,X274),X277),s(X276,X278))))),s(X273,X279))),s(sK1,X280)))) ) | $spl2),
  inference(forward_demodulation,[],[f2080,f91])).
fof(f2080,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X274] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,happ(s(fun(X276,sK0),'_dst_x0x1x0_3'(s(fun(X274,fun(X273,sK0)),X275),s(fun(X276,X274),X277),s(X273,X279))),s(X276,X278))))),s(sK1,X280))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X273,sK0),'const.bool.LET_2'(s(fun(X274,fun(X273,sK0)),X275),s(X274,'const.bool.LET_2'(s(fun(X276,X274),X277),s(X276,X278))))),s(X273,X279))),s(sK1,X280)))) ) | $spl2),
  inference(superposition,[],[f1469,f145])).
fof(f2128,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272,X274] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X274),s(sK1,'const.bool.LET_2'(s(fun(X267,sK1),'_dst_x0x1x0_3'(s(fun(X268,fun(X269,sK1)),X270),s(fun(X267,X268),X271),s(X269,X272))),s(X267,X273))))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X274))),s(sK1,'const.bool.LET_2'(s(fun(X269,sK1),'const.bool.LET_2'(s(fun(X268,fun(X269,sK1)),X270),s(X268,'const.bool.LET_2'(s(fun(X267,X268),X271),s(X267,X273))))),s(X269,X272)))))) ) | $spl2),
  inference(forward_demodulation,[],[f2127,f91])).
fof(f2127,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272,X274] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X274))),s(sK1,'const.bool.LET_2'(s(fun(X269,sK1),'const.bool.LET_2'(s(fun(X268,fun(X269,sK1)),X270),s(X268,'const.bool.LET_2'(s(fun(X267,X268),X271),s(X267,X273))))),s(X269,X272))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X274),s(sK1,happ(s(fun(X267,sK1),'_dst_x0x1x0_3'(s(fun(X268,fun(X269,sK1)),X270),s(fun(X267,X268),X271),s(X269,X272))),s(X267,X273)))))) ) | $spl2),
  inference(forward_demodulation,[],[f2048,f91])).
fof(f2048,plain,(
  ( ! [X269,X271,X267,X273,X268,X270,X272,X274] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X274))),s(sK1,happ(s(fun(X269,sK1),'const.bool.LET_2'(s(fun(X268,fun(X269,sK1)),X270),s(X268,'const.bool.LET_2'(s(fun(X267,X268),X271),s(X267,X273))))),s(X269,X272))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X274),s(sK1,happ(s(fun(X267,sK1),'_dst_x0x1x0_3'(s(fun(X268,fun(X269,sK1)),X270),s(fun(X267,X268),X271),s(X269,X272))),s(X267,X273)))))) ) | $spl2),
  inference(superposition,[],[f1469,f145])).
fof(f2126,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264,X266] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X259,sK0),'_dst_x0x1x0_3'(s(fun(X260,fun(X261,sK0)),X262),s(fun(X259,X260),X263),s(X261,X264))),s(X259,X265))),s(sK1,X266))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X261,sK0),'const.bool.LET_2'(s(fun(X260,fun(X261,sK0)),X262),s(X260,'const.bool.LET_2'(s(fun(X259,X260),X263),s(X259,X265))))),s(X261,X264))))),s(sK1,X266)))) ) | $spl2),
  inference(forward_demodulation,[],[f2125,f91])).
fof(f2125,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264,X266] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X261,sK0),'const.bool.LET_2'(s(fun(X260,fun(X261,sK0)),X262),s(X260,'const.bool.LET_2'(s(fun(X259,X260),X263),s(X259,X265))))),s(X261,X264))))),s(sK1,X266))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X259,sK0),'_dst_x0x1x0_3'(s(fun(X260,fun(X261,sK0)),X262),s(fun(X259,X260),X263),s(X261,X264))),s(X259,X265))),s(sK1,X266)))) ) | $spl2),
  inference(forward_demodulation,[],[f2047,f91])).
fof(f2047,plain,(
  ( ! [X261,X263,X259,X265,X260,X262,X264,X266] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,happ(s(fun(X261,sK0),'const.bool.LET_2'(s(fun(X260,fun(X261,sK0)),X262),s(X260,'const.bool.LET_2'(s(fun(X259,X260),X263),s(X259,X265))))),s(X261,X264))))),s(sK1,X266))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X259,sK0),'_dst_x0x1x0_3'(s(fun(X260,fun(X261,sK0)),X262),s(fun(X259,X260),X263),s(X261,X264))),s(X259,X265))),s(sK1,X266)))) ) | $spl2),
  inference(superposition,[],[f1469,f145])).
fof(f2098,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK2,sK4))),s(sK1,sK9(sK5,sK1,sK0))))) | $spl2),
  inference(forward_demodulation,[],[f2096,f91])).
fof(f2096,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),'_dst_x0x1x0_3'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK2,sK4))),s(sK1,sK9(sK5,sK1,sK0))))) | $spl2),
  inference(backward_demodulation,[],[f2011,f1932])).
fof(f1931,plain,(
  s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))) | $spl2),
  inference(superposition,[],[f1517,f91])).
fof(f1933,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0))))))) | $spl2),
  inference(forward_demodulation,[],[f1929,f91])).
fof(f1929,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1517,f141])).
fof(f1932,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))))),s(sK2,sK4)))) | $spl2),
  inference(forward_demodulation,[],[f1930,f91])).
fof(f1930,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))))),s(sK2,sK4)))) | $spl2),
  inference(unit_resulting_resolution,[],[f1517,f141])).
fof(f1822,plain,(
  ( ! [X0,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,X0)),X1),s('type.pair.prod'(sK0,sK1),sK5))) = s(X0,happ(s(fun(sK1,X0),'const.bool.LET_2'(s(fun(sK0,fun(sK1,X0)),X1),s(sK0,sK6))),s(sK1,sK7)))) ) | $spl6),
  inference(superposition,[],[f143,f136])).
fof(f1522,plain,(
  ( ! [X0,X1] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X0))),s(sK1,X1))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X0),s(sK1,X1)))) ) | $spl2),
  inference(superposition,[],[f1469,f91])).
fof(f1530,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X7))),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),X5),s(X4,X6))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X7),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),X5),s(X4,X6)))))) ) | $spl2),
  inference(forward_demodulation,[],[f1529,f889])).
fof(f1529,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X7),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),'const.bool.LET_2'(s(fun(fun(X4,sK1),fun(X4,sK1)),'const.bool.LET_0'),s(fun(X4,sK1),X5))),s(X4,X6))))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X7))),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),X5),s(X4,X6)))))) ) | $spl2),
  inference(forward_demodulation,[],[f1528,f91])).
fof(f1528,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X7))),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),X5),s(X4,X6))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X7),s(sK1,happ(s(fun(X4,sK1),'const.bool.LET_2'(s(fun(fun(X4,sK1),fun(X4,sK1)),'const.bool.LET_0'),s(fun(X4,sK1),X5))),s(X4,X6)))))) ) | $spl2),
  inference(forward_demodulation,[],[f1521,f91])).
fof(f1521,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X7))),s(sK1,'const.bool.LET_2'(s(fun(X4,sK1),X5),s(X4,X6))))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X7),s(sK1,happ(s(fun(X4,sK1),'const.bool.LET_2'(s(fun(fun(X4,sK1),fun(X4,sK1)),'const.bool.LET_0'),s(fun(X4,sK1),X5))),s(X4,X6)))))) ) | $spl2),
  inference(superposition,[],[f1469,f137])).
fof(f1526,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),X5),s(X4,X6))))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),X5),s(X4,X6))),s(sK1,X7)))) ) | $spl2),
  inference(forward_demodulation,[],[f1525,f889])).
fof(f1525,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),'const.bool.LET_2'(s(fun(fun(X4,sK0),fun(X4,sK0)),'const.bool.LET_0'),s(fun(X4,sK0),X5))),s(X4,X6))),s(sK1,X7))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),X5),s(X4,X6))))),s(sK1,X7)))) ) | $spl2),
  inference(forward_demodulation,[],[f1524,f91])).
fof(f1524,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),X5),s(X4,X6))))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X4,sK0),'const.bool.LET_2'(s(fun(fun(X4,sK0),fun(X4,sK0)),'const.bool.LET_0'),s(fun(X4,sK0),X5))),s(X4,X6))),s(sK1,X7)))) ) | $spl2),
  inference(forward_demodulation,[],[f1519,f91])).
fof(f1519,plain,(
  ( ! [X6,X4,X7,X5] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,'const.bool.LET_2'(s(fun(X4,sK0),X5),s(X4,X6))))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,happ(s(fun(X4,sK0),'const.bool.LET_2'(s(fun(fun(X4,sK0),fun(X4,sK0)),'const.bool.LET_0'),s(fun(X4,sK0),X5))),s(X4,X6))),s(sK1,X7)))) ) | $spl2),
  inference(superposition,[],[f1469,f137])).
fof(f1517,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK8(sK5,sK1,sK0)))),s(sK1,sK9(sK5,sK1,sK0)))) | $spl2),
  inference(unit_resulting_resolution,[],[f90,f1469])).
fof(f1469,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7)))) ) | $spl2),
  inference(forward_demodulation,[],[f132,f91])).
fof(f1419,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X3))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl4),
  inference(superposition,[],[f1197,f156])).
fof(f1418,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X2))),s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | $spl4),
  inference(superposition,[],[f1197,f724])).
fof(f1455,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X24,bool),X25),s(X24,X26))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),'const.bool.LET_2'(s(fun(fun(X24,bool),fun(X24,bool)),'const.bool.LET_0'),s(fun(X24,bool),X25))),s(X24,X26)))) ) | $spl4),
  inference(forward_demodulation,[],[f1417,f696])).
fof(f1417,plain,(
  ( ! [X26,X24,X25] : (s(bool,'const.bool.LET_2'(s(fun(X24,bool),X25),s(X24,X26))) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),'const.bool.LET_2'(s(fun(fun(X24,bool),fun(X24,bool)),'const.bool.LET_0'),s(fun(X24,bool),X25))),s(X24,X26)))) ) | $spl4),
  inference(superposition,[],[f1197,f137])).
fof(f1454,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.bool.LET_2'(s(fun(X21,bool),X22),s(X21,X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) ) | $spl4),
  inference(forward_demodulation,[],[f1416,f696])).
fof(f1416,plain,(
  ( ! [X23,X21,X22] : (s(bool,'const.bool.LET_2'(s(fun(X21,bool),X22),s(X21,X23))) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) ) | $spl4),
  inference(superposition,[],[f1197,f91])).
fof(f1400,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f91,f1197])).
fof(f1401,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.bool.LET_2'(s(fun(fun(bool,bool),fun(bool,bool)),'const.bool.LET_0'),s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))),s(bool,'T_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f137,f1197])).
fof(f1197,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1189])).
fof(f1189,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(equality_factoring,[],[f724])).
fof(f1183,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(equality_factoring,[],[f724])).
fof(f1203,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0)))) ) | $spl4),
  inference(trivial_inequality_removal,[],[f1181])).
fof(f1181,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,X0) | s(bool,X0) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0)))) ) | $spl4),
  inference(equality_factoring,[],[f724])).
fof(f1180,plain,(
  ( ! [X101,X99,X100,X98] : (s(X100,f(s(bool,'T_0'))) != s(X100,g(s(bool,'T_0'))) | s(fun(bool,X100),X98) = s(fun(bool,X100),X99) | s(bool,X101) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X101))),s(bool,X101))) | s(bool,X101) = s(bool,sK10(X98,X99,X100,bool))) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1179,plain,(
  ( ! [X97,X95,X96] : (s('type.pair.prod'(X96,bool),X95) = s('type.pair.prod'(X96,bool),'const.pair.,_2'(s(X96,sK8(X95,bool,X96)),s(bool,'T_0'))) | s(bool,X97) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X97))),s(bool,X97))) | s(bool,X97) = s(bool,sK9(X95,bool,X96))) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1178,plain,(
  ( ! [X94,X92,X93] : (s('type.pair.prod'(bool,X93),X92) = s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'T_0'),s(X93,sK9(X92,X93,bool)))) | s(bool,X94) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X94))),s(bool,X94))) | s(bool,X94) = s(bool,sK8(X92,X93,bool))) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1175,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,f(s(X80,sK10(X81,X82,bool,X80)))) | s(fun(X80,bool),X81) = s(fun(X80,bool),X82) | s(bool,X83) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X83))),s(bool,X83))) | s(bool,X83) = s(bool,g(s(X80,sK10(X81,X82,bool,X80))))) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1173,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X71)))) | s(bool,X70) != s(bool,X71) | s(bool,X72) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X72))),s(bool,X72))) | s(bool,X70) = s(bool,X72)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1172,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,'T_0')))) | s(bool,X67) != s(bool,X68) | s(bool,X69) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X69))) | s(bool,X67) = s(bool,X69)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1171,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X65)))) | s(bool,X64) = s(bool,X65) | s(bool,X66) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X66))),s(bool,X66))) | s(bool,X64) = s(bool,X66)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1170,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'T_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X63) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X63))),s(bool,X63))) | s(bool,X61) = s(bool,X63)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1343,plain,(
  ( ! [X47,X48,X46,X49] : (s(X46,X47) = s(X46,X48) | s(bool,X49) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X49))),s(bool,X49))) | s(bool,X49) = s(bool,'const.bool.LET_2'(s(fun(X46,bool),'const.bool.LET_2'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,X47))),s(X46,X48)))) ) | $spl4),
  inference(subsumption_resolution,[],[f1166,f78])).
fof(f1166,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | s(X46,X47) = s(X46,X48) | s(bool,X49) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X49))),s(bool,X49))) | s(bool,X49) = s(bool,'const.bool.LET_2'(s(fun(X46,bool),'const.bool.LET_2'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,X47))),s(X46,X48)))) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1342,plain,(
  ( ! [X42] : (s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42))) | s(bool,'F_0') = s(bool,X42)) ) | $spl4),
  inference(subsumption_resolution,[],[f1163,f78])).
fof(f1163,plain,(
  ( ! [X42] : (~p(s(bool,'T_0')) | s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42))) | s(bool,'F_0') = s(bool,X42)) ) | $spl4),
  inference(superposition,[],[f80,f724])).
fof(f1340,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,X13) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X13))),s(bool,X13))) | s(bool,X11) = s(bool,X13)) ) | $spl4),
  inference(subsumption_resolution,[],[f1151,f78])).
fof(f1151,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | s(bool,X11) != s(bool,X12) | p(s(bool,X12)) | s(bool,X13) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X13))),s(bool,X13))) | s(bool,X11) = s(bool,X13)) ) | $spl4),
  inference(superposition,[],[f128,f724])).
fof(f1339,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10))) | s(bool,X8) = s(bool,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f1150,f78])).
fof(f1150,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10))) | s(bool,X8) = s(bool,X10)) ) | $spl4),
  inference(superposition,[],[f89,f724])).
fof(f1145,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X56))),s(bool,X57)))) | s(bool,'T_0') != s(bool,X57) | s(bool,X58) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,X58))) | s(bool,X56) = s(bool,X58)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1144,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X54))),s(bool,X53)))) | s(bool,'T_0') != s(bool,X54) | s(bool,X55) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X55))),s(bool,X55))) | s(bool,X53) = s(bool,X55)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1126,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(superposition,[],[f696,f724])).
fof(f1125,plain,(
  ( ! [X109,X107,X108,X106] : (s(X108,f(s(bool,X109))) != s(X108,g(s(bool,X109))) | s(fun(bool,X108),X106) = s(fun(bool,X108),X107) | s(bool,X109) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X109))),s(bool,X109))) | s(bool,'T_0') = s(bool,sK10(X106,X107,X108,bool))) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1124,plain,(
  ( ! [X103,X105,X104] : (s('type.pair.prod'(X104,bool),X103) = s('type.pair.prod'(X104,bool),'const.pair.,_2'(s(X104,sK8(X103,bool,X104)),s(bool,X105))) | s(bool,X105) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X105))),s(bool,X105))) | s(bool,'T_0') = s(bool,sK9(X103,bool,X104))) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1123,plain,(
  ( ! [X101,X102,X100] : (s('type.pair.prod'(bool,X101),X100) = s('type.pair.prod'(bool,X101),'const.pair.,_2'(s(bool,X102),s(X101,sK9(X100,X101,bool)))) | s(bool,X102) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X102))),s(bool,X102))) | s(bool,'T_0') = s(bool,sK8(X100,X101,bool))) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1120,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X91) != s(bool,f(s(X88,sK10(X89,X90,bool,X88)))) | s(fun(X88,bool),X89) = s(fun(X88,bool),X90) | s(bool,X91) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X91))) | s(bool,'T_0') = s(bool,g(s(X88,sK10(X89,X90,bool,X88))))) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1118,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,X77) != s(bool,X79) | s(bool,X78) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X78))) | s(bool,'T_0') = s(bool,X77)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1117,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X76))),s(bool,X75)))) | s(bool,X74) != s(bool,X76) | s(bool,X75) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X75))),s(bool,X75))) | s(bool,'T_0') = s(bool,X74)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1116,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X72))),s(bool,X73)))) | s(bool,X71) = s(bool,X73) | s(bool,X72) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X72))),s(bool,X72))) | s(bool,'T_0') = s(bool,X71)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1115,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X70))),s(bool,X69)))) | s(bool,X68) = s(bool,X70) | s(bool,X69) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X69))) | s(bool,'T_0') = s(bool,X68)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1087,plain,(
  ( ! [X109,X107,X108,X106] : (s(X108,f(s(bool,X109))) != s(X108,g(s(bool,X109))) | s(fun(bool,X108),X106) = s(fun(bool,X108),X107) | s(bool,sK10(X106,X107,X108,bool)) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK10(X106,X107,X108,bool)))),s(bool,sK10(X106,X107,X108,bool)))) | s(bool,'T_0') = s(bool,X109)) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1086,plain,(
  ( ! [X103,X105,X104] : (s('type.pair.prod'(X104,bool),X103) = s('type.pair.prod'(X104,bool),'const.pair.,_2'(s(X104,sK8(X103,bool,X104)),s(bool,X105))) | s(bool,sK9(X103,bool,X104)) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK9(X103,bool,X104)))),s(bool,sK9(X103,bool,X104)))) | s(bool,'T_0') = s(bool,X105)) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1085,plain,(
  ( ! [X101,X102,X100] : (s('type.pair.prod'(bool,X101),X100) = s('type.pair.prod'(bool,X101),'const.pair.,_2'(s(bool,X102),s(X101,sK9(X100,X101,bool)))) | s(bool,sK8(X100,X101,bool)) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,sK8(X100,X101,bool)))),s(bool,sK8(X100,X101,bool)))) | s(bool,'T_0') = s(bool,X102)) ) | $spl4),
  inference(superposition,[],[f90,f724])).
fof(f1082,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X91) != s(bool,f(s(X88,sK10(X89,X90,bool,X88)))) | s(fun(X88,bool),X89) = s(fun(X88,bool),X90) | s(bool,g(s(X88,sK10(X89,X90,bool,X88)))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,g(s(X88,sK10(X89,X90,bool,X88)))))),s(bool,g(s(X88,sK10(X89,X90,bool,X88)))))) | s(bool,'T_0') = s(bool,X91)) ) | $spl4),
  inference(superposition,[],[f93,f724])).
fof(f1080,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,X77) != s(bool,X79) | s(bool,X77) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X77))),s(bool,X77))) | s(bool,'T_0') = s(bool,X78)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1079,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X76))),s(bool,X75)))) | s(bool,X74) != s(bool,X76) | s(bool,X74) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X74))),s(bool,X74))) | s(bool,'T_0') = s(bool,X75)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1078,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X72))),s(bool,X73)))) | s(bool,X71) = s(bool,X73) | s(bool,X71) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X71))),s(bool,X71))) | s(bool,'T_0') = s(bool,X72)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1077,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X70))),s(bool,X69)))) | s(bool,X68) = s(bool,X70) | s(bool,X68) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,X68))) | s(bool,'T_0') = s(bool,X69)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1048,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X57))),s(bool,X58)))) | s(bool,X58) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X57))),s(bool,X57))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1047,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X55))),s(bool,X54)))) | s(bool,X55) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X54))),s(bool,X54))) | s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1046,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X51))),s(bool,X52)))) | s(bool,X52) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X51))),s(bool,X51))) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1045,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X49))),s(bool,X48)))) | s(bool,X49) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X48))) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl4),
  inference(superposition,[],[f141,f724])).
fof(f1303,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X17))),s(bool,X17))) | p(s(bool,X18)) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl4),
  inference(subsumption_resolution,[],[f1034,f180])).
fof(f1034,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,X18) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X17))),s(bool,X17))) | p(s(bool,X18)) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl4),
  inference(superposition,[],[f128,f724])).
fof(f1302,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X14))),s(bool,X14))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl4),
  inference(subsumption_resolution,[],[f1033,f180])).
fof(f1033,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X15) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X14))),s(bool,X14))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl4),
  inference(superposition,[],[f89,f724])).
fof(f1026,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,X32) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X32))),s(bool,X32))) | s(bool,X32) = s(bool,happ(s(fun(X29,bool),'const.bool.LET_2'(s(fun(fun(X29,bool),fun(X29,bool)),'const.bool.LET_0'),s(fun(X29,bool),X30))),s(X29,X31)))) ) | $spl4),
  inference(superposition,[],[f724,f137])).
fof(f1025,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X25,bool),X26),s(X25,X27))) | s(bool,X28) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X28))),s(bool,X28))) | s(bool,X28) = s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) ) | $spl4),
  inference(superposition,[],[f724,f91])).
fof(f1009,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X6))),s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl4),
  inference(superposition,[],[f724,f156])).
fof(f1008,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X2))),s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X3))),s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f1007,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X43) = s(bool,'const.bool.LET_2'(s(fun(X40,bool),X41),s(X40,X42))) | s(bool,X43) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X43))),s(bool,X43))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),'const.bool.LET_2'(s(fun(fun(X40,bool),fun(X40,bool)),'const.bool.LET_0'),s(fun(X40,bool),X41))),s(X40,X42)))) ) | $spl4),
  inference(superposition,[],[f724,f137])).
fof(f1006,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X39) = s(bool,'const.bool.LET_2'(s(fun(X36,bool),X37),s(X36,X38))) | s(bool,X39) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X39))),s(bool,X39))) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) ) | $spl4),
  inference(superposition,[],[f724,f91])).
fof(f1275,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))))),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0)))))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1274])).
fof(f1274,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))))),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(condensation,[],[f971])).
fof(f971,plain,(
  ( ! [X17,X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,X15))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,X15))))),s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,X15))))) | s(bool,'T_0') = s(bool,X16)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f965,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,'const.bool.LET_2'(s(fun(X43,bool),X44),s(X43,X45))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X43,bool),X44),s(X43,X45))))),s(bool,'const.bool.LET_2'(s(fun(X43,bool),X44),s(X43,X45))))) | s(bool,X46) = s(bool,happ(s(fun(X43,bool),'const.bool.LET_2'(s(fun(fun(X43,bool),fun(X43,bool)),'const.bool.LET_0'),s(fun(X43,bool),X44))),s(X43,X45))) | s(bool,'T_0') = s(bool,X46)) ) | $spl4),
  inference(superposition,[],[f724,f137])).
fof(f824,plain,(
  ( ! [X14,X13] : (s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X13))),s(bool,X13)))) ) | $spl4),
  inference(subsumption_resolution,[],[f805,f80])).
fof(f805,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X13))),s(bool,X13)))) ) | $spl4),
  inference(superposition,[],[f697,f81])).
fof(f724,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(superposition,[],[f696,f156])).
fof(f697,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0)))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(superposition,[],[f626,f156])).
fof(f696,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f626,f158])).
fof(f626,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl4),
  inference(forward_demodulation,[],[f598,f581])).
fof(f598,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4))))),s(bool,'T_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f581,f139])).
fof(f581,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4))) | $spl4),
  inference(unit_resulting_resolution,[],[f547,f158])).
fof(f547,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4)))) | $spl4),
  inference(forward_demodulation,[],[f513,f466])).
fof(f513,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))))),s(sK2,sK4)))) | $spl4),
  inference(unit_resulting_resolution,[],[f466,f139])).
fof(f466,plain,(
  s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(superposition,[],[f91,f155])).
fof(f155,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(forward_demodulation,[],[f134,f91])).
fof(f4670,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(global_subsumption,[],[f77,f76,f75,f78,f81,f89,f86,f90,f91,f137,f93,f102,f142,f143,f128,f145,f154,f155,f156,f158,f161,f177,f180,f181,f149,f153,f139,f141,f80,f213,f214,f239,f271,f272,f438,f440,f441,f307,f308,f447,f461,f463,f465,f466,f546,f547,f550,f516,f517,f519,f521,f522,f529,f530,f532,f534,f535,f574,f581,f626,f641,f643,f645,f646,f648,f655,f659,f662,f688,f696,f697,f724,f778,f780,f782,f783,f785,f787,f824,f920,f921,f889,f890,f892,f894,f1284,f965,f1275,f1268,f1257,f1252,f1006,f1007,f1008,f1009,f1243,f1242,f1025,f1026,f1302,f1303,f1045,f1046,f1047,f1048,f1077,f1078,f1079,f1080,f1082,f1085,f1086,f1087,f1220,f1328,f1330,f1215,f1115,f1116,f1117,f1118,f1120,f1123,f1124,f1125,f1126,f1212,f1144,f1145,f1339,f1340,f1342,f1343,f1170,f1171,f1172,f1173,f1175,f1178,f1179,f1180,f1203,f1183,f1197,f1385,f1386,f1375,f1387,f1389,f1401,f1400,f1445,f1454,f1455,f1418,f1419,f1469,f1517,f1526,f1530,f1522,f1655,f1656,f1624,f1661,f1662,f1659,f1625,f1626,f1627,f1876,f1877,f1777,f1822,f1823,f1878,f1879,f1880,f1883,f1884,f1885,f1886,f1887,f1888,f1889,f1890,f1932,f1933,f1931,f2092,f2094,f1938,f1939,f1940,f1941,f1983,f1984,f1985,f2010,f2011,f2098,f2012,f2013,f2100,f2102,f2104,f2106,f2108,f2112,f2114,f2116,f2120,f2126,f2128,f2130,f2132,f2134,f2136,f2138,f2142,f2144,f2146,f2150,f2156,f2158,f2082,f2396,f2397,f2351,f2398,f2399,f2400,f2401,f2402,f2403,f2404,f2405,f2406,f2600,f2601,f2602,f2606,f2608,f2610,f2613,f2542,f2543,f2544,f2616,f2617,f2621,f2623,f2625,f2628,f2632,f2633,f2634,f2638,f2640,f2642,f2645,f2649,f2651,f2653,f2656,f2574,f2575,f2576,f2579,f2581,f2661,f2662,f2592,f2593,f2594,f2816,f2817,f2821,f2823,f2825,f2828,f2832,f2833,f2834,f2838,f2840,f2842,f2845,f2770,f2771,f2772,f2852,f2853,f2854,f2858,f2860,f2862,f2865,f2869,f2871,f2873,f2876,f2793,f2794,f2795,f2798,f2800,f2881,f2882,f2886,f2960,f2961,f2899,f2901,f2902,f2903,f3011,f3013,f3108,f3110,f3122,f3124,f3145,f3146,f3147,f3213,f3214,f3152,f3183,f3233,f3249,f3251,f3360,f3361,f3362,f3366,f3368,f3370,f3376,f3382,f3384,f3386,f3391,f3332,f3333,f3334,f3337,f3339,f3398,f3399,f3349,f3407,f3409,f3584,f3585,f3589,f3591,f3593,f3599,f3605,f3607,f3609,f3614,f3553,f3554,f3555,f3558,f3560,f3630,f3631,f3864,f3865,f3869,f3871,f3873,f3879,f3885,f3887,f3889,f3894,f3825,f3826,f3827,f3830,f3832,f3902,f3903,f136,f4643,f4644])).
fof(f4669,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4645,f4643])).
fof(f4645,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f136,f1469])).
fof(f1301,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f1301_D])).
fof(f1301_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4478,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f4276])).
fof(f4276,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f81,f1695])).
fof(f1695,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(superposition,[],[f1586,f156])).
fof(f1586,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f1566,f158])).
fof(f1566,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl0),
  inference(forward_demodulation,[],[f1534,f1497])).
fof(f1497,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4))) | $spl0),
  inference(unit_resulting_resolution,[],[f1483,f158])).
fof(f1483,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4)))) | $spl0),
  inference(forward_demodulation,[],[f1476,f130])).
fof(f1476,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))))),s(sK2,sK4)))) | $spl0),
  inference(unit_resulting_resolution,[],[f130,f139])).
fof(f1534,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,sK4))))),s(bool,'T_0')))) | $spl0),
  inference(unit_resulting_resolution,[],[f1497,f139])).
fof(f4598,plain,(
  ( ! [X28,X29] : (p(s(bool,X29)) | s(bool,'F_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X28))),s(bool,X28)))) ) | ($spl0 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f4225,f439])).
fof(f4225,plain,(
  ( ! [X28,X29] : (p(s(bool,X29)) | s(bool,'F_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X28))),s(bool,X28))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f179,f1695])).
fof(f4502,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f4168])).
fof(f4168,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl0),
  inference(superposition,[],[f1695,f81])).
fof(f4517,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f4132])).
fof(f4132,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1695,f81])).
fof(f3992,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl40),
  inference(cnf_transformation,[],[f3992_D])).
fof(f3992_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3991,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl5 | $spl6 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f3964,f439])).
fof(f3964,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl5 | $spl6 | $spl11)),
  inference(superposition,[],[f3792,f81])).
fof(f3792,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl5 | $spl6 | $spl11)),
  inference(subsumption_resolution,[],[f3730,f78])).
fof(f3730,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X12)) ) | ($spl5 | $spl6 | $spl11)),
  inference(superposition,[],[f128,f3456])).
fof(f3456,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl6 | $spl11)),
  inference(superposition,[],[f3427,f81])).
fof(f3427,plain,(
  s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'F_0'))) | ($spl5 | $spl6 | $spl11)),
  inference(unit_resulting_resolution,[],[f3299,f158])).
fof(f3299,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'F_0')))) | ($spl5 | $spl6 | $spl11)),
  inference(forward_demodulation,[],[f3256,f3029])).
fof(f3029,plain,(
  s(bool,'F_0') = s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5))))) | ($spl5 | $spl6 | $spl11)),
  inference(unit_resulting_resolution,[],[f164,f2963,f180])).
fof(f2963,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5)))))) | ($spl5 | $spl6)),
  inference(backward_demodulation,[],[f2899,f1474])).
fof(f1474,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.bool.LET_2'(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7)))))) | $spl5),
  inference(forward_demodulation,[],[f1470,f91])).
fof(f1470,plain,(
  ~p(s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f1468,f141])).
fof(f1468,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),'const.bool.LET_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl5),
  inference(forward_demodulation,[],[f133,f91])).
fof(f133,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl5),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f3256,plain,(
  p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'const.bool.LET_2'(s(fun(sK2,bool),'const.bool.LET_2'(s(fun(sK2,fun(sK2,bool)),'=_0'),s(sK2,sK4))),s(sK2,'const.pair.UNCURRY_2'(s(fun(sK0,fun(sK1,sK2)),sK3),s('type.pair.prod'(sK0,sK1),sK5)))))))) | ($spl5 | $spl6 | $spl11)),
  inference(unit_resulting_resolution,[],[f3029,f139])).
fof(f2735,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) ) | $spl38),
  inference(cnf_transformation,[],[f2735_D])).
fof(f2735_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2734,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) ) | ($spl14 | $spl16 | $spl26)),
  inference(subsumption_resolution,[],[f2714,f439])).
fof(f2714,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl14 | $spl26)),
  inference(superposition,[],[f2225,f81])).
fof(f2225,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl14 | $spl26)),
  inference(subsumption_resolution,[],[f2188,f179])).
fof(f2188,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,X10) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,X9)) ) | $spl26),
  inference(superposition,[],[f128,f1301])).
fof(f1453,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12)))) ) | $spl36),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1445,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1408])).
fof(f1408,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | $spl4),
  inference(superposition,[],[f1197,f724])).
fof(f1338,plain,(
  ( ! [X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,X40)))) | s(bool,X40) = s(bool,X42) | s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42)))) ) | $spl34),
  inference(cnf_transformation,[],[f1338_D])).
fof(f1338_D,plain,(
  ( ! [X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,X40)))) | s(bool,X40) = s(bool,X42) | s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1212,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,X40)))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1140])).
fof(f1140,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X40))),s(bool,X40)))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | $spl4),
  inference(superposition,[],[f697,f724])).
fof(f1329,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) | $spl32),
  inference(cnf_transformation,[],[f1329_D])).
fof(f1329_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1215,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,X27)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1097])).
fof(f1097,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X27))),s(bool,X27))) | s(bool,'T_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f158,f724])).
fof(f1330,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | $spl4),
  inference(subsumption_resolution,[],[f1092,f180])).
fof(f1092,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | p(s(bool,X13)) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,'T_0') = s(bool,X11)) ) | $spl4),
  inference(superposition,[],[f128,f724])).
fof(f1328,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) | $spl4),
  inference(condensation,[],[f1327])).
fof(f1327,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X10) | s(bool,X9) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(subsumption_resolution,[],[f1091,f180])).
fof(f1091,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s(bool,X9) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X9))) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f89,f724])).
fof(f1220,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1088])).
fof(f1088,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f81,f724])).
fof(f1318,plain,(
  ( ! [X57,X56,X55] : (s(X55,X56) != s(X55,X57) | s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))))),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57)))))) ) | $spl30),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  ( ! [X57,X56,X55] : (s(X55,X56) != s(X55,X57) | s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))))),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1317,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X58)) | s(X55,X56) != s(X55,X57) | s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))))),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57)))))) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f1072,f439])).
fof(f1072,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X58)) | s(X55,X56) != s(X55,X57) | s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))))),s(bool,'const.bool.LET_2'(s(fun(X55,bool),'const.bool.LET_2'(s(fun(X55,fun(X55,bool)),'=_0'),s(X55,X56))),s(X55,X57))))) | s(bool,'T_0') = s(bool,X58)) ) | $spl4),
  inference(superposition,[],[f139,f724])).
fof(f1314,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) | $spl28),
  inference(cnf_transformation,[],[f1314_D])).
fof(f1314_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1313,plain,(
  ( ! [X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33)))) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f1062,f439])).
fof(f1062,plain,(
  ( ! [X33,X34] : (p(s(bool,X34)) | s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X33))),s(bool,X33))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f439,f724])).
fof(f1311,plain,(
  ( ! [X28,X29] : (p(s(bool,X29)) | s(bool,'F_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X28))),s(bool,X28)))) ) | ($spl4 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f1060,f439])).
fof(f1060,plain,(
  ( ! [X28,X29] : (p(s(bool,X29)) | s(bool,'F_0') = s(bool,X28) | s(bool,X28) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X28))),s(bool,X28))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl4 | $spl14)),
  inference(superposition,[],[f179,f724])).
fof(f1242,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1023])).
fof(f1023,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X21))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl4),
  inference(superposition,[],[f724,f81])).
fof(f1292,plain,(
  ( ! [X23,X25] : (s(bool,X23) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,X23))) | s(bool,X23) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X25)))) ) | $spl24),
  inference(cnf_transformation,[],[f1292_D])).
fof(f1292_D,plain,(
  ( ! [X23,X25] : (s(bool,X23) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,X23))) | s(bool,X23) = s(bool,X25) | s(bool,X25) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X25)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1243,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1021])).
fof(f1021,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X12))),s(bool,X12))) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f1252,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X19))),s(bool,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X18))),s(bool,X18)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f997])).
fof(f997,plain,(
  ( ! [X19,X17,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X19))),s(bool,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X18))),s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f1257,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f991])).
fof(f991,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f724,f81])).
fof(f1268,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X20))),s(bool,X20))) | s(bool,X22) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X22))),s(bool,X22))) | s(bool,X20) = s(bool,X22)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f981])).
fof(f981,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X20))),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X22))),s(bool,X22))) | s(bool,X20) = s(bool,X22)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f1284,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X25))) | s(bool,X23) = s(bool,X25)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f957])).
fof(f957,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X25))),s(bool,X25))) | s(bool,X23) = s(bool,X25)) ) | $spl4),
  inference(superposition,[],[f724,f724])).
fof(f866,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f826,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f572,f141])).
fof(f826,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f80,f572,f128])).
fof(f865,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f835,f78])).
fof(f835,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f164,f572,f89])).
fof(f864,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f841,f78])).
fof(f841,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f164,f572,f89])).
fof(f848,plain,(
  $false | ($spl1 | $spl20)),
  inference(unit_resulting_resolution,[],[f131,f572,f141])).
fof(f846,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f164,f572,f158])).
fof(f844,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f572,f164,f572,f89])).
fof(f842,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f78,f164,f572,f89])).
fof(f838,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f572,f164,f572,f89])).
fof(f836,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f78,f164,f572,f89])).
fof(f825,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f80,f572])).
fof(f822,plain,(
  ( ! [X7,X9] : (s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X7))),s(bool,X7))) | s(bool,'T_0') = s(bool,X9)) ) | $spl22),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  ( ! [X7,X9] : (s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X7))),s(bool,X7))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f823,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10)))) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f804,f439])).
fof(f804,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X10))),s(bool,X10)))) ) | $spl4),
  inference(superposition,[],[f697,f156])).
fof(f821,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X8)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X7))),s(bool,X7)))) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f803,f439])).
fof(f803,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X8)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(bool,bool),'const.bool.LET_2'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X7))),s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f697,f156])).
fof(f571,plain,(
  ( ! [X10,X8,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),'const.bool.LET_2'(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) | $spl18),
  inference(cnf_transformation,[],[f571_D])).
fof(f571_D,plain,(
  ( ! [X10,X8,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),'const.bool.LET_2'(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f573,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(X12,X13) != s(X12,X14) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X12,bool),'const.bool.LET_2'(s(fun(X12,fun(X12,bool)),'=_0'),s(X12,X13))),s(X12,X14)))) ) | $spl16),
  inference(subsumption_resolution,[],[f544,f439])).
fof(f544,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(X12,X13) != s(X12,X14) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X12,bool),'const.bool.LET_2'(s(fun(X12,fun(X12,bool)),'=_0'),s(X12,X13))),s(X12,X14)))) )),
  inference(superposition,[],[f139,f156])).
fof(f570,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),'const.bool.LET_2'(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10)))) ) | $spl16),
  inference(subsumption_resolution,[],[f543,f439])).
fof(f543,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,'const.bool.LET_2'(s(fun(X8,bool),'const.bool.LET_2'(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f139,f156])).
fof(f439,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f447,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f446])).
fof(f446,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,X14)) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f313,f180])).
fof(f313,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X14)) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f128,f156])).
fof(f438,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f437])).
fof(f437,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,X14)) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f291,f180])).
fof(f291,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X14)) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f128,f156])).
fof(f187,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f182,f78])).
fof(f182,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f178,f80])).
fof(f179,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f181,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f175,f158])).
fof(f175,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f86,f81])).
fof(f177,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f173,f158])).
fof(f173,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f86,f81])).
fof(f162,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f164,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f161,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f81])).
fof(f130,plain,(
  s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) | $spl0),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f136,plain,(
  s('type.pair.prod'(sK0,sK1),sK5) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | $spl6),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  s('type.pair.prod'(sK0,sK1),sK5) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f75,plain,(
  s('type.pair.prod'(sK0,sK1),sK5) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) | s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5)))),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  (s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) | (s('type.pair.prod'(sK0,sK1),sK5) = s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK6),s(sK1,sK7))) & s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))))) & (s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) | ! [X6,X7] : (s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7))) | s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X2,X4) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) | ? [X8,X9] : (s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X8),s(X1,X9))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X8))),s(X1,X9))))) & (s(X2,X4) != s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) | ! [X6,X7] : (s('type.pair.prod'(X0,X1),X5) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) | s(X2,X4) != s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7))))))),
  inference(rectify,[],[f62])).
fof(f62,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X2,X4) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) | ? [X6,X7] : (s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7))))) & (s(X2,X4) != s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) | ! [X6,X7] : (s('type.pair.prod'(X0,X1),X5) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) | s(X2,X4) != s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7))))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X2,X4) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) <~> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X2,X4) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : (s(X2,X4) = s(X2,'const.bool.LET_2'(s(fun('type.pair.prod'(X0,X1),X2),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,X2)),X3))),s('type.pair.prod'(X0,X1),X5))) <=> ? [X6,X7] : (s('type.pair.prod'(X0,X1),X5) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))) & s(X2,X4) = s(X2,happ(s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),X3),s(X0,X6))),s(X1,X7)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : (s(X8,X10) = s(X8,'const.bool.LET_2'(s(fun('type.pair.prod'(X6,X7),X8),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,X8)),X9))),s('type.pair.prod'(X6,X7),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X6,X7),X11) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13))) & s(X8,X10) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X9),s(X6,X12))),s(X7,X13)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : (s(X8,X10) = s(X8,'const.bool.LET_2'(s(fun('type.pair.prod'(X6,X7),X8),'const.pair.UNCURRY_1'(s(fun(X6,fun(X7,X8)),X9))),s('type.pair.prod'(X6,X7),X11))) <=> ? [X12,X13] : (s('type.pair.prod'(X6,X7),X11) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13))) & s(X8,X10) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X9),s(X6,X12))),s(X7,X13)))))),
  file('thm.evalProps.evalProps_43_prove_ltb','thm.evalProps.evalProps_43_prove_ltb')).
fof(f134,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | $spl4),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f76,plain,(
  s(sK2,sK4) = s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,sK6))),s(sK1,sK7))) | s(sK2,sK4) = s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5)))),
  inference(cnf_transformation,[],[f64])).
fof(f131,plain,(
  s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) | $spl1),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f132,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f77,plain,(
  ( ! [X6,X7] : (s(sK2,sK4) != s(sK2,happ(s(fun(sK1,sK2),happ(s(fun(sK0,fun(sK1,sK2)),sK3),s(sK0,X6))),s(sK1,X7))) | s('type.pair.prod'(sK0,sK1),sK5) != s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,X6),s(sK1,X7))) | s(sK2,sK4) != s(sK2,'const.bool.LET_2'(s(fun('type.pair.prod'(sK0,sK1),sK2),'const.pair.UNCURRY_1'(s(fun(sK0,fun(sK1,sK2)),sK3))),s('type.pair.prod'(sK0,sK1),sK5)))) )),
  inference(cnf_transformation,[],[f64])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 194
% Passive clauses: 2144
% Generated clauses: 18979
% Final active clauses: 96
% Final passive clauses: 1007
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 700
% Trivial inequalities: 196
% Fw subsumption resolutions: 798
% Fw demodulations: 4651
% Bw demodulations: 524
% Condensations: 944
% Global subsumptions: 4
% 
% Simple tautologies: 172
% Equational tautologies: 108
% Forward subsumptions: 9191
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 54
% 
% Unit resulting resolution: 755
% Factoring: 2
% Forward superposition: 5715
% Backward superposition: 3923
% Self superposition: 174
% Equality factoring: 70
% Equality resolution: 9
% 
% Split clauses: 47
% Split components: 100
% SAT solver clauses: 2276
% SAT solver unit clauses: 70
% SAT solver binary clauses: 402
% 
% Sat splits: 47
% Sat splitting refutations: 32
% 
% Memory used [KB]: 22771
% Time elapsed: 1.227 s
% ------------------------------
% ------------------------------
% Success in time 2.57 s
