Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closSem
% (50894)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 3760
% Generated clauses: 14941
% Final active clauses: 89
% Final passive clauses: 3542
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 693
% Trivial inequalities: 182
% Fw subsumption resolutions: 1726
% Fw demodulations: 993
% Bw demodulations: 2
% 
% Simple tautologies: 232
% Equational tautologies: 79
% Forward subsumptions: 6205
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 189
% Factoring: 3
% Forward superposition: 5113
% Backward superposition: 5397
% Self superposition: 279
% Equality factoring: 200
% Equality resolution: 15
% 
% Split clauses: 341
% Split components: 705
% SAT solver clauses: 1023
% SAT solver unit clauses: 18
% SAT solver binary clauses: 185
% 
% Sat splits: 341
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1022
% TWLsolver calls for satisfiability: 140
% 
% Memory used [KB]: 13304
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closSem
% SZS output start Proof for thm.closSem
fof(f3795,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3794,f2529,f2530,f92,f120,f70,f98,f96,f2672,f3739,f3332,f3333,f82,f112,f75,f105,f74,f104,f69,f97,f95])).
fof(f95,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) | $spl0),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f97,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),sK2) | $spl2),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f69,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),sK2) | s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2)))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  (s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) | s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),sK2)) & (s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) != s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) | s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2] : ((s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) & (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) != s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ? [X0,X1,X2] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) <~> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X2))) <=> s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X7))) = s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X8))) <=> s('type.list.list'(X6),X7) = s('type.list.list'(X6),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X7))) = s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X8))) <=> s('type.list.list'(X6),X7) = s('type.list.list'(X6),X8))),
  file('thm.closSem.ref_|49||49||split|0','thm.closSem.ref_|49||49||split|0')).
fof(f104,plain,(
  ( ! [X0,X1] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),happ(s(fun('type.list.list'(X0),'type.closSem.ref'(X0)),'const.closSem.ValueArray_0'),s('type.list.list'(X0),X1)))) ) | $spl10),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  ( ! [X0,X1] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),happ(s(fun('type.list.list'(X0),'type.closSem.ref'(X0)),'const.closSem.ValueArray_0'),s('type.list.list'(X0),X1)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f74,plain,(
  ( ! [X0,X1] : (s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),happ(s(fun('type.list.list'(X0),'type.closSem.ref'(X0)),'const.closSem.ValueArray_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),happ(s(fun('type.list.list'(X0),'type.closSem.ref'(X0)),'const.closSem.ValueArray_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s('type.closSem.ref'(X0),'const.closSem.ValueArray_1'(s('type.list.list'(X0),X1))) = s('type.closSem.ref'(X0),happ(s(fun('type.list.list'(X0),'type.closSem.ref'(X0)),'const.closSem.ValueArray_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9] : s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X9))) = s('type.closSem.ref'(X6),happ(s(fun('type.list.list'(X6),'type.closSem.ref'(X6)),'const.closSem.ValueArray_0'),s('type.list.list'(X6),X9)))),
  file('thm.closSem.ref_|49||49||split|0','ARITY_0_1_const.closSem.ValueArray')).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X11] : s(X6,X7) = s(X6,'const.ind_type.FCONS_3'(s(X6,X7),s(fun('type.num.num',X6),X11),s('type.num.num','const.num.|48|_0')))),
  file('thm.closSem.ref_|49||49||split|0','thm.ind_type.FCONS|split|0')).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) | $spl26),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X9,X20,X21] : s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X20),s('type.num.num',X21))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X9))),s(fun('type.num.num',X6),X20))),s('type.num.num',X21)))),
  file('thm.closSem.ref_|49||49||split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f3333,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK1),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5))) = s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK2),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5)))) ) | $spl666),
  inference(cnf_transformation,[],[f3333_D])).
fof(f3333_D,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK1),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5))) = s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK2),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5)))) ) <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f3332,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK1),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5))) = s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK2),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5)))) ) | ($spl26 | $spl552)),
  inference(forward_demodulation,[],[f3219,f112])).
fof(f3219,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK2),s(fun('type.num.num','type.list.list'(sK0)),X4),s('type.num.num',X5))) = s('type.list.list'(sK0),happ(s(fun('type.num.num','type.list.list'(sK0)),happ(s(fun(fun('type.num.num','type.list.list'(sK0)),fun('type.num.num','type.list.list'(sK0))),happ(s(fun('type.list.list'(sK0),fun(fun('type.num.num','type.list.list'(sK0)),fun('type.num.num','type.list.list'(sK0)))),'const.ind_type.FCONS_0'),s('type.list.list'(sK0),sK1))),s(fun('type.num.num','type.list.list'(sK0)),X4))),s('type.num.num',X5)))) ) | ($spl26 | $spl552)),
  inference(superposition,[],[f112,f2530])).
fof(f3739,plain,(
  $false | ($spl3 | $spl12 | $spl666)),
  inference(subsumption_resolution,[],[f3738,f98])).
fof(f3738,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),sK2) | ($spl12 | $spl666)),
  inference(forward_demodulation,[],[f3719,f105])).
fof(f3719,plain,(
  ( ! [X2] : (s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.ind_type.FCONS_3'(s('type.list.list'(sK0),sK1),s(fun('type.num.num','type.list.list'(sK0)),X2),s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl666)),
  inference(superposition,[],[f3333,f105])).
fof(f2672,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) = s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) | ($spl10 | $spl552)),
  inference(forward_demodulation,[],[f2633,f104])).
fof(f2633,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) = s('type.closSem.ref'(sK0),happ(s(fun('type.list.list'(sK0),'type.closSem.ref'(sK0)),'const.closSem.ValueArray_0'),s('type.list.list'(sK0),sK1))) | ($spl10 | $spl552)),
  inference(superposition,[],[f104,f2530])).
fof(f96,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) != s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) | $spl1),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) != s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f98,plain,(
  s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),sK2) | $spl3),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f70,plain,(
  s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),sK2) | s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) != s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK2)))),
  inference(cnf_transformation,[],[f65])).
fof(f120,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),X3),s('type.list.list'(X1),X2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2))),s(fun('type.list.list'(X1),X0),X3),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X4)))) ) | $spl42),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),X3),s('type.list.list'(X1),X2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2))),s(fun('type.list.list'(X1),X0),X3),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X4)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f92,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),X3),s('type.list.list'(X1),X2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2))),s(fun('type.list.list'(X1),X0),X3),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X4)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.list.list'(X1),X0),X3),s('type.list.list'(X1),X2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2))),s(fun('type.list.list'(X1),X0),X3),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X4)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun('type.list.list'(X1),X0),X3),s('type.list.list'(X1),X2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X1),'const.closSem.ValueArray_1'(s('type.list.list'(X1),X2))),s(fun('type.list.list'(X1),X0),X3),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X6] : ! [X7,X11,X12] : s(X10,happ(s(fun('type.list.list'(X6),X10),X11),s('type.list.list'(X6),X7))) = s(X10,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(X6),'const.closSem.ValueArray_1'(s('type.list.list'(X6),X7))),s(fun('type.list.list'(X6),X10),X11),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X10),X12)))),
  file('thm.closSem.ref_|49||49||split|0','thm.closSem.ref_case_def|split|0')).
fof(f2530,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK2)))) ) | $spl552),
  inference(cnf_transformation,[],[f2530_D])).
fof(f2530_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK2)))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f2529,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK2)))) ) | ($spl0 | $spl42)),
  inference(forward_demodulation,[],[f2448,f120])).
fof(f2448,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.list.list'(sK0),X0),X1),s('type.list.list'(sK0),sK2))) = s(X0,'const.closSem.ref_CASE_3'(s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))),s(fun('type.list.list'(sK0),X0),X1),s(fun('type.list.list'('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))),X0),X2)))) ) | ($spl0 | $spl42)),
  inference(superposition,[],[f120,f95])).
fof(f3794,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f3793])).
fof(f3793,plain,(
  s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) != s('type.closSem.ref'(sK0),'const.closSem.ValueArray_1'(s('type.list.list'(sK0),sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f96,f97])).
% SZS output end Proof for thm.closSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 51
% Passive clauses: 360
% Generated clauses: 3673
% Final active clauses: 41
% Final passive clauses: 270
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 178
% Trivial inequalities: 15
% Fw subsumption resolutions: 103
% Fw demodulations: 246
% Bw demodulations: 42
% Condensations: 322
% 
% Simple tautologies: 49
% Equational tautologies: 29
% Forward subsumptions: 1733
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 87
% Factoring: 1
% Forward superposition: 1095
% Backward superposition: 1044
% Self superposition: 83
% Equality factoring: 52
% Equality resolution: 1
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 664
% SAT solver unit clauses: 39
% SAT solver binary clauses: 25
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 10
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 662
% TWLsolver calls for satisfiability: 114
% 
% Memory used [KB]: 7291
% Time elapsed: 0.207 s
% ------------------------------
% ------------------------------
% Success in time 0.84 s
