Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stack_removeProof
% (105995)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 956
% Passive clauses: 6360
% Generated clauses: 29045
% Final active clauses: 828
% Final passive clauses: 4858
% Input formulas: 260
% Initial clauses: 530
% 
% Duplicate literals: 186
% Trivial inequalities: 98
% Fw subsumption resolutions: 2069
% Fw demodulations: 2993
% Bw demodulations: 75
% 
% Simple tautologies: 272
% Equational tautologies: 36
% Forward subsumptions: 16642
% 
% Binary resolution: 6223
% Factoring: 559
% Forward superposition: 6768
% Backward superposition: 9288
% Self superposition: 34
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 330
% Split components: 662
% SAT solver clauses: 506
% SAT solver unit clauses: 15
% SAT solver binary clauses: 107
% 
% Sat splits: 330
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 501
% TWLsolver calls for satisfiability: 296
% 
% Memory used [KB]: 12665
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stack_removeProof
% (106016)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 665
% Passive clauses: 3722
% Generated clauses: 24125
% Final active clauses: 598
% Final passive clauses: 2584
% Input formulas: 260
% Initial clauses: 530
% 
% Duplicate literals: 102
% Trivial inequalities: 74
% Fw subsumption resolutions: 932
% Fw demodulations: 3531
% Bw demodulations: 519
% Condensations: 82
% Global subsumptions: 38
% 
% Simple tautologies: 272
% Equational tautologies: 15
% Forward subsumptions: 11101
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2923
% Factoring: 150
% Forward superposition: 6668
% Backward superposition: 4680
% Self superposition: 50
% Equality factoring: 2
% Equality resolution: 18
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 4723
% SAT solver unit clauses: 691
% SAT solver binary clauses: 184
% SAT solver learnt clauses: 77
% SAT solver learnt literals: 19
% 
% Sat splits: 17
% Sat splitting refutations: 25
% 
% TWLsolver clauses: 4571
% TWLsolver calls for satisfiability: 3593
% 
% Memory used [KB]: 20340
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stack_removeProof
% (106046)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2650
% Passive clauses: 22271
% Generated clauses: 138033
% Final active clauses: 2247
% Final passive clauses: 5830
% Input formulas: 260
% Initial clauses: 530
% Discarded non-redundant clauses: 26625
% 
% Duplicate literals: 10
% Trivial inequalities: 62
% Fw subsumption resolutions: 773
% Fw demodulations: 29094
% Bw demodulations: 1112
% Condensations: 42
% Global subsumptions: 3
% 
% Simple tautologies: 148
% Equational tautologies: 115
% Forward subsumptions: 78278
% Fw demodulations to eq. taut.: 489
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 55195
% Factoring: 3
% Forward superposition: 30397
% Backward superposition: 20150
% Self superposition: 241
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 22306
% SAT solver unit clauses: 12401
% SAT solver binary clauses: 7506
% 
% Sat splits: 11
% Sat splitting refutations: 7
% 
% Memory used [KB]: 105925
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.stack_removeProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stack_removeProof
% SZS output start Proof for thm.stack_removeProof
fof(f6829,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1014,f1019,f1020,f1021,f1165,f1022,f1163,f1023,f1169,f1167,f1024,f1025,f1026,f1027,f1028,f1029,f1031,f1032,f1120,f1186,f1187,f1188,f1363,f1189,f1365,f1203,f1370,f1371,f1373,f1164,f1205,f1206,f1207,f1208,f1209,f1211,f1375,f1212,f1300,f1374,f1601,f1513,f1673,f1674,f1675,f1676,f1690,f1691,f1705,f1706,f1707,f1708,f1709,f1710,f1711,f1714,f1724,f1725,f1726,f1727,f1728,f1729,f1731,f1821,f1872,f1873,f1874,f1875,f1876,f1878,f1879,f1967,f2003,f2018,f2019,f2020,f2021,f2022,f1858,f2023,f2025,f2115,f1857,f2164,f2165,f2166,f2167,f2168,f2169,f2171,f2261,f2310,f2311,f2312,f2313,f2329,f2330,f2331,f2332,f2333,f2335,f2336,f2424,f2589,f2590,f2593,f2595,f2596,f2597,f2598,f2599,f2600,f2602,f2692,f3019,f3286,f3287,f3288,f3289,f3290,f2298,f3291,f3293,f3383,f2297,f3426,f1157,f1155,f3431,f3436,f1159,f1161,f3437,f3438,f3439,f3440,f3441,f3443,f1172,f3444,f3532,f1170,f3591,f3592,f3593,f2485,f3594,f2487,f1361,f3610,f3611,f3612,f3613,f3614,f3616,f2492,f3617,f3705,f2491,f3863,f3864,f2742,f2740,f3867,f2746,f2744,f3869,f3870,f3871,f3872,f3873,f2749,f3874,f3876,f3966,f2748,f4151,f4150,f4149,f4417,f4394,f4381,f4380,f4520,f4521,f4522,f4523,f4524,f4525,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4962,f4965,f4972,f4979,f5489,f5492,f5499,f5506,f4154,f5634,f5635,f5636,f5637,f1171,f5638,f1173,f5773,f5639,f1175,f5641,f1174,f5731,f1176,f5772,f6813,f6816])).
fof(f6816,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f6815,f666])).
fof(f666,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14))))),
  inference(cnf_transformation,[],[f533])).
fof(f533,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) & ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f448])).
fof(f448,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X1)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) & ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))),s('type.fcp.cart'(bool,X0),X2)))))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X1)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) & ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))),s('type.fcp.cart'(bool,X0),X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X1)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))),s('type.fcp.cart'(bool,X0),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X1)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))),s('type.fcp.cart'(bool,X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X7)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X8)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X7),s('type.fcp.cart'(bool,X6),X8))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X8),s('type.fcp.cart'(bool,X6),X7))),s('type.fcp.cart'(bool,X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X7)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X8)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X7),s('type.fcp.cart'(bool,X6),X8))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X8),s('type.fcp.cart'(bool,X6),X7))),s('type.fcp.cart'(bool,X6),X8)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.stack_removeProof.word_sub_lt')).
fof(f6815,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl6),
  inference(subsumption_resolution,[],[f6666,f669])).
fof(f669,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(cnf_transformation,[],[f533])).
fof(f6666,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl6),
  inference(resolution,[],[f1161,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(ennf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9,X10] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X10)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X10),s('type.fcp.cart'(bool,X6),X9))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X9),s('type.fcp.cart'(bool,X6),X10))),s('type.fcp.cart'(bool,X6),X9)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LT_SUB_UPPER')).
fof(f6813,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f6812,f666])).
fof(f6812,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl6),
  inference(subsumption_resolution,[],[f6664,f669])).
fof(f6664,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl6),
  inference(resolution,[],[f1161,f675])).
fof(f675,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1))))))),
  inference(flattening,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) | (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1))))))),
  inference(ennf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1))))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))),s('type.fcp.cart'(bool,X0),X1))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X10] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X10)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X10),s('type.fcp.cart'(bool,X6),X9))))) => (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X9),s('type.fcp.cart'(bool,X6),X10)))))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X9),s('type.fcp.cart'(bool,X6),X10))),s('type.fcp.cart'(bool,X6),X9))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_SUB_LT')).
fof(f5772,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl76),
  inference(cnf_transformation,[],[f5772_D])).
fof(f5772_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1176,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl28),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5731,plain,(
  ( ! [X175,X176] : (p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X176) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f1017,f1006])).
fof(f1006,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f708])).
fof(f708,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f545,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f544])).
fof(f544,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9,X17,X10,X18] : ((s(bool,X9) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X10) = s(bool,X18))) => ((p(s(bool,X9)) => p(s(bool,X10))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.IMP_CONG')).
fof(f1017,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))),
  inference(resolution,[],[f668,f685])).
fof(f685,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) )),
  inference(cnf_transformation,[],[f536])).
fof(f536,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) & (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(nnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0] : ! [X1,X2] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X12,X11] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_NOT_LESS_EQUAL')).
fof(f668,plain,(
  p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(cnf_transformation,[],[f533])).
fof(f1174,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9))) ) | $spl24),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5641,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f1017,f711])).
fof(f711,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f1175,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) | $spl26),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5639,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f1017,f710])).
fof(f710,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f5773,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl78),
  inference(cnf_transformation,[],[f5773_D])).
fof(f5773_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1173,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X6))) ) | $spl22),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X6))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5638,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f1017,f709])).
fof(f709,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f1171,plain,(
  ( ! [X4,X2] : (s(bool,X2) != s(bool,X4) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  ( ! [X4,X2] : (s(bool,X2) != s(bool,X4) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5637,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f1017,f709])).
fof(f5636,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f1017,f707])).
fof(f707,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f5635,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f1017,f705])).
fof(f705,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f545])).
fof(f5634,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f1017,f705])).
fof(f4154,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f4154_D])).
fof(f4154_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5506,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X106) != s(bool,X107)) ) | $spl15),
  inference(subsumption_resolution,[],[f5241,f1006])).
fof(f5241,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,X107)) | s(bool,X106) != s(bool,X107)) ) | $spl15),
  inference(superposition,[],[f711,f2253])).
fof(f2253,plain,(
  s(bool,'F_0') = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | $spl15),
  inference(resolution,[],[f1169,f898])).
fof(f898,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f630])).
fof(f630,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f498])).
fof(f498,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X13))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.IMP_F_EQ_F')).
fof(f5499,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,X75) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X76) != s(bool,X77)) ) | $spl15),
  inference(subsumption_resolution,[],[f5231,f711])).
fof(f5231,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X76) != s(bool,X77)) ) | $spl15),
  inference(superposition,[],[f709,f2253])).
fof(f5492,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X50)) ) | $spl15),
  inference(subsumption_resolution,[],[f5222,f1006])).
fof(f5222,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X50)) ) | $spl15),
  inference(superposition,[],[f707,f2253])).
fof(f5489,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X30) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X32)) ) | $spl15),
  inference(subsumption_resolution,[],[f5216,f711])).
fof(f5216,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X30) | p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X32)) ) | $spl15),
  inference(superposition,[],[f705,f2253])).
fof(f4979,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~p(s(bool,X107)) | s(bool,X106) != s(bool,X107)) ) | $spl11),
  inference(subsumption_resolution,[],[f4714,f1006])).
fof(f4714,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X106)) | ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | ~p(s(bool,X107)) | s(bool,X106) != s(bool,X107)) ) | $spl11),
  inference(superposition,[],[f711,f2107])).
fof(f2107,plain,(
  s(bool,'F_0') = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | $spl11),
  inference(resolution,[],[f1164,f898])).
fof(f4972,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,X75) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X76) != s(bool,X77)) ) | $spl11),
  inference(subsumption_resolution,[],[f4704,f711])).
fof(f4704,plain,(
  ( ! [X76,X77,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X76) != s(bool,X77)) ) | $spl11),
  inference(superposition,[],[f709,f2107])).
fof(f4965,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X50)) ) | $spl11),
  inference(subsumption_resolution,[],[f4695,f1006])).
fof(f4695,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X48) | p(s(bool,X49)) | ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X50)) ) | $spl11),
  inference(superposition,[],[f707,f2107])).
fof(f4962,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X30) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X32)) ) | $spl11),
  inference(subsumption_resolution,[],[f4689,f711])).
fof(f4689,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X30) | p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X31) != s(bool,X32)) ) | $spl11),
  inference(superposition,[],[f705,f2107])).
fof(f4535,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4418,f4237])).
fof(f4237,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(equality_resolution,[],[f1174])).
fof(f4418,plain,(
  ( ! [X420,X419] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X419),s('type.num.num',X420))))) ) | $spl24),
  inference(resolution,[],[f4237,f868])).
fof(f868,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.numeral.numeral_lte|split|5')).
fof(f4534,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4412,f4237])).
fof(f4412,plain,(
  ( ! [X410,X409] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X409))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X410))))))) ) | $spl24),
  inference(resolution,[],[f4237,f868])).
fof(f4533,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4411,f4237])).
fof(f4411,plain,(
  ( ! [X408,X407] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X407),s('type.num.num',X408))))) ) | $spl24),
  inference(resolution,[],[f4237,f827])).
fof(f827,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f258])).
fof(f258,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f4532,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4410,f4237])).
fof(f4410,plain,(
  ( ! [X405,X406] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X405),s('type.num.num',X406))))) ) | $spl24),
  inference(resolution,[],[f4237,f825])).
fof(f825,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f596])).
fof(f596,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.arithmetic.NOT_LESS')).
fof(f4531,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4405,f4237])).
fof(f4405,plain,(
  ( ! [X391,X392] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X391),s('type.num.num',X392))))) ) | $spl24),
  inference(resolution,[],[f4237,f866])).
fof(f866,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f236])).
fof(f236,axiom,(
  ! [X7,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.numeral.numeral_lt|split|5')).
fof(f4530,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4399,f4237])).
fof(f4399,plain,(
  ( ! [X377,X378] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X377))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X378))))))) ) | $spl24),
  inference(resolution,[],[f4237,f866])).
fof(f4529,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4398,f4237])).
fof(f4398,plain,(
  ( ! [X376,X375] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X375),s('type.num.num',X376))))) ) | $spl24),
  inference(resolution,[],[f4237,f827])).
fof(f4528,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4397,f4237])).
fof(f4397,plain,(
  ( ! [X374,X373] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X373),s('type.num.num',X374))))) ) | $spl24),
  inference(resolution,[],[f4237,f825])).
fof(f4527,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4393,f4237])).
fof(f4393,plain,(
  ( ! [X364,X366,X365] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X364),X365),s('type.fcp.cart'(bool,X364),X366))))) ) | $spl24),
  inference(resolution,[],[f4237,f703])).
fof(f703,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X12,X11] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LESS_CASES')).
fof(f4526,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4392,f4237])).
fof(f4392,plain,(
  ( ! [X362,X361,X363] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X361),X362),s('type.fcp.cart'(bool,X361),X363))))) ) | $spl24),
  inference(resolution,[],[f4237,f702])).
fof(f702,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f200])).
fof(f200,axiom,(
  ! [X6] : ! [X12,X11] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LESS_EQ_CASES')).
fof(f4525,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4391,f4237])).
fof(f4391,plain,(
  ( ! [X358,X360,X359] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X358),X359),s('type.fcp.cart'(bool,X358),X360))))) ) | $spl24),
  inference(resolution,[],[f4237,f702])).
fof(f4524,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4384,f4237])).
fof(f4384,plain,(
  ( ! [X337,X339,X338] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X337),X338),s('type.fcp.cart'(bool,X337),X339))))) ) | $spl24),
  inference(resolution,[],[f4237,f686])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f537])).
fof(f537,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) & (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(nnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0] : ! [X1,X2] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f177])).
fof(f177,axiom,(
  ! [X6] : ! [X12,X11] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) <=> p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_NOT_LESS')).
fof(f4523,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4383,f4237])).
fof(f4383,plain,(
  ( ! [X335,X334,X336] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X334),X335),s('type.fcp.cart'(bool,X334),X336))))) ) | $spl24),
  inference(resolution,[],[f4237,f684])).
fof(f684,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f536])).
fof(f4522,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4376,f4237])).
fof(f4376,plain,(
  ( ! [X318,X320,X319] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X318),X319),s('type.fcp.cart'(bool,X318),X320))))) ) | $spl24),
  inference(resolution,[],[f4237,f703])).
fof(f4521,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4373,f4237])).
fof(f4373,plain,(
  ( ! [X310,X309,X311] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X309),X310),s('type.fcp.cart'(bool,X309),X311))))) ) | $spl24),
  inference(resolution,[],[f4237,f686])).
fof(f4520,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f4372,f4237])).
fof(f4372,plain,(
  ( ! [X308,X306,X307] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X306),X307),s('type.fcp.cart'(bool,X306),X308))))) ) | $spl24),
  inference(resolution,[],[f4237,f684])).
fof(f4380,plain,(
  $false | $spl24),
  inference(resolution,[],[f4237,f666])).
fof(f4381,plain,(
  $false | $spl24),
  inference(resolution,[],[f4237,f667])).
fof(f667,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(cnf_transformation,[],[f533])).
fof(f4394,plain,(
  $false | $spl24),
  inference(resolution,[],[f4237,f668])).
fof(f4417,plain,(
  $false | $spl24),
  inference(resolution,[],[f4237,f944])).
fof(f944,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))),
  file('thm.stack_removeProof.word_sub_lt','thm.arithmetic.ZERO_LESS_EQ')).
fof(f4149,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4085])).
fof(f4085,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl18),
  inference(superposition,[],[f1171,f932])).
fof(f932,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.EQ_CLAUSES|split|3')).
fof(f4150,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4084])).
fof(f4084,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl18),
  inference(superposition,[],[f1171,f930])).
fof(f930,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f649])).
fof(f649,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.EQ_CLAUSES|split|2')).
fof(f4151,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f4082])).
fof(f4082,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl18),
  inference(superposition,[],[f1171,f898])).
fof(f2748,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl70),
  inference(cnf_transformation,[],[f2748_D])).
fof(f2748_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3966,plain,(
  ( ! [X175,X176] : (p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X176) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f1006])).
fof(f3876,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f669,f711])).
fof(f3874,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f710])).
fof(f2749,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl72),
  inference(cnf_transformation,[],[f2749_D])).
fof(f2749_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3873,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f669,f709])).
fof(f3872,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f709])).
fof(f3871,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f669,f707])).
fof(f3870,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f669,f705])).
fof(f3869,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f705])).
fof(f2744,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))) | $spl67),
  inference(cnf_transformation,[],[f2744_D])).
fof(f2744_D,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f2746,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) | $spl69),
  inference(cnf_transformation,[],[f2746_D])).
fof(f2746_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f3867,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))),
  inference(resolution,[],[f669,f950])).
fof(f950,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2)))))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))) )),
  inference(cnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | (s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) & (((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(flattening,[],[f651])).
fof(f651,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(nnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : ! [X12,X11] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X11))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X12))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X6),X12))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X6),X11))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X12)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X11)))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LT')).
fof(f2740,plain,(
  ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) | $spl63),
  inference(cnf_transformation,[],[f2740_D])).
fof(f2740_D,plain,(
  ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f2742,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | $spl64),
  inference(cnf_transformation,[],[f2742_D])).
fof(f2742_D,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3864,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f669,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)) & ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) & s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2)) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2))) & ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) & s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2)) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(nnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X12,X11] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) <=> (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) | s('type.fcp.cart'(bool,X6),X11) = s('type.fcp.cart'(bool,X6),X12)))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LESS_OR_EQ')).
fof(f3863,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f669,f688])).
fof(f688,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) )),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) & (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) & s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2))))),
  inference(flattening,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) & (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) & s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2))))),
  inference(nnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) & s('type.fcp.cart'(bool,X0),X1) != s('type.fcp.cart'(bool,X0),X2)) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) & ~s('type.fcp.cart'(bool,X0),X1) = s('type.fcp.cart'(bool,X0),X2)) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X12,X11] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))) & ~s('type.fcp.cart'(bool,X6),X11) = s('type.fcp.cart'(bool,X6),X12)) <=> p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X11),s('type.fcp.cart'(bool,X6),X12)))))),
  file('thm.stack_removeProof.word_sub_lt','thm.address.address_83_prove_ltb')).
fof(f2491,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl58),
  inference(cnf_transformation,[],[f2491_D])).
fof(f2491_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3705,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f1006])).
fof(f3617,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f667,f711])).
fof(f2492,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl60),
  inference(cnf_transformation,[],[f2492_D])).
fof(f2492_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3616,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f711])).
fof(f3614,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f710])).
fof(f3613,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f709])).
fof(f3612,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f667,f707])).
fof(f3611,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f707])).
fof(f3610,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f705])).
fof(f1361,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))) | $spl31),
  inference(cnf_transformation,[],[f1361_D])).
fof(f1361_D,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f2487,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) | $spl56),
  inference(cnf_transformation,[],[f2487_D])).
fof(f2487_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3594,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f678])).
fof(f2485,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl54),
  inference(cnf_transformation,[],[f2485_D])).
fof(f2485_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3593,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f677])).
fof(f677,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(ennf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X9,X10] : ((p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X10)))) & p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X10),s('type.fcp.cart'(bool,X6),X9))))) => p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X9),s('type.fcp.cart'(bool,X6),X10)))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.words_7|48|5_prove_ltb')).
fof(f3592,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f675])).
fof(f3591,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f674])).
fof(f674,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f450])).
fof(f1170,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl16),
  inference(cnf_transformation,[],[f1170_D])).
fof(f1170_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3532,plain,(
  ( ! [X177,X176] : (p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X177) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f1006])).
fof(f3444,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f668,f711])).
fof(f1172,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl20),
  inference(cnf_transformation,[],[f1172_D])).
fof(f1172_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3443,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f711])).
fof(f3441,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f710])).
fof(f3440,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f709])).
fof(f3439,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f668,f707])).
fof(f3438,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f707])).
fof(f3437,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f705])).
fof(f1161,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) | $spl6),
  inference(cnf_transformation,[],[f1161_D])).
fof(f1161_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1159,plain,(
  s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15) | $spl4),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3436,plain,(
  s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f668,f697])).
fof(f3431,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) | s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15)),
  inference(resolution,[],[f668,f688])).
fof(f1155,plain,(
  ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1157,plain,(
  p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))) | $spl2),
  inference(cnf_transformation,[],[f1157_D])).
fof(f1157_D,plain,(
  p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3426,plain,(
  p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14))))),
  inference(resolution,[],[f668,f676])).
fof(f676,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(ennf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),X2)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) => p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : ! [X9,X10] : ((p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),X10)))) & p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X10),s('type.fcp.cart'(bool,X6),X9))))) => p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X9),s('type.fcp.cart'(bool,X6),X10)))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.words_7|48|4_prove_ltb')).
fof(f2297,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl50),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3383,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f1006])).
fof(f3293,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X27) != s(bool,X29)) ) | $spl15),
  inference(resolution,[],[f1169,f711])).
fof(f3291,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f710])).
fof(f2298,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl52),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3290,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X18) != s(bool,X19)) ) | $spl15),
  inference(resolution,[],[f1169,f709])).
fof(f3289,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f709])).
fof(f3288,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X12) != s(bool,X14)) ) | $spl15),
  inference(resolution,[],[f1169,f707])).
fof(f3287,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X10)) ) | $spl15),
  inference(resolution,[],[f1169,f705])).
fof(f3286,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f705])).
fof(f3019,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f3018,f666])).
fof(f3018,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl4),
  inference(forward_demodulation,[],[f2815,f681])).
fof(f681,plain,(
  ( ! [X0,X1] : (s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))) = s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X1)))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))) = s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X1)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0] : ! [X1] : s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))) = s('type.fcp.cart'(bool,X0),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X1)))),
  inference(rectify,[],[f204])).
fof(f204,axiom,(
  ! [X6] : ! [X31] : s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))) = s('type.fcp.cart'(bool,X6),'const.words.word_sub_2'(s('type.fcp.cart'(bool,X6),X31),s('type.fcp.cart'(bool,X6),X31)))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_SUB_REFL')).
fof(f2815,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl4),
  inference(backward_demodulation,[],[f1159,f669])).
fof(f2692,plain,(
  ( ! [X175,X176] : (p(s(bool,X175)) | ~p(s(bool,X176)) | s(bool,X176) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f1006])).
fof(f2602,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f669,f711])).
fof(f2600,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f710])).
fof(f2599,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f669,f709])).
fof(f2598,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f709])).
fof(f2597,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f669,f707])).
fof(f2596,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f669,f705])).
fof(f2595,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f669,f705])).
fof(f2593,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))),
  inference(resolution,[],[f669,f950])).
fof(f2590,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f669,f697])).
fof(f2589,plain,(
  s('type.fcp.cart'(bool,sK13),sK15) = s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14))),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f669,f688])).
fof(f2424,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f1006])).
fof(f2336,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f667,f711])).
fof(f2335,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f711])).
fof(f2333,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f710])).
fof(f2332,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f709])).
fof(f2331,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f667,f707])).
fof(f2330,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f707])).
fof(f2329,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f667,f705])).
fof(f2313,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f678])).
fof(f2312,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f677])).
fof(f2311,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK15)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f675])).
fof(f2310,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f667,f674])).
fof(f2261,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f1006])).
fof(f2171,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X27) != s(bool,X29)) ) | $spl15),
  inference(resolution,[],[f1169,f711])).
fof(f2169,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f710])).
fof(f2168,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X18) != s(bool,X19)) ) | $spl15),
  inference(resolution,[],[f1169,f709])).
fof(f2167,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f709])).
fof(f2166,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X12) != s(bool,X14)) ) | $spl15),
  inference(resolution,[],[f1169,f707])).
fof(f2165,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X10)) ) | $spl15),
  inference(resolution,[],[f1169,f705])).
fof(f2164,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) ) | $spl15),
  inference(resolution,[],[f1169,f705])).
fof(f1857,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl44),
  inference(cnf_transformation,[],[f1857_D])).
fof(f1857_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2115,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f1006])).
fof(f2025,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X27) != s(bool,X29)) ) | $spl11),
  inference(resolution,[],[f1164,f711])).
fof(f2023,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f710])).
fof(f1858,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl46),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2022,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X18) != s(bool,X19)) ) | $spl11),
  inference(resolution,[],[f1164,f709])).
fof(f2021,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f709])).
fof(f2020,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X12) != s(bool,X14)) ) | $spl11),
  inference(resolution,[],[f1164,f707])).
fof(f2019,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X9) != s(bool,X10)) ) | $spl11),
  inference(resolution,[],[f1164,f705])).
fof(f2018,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f705])).
fof(f2003,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl48),
  inference(cnf_transformation,[],[f2003_D])).
fof(f2003_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1967,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f1006])).
fof(f1879,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X28) != s(bool,X29)) ) | $spl10),
  inference(resolution,[],[f1165,f711])).
fof(f1878,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X25) | s(bool,X26) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f711])).
fof(f1876,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f710])).
fof(f1875,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f709])).
fof(f1874,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl10),
  inference(resolution,[],[f1165,f707])).
fof(f1873,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X10) | s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f707])).
fof(f1872,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl10),
  inference(resolution,[],[f1165,f705])).
fof(f1821,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f1006])).
fof(f1731,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X27) != s(bool,X29)) ) | $spl11),
  inference(resolution,[],[f1164,f711])).
fof(f1729,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f710])).
fof(f1728,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X20) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X18) != s(bool,X19)) ) | $spl11),
  inference(resolution,[],[f1164,f709])).
fof(f1727,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f709])).
fof(f1726,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X12) != s(bool,X14)) ) | $spl11),
  inference(resolution,[],[f1164,f707])).
fof(f1725,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X9) != s(bool,X10)) ) | $spl11),
  inference(resolution,[],[f1164,f705])).
fof(f1724,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl11),
  inference(resolution,[],[f1164,f705])).
fof(f1714,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1589,f1176])).
fof(f1589,plain,(
  ( ! [X512,X513] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X512),s('type.num.num',X513))))) ) | $spl28),
  inference(resolution,[],[f1176,f867])).
fof(f867,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f615])).
fof(f1711,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1584,f1176])).
fof(f1584,plain,(
  ( ! [X502,X501] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X501))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X502))))))) ) | $spl28),
  inference(resolution,[],[f1176,f867])).
fof(f1710,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1583,f1176])).
fof(f1583,plain,(
  ( ! [X499,X500] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X499),s('type.num.num',X500))))) ) | $spl28),
  inference(resolution,[],[f1176,f828])).
fof(f828,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f597])).
fof(f1709,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1582,f1176])).
fof(f1582,plain,(
  ( ! [X498,X497] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X497),s('type.num.num',X498))))) ) | $spl28),
  inference(resolution,[],[f1176,f826])).
fof(f826,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f596])).
fof(f1708,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1573,f1176])).
fof(f1573,plain,(
  ( ! [X476,X475] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X475),s('type.num.num',X476))))) ) | $spl28),
  inference(resolution,[],[f1176,f865])).
fof(f865,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f1707,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1565,f1176])).
fof(f1565,plain,(
  ( ! [X455,X456] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X455))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X456))))))) ) | $spl28),
  inference(resolution,[],[f1176,f865])).
fof(f1706,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1564,f1176])).
fof(f1564,plain,(
  ( ! [X453,X454] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X453),s('type.num.num',X454))))) ) | $spl28),
  inference(resolution,[],[f1176,f828])).
fof(f1705,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1563,f1176])).
fof(f1563,plain,(
  ( ! [X451,X452] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X451),s('type.num.num',X452))))) ) | $spl28),
  inference(resolution,[],[f1176,f826])).
fof(f1691,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1547,f1176])).
fof(f1547,plain,(
  ( ! [X405,X403,X404] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X403),X404),s('type.fcp.cart'(bool,X403),X405))))) ) | $spl28),
  inference(resolution,[],[f1176,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X2),s('type.fcp.cart'(bool,X0),X1))))) )),
  inference(cnf_transformation,[],[f537])).
fof(f1690,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1546,f1176])).
fof(f1546,plain,(
  ( ! [X401,X400,X402] : (~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X400),X401),s('type.fcp.cart'(bool,X400),X402))))) ) | $spl28),
  inference(resolution,[],[f1176,f685])).
fof(f1676,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1521,f1176])).
fof(f1521,plain,(
  ( ! [X321,X323,X322] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X321),X322),s('type.fcp.cart'(bool,X321),X323))))) ) | $spl28),
  inference(resolution,[],[f1176,f687])).
fof(f1675,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1520,f1176])).
fof(f1520,plain,(
  ( ! [X318,X320,X319] : (~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X318),X319),s('type.fcp.cart'(bool,X318),X320))))) ) | $spl28),
  inference(resolution,[],[f1176,f685])).
fof(f1674,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1511,f1176])).
fof(f1511,plain,(
  ( ! [X296] : (~p(s(bool,X296))) ) | $spl28),
  inference(resolution,[],[f1176,f940])).
fof(f940,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f388])).
fof(f388,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) => ~p(s(bool,X13)))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.IMP_F')).
fof(f1673,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f1501,f1176])).
fof(f1501,plain,(
  ( ! [X286] : (~p(s(bool,X286))) ) | $spl28),
  inference(resolution,[],[f1176,f896])).
fof(f896,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f628])).
fof(f628,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.stack_removeProof.word_sub_lt','thm.bool.IMP_CLAUSES|split|4')).
fof(f1513,plain,(
  $false | $spl28),
  inference(resolution,[],[f1176,f945])).
fof(f945,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.stack_removeProof.word_sub_lt',aHL_FALSITY)).
fof(f1601,plain,(
  $false | $spl28),
  inference(resolution,[],[f1176,f964])).
fof(f964,plain,(
  ( ! [X0] : (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0] : ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f413])).
fof(f413,plain,(
  ! [X0] : ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6] : ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_|48|_POS')).
fof(f1374,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl40),
  inference(cnf_transformation,[],[f1374_D])).
fof(f1374_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1300,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f1006])).
fof(f1212,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f666,f711])).
fof(f1375,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) ) | $spl42),
  inference(cnf_transformation,[],[f1375_D])).
fof(f1375_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1211,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f711])).
fof(f1209,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f710])).
fof(f1208,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f709])).
fof(f1207,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f666,f707])).
fof(f1206,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f707])).
fof(f1205,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14)))) )),
  inference(resolution,[],[f666,f705])).
fof(f1164,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | $spl11),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1373,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14)))))) | $spl38),
  inference(cnf_transformation,[],[f1373_D])).
fof(f1373_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1371,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14)))))) | ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))))),
  inference(forward_demodulation,[],[f1204,f965])).
fof(f965,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f229])).
fof(f229,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X6),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.word_|48|_n2w')).
fof(f1204,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14))))))),
  inference(resolution,[],[f666,f949])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f652])).
fof(f1370,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))) | $spl36),
  inference(cnf_transformation,[],[f1370_D])).
fof(f1370_D,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1203,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(resolution,[],[f666,f948])).
fof(f948,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f652])).
fof(f1365,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK14)))) | $spl34),
  inference(cnf_transformation,[],[f1365_D])).
fof(f1365_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK14)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1189,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK14)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f666,f678])).
fof(f1363,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl32),
  inference(cnf_transformation,[],[f1363_D])).
fof(f1363_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1188,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f666,f677])).
fof(f1187,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))),s('type.fcp.cart'(bool,sK13),sK14)))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f666,f675])).
fof(f1186,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f666,f674])).
fof(f1120,plain,(
  ( ! [X177,X176] : (p(s(bool,X176)) | p(s(bool,X177)) | s(bool,X177) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f1006])).
fof(f1032,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X24) != s(bool,X25)) )),
  inference(resolution,[],[f668,f711])).
fof(f1031,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f711])).
fof(f1029,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f710])).
fof(f1028,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f709])).
fof(f1027,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))) | s(bool,X9) != s(bool,X10)) )),
  inference(resolution,[],[f668,f707])).
fof(f1026,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f707])).
fof(f1025,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) )),
  inference(resolution,[],[f668,f705])).
fof(f1024,plain,(
  s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15))))),
  inference(resolution,[],[f668,f697])).
fof(f1167,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) | $spl12),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15)))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1169,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) | $spl15),
  inference(cnf_transformation,[],[f1169_D])).
fof(f1169_D,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1023,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15))))))),
  inference(resolution,[],[f668,f694])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | (s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) & (((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(flattening,[],[f540])).
fof(f540,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) | ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) != s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) & (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))) | p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))))),
  inference(nnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : ! [X12,X11] : (p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X6),X12),s('type.fcp.cart'(bool,X6),X11)))) <=> ((s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X11))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X12))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X6),X12))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X6),X11))))))) | (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X12)))) & ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X6),X11)))))))),
  file('thm.stack_removeProof.word_sub_lt','thm.words.WORD_LE')).
fof(f1163,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) | $spl8),
  inference(cnf_transformation,[],[f1163_D])).
fof(f1163_D,plain,(
  s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1022,plain,(
  ~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))),
  inference(resolution,[],[f668,f693])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f1165,plain,(
  p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | $spl10),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1021,plain,(
  p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK14))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,sK13),sK15))))))),
  inference(resolution,[],[f668,f692])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X1))),s('type.num.num','const.words.w2n_1'(s('type.fcp.cart'(bool,X0),X2)))))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f1020,plain,(
  p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK14))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,sK13),sK15)))),
  inference(resolution,[],[f668,f691])).
fof(f691,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1)))) | s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X1))) = s(bool,'const.words.word_msb_1'(s('type.fcp.cart'(bool,X0),X2))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,X0),X1),s('type.fcp.cart'(bool,X0),X2))))) )),
  inference(cnf_transformation,[],[f541])).
fof(f1019,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK14),s('type.fcp.cart'(bool,sK13),sK15)))) | s('type.fcp.cart'(bool,sK13),sK14) = s('type.fcp.cart'(bool,sK13),sK15)),
  inference(resolution,[],[f668,f688])).
fof(f1014,plain,(
  p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),'const.words.word_sub_2'(s('type.fcp.cart'(bool,sK13),sK15),s('type.fcp.cart'(bool,sK13),sK14)))))) | ~p(s(bool,'const.words.word_le_2'(s('type.fcp.cart'(bool,sK13),'const.words.n2w_1'(s('type.num.num','const.num.|48|_0'))),s('type.fcp.cart'(bool,sK13),sK14))))),
  inference(resolution,[],[f668,f676])).
% SZS output end Proof for thm.stack_removeProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 346
% Passive clauses: 2537
% Generated clauses: 5802
% Final active clauses: 333
% Final passive clauses: 2113
% Input formulas: 260
% Initial clauses: 340
% 
% Selected by SInE selection: 123
% SInE iterations: 1
% 
% Duplicate literals: 134
% Trivial inequalities: 6
% Fw subsumption resolutions: 571
% Fw demodulations: 477
% Bw demodulations: 6
% 
% Simple tautologies: 330
% Equational tautologies: 1
% Forward subsumptions: 1469
% 
% Binary resolution: 2397
% Forward superposition: 1460
% Backward superposition: 711
% Equality resolution: 2
% 
% Split clauses: 145
% Split components: 303
% SAT solver clauses: 185
% SAT solver unit clauses: 39
% SAT solver binary clauses: 86
% 
% Sat splits: 145
% Sat splitting refutations: 39
% 
% Memory used [KB]: 3582
% Time elapsed: 0.109 s
% ------------------------------
% ------------------------------
% Success in time 11.86 s
