Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% (118611)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 157
% Passive clauses: 5064
% Generated clauses: 16632
% Final active clauses: 145
% Final passive clauses: 4681
% Input formulas: 33
% Initial clauses: 85
% 
% Duplicate literals: 653
% Trivial inequalities: 116
% Fw subsumption resolutions: 2586
% Fw demodulations: 1272
% Bw demodulations: 9
% 
% Simple tautologies: 228
% Equational tautologies: 83
% Forward subsumptions: 5407
% 
% Binary resolution: 1160
% Factoring: 2
% Forward superposition: 5512
% Backward superposition: 4602
% Self superposition: 234
% Equality factoring: 194
% Equality resolution: 25
% 
% Split clauses: 539
% Split components: 1078
% SAT solver clauses: 1385
% SAT solver unit clauses: 4
% SAT solver binary clauses: 180
% 
% Sat splits: 539
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1385
% TWLsolver calls for satisfiability: 201
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_numberProof
% (118639)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 942
% Generated clauses: 10227
% Final active clauses: 112
% Final passive clauses: 799
% Input formulas: 33
% Initial clauses: 85
% 
% Duplicate literals: 733
% Trivial inequalities: 107
% Fw subsumption resolutions: 841
% Fw demodulations: 553
% Bw demodulations: 31
% Condensations: 1142
% Global subsumptions: 4
% 
% Simple tautologies: 213
% Equational tautologies: 49
% Forward subsumptions: 3536
% 
% Binary resolution: 235
% Factoring: 1
% Forward superposition: 2967
% Backward superposition: 2499
% Self superposition: 138
% Equality factoring: 124
% Equality resolution: 26
% 
% Split clauses: 128
% Split components: 256
% SAT solver clauses: 1530
% SAT solver unit clauses: 89
% SAT solver binary clauses: 73
% SAT solver learnt clauses: 8
% 
% Sat splits: 128
% 
% TWLsolver clauses: 1514
% TWLsolver calls for satisfiability: 687
% 
% Memory used [KB]: 10618
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_numberProof
% (118655)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 861
% Passive clauses: 12151
% Generated clauses: 208019
% Final active clauses: 107
% Final passive clauses: 1720
% Input formulas: 33
% Initial clauses: 85
% Discarded non-redundant clauses: 1449
% 
% Duplicate literals: 8819
% Trivial inequalities: 1523
% Fw subsumption resolutions: 20004
% Fw demodulations: 36873
% Bw demodulations: 3063
% Condensations: 14255
% Global subsumptions: 103
% 
% Simple tautologies: 1929
% Equational tautologies: 763
% Forward subsumptions: 110949
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 790
% 
% Unit resulting resolution: 2989
% Factoring: 1
% Forward superposition: 44879
% Backward superposition: 60702
% Self superposition: 1271
% Equality factoring: 796
% Equality resolution: 157
% 
% Split clauses: 809
% Split components: 1644
% SAT solver clauses: 13927
% SAT solver unit clauses: 37
% SAT solver binary clauses: 1245
% 
% Sat splits: 809
% Sat splitting refutations: 53
% 
% Memory used [KB]: 109124
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f30430,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f83,f129,f127,f82,f128,f126,f235,f236,f287,f284,f283,f395,f398,f403,f404,f423,f424,f394,f396,f392,f443,f444,f468,f663,f666,f667,f671,f672,f673,f674,f675,f662,f604,f676,f605,f677,f606,f607,f679,f608,f610,f681,f611,f612,f683,f613,f685,f687,f688,f689,f691,f693,f694,f695,f621,f622,f623,f624,f625,f626,f627,f628,f629,f706,f708,f710,f712,f668,f714,f715,f665,f716,f717,f660,f428,f1112,f1113,f1114,f1115,f1117,f1119,f1120,f1121,f1122,f1123,f1125,f1127,f1130,f1131,f1132,f1133,f1135,f690,f1137,f696,f1138,f1139,f1140,f1141,f1143,f1145,f1237,f1238,f1460,f1461,f1715,f1716,f1840,f1841,f1862,f1864,f2003,f2004,f2027,f2029,f2210,f2211,f2372,f2373,f2493,f2494,f2692,f2693,f2716,f2718,f2883,f2884,f3243,f3244,f3544,f3546,f3749,f3750,f259,f3769,f3770,f427,f3982,f3994,f3996,f4033,f4035,f669,f4075,f4194,f4217,f4227,f4224,f4223,f5137,f5138,f5139,f5144,f678,f5145,f680,f5146,f2717,f5408,f5409,f684,f5450,f5451,f5452,f5454,f5453,f5467,f5469,f5468,f5481,f5701,f5702,f682,f5842,f5843,f5981,f5982,f6126,f6128,f6169,f6171,f6285,f6287,f6288,f6170,f6387,f6388,f6582,f6583,f6706,f6707,f6286,f6708,f6127,f6794,f6795,f6917,f6918,f7446,f7448,f7874,f7876,f7878,f7868,f7867,f7881,f7882,f7862,f7859,f7883,f7884,f7887,f7889,f7890,f7891,f7892,f7894,f7895,f7897,f7875,f7555,f7898,f7899,f7900,f7904,f7905,f7907,f7909,f7906,f7910,f7908,f7911,f7912,f7913,f7914,f7915,f7917,f7916,f7918,f7919,f7920,f7921,f7922,f7924,f7925,f7926,f7853,f7852,f7927,f7929,f7930,f7931,f7932,f7654,f7934,f7655,f692,f7656,f7935,f7657,f7936,f7937,f7938,f7660,f7939,f7661,f7940,f7941,f7942,f7664,f7943,f7668,f686,f7670,f7946,f7947,f7948,f7674,f7949,f7950,f7951,f7678,f7952,f7679,f7953,f7680,f7954,f7681,f698,f7682,f699,f7683,f701,f7684,f700,f7685,f7955,f7686,f7956,f7687,f697,f7689,f703,f7690,f702,f7691,f705,f7692,f704,f7693,f711,f7694,f707,f7695,f709,f7696,f713,f7957,f7958,f7923,f7705,f7706,f7959,f7960,f7847,f7965,f7966,f7967,f7969,f7970,f7972,f7973,f7978,f7877,f7775,f7447,f7980,f7843,f7839,f7928,f8142,f8144,f8143,f397,f8330,f8332,f8331,f8448,f8450,f8451,f8886,f8888,f8889,f8990,f8992,f8991,f8993,f8995,f8997,f9000,f9004,f9005,f9040,f9041,f1863,f9076,f9077,f2028,f9169,f8998,f9279,f9001,f9371,f9373,f9375,f9374,f9492,f9494,f9493,f10061,f10063,f10064,f10197,f10199,f10200,f10354,f10356,f10355,f10446,f10448,f10447,f10449,f10554,f10556,f10555,f10557,f10662,f10664,f10663,f10665,f10770,f10772,f10771,f10773,f10855,f10857,f10856,f11232,f11233,f11234,f11244,f11396,f11375,f11398,f11400,f11399,f11402,f11404,f11403,f11405,f11552,f11554,f11555,f11556,f11752,f11553,f12055,f12022,f12080,f12021,f12081,f12084,f12092,f12093,f12085,f12179,f12180,f12192,f12201,f12203,f12205,f12207,f12209,f12221,f12228,f12230,f12518,f12520,f12523,f12526,f12524,f12732,f12519,f13140,f13105,f13165,f13104,f13166,f13169,f13177,f13178,f13170,f13280,f13289,f13291,f13293,f13295,f13297,f13309,f13316,f13318,f13895,f13897,f13896,f14061,f14063,f14062,f14809,f11374,f14810,f11377,f14811,f11379,f14830,f15007,f14990,f15009,f15011,f15010,f15017,f15019,f15018,f15024,f15027,f15031,f15137,f15139,f15144,f15147,f15151,f15145,f15350,f15138,f15785,f15748,f15812,f15747,f15813,f15816,f15824,f15825,f15817,f15946,f15947,f15959,f15968,f15970,f15972,f15974,f15976,f15978,f15990,f15997,f15999,f16392,f16394,f16401,f16404,f16408,f16412,f16402,f16646,f16393,f17126,f17087,f17153,f17086,f17154,f17157,f17165,f17166,f17158,f17305,f8452,f17314,f17316,f17318,f17320,f17322,f17324,f17336,f8890,f17343,f10065,f17345,f18415,f18417,f18416,f18852,f18854,f18853,f18864,f18866,f18865,f18875,f18878,f18882,f18886,f18890,f19142,f19144,f19143,f19154,f19156,f19155,f19165,f19168,f19172,f19176,f19180,f19491,f19493,f19495,f19494,f19497,f19499,f19501,f19503,f19524,f19527,f19526,f19525,f20183,f20185,f20184,f20510,f20511,f20567,f20568,f20560,f20741,f20743,f20742,f20752,f20755,f20759,f20763,f20767,f20753,f20802,f20804,f20803,f20844,f20846,f21025,f21026,f21100,f21102,f21101,f21093,f21092,f21103,f21341,f21342,f21345,f21346,f21027,f21677,f21679,f21678,f21688,f21691,f21695,f21699,f21703,f21709,f21711,f21710,f21944,f21952,f21943,f21953,f21942,f21954,f21941,f21955,f21984,f22050,f22058,f22049,f22059,f22048,f22060,f22047,f22061,f22090,f22416,f22424,f22415,f22425,f22414,f22426,f22413,f22427,f22456,f22584,f22592,f22583,f22593,f22582,f22594,f22581,f22595,f22624,f20845,f22700,f22702,f22701,f670,f23507,f23508,f8994,f23936,f23946,f23954,f8449,f23882,f23980,f24070,f24071,f24481,f24483,f24493,f24495,f24504,f24507,f24511,f24515,f24519,f24007,f24006,f24063,f24984,f24994,f25002,f24953,f25031,f25152,f25154,f25200,f25202,f25276,f25277,f8887,f25446,f25448,f25457,f25459,f25463,f25467,f25471,f25515,f25516,f25770,f25772,f25771,f25782,f25784,f25783,f25793,f25796,f25800,f25804,f25808,f25975,f25977,f25976,f25986,f25989,f25992,f25995,f25998,f26384,f26386,f26387,f26388,f26389,f26390,f26391,f26398,f26124,f26403,f10201,f26404,f26405,f26410,f26152,f26156,f26402,f26421,f26431,f26440,f10062,f26449,f26450,f26452,f26453,f26455,f26456,f26458,f26459,f26331,f26462,f26642,f26644,f26643,f26690,f26691,f26900,f26901,f26911,f26912,f24494,f26922,f26923,f26925,f26929,f26933,f26937,f24505,f27112,f27113,f27450,f24482,f27457,f27458,f27470,f27476,f27482,f27488,f27532,f27533,f27808,f27809,f27819,f27820,f27822,f27826,f27830,f27834,f27879,f27880,f27944,f11381,f28285,f10450,f28287,f10558,f28288,f10666,f28289,f10774,f28290,f11406,f28291,f25794,f28292,f15025,f28299,f25987,f28013,f26401,f28302,f18876,f28303,f19166,f28308,f21689,f28041,f28045,f28270,f26411,f28309,f28318,f9372,f28328,f19492,f28337,f10198,f28346,f28347,f28349,f28350,f28352,f28353,f28355,f28356,f28227,f23981,f28359,f27432,f27430,f27437,f27436,f27431,f258,f28694,f28695,f25153,f29349,f29350,f25447,f29362,f29368,f29374,f29380,f25201,f29868,f29870,f29869,f29880,f29882,f29881,f29893,f29895,f29896,f29898,f29902,f29906,f29910,f29894,f1303,f30331,f30333,f30335,f30334,f30337,f30339,f30341,f30343,f664,f30344,f28286,f30345,f26385,f30364,f24985,f30367,f30332,f23937,f30424,f30426,f30428,f30429])).
fof(f30429,plain,(
  $false | ($spl1 | $spl2 | $spl86)),
  inference(subsumption_resolution,[],[f30419,f175])).
fof(f175,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f127,f128])).
fof(f30419,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl86),
  inference(equality_resolution,[],[f30404])).
fof(f30404,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0))))) ) | $spl86),
  inference(forward_demodulation,[],[f30394,f117])).
fof(f117,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f16])).
fof(f16,axiom,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.closSem.Unit_def')).
fof(f30394,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))),s('type.closSem.v',X0))))) ) | $spl86),
  inference(superposition,[],[f7109,f117])).
fof(f7109,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X1) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))),s('type.closSem.v',X1))))) ) | $spl86),
  inference(subsumption_resolution,[],[f7097,f124])).
fof(f124,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f96])).
fof(f96,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ! [X6] : (s(X10,X6) = s(X10,X6) <=> p(s(bool,'T_0')))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.REFL_CLAUSE')).
fof(f7097,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))),s('type.closSem.v',X1)))) | s('type.closSem.v',X1) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) ) | $spl86),
  inference(superposition,[],[f6970,f2717])).
fof(f6970,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3)))) )),
  inference(forward_demodulation,[],[f120,f92])).
fof(f92,plain,(
  ( ! [X0,X1] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v',X1))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v',X1)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v',X1))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v',X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X47,X48] : s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X47),s('type.closSem.v',X48))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X47))),s('type.closSem.v',X48)))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','ARITY_0_2_const.clos_numberProof.v_rel')).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0)))) | (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0))))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0)))) | ? [X4] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X4))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0))))))),
  inference(rectify,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0)))) | ? [X3] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3)))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0))))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0)))) <=> ? [X3] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X3))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13,X16,X17] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X17))))),s('type.closSem.v',X13)))) <=> ? [X18] : (s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X18))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X18))))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|2')).
fof(f30428,plain,(
  $false | ($spl1 | $spl2 | $spl86)),
  inference(subsumption_resolution,[],[f30427,f175])).
fof(f30427,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ($spl2 | $spl86)),
  inference(forward_demodulation,[],[f30420,f128])).
fof(f30420,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | ($spl2 | $spl86)),
  inference(trivial_inequality_removal,[],[f30415])).
fof(f30415,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Unit_0') | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | ($spl2 | $spl86)),
  inference(superposition,[],[f30404,f128])).
fof(f30426,plain,(
  $false | ($spl1 | $spl2 | $spl86 | $spl88)),
  inference(subsumption_resolution,[],[f30425,f175])).
fof(f30425,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ($spl86 | $spl88)),
  inference(forward_demodulation,[],[f30421,f3546])).
fof(f30421,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))))))))) | ($spl86 | $spl88)),
  inference(trivial_inequality_removal,[],[f30410])).
fof(f30410,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Unit_0') | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))))))))) | ($spl86 | $spl88)),
  inference(superposition,[],[f30404,f3546])).
fof(f30424,plain,(
  $false | ($spl1 | $spl2 | $spl86)),
  inference(subsumption_resolution,[],[f30423,f175])).
fof(f30423,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl86),
  inference(forward_demodulation,[],[f30422,f117])).
fof(f30422,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) | $spl86),
  inference(trivial_inequality_removal,[],[f30409])).
fof(f30409,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Unit_0') | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) | $spl86),
  inference(superposition,[],[f30404,f117])).
fof(f23937,plain,(
  ( ! [X152,X151,X150] : (s(X150,X152) = s(X150,sK1(X152,X151,X150))) ) | $spl386),
  inference(cnf_transformation,[],[f23937_D])).
fof(f23937_D,plain,(
  ( ! [X152,X151,X150] : (s(X150,X152) = s(X150,sK1(X152,X151,X150))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f30332,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,X18)) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,sK1(X18,X17,bool)))))) ) | $spl462),
  inference(cnf_transformation,[],[f30332_D])).
fof(f30332_D,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,X18)) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,sK1(X18,X17,bool)))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f30367,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(bool,bool),X125),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X125),s(bool,sK1(X126,X125,bool))))) | p(s(bool,X126)) | s(X122,X124) = s(X122,sK1(X124,X123,X122))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f30308,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK1(X2,X1,X0)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK1(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : ! [X11,X12] : (? [X6] : (s(X10,X6) = s(X10,X12) & p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X6))))) <=> p(s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.UNWIND_THM2')).
fof(f30308,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(bool,bool),X125),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X125),s(bool,sK1(X126,X125,bool))))) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X124)))) | p(s(bool,X126)) | s(X122,X124) = s(X122,sK1(X124,X123,X122))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2458,f4403])).
fof(f4403,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(X12,X14) = s(X12,sK1(X14,X13,X12))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f4043,f94])).
fof(f4043,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f175,f88])).
fof(f88,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f2458,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,sK1(X4,X6,bool))))) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(superposition,[],[f95,f88])).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f24985,plain,(
  ( ! [X46] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X46,'type.closSem.v'))) ) | $spl400),
  inference(cnf_transformation,[],[f24985_D])).
fof(f24985_D,plain,(
  ( ! [X46] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X46,'type.closSem.v'))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f30364,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X113),s(bool,sK1(X114,X113,bool))))) | p(s(bool,X114)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X112,'type.closSem.v'))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f30305,f4041])).
fof(f4041,plain,(
  ( ! [X39] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),X39),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X39,'type.closSem.v'))) ) | $spl2),
  inference(forward_demodulation,[],[f4019,f128])).
fof(f4019,plain,(
  ( ! [X39] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),X39),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',sK0) = s('type.closSem.v',sK1(sK0,X39,'type.closSem.v'))) ) | $spl2),
  inference(superposition,[],[f94,f128])).
fof(f30305,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,happ(s(fun(bool,bool),X113),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X113),s(bool,sK1(X114,X113,bool))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),X112),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,X114)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X112,'type.closSem.v'))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2458,f5283])).
fof(f5283,plain,(
  ( ! [X2] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X2),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X2,'type.closSem.v'))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f4041,f4043])).
fof(f26385,plain,(
  ( ! [X98] : (s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X98))) ) | $spl420),
  inference(cnf_transformation,[],[f26385_D])).
fof(f26385_D,plain,(
  ( ! [X98] : (s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X98))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f30345,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X63),s(bool,sK1(X64,X63,bool))))) | p(s(bool,X64)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f30290,f8638])).
fof(f8638,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X0))) )),
  inference(superposition,[],[f3956,f117])).
fof(f3956,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))),s('type.closSem.v',X1)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0',X0,X1))) )),
  inference(resolution,[],[f3923,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3)))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) & (s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3))))))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3)))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3)))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10,X38] : ! [X11,X18] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X38,bool)),X11),s('type.list.list'(X10),'const.list.NIL_0'),s('type.list.list'(X38),X18)))) <=> s('type.list.list'(X38),'const.list.NIL_0') = s('type.list.list'(X38),X18))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.list.EVERY2_THM|split|0')).
fof(f3923,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.v',X0))))) )),
  inference(forward_demodulation,[],[f119,f92])).
fof(f119,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0))))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f30290,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),X63),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X63),s(bool,sK1(X64,X63,bool))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X62)))) | p(s(bool,X64)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2458,f8674])).
fof(f8674,plain,(
  ( ! [X10] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X10))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X10))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f8638,f4043])).
fof(f28286,plain,(
  ( ! [X102] : (s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X102))) ) | $spl444),
  inference(cnf_transformation,[],[f28286_D])).
fof(f28286_D,plain,(
  ( ! [X102] : (s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X102))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f30344,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun(bool,bool),X60),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK1(X61,X60,bool))))) | p(s(bool,X61)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X59))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f30289,f8750])).
fof(f8750,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X0))) )),
  inference(superposition,[],[f4127,f117])).
fof(f4127,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0',X1,X0))) )),
  inference(resolution,[],[f4051,f84])).
fof(f84,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) & (s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0'))))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X42,X43] : ! [X11,X32] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X43,fun(X42,bool)),X11),s('type.list.list'(X43),X32),s('type.list.list'(X42),'const.list.NIL_0')))) <=> s('type.list.list'(X43),'const.list.NIL_0') = s('type.list.list'(X43),X32))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.list.EVERY2_THM|split|2')).
fof(f4051,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0)),s('type.list.list'('type.closSem.v'),X2)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))) )),
  inference(forward_demodulation,[],[f122,f92])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0)),s('type.list.list'('type.closSem.v'),X2)))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0)),s('type.list.list'('type.closSem.v'),X2)))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | ? [X4] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X4))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X4),s('type.list.list'('type.closSem.v'),X2)))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))))),
  inference(rectify,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | ? [X3] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))))) & (! [X3] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) <=> ? [X3] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13,X16,X17] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X13))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X17)))))) <=> ? [X32] : (s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X16),s('type.list.list'('type.closSem.v'),X32))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X32),s('type.list.list'('type.closSem.v'),X17))))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|7')).
fof(f30289,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun(bool,bool),X60),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK1(X61,X60,bool))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X59),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,X61)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X59))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2458,f8773])).
fof(f8773,plain,(
  ( ! [X10] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X10),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X10))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f8750,f4043])).
fof(f664,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f664_D])).
fof(f664_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f30343,plain,(
  ( ! [X47,X50,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X49),s(bool,sK1(X50,X49,bool))))) | p(s(bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl292)),
  inference(forward_demodulation,[],[f30342,f16150])).
fof(f16150,plain,(
  ( ! [X27] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X27))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl292)),
  inference(resolution,[],[f16131,f4043])).
fof(f16131,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))),s('type.closSem.v','const.closSem.Unit_0'))))) ) | $spl292),
  inference(equality_resolution,[],[f16012])).
fof(f16012,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X1) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))),s('type.closSem.v',X1))))) ) | $spl292),
  inference(superposition,[],[f15947,f115])).
fof(f115,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v',X0))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v',X0))))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v',X0)))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13,X14] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14))),s('type.closSem.v',X13)))) <=> s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|0')).
fof(f30342,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X49),s(bool,sK1(X50,X49,bool))))) | p(s(bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl8 | $spl292)),
  inference(subsumption_resolution,[],[f30286,f396])).
fof(f30286,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X49),s(bool,sK1(X50,X49,bool))))) | p(s(bool,X47)) | p(s(bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2458,f16151])).
fof(f16151,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl292)),
  inference(resolution,[],[f16131,f4239])).
fof(f4239,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(resolution,[],[f396,f88])).
fof(f30341,plain,(
  ( ! [X45,X43,X46] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,sK1(X46,X45,bool))))) | p(s(bool,X46)) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f30340,f15070])).
fof(f15070,plain,(
  ( ! [X103] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f14827,f175])).
fof(f14827,plain,(
  ( ! [X66,X67] : (p(s(bool,X66)) | s(bool,X66) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X67)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f2527,f175])).
fof(f2527,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X9),s('type.closSem.v',X9)))) | p(s(bool,X11)) | s(bool,X11) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X9),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10)))))) )),
  inference(resolution,[],[f2013,f88])).
fof(f2013,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1))))) )),
  inference(superposition,[],[f1999,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13,X14] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14)))))) <=> s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|6')).
fof(f1999,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))))))) )),
  inference(equality_resolution,[],[f112])).
fof(f112,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f30340,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44)))))))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,sK1(X46,X45,bool))))) | p(s(bool,X46)) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f30285,f396])).
fof(f30285,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44)))))))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,sK1(X46,X45,bool))))) | p(s(bool,X43)) | p(s(bool,X46)) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2458,f15039])).
fof(f15039,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11))))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl2 | $spl8)),
  inference(resolution,[],[f14827,f396])).
fof(f30339,plain,(
  ( ! [X39,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X41),s(bool,sK1(X42,X41,bool))))) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl8 | $spl258)),
  inference(forward_demodulation,[],[f30338,f12328])).
fof(f12328,plain,(
  ( ! [X24] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl258)),
  inference(resolution,[],[f12309,f4043])).
fof(f12309,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))),s('type.closSem.v','const.closSem.Unit_0'))))) ) | $spl258),
  inference(equality_resolution,[],[f12239])).
fof(f12239,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X1) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))),s('type.closSem.v',X1))))) ) | $spl258),
  inference(superposition,[],[f12180,f113])).
fof(f113,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v',X0))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v',X0))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v',X0)))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13,X19] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v',X13)))) <=> s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|3')).
fof(f30338,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X41),s(bool,sK1(X42,X41,bool))))) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X39)) ) | ($spl8 | $spl258)),
  inference(subsumption_resolution,[],[f30284,f396])).
fof(f30284,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X41),s(bool,sK1(X42,X41,bool))))) | p(s(bool,X39)) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X39)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2458,f12329])).
fof(f12329,plain,(
  ( ! [X26,X25] : (s(bool,X25) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl8 | $spl258)),
  inference(resolution,[],[f12309,f4239])).
fof(f30337,plain,(
  ( ! [X37,X35,X38] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,sK1(X38,X37,bool))))) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f30336,f11496])).
fof(f11496,plain,(
  ( ! [X69] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X69)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f11243,f175])).
fof(f11243,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f2382,f175])).
fof(f2382,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X9),s('type.closSem.v',X9)))) | p(s(bool,X11)) | s(bool,X11) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X9),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10)))))) )),
  inference(resolution,[],[f1850,f88])).
fof(f1850,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0)))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1))))) )),
  inference(superposition,[],[f1836,f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13,X19] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19)))))) <=> s('type.closSem.v',X13) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_simp|split|8')).
fof(f1836,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))))))) )),
  inference(equality_resolution,[],[f110])).
fof(f110,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f72])).
fof(f30336,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36)))))))) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,sK1(X38,X37,bool))))) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f30283,f396])).
fof(f30283,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36)))))))) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,sK1(X38,X37,bool))))) | p(s(bool,X35)) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2458,f11475])).
fof(f11475,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X11))))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl2 | $spl8)),
  inference(resolution,[],[f11243,f396])).
fof(f30334,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,sK1(X26,X25,bool)))))) ) | $spl464),
  inference(cnf_transformation,[],[f30334_D])).
fof(f30334_D,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,sK1(X26,X25,bool)))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f30335,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,sK1(X30,X29,bool))))) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f30281,f396])).
fof(f30281,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | p(s(bool,happ(s(fun(bool,bool),X29),s(bool,sK1(X30,X29,bool))))) | p(s(bool,X27)) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f2458,f4721])).
fof(f4721,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f259,f259])).
fof(f30333,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,sK1(X26,X25,bool))))) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f30280,f396])).
fof(f30280,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,sK1(X26,X25,bool))))) | p(s(bool,X23)) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2458,f4721])).
fof(f30331,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,sK1(X18,X17,bool))))) | p(s(bool,X18)) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f30278,f396])).
fof(f30278,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,sK1(X18,X17,bool))))) | p(s(bool,X16)) | p(s(bool,X18)) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f2458,f259])).
fof(f1303,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl80),
  inference(cnf_transformation,[],[f1303_D])).
fof(f1303_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f29894,plain,(
  ( ! [X66,X67] : (p(s(bool,sK1(sK1('T_0',X66,bool),X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67)))) ) | $spl460),
  inference(cnf_transformation,[],[f29894_D])).
fof(f29894_D,plain,(
  ( ! [X66,X67] : (p(s(bool,sK1(sK1('T_0',X66,bool),X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67)))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f29910,plain,(
  ( ! [X103,X102,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X103),s(bool,X104))) | p(s(bool,X102)) | p(s(bool,sK1(sK1('T_0',X103,bool),X103,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258 | $spl344 | $spl426)),
  inference(forward_demodulation,[],[f29909,f12328])).
fof(f29909,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | p(s(bool,sK1(sK1('T_0',X103,bool),X103,bool))) | s(bool,happ(s(fun(bool,bool),X103),s(bool,X104))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X101))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29854,f4107])).
fof(f4107,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f4103,f175])).
fof(f4103,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))))) ) | ($spl2 | $spl15)),
  inference(superposition,[],[f4094,f88])).
fof(f4094,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f467,f128])).
fof(f467,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) | $spl15),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f29854,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | p(s(bool,sK1(sK1('T_0',X103,bool),X103,bool))) | s(bool,happ(s(fun(bool,bool),X103),s(bool,X104))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X101))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl6 | $spl8 | $spl258 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f12329])).
fof(f29008,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK1(sK1('T_0',X0,bool),X0,bool))) | p(s(bool,X1)) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) ) | ($spl6 | $spl344 | $spl426)),
  inference(resolution,[],[f28992,f2432])).
fof(f2432,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK1(X9,X8,X7))))) | p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) )),
  inference(resolution,[],[f95,f88])).
fof(f28992,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | p(s(bool,sK1(sK1('T_0',X5,bool),X5,bool)))) ) | ($spl6 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f28964,f21356])).
fof(f21356,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,sK1(sK1('T_0',X0,bool),X0,bool)))) ) | $spl344),
  inference(resolution,[],[f21342,f93])).
fof(f93,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f64])).
fof(f28964,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | p(s(bool,sK1(sK1('T_0',X5,bool),X5,bool))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl6 | $spl426)),
  inference(superposition,[],[f26411,f259])).
fof(f29906,plain,(
  ( ! [X94,X92,X93] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X93),s(bool,X94))) | p(s(bool,X92)) | p(s(bool,sK1(sK1('T_0',X93,bool),X93,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292 | $spl344 | $spl426)),
  inference(forward_demodulation,[],[f29905,f16150])).
fof(f29905,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | p(s(bool,sK1(sK1('T_0',X93,bool),X93,bool))) | s(bool,happ(s(fun(bool,bool),X93),s(bool,X94))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X91))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29851,f4107])).
fof(f29851,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | p(s(bool,sK1(sK1('T_0',X93,bool),X93,bool))) | s(bool,happ(s(fun(bool,bool),X93),s(bool,X94))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X91))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl6 | $spl8 | $spl292 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f16151])).
fof(f29902,plain,(
  ( ! [X83,X84,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X84))) | p(s(bool,X82)) | p(s(bool,sK1(sK1('T_0',X83,bool),X83,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl344 | $spl426)),
  inference(forward_demodulation,[],[f29901,f11496])).
fof(f29901,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X82)) | p(s(bool,sK1(sK1('T_0',X83,bool),X83,bool))) | s(bool,happ(s(fun(bool,bool),X83),s(bool,X84))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X81)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29848,f4107])).
fof(f29848,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X82)) | p(s(bool,sK1(sK1('T_0',X83,bool),X83,bool))) | s(bool,happ(s(fun(bool,bool),X83),s(bool,X84))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X81))))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f11475])).
fof(f29898,plain,(
  ( ! [X76,X77,X75] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X76),s(bool,X77))) | p(s(bool,X75)) | p(s(bool,sK1(sK1('T_0',X76,bool),X76,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl344 | $spl426)),
  inference(forward_demodulation,[],[f29897,f15070])).
fof(f29897,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | p(s(bool,sK1(sK1('T_0',X76,bool),X76,bool))) | s(bool,happ(s(fun(bool,bool),X76),s(bool,X77))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X74)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29846,f4107])).
fof(f29846,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | p(s(bool,sK1(sK1('T_0',X76,bool),X76,bool))) | s(bool,happ(s(fun(bool,bool),X76),s(bool,X77))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X74))))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f15039])).
fof(f29896,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,sK1(sK1('T_0',X72,bool),X72,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X72),s(bool,X73)))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29845,f4107])).
fof(f29845,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,sK1(sK1('T_0',X72,bool),X72,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X72),s(bool,X73))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f259])).
fof(f29895,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,sK1(sK1('T_0',X69,bool),X69,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X69),s(bool,X70))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X68))) ) | ($spl1 | $spl2 | $spl6 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29844,f8638])).
fof(f29844,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X68)))) | p(s(bool,sK1(sK1('T_0',X69,bool),X69,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X69),s(bool,X70))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X68))) ) | ($spl1 | $spl2 | $spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f8674])).
fof(f29893,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,sK1(sK1('T_0',X66,bool),X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X65))) ) | ($spl1 | $spl2 | $spl6 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29843,f8750])).
fof(f29843,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X65),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,sK1(sK1('T_0',X66,bool),X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X65))) ) | ($spl1 | $spl2 | $spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f8773])).
fof(f29881,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,sK1(sK1('T_0',X25,bool),X25,bool))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) | $spl458),
  inference(cnf_transformation,[],[f29881_D])).
fof(f29881_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,sK1(sK1('T_0',X25,bool),X25,bool))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f29882,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1(sK1('T_0',X29,bool),X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29833,f4107])).
fof(f29833,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1(sK1('T_0',X29,bool),X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f4721])).
fof(f29880,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1(sK1('T_0',X25,bool),X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29832,f4107])).
fof(f29832,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1(sK1('T_0',X25,bool),X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f4721])).
fof(f29869,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,X25)) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X23,bool),X23,bool)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) | $spl456),
  inference(cnf_transformation,[],[f29869_D])).
fof(f29869_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,X25)) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X23,bool),X23,bool)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f29870,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X27,bool),X27,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29819,f4107])).
fof(f29819,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X27,bool),X27,bool)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f4721])).
fof(f29868,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X23,bool),X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl344 | $spl426)),
  inference(subsumption_resolution,[],[f29818,f4107])).
fof(f29818,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X23,bool),X23,bool))) ) | ($spl6 | $spl344 | $spl426)),
  inference(superposition,[],[f29008,f4721])).
fof(f25201,plain,(
  ( ! [X19,X17] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | $spl406),
  inference(cnf_transformation,[],[f25201_D])).
fof(f25201_D,plain,(
  ( ! [X19,X17] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f29380,plain,(
  ( ! [X401,X403,X402] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X401),s(bool,X402))) | s(bool,'T_0') = s(bool,X403) | s(bool,'T_0') = s(bool,sK1('T_0',X401,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl292 | $spl404)),
  inference(forward_demodulation,[],[f29379,f16150])).
fof(f29379,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,'T_0') = s(bool,X403) | s(bool,happ(s(fun(bool,bool),X401),s(bool,X402))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X404))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X401,bool))) ) | ($spl6 | $spl96 | $spl292 | $spl404)),
  inference(subsumption_resolution,[],[f29282,f4721])).
fof(f29282,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,X403) != s(bool,sK1('T_0',X401,bool)) | s(bool,'T_0') = s(bool,X403) | s(bool,happ(s(fun(bool,bool),X401),s(bool,X402))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X404))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X401,bool))) ) | ($spl96 | $spl292 | $spl404)),
  inference(superposition,[],[f16156,f25153])).
fof(f16156,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,X40) | s(bool,'T_0') = s(bool,X39) | s(bool,X40) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X41))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl96 | $spl292)),
  inference(resolution,[],[f16131,f5453])).
fof(f29374,plain,(
  ( ! [X380,X382,X381] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X380),s(bool,X381))) | s(bool,'T_0') = s(bool,X382) | s(bool,'T_0') = s(bool,sK1('T_0',X380,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(forward_demodulation,[],[f29373,f15070])).
fof(f29373,plain,(
  ( ! [X383,X380,X382,X381] : (s(bool,'T_0') = s(bool,X382) | s(bool,happ(s(fun(bool,bool),X380),s(bool,X381))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X383))))) | s(bool,'T_0') = s(bool,sK1('T_0',X380,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(subsumption_resolution,[],[f29276,f4721])).
fof(f29276,plain,(
  ( ! [X383,X380,X382,X381] : (s(bool,X382) != s(bool,sK1('T_0',X380,bool)) | s(bool,'T_0') = s(bool,X382) | s(bool,happ(s(fun(bool,bool),X380),s(bool,X381))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X383))))) | s(bool,'T_0') = s(bool,sK1('T_0',X380,bool))) ) | ($spl1 | $spl2 | $spl96 | $spl404)),
  inference(superposition,[],[f15349,f25153])).
fof(f15349,plain,(
  ( ! [X116,X117,X115] : (s(bool,X116) != s(bool,X117) | s(bool,'T_0') = s(bool,X116) | s(bool,X117) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X115)))))) ) | ($spl1 | $spl2 | $spl96)),
  inference(subsumption_resolution,[],[f15275,f175])).
fof(f15275,plain,(
  ( ! [X116,X117,X115] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X116) | s(bool,X116) != s(bool,X117) | s(bool,X117) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X115)))))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f5453,f15070])).
fof(f29368,plain,(
  ( ! [X360,X359,X361] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X359),s(bool,X360))) | s(bool,'T_0') = s(bool,X361) | s(bool,'T_0') = s(bool,sK1('T_0',X359,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl258 | $spl404)),
  inference(forward_demodulation,[],[f29367,f12328])).
fof(f29367,plain,(
  ( ! [X360,X362,X359,X361] : (s(bool,'T_0') = s(bool,X361) | s(bool,happ(s(fun(bool,bool),X359),s(bool,X360))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X362))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X359,bool))) ) | ($spl6 | $spl96 | $spl258 | $spl404)),
  inference(subsumption_resolution,[],[f29270,f4721])).
fof(f29270,plain,(
  ( ! [X360,X362,X359,X361] : (s(bool,X361) != s(bool,sK1('T_0',X359,bool)) | s(bool,'T_0') = s(bool,X361) | s(bool,happ(s(fun(bool,bool),X359),s(bool,X360))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X362))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X359,bool))) ) | ($spl96 | $spl258 | $spl404)),
  inference(superposition,[],[f12334,f25153])).
fof(f12334,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl96 | $spl258)),
  inference(resolution,[],[f12309,f5453])).
fof(f29362,plain,(
  ( ! [X339,X340,X338] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X338),s(bool,X339))) | s(bool,'T_0') = s(bool,X340) | s(bool,'T_0') = s(bool,sK1('T_0',X338,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(forward_demodulation,[],[f29361,f11496])).
fof(f29361,plain,(
  ( ! [X341,X339,X340,X338] : (s(bool,'T_0') = s(bool,X340) | s(bool,happ(s(fun(bool,bool),X338),s(bool,X339))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X341))))) | s(bool,'T_0') = s(bool,sK1('T_0',X338,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(subsumption_resolution,[],[f29264,f4721])).
fof(f29264,plain,(
  ( ! [X341,X339,X340,X338] : (s(bool,X340) != s(bool,sK1('T_0',X338,bool)) | s(bool,'T_0') = s(bool,X340) | s(bool,happ(s(fun(bool,bool),X338),s(bool,X339))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X341))))) | s(bool,'T_0') = s(bool,sK1('T_0',X338,bool))) ) | ($spl1 | $spl2 | $spl96 | $spl404)),
  inference(superposition,[],[f11751,f25153])).
fof(f11751,plain,(
  ( ! [X103,X105,X104] : (s(bool,X104) != s(bool,X105) | s(bool,'T_0') = s(bool,X104) | s(bool,X105) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103)))))) ) | ($spl1 | $spl2 | $spl96)),
  inference(subsumption_resolution,[],[f11687,f175])).
fof(f11687,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X104) | s(bool,X104) != s(bool,X105) | s(bool,X105) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103)))))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f5453,f11496])).
fof(f25447,plain,(
  ( ! [X26,X23,X25] : (s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) | $spl408),
  inference(cnf_transformation,[],[f25447_D])).
fof(f25447_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f29350,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,X264) = s(bool,happ(s(fun(bool,bool),X261),s(bool,X262))) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X263) | s(bool,'T_0') = s(bool,sK1('T_0',X261,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(subsumption_resolution,[],[f29242,f4721])).
fof(f29242,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,X263) != s(bool,sK1('T_0',X261,bool)) | s(bool,X264) = s(bool,happ(s(fun(bool,bool),X261),s(bool,X262))) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,X263) | s(bool,'T_0') = s(bool,sK1('T_0',X261,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(superposition,[],[f6249,f25153])).
fof(f6249,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(resolution,[],[f4936,f396])).
fof(f4936,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4107,f4721])).
fof(f29349,plain,(
  ( ! [X255,X257,X256] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X255),s(bool,X256))) | s(bool,'T_0') = s(bool,X257) | s(bool,'T_0') = s(bool,sK1('T_0',X255,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(subsumption_resolution,[],[f29240,f4721])).
fof(f29240,plain,(
  ( ! [X255,X257,X256] : (s(bool,X257) != s(bool,sK1('T_0',X255,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X255),s(bool,X256))) | s(bool,'T_0') = s(bool,X257) | s(bool,'T_0') = s(bool,sK1('T_0',X255,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl404)),
  inference(superposition,[],[f6091,f25153])).
fof(f6091,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(resolution,[],[f4394,f396])).
fof(f4394,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X0) != s(bool,X1)) ) | ($spl1 | $spl2 | $spl15)),
  inference(resolution,[],[f4043,f4124])).
fof(f4124,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X5)) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f4111,f124])).
fof(f4111,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,'T_0')) | ~p(s(bool,X4))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4107,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f57])).
fof(f25153,plain,(
  ( ! [X19,X17] : (s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | $spl404),
  inference(cnf_transformation,[],[f25153_D])).
fof(f25153_D,plain,(
  ( ! [X19,X17] : (s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f28695,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f28684,f4107])).
fof(f28684,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f24430,plain,(
  ( ! [X0,X1] : (p(s(bool,sK1('T_0',X0,bool))) | s(bool,sK1('T_0',X0,bool)) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(factoring,[],[f24208])).
fof(f24208,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,sK1('T_0',X7,bool))) | p(s(bool,X9)) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(resolution,[],[f24129,f88])).
fof(f24129,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,sK1('T_0',X1,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24103,f21004])).
fof(f21004,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X29)))) | p(s(bool,sK1('T_0',X28,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20892,f4107])).
fof(f20892,plain,(
  ( ! [X28,X29] : (p(s(bool,sK1('T_0',X28,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X29)))) | s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,sK1('T_0',X28,bool))) ) | ($spl8 | $spl106)),
  inference(superposition,[],[f93,f20778])).
fof(f20778,plain,(
  ( ! [X11] : (s(bool,sK1('T_0',X11,bool)) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f20688,f396])).
fof(f20688,plain,(
  ( ! [X0] : (p(s(bool,sK1('T_0',X0,bool))) | s(bool,sK1('T_0',X0,bool)) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl106),
  inference(factoring,[],[f20508])).
fof(f20508,plain,(
  ( ! [X92,X91] : (p(s(bool,sK1('T_0',X91,bool))) | p(s(bool,X92)) | s(bool,X92) = s(bool,happ(s(fun(bool,bool),X91),s(bool,'T_0')))) ) | $spl106),
  inference(resolution,[],[f8010,f108])).
fof(f108,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.v_rel_Unit|split|1',aHL_TRUTH)).
fof(f8010,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,sK1(X3,X4,bool))) | s(bool,X5) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))) | p(s(bool,X5))) ) | $spl106),
  inference(resolution,[],[f7447,f88])).
fof(f24103,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,sK1('T_0',X1,bool))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl390)),
  inference(superposition,[],[f23981,f259])).
fof(f28694,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f28683,f4107])).
fof(f28683,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f258,plain,(
  ( ! [X11] : (~p(s(bool,X11))) ) | $spl4),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  ( ! [X11] : (~p(s(bool,X11))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f27431,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X12)))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | $spl440),
  inference(cnf_transformation,[],[f27431_D])).
fof(f27431_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X12)))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f27436,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,happ(s(fun(bool,bool),X103),s(bool,X104)))) | ~p(s(bool,X105)) | s(bool,'T_0') = s(bool,sK1('T_0',X103,bool))) ) | ($spl176 | $spl404)),
  inference(subsumption_resolution,[],[f27310,f8143])).
fof(f27310,plain,(
  ( ! [X103,X105,X104] : (s(bool,X105) = s(bool,sK1('T_0',X103,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X103),s(bool,X104)))) | ~p(s(bool,X105)) | s(bool,'T_0') = s(bool,sK1('T_0',X103,bool))) ) | $spl404),
  inference(superposition,[],[f91,f25153])).
fof(f27437,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X108)) | ~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X107)))) | s(bool,'T_0') = s(bool,sK1('T_0',X106,bool))) ) | ($spl176 | $spl404)),
  inference(subsumption_resolution,[],[f27311,f8143])).
fof(f27311,plain,(
  ( ! [X107,X108,X106] : (s(bool,X108) = s(bool,sK1('T_0',X106,bool)) | ~p(s(bool,X108)) | ~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X107)))) | s(bool,'T_0') = s(bool,sK1('T_0',X106,bool))) ) | $spl404),
  inference(superposition,[],[f91,f25153])).
fof(f27430,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,sK1('T_0',X11,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X12)))) | ~p(s(bool,X13))) ) | ($spl176 | $spl404)),
  inference(subsumption_resolution,[],[f27257,f8143])).
fof(f27257,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,sK1('T_0',X11,bool)) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X12)))) | ~p(s(bool,X13))) ) | $spl404),
  inference(superposition,[],[f25153,f91])).
fof(f27432,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | ~p(s(bool,X16)) | ~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15))))) ) | ($spl176 | $spl404)),
  inference(subsumption_resolution,[],[f27258,f8143])).
fof(f27258,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) = s(bool,sK1('T_0',X14,bool)) | s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | ~p(s(bool,X16)) | ~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15))))) ) | $spl404),
  inference(superposition,[],[f25153,f91])).
fof(f28359,plain,(
  ( ! [X337,X339] : (s(bool,X339) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X337),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X337)) | ~p(s(bool,X339))) ) | $spl454),
  inference(cnf_transformation,[],[f28359_D])).
fof(f28359_D,plain,(
  ( ! [X337,X339] : (s(bool,X339) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X337),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X337)) | ~p(s(bool,X339))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f23981,plain,(
  ( ! [X515] : (~p(s(bool,happ(s(fun(bool,bool),X515),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1('T_0',X515,bool)))) ) | $spl390),
  inference(cnf_transformation,[],[f23981_D])).
fof(f23981_D,plain,(
  ( ! [X515] : (~p(s(bool,happ(s(fun(bool,bool),X515),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1('T_0',X515,bool)))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f28227,plain,(
  ( ! [X337,X339,X338] : (~p(s(bool,happ(s(fun(bool,bool),X338),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X339) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X337),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,sK1('T_0',X338,bool))) | ~p(s(bool,X339)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X337))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f8773])).
fof(f21098,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | s(bool,X14) != s(bool,X15) | p(s(bool,sK1('T_0',X16,bool))) | ~p(s(bool,X14))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21052,f124])).
fof(f21052,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | p(s(bool,sK1('T_0',X16,bool))) | ~p(s(bool,'T_0')) | ~p(s(bool,X14))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f91])).
fof(f28356,plain,(
  ( ! [X315,X317] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X315),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X317))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X315))) ) | $spl452),
  inference(cnf_transformation,[],[f28356_D])).
fof(f28356_D,plain,(
  ( ! [X315,X317] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X315),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X317))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X315))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f28355,plain,(
  ( ! [X315,X316,X317] : (s(bool,'T_0') = s(bool,X316) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X315),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X317))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X315))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl292)),
  inference(subsumption_resolution,[],[f28218,f259])).
fof(f28218,plain,(
  ( ! [X315,X316,X317] : (s(bool,X316) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X316) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X315),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X317))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X315))) ) | ($spl1 | $spl2 | $spl96 | $spl292)),
  inference(superposition,[],[f16156,f8773])).
fof(f28353,plain,(
  ( ! [X300,X302] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X300),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X302))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X300))) ) | $spl450),
  inference(cnf_transformation,[],[f28353_D])).
fof(f28353_D,plain,(
  ( ! [X300,X302] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X300),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X302))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X300))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f28352,plain,(
  ( ! [X300,X302,X301] : (s(bool,'T_0') = s(bool,X301) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X300),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X302))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X300))) ) | ($spl1 | $spl2 | $spl6 | $spl96)),
  inference(subsumption_resolution,[],[f28212,f259])).
fof(f28212,plain,(
  ( ! [X300,X302,X301] : (s(bool,X301) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X301) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X300),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X302))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X300))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f15349,f8773])).
fof(f28350,plain,(
  ( ! [X285,X287] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X285),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X287))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X285))) ) | $spl448),
  inference(cnf_transformation,[],[f28350_D])).
fof(f28350_D,plain,(
  ( ! [X285,X287] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X285),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X287))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X285))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f28349,plain,(
  ( ! [X285,X287,X286] : (s(bool,'T_0') = s(bool,X286) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X285),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X287))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X285))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl258)),
  inference(subsumption_resolution,[],[f28206,f259])).
fof(f28206,plain,(
  ( ! [X285,X287,X286] : (s(bool,X286) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X286) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X285),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X287))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X285))) ) | ($spl1 | $spl2 | $spl96 | $spl258)),
  inference(superposition,[],[f12334,f8773])).
fof(f28347,plain,(
  ( ! [X270,X272] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X270),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X272))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X270))) ) | $spl446),
  inference(cnf_transformation,[],[f28347_D])).
fof(f28347_D,plain,(
  ( ! [X270,X272] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X270),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X272))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X270))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f28346,plain,(
  ( ! [X271,X270,X272] : (s(bool,'T_0') = s(bool,X271) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X270),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X272))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X270))) ) | ($spl1 | $spl2 | $spl6 | $spl96)),
  inference(subsumption_resolution,[],[f28200,f259])).
fof(f28200,plain,(
  ( ! [X271,X270,X272] : (s(bool,X271) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X271) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X270),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X272))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X270))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f11751,f8773])).
fof(f10198,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X19),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19))) ) | $spl208),
  inference(cnf_transformation,[],[f10198_D])).
fof(f10198_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X19),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f28337,plain,(
  ( ! [X216,X215,X217] : (s(bool,X217) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X215),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X217) | s(bool,'T_0') = s(bool,X216) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X215))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f28178,f259])).
fof(f28178,plain,(
  ( ! [X216,X215,X217] : (s(bool,X216) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X217) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X215),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X217) | s(bool,'T_0') = s(bool,X216) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X215))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f8773])).
fof(f19492,plain,(
  ( ! [X17,X18] : (s(bool,sK1(X17,X18,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17)))) | ~p(s(bool,X17))) ) | $spl318),
  inference(cnf_transformation,[],[f19492_D])).
fof(f19492_D,plain,(
  ( ! [X17,X18] : (s(bool,sK1(X17,X18,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17)))) | ~p(s(bool,X17))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f28328,plain,(
  ( ! [X182,X180,X181] : (s(bool,sK1(X181,X182,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X181)) | ~p(s(bool,happ(s(fun(bool,bool),X182),s(bool,X181)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X180))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f28163,f8750])).
fof(f28163,plain,(
  ( ! [X182,X180,X181] : (s(bool,sK1(X181,X182,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X180),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,X181)) | ~p(s(bool,happ(s(fun(bool,bool),X182),s(bool,X181)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X180))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f8773])).
fof(f4312,plain,(
  ( ! [X17,X15,X16] : (s(bool,X17) != s(bool,sK1(X15,X16,bool)) | p(s(bool,X17)) | ~p(s(bool,X15)) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15))))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4124,f94])).
fof(f9372,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,X18)) | s(bool,X18) = s(bool,sK1(X18,X17,bool))) ) | $spl198),
  inference(cnf_transformation,[],[f9372_D])).
fof(f9372_D,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,X18)) | s(bool,X18) = s(bool,sK1(X18,X17,bool))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f28318,plain,(
  ( ! [X127,X128,X126] : (~p(s(bool,happ(s(fun(bool,bool),X127),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X128) = s(bool,sK1(X128,X127,bool)) | p(s(bool,X128)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X126))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f28143,f8750])).
fof(f28143,plain,(
  ( ! [X127,X128,X126] : (~p(s(bool,happ(s(fun(bool,bool),X127),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X128) = s(bool,sK1(X128,X127,bool)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X126),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,X128)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X126))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1730,f8773])).
fof(f1730,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X5)))) | s(bool,X4) = s(bool,sK1(X4,X6,bool)) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(superposition,[],[f94,f88])).
fof(f28309,plain,(
  ( ! [X313,X315,X314] : (s(bool,'T_0') = s(bool,sK1('T_0',X314,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X314),s(bool,X315))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X313))) ) | ($spl1 | $spl2 | $spl394)),
  inference(subsumption_resolution,[],[f28048,f8750])).
fof(f28048,plain,(
  ( ! [X313,X315,X314] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X313),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,sK1('T_0',X314,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X314),s(bool,X315))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X313))) ) | ($spl1 | $spl2 | $spl394)),
  inference(superposition,[],[f24482,f8773])).
fof(f26411,plain,(
  ( ! [X302] : (~p(s(bool,happ(s(fun(bool,bool),X302),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1(sK1('T_0',X302,bool),X302,bool)))) ) | $spl426),
  inference(cnf_transformation,[],[f26411_D])).
fof(f26411_D,plain,(
  ( ! [X302] : (~p(s(bool,happ(s(fun(bool,bool),X302),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1(sK1('T_0',X302,bool),X302,bool)))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f28270,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | $spl442),
  inference(cnf_transformation,[],[f28270_D])).
fof(f28270_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f28045,plain,(
  ( ! [X306,X305] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X305),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X306),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1(sK1('T_0',X306,bool),X306,bool))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X305))) ) | ($spl1 | $spl2 | $spl344)),
  inference(superposition,[],[f21356,f8773])).
fof(f28041,plain,(
  ( ! [X296,X295] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X295),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,happ(s(fun(bool,bool),X296),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1('T_0',X296,bool))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X295))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f8773])).
fof(f21689,plain,(
  ( ! [X45] : (p(s(bool,sK1(sK1('T_0',X45,bool),X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) | $spl348),
  inference(cnf_transformation,[],[f21689_D])).
fof(f21689_D,plain,(
  ( ! [X45] : (p(s(bool,sK1(sK1('T_0',X45,bool),X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f28308,plain,(
  ( ! [X294,X293] : (p(s(bool,sK1(sK1('T_0',X294,bool),X294,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X294),s(bool,'T_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f28040,f8750])).
fof(f28040,plain,(
  ( ! [X294,X293] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X293),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,sK1(sK1('T_0',X294,bool),X294,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X294),s(bool,'T_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f8773])).
fof(f20737,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | p(s(bool,sK1(sK1('T_0',X4,bool),X4,bool))) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20573,f88])).
fof(f20573,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))) | p(s(bool,sK1(sK1('T_0',X4,bool),X4,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(resolution,[],[f20508,f19531])).
fof(f19531,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK1(X0,X1,bool))) | p(s(bool,sK1(sK1(X0,X1,bool),X1,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) ) | ($spl1 | $spl2 | $spl15)),
  inference(resolution,[],[f19487,f95])).
fof(f19487,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | ~p(s(bool,X0)) | p(s(bool,sK1(X0,X1,bool)))) ) | ($spl1 | $spl2 | $spl15)),
  inference(equality_resolution,[],[f4312])).
fof(f19166,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X61))),s('type.closSem.v',X60)))) ) | $spl316),
  inference(cnf_transformation,[],[f19166_D])).
fof(f19166_D,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X61))),s('type.closSem.v',X60)))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f28303,plain,(
  ( ! [X218,X219,X217] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X218),s('type.closSem.v',X218)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X219))),s('type.closSem.v',X218))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X217))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f28020,f8750])).
fof(f28020,plain,(
  ( ! [X218,X219,X217] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X217),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X218),s('type.closSem.v',X218)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X219))),s('type.closSem.v',X218))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X217))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f9506,f8773])).
fof(f9506,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X6),s('type.closSem.v',X6)))) | p(s(bool,X8)) | s(bool,X8) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X7))),s('type.closSem.v',X6)))) )),
  inference(resolution,[],[f9058,f88])).
fof(f9058,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0))),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1))))) )),
  inference(superposition,[],[f1999,f115])).
fof(f18876,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X61))),s('type.closSem.v',X60)))) ) | $spl310),
  inference(cnf_transformation,[],[f18876_D])).
fof(f18876_D,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X61))),s('type.closSem.v',X60)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f28302,plain,(
  ( ! [X216,X215,X214] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X215),s('type.closSem.v',X215)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X216))),s('type.closSem.v',X215))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X214))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f28019,f8750])).
fof(f28019,plain,(
  ( ! [X216,X215,X214] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X214),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X215),s('type.closSem.v',X215)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X216))),s('type.closSem.v',X215))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X214))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f9393,f8773])).
fof(f9393,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X6),s('type.closSem.v',X6)))) | p(s(bool,X8)) | s(bool,X8) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X7))),s('type.closSem.v',X6)))) )),
  inference(resolution,[],[f9022,f88])).
fof(f9022,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))),s('type.closSem.v',X1)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1))))) )),
  inference(superposition,[],[f1836,f113])).
fof(f26401,plain,(
  ( ! [X191,X190] : (~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | ~p(s(bool,X190)) | p(s(bool,sK1(X190,X191,bool)))) ) | $spl422),
  inference(cnf_transformation,[],[f26401_D])).
fof(f26401_D,plain,(
  ( ! [X191,X190] : (~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | ~p(s(bool,X190)) | p(s(bool,sK1(X190,X191,bool)))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f28013,plain,(
  ( ! [X194,X195,X193] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X193),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,sK1(X194,X195,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X195),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | ~p(s(bool,X194)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X193))) ) | ($spl1 | $spl2 | $spl106)),
  inference(superposition,[],[f8009,f8773])).
fof(f8009,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,sK1(X0,X1,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,X0))) ) | $spl106),
  inference(resolution,[],[f7447,f93])).
fof(f25987,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X60) = s(bool,X61) | s(bool,X61) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl418),
  inference(cnf_transformation,[],[f25987_D])).
fof(f25987_D,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X60) = s(bool,X61) | s(bool,X61) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f28299,plain,(
  ( ! [X173,X174,X172] : (s(bool,X173) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X174) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X173) = s(bool,X174) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X172))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f28005,f8750])).
fof(f28005,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X172),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X173) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X174) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X173) = s(bool,X174) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X172))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f5987,f8773])).
fof(f5987,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f4398,f88])).
fof(f4398,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f4043,f91])).
fof(f15025,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v',X52)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X53)))))) ) | $spl280),
  inference(cnf_transformation,[],[f15025_D])).
fof(f15025_D,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v',X52)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X53)))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f28292,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X126),s('type.closSem.v',X126)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X126),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X127))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X125))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27984,f8750])).
fof(f27984,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X125),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X126),s('type.closSem.v',X126)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X126),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X127))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X125))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2527,f8773])).
fof(f25794,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK1(X77,X76,X75))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,X77)))) ) | $spl414),
  inference(cnf_transformation,[],[f25794_D])).
fof(f25794_D,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK1(X77,X76,X75))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,X77)))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f28291,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,happ(s(fun(X122,bool),X123),s(X122,sK1(X124,X123,X122))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X122,bool),X123),s(X122,X124))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X121))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27983,f8750])).
fof(f27983,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,sK1(X124,X123,X122))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X122,bool),X123),s(X122,X124))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X121))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2432,f8773])).
fof(f11406,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v',X44)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))))) ) | $spl246),
  inference(cnf_transformation,[],[f11406_D])).
fof(f11406_D,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v',X44)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f28290,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X116))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X114))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27981,f8750])).
fof(f27981,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X114),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X116))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X114))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2382,f8773])).
fof(f10774,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44)))) ) | $spl228),
  inference(cnf_transformation,[],[f10774_D])).
fof(f10774_D,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f28289,plain,(
  ( ! [X111,X112,X113] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X112))),s('type.closSem.v',X113))) | s('type.closSem.v',X113) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X112))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X111))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27980,f8750])).
fof(f27980,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X112))),s('type.closSem.v',X113))) | s('type.closSem.v',X113) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X112))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X111))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2217,f8773])).
fof(f2217,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X8))),s('type.closSem.v',X7))) | s('type.closSem.v',X7) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X8)))) )),
  inference(resolution,[],[f115,f88])).
fof(f10666,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44)))) ) | $spl224),
  inference(cnf_transformation,[],[f10666_D])).
fof(f10666_D,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f28288,plain,(
  ( ! [X109,X110,X108] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))),s('type.closSem.v',X110))) | s('type.closSem.v',X110) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X108))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27979,f8750])).
fof(f27979,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X108),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))),s('type.closSem.v',X110))) | s('type.closSem.v',X110) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X108))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2079,f8773])).
fof(f2079,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X8))),s('type.closSem.v',X7))) | s('type.closSem.v',X7) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X8)))) )),
  inference(resolution,[],[f113,f88])).
fof(f10558,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45)))) ) | $spl220),
  inference(cnf_transformation,[],[f10558_D])).
fof(f10558_D,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f28287,plain,(
  ( ! [X107,X105,X106] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X106),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))))) | s('type.closSem.v',X106) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X105))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27978,f8750])).
fof(f27978,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X106),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))))) | s('type.closSem.v',X106) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X105))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1911,f8773])).
fof(f1911,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X3),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X4))))) | s('type.closSem.v',X3) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X4)))) )),
  inference(resolution,[],[f111,f88])).
fof(f10450,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))) ) | $spl216),
  inference(cnf_transformation,[],[f10450_D])).
fof(f10450_D,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f28285,plain,(
  ( ! [X103,X102,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X103),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X104))))) | s('type.closSem.v',X103) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X104))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X102))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f27977,f8750])).
fof(f27977,plain,(
  ( ! [X103,X102,X104] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X102),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X103),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X104))))) | s('type.closSem.v',X103) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X104))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X102))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1773,f8773])).
fof(f1773,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X3),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X4))))) | s('type.closSem.v',X3) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X4)))) )),
  inference(resolution,[],[f109,f88])).
fof(f11381,plain,(
  s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')) | $spl240),
  inference(cnf_transformation,[],[f11381_D])).
fof(f11381_D,plain,(
  s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f27944,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f26043,f8773])).
fof(f26043,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X3))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X3))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4094,f8674])).
fof(f27880,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl396)),
  inference(subsumption_resolution,[],[f27869,f4107])).
fof(f27869,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl396)),
  inference(superposition,[],[f24494,f4721])).
fof(f27879,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl396)),
  inference(subsumption_resolution,[],[f27868,f4107])).
fof(f27868,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl6 | $spl396)),
  inference(superposition,[],[f24494,f4721])).
fof(f27834,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) | p(s(bool,X96)) | s(bool,'T_0') = s(bool,sK1('T_0',X97,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292 | $spl394)),
  inference(forward_demodulation,[],[f27833,f16150])).
fof(f27833,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | s(bool,'T_0') = s(bool,sK1('T_0',X97,bool)) | s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X95))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292 | $spl394)),
  inference(subsumption_resolution,[],[f27797,f4107])).
fof(f27797,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | s(bool,'T_0') = s(bool,sK1('T_0',X97,bool)) | s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X95))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X96)) ) | ($spl8 | $spl292 | $spl394)),
  inference(superposition,[],[f24482,f16151])).
fof(f27830,plain,(
  ( ! [X88,X87,X86] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) | p(s(bool,X86)) | s(bool,'T_0') = s(bool,sK1('T_0',X87,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl394)),
  inference(forward_demodulation,[],[f27829,f12328])).
fof(f27829,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,'T_0') = s(bool,sK1('T_0',X87,bool)) | s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl394)),
  inference(subsumption_resolution,[],[f27794,f4107])).
fof(f27794,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,'T_0') = s(bool,sK1('T_0',X87,bool)) | s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X86)) ) | ($spl8 | $spl258 | $spl394)),
  inference(superposition,[],[f24482,f12329])).
fof(f27826,plain,(
  ( ! [X78,X76,X77] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) | p(s(bool,X76)) | s(bool,'T_0') = s(bool,sK1('T_0',X77,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(forward_demodulation,[],[f27825,f11496])).
fof(f27825,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | s(bool,'T_0') = s(bool,sK1('T_0',X77,bool)) | s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f27791,f4107])).
fof(f27791,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | s(bool,'T_0') = s(bool,sK1('T_0',X77,bool)) | s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75))))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl1 | $spl2 | $spl8 | $spl394)),
  inference(superposition,[],[f24482,f11475])).
fof(f27822,plain,(
  ( ! [X70,X71,X69] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) | p(s(bool,X69)) | s(bool,'T_0') = s(bool,sK1('T_0',X70,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(forward_demodulation,[],[f27821,f15070])).
fof(f27821,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,sK1('T_0',X70,bool)) | s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f27789,f4107])).
fof(f27789,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,sK1('T_0',X70,bool)) | s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68))))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl1 | $spl2 | $spl8 | $spl394)),
  inference(superposition,[],[f24482,f15039])).
fof(f27820,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | s(bool,'T_0') = s(bool,sK1('T_0',X66,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67)))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f27788,f4107])).
fof(f27788,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | s(bool,'T_0') = s(bool,sK1('T_0',X66,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f259])).
fof(f27819,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,sK1('T_0',X63,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,X64))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2 | $spl394)),
  inference(subsumption_resolution,[],[f27787,f8638])).
fof(f27787,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X62)))) | s(bool,'T_0') = s(bool,sK1('T_0',X63,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,X64))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2 | $spl394)),
  inference(superposition,[],[f24482,f8674])).
fof(f27809,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,sK1('T_0',X29,bool)) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f27778,f4107])).
fof(f27778,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,sK1('T_0',X29,bool)) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f4721])).
fof(f27808,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f27777,f4107])).
fof(f27777,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f4721])).
fof(f27533,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl398)),
  inference(subsumption_resolution,[],[f27522,f4107])).
fof(f27522,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl398)),
  inference(superposition,[],[f24505,f4721])).
fof(f27532,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl398)),
  inference(subsumption_resolution,[],[f27521,f4107])).
fof(f27521,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl6 | $spl398)),
  inference(superposition,[],[f24505,f4721])).
fof(f27488,plain,(
  ( ! [X397,X396,X398] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X396),s(bool,X397))) | s(bool,'T_0') = s(bool,X398) | s(bool,'T_0') = s(bool,sK1('T_0',X396,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl292 | $spl404)),
  inference(forward_demodulation,[],[f27487,f16150])).
fof(f27487,plain,(
  ( ! [X397,X399,X396,X398] : (s(bool,'T_0') = s(bool,X398) | s(bool,happ(s(fun(bool,bool),X396),s(bool,X397))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X399))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X396,bool))) ) | ($spl6 | $spl96 | $spl292 | $spl404)),
  inference(subsumption_resolution,[],[f27395,f4721])).
fof(f27395,plain,(
  ( ! [X397,X399,X396,X398] : (s(bool,X398) != s(bool,sK1('T_0',X396,bool)) | s(bool,'T_0') = s(bool,X398) | s(bool,happ(s(fun(bool,bool),X396),s(bool,X397))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X399))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X396,bool))) ) | ($spl96 | $spl292 | $spl404)),
  inference(superposition,[],[f16156,f25153])).
fof(f27482,plain,(
  ( ! [X377,X376,X375] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X375),s(bool,X376))) | s(bool,'T_0') = s(bool,X377) | s(bool,'T_0') = s(bool,sK1('T_0',X375,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(forward_demodulation,[],[f27481,f15070])).
fof(f27481,plain,(
  ( ! [X377,X376,X378,X375] : (s(bool,'T_0') = s(bool,X377) | s(bool,happ(s(fun(bool,bool),X375),s(bool,X376))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X378))))) | s(bool,'T_0') = s(bool,sK1('T_0',X375,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(subsumption_resolution,[],[f27389,f4721])).
fof(f27389,plain,(
  ( ! [X377,X376,X378,X375] : (s(bool,X377) != s(bool,sK1('T_0',X375,bool)) | s(bool,'T_0') = s(bool,X377) | s(bool,happ(s(fun(bool,bool),X375),s(bool,X376))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X378))))) | s(bool,'T_0') = s(bool,sK1('T_0',X375,bool))) ) | ($spl1 | $spl2 | $spl96 | $spl404)),
  inference(superposition,[],[f15349,f25153])).
fof(f27476,plain,(
  ( ! [X356,X354,X355] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X354),s(bool,X355))) | s(bool,'T_0') = s(bool,X356) | s(bool,'T_0') = s(bool,sK1('T_0',X354,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl258 | $spl404)),
  inference(forward_demodulation,[],[f27475,f12328])).
fof(f27475,plain,(
  ( ! [X356,X354,X357,X355] : (s(bool,'T_0') = s(bool,X356) | s(bool,happ(s(fun(bool,bool),X354),s(bool,X355))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X357))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X354,bool))) ) | ($spl6 | $spl96 | $spl258 | $spl404)),
  inference(subsumption_resolution,[],[f27383,f4721])).
fof(f27383,plain,(
  ( ! [X356,X354,X357,X355] : (s(bool,X356) != s(bool,sK1('T_0',X354,bool)) | s(bool,'T_0') = s(bool,X356) | s(bool,happ(s(fun(bool,bool),X354),s(bool,X355))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X357))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X354,bool))) ) | ($spl96 | $spl258 | $spl404)),
  inference(superposition,[],[f12334,f25153])).
fof(f27470,plain,(
  ( ! [X333,X335,X334] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X333),s(bool,X334))) | s(bool,'T_0') = s(bool,X335) | s(bool,'T_0') = s(bool,sK1('T_0',X333,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(forward_demodulation,[],[f27469,f11496])).
fof(f27469,plain,(
  ( ! [X333,X335,X334,X336] : (s(bool,'T_0') = s(bool,X335) | s(bool,happ(s(fun(bool,bool),X333),s(bool,X334))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X336))))) | s(bool,'T_0') = s(bool,sK1('T_0',X333,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl404)),
  inference(subsumption_resolution,[],[f27377,f4721])).
fof(f27377,plain,(
  ( ! [X333,X335,X334,X336] : (s(bool,X335) != s(bool,sK1('T_0',X333,bool)) | s(bool,'T_0') = s(bool,X335) | s(bool,happ(s(fun(bool,bool),X333),s(bool,X334))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X336))))) | s(bool,'T_0') = s(bool,sK1('T_0',X333,bool))) ) | ($spl1 | $spl2 | $spl96 | $spl404)),
  inference(superposition,[],[f11751,f25153])).
fof(f27458,plain,(
  ( ! [X257,X259,X256,X258] : (s(bool,X259) = s(bool,happ(s(fun(bool,bool),X256),s(bool,X257))) | s(bool,'T_0') = s(bool,X259) | s(bool,'T_0') = s(bool,X258) | s(bool,'T_0') = s(bool,sK1('T_0',X256,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(subsumption_resolution,[],[f27355,f4721])).
fof(f27355,plain,(
  ( ! [X257,X259,X256,X258] : (s(bool,X258) != s(bool,sK1('T_0',X256,bool)) | s(bool,X259) = s(bool,happ(s(fun(bool,bool),X256),s(bool,X257))) | s(bool,'T_0') = s(bool,X259) | s(bool,'T_0') = s(bool,X258) | s(bool,'T_0') = s(bool,sK1('T_0',X256,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(superposition,[],[f6249,f25153])).
fof(f27457,plain,(
  ( ! [X251,X252,X250] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X250),s(bool,X251))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,sK1('T_0',X250,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl404)),
  inference(subsumption_resolution,[],[f27353,f4721])).
fof(f27353,plain,(
  ( ! [X251,X252,X250] : (s(bool,X252) != s(bool,sK1('T_0',X250,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X250),s(bool,X251))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,sK1('T_0',X250,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl404)),
  inference(superposition,[],[f6091,f25153])).
fof(f24482,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,X25)) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) | $spl394),
  inference(cnf_transformation,[],[f24482_D])).
fof(f24482_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,X25)) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26)))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f27450,plain,(
  ( ! [X204,X202,X205,X203] : (p(s(bool,X204)) | s(bool,X205) = s(bool,happ(s(fun(bool,bool),X202),s(bool,X203))) | p(s(bool,X205)) | s(bool,'T_0') = s(bool,sK1('T_0',X202,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl404)),
  inference(subsumption_resolution,[],[f27339,f4239])).
fof(f27339,plain,(
  ( ! [X204,X202,X205,X203] : (s(bool,X204) != s(bool,sK1('T_0',X202,bool)) | p(s(bool,X204)) | s(bool,X205) = s(bool,happ(s(fun(bool,bool),X202),s(bool,X203))) | p(s(bool,X205)) | s(bool,'T_0') = s(bool,sK1('T_0',X202,bool))) ) | ($spl1 | $spl2 | $spl15 | $spl404)),
  inference(superposition,[],[f4300,f25153])).
fof(f4300,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) ) | ($spl1 | $spl2 | $spl15)),
  inference(resolution,[],[f4124,f88])).
fof(f27113,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f27102,f4107])).
fof(f27102,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26477,f4721])).
fof(f26477,plain,(
  ( ! [X50,X51] : (p(s(bool,sK1('T_0',X50,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X50),s(bool,X51)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(resolution,[],[f24129,f4043])).
fof(f27112,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f27101,f4107])).
fof(f27101,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26477,f4721])).
fof(f24505,plain,(
  ( ! [X61,X60] : (p(s(bool,sK1('T_0',X60,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61)))) ) | $spl398),
  inference(cnf_transformation,[],[f24505_D])).
fof(f24505_D,plain,(
  ( ! [X61,X60] : (p(s(bool,sK1('T_0',X60,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61)))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f26937,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) | p(s(bool,X96)) | p(s(bool,sK1('T_0',X97,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(forward_demodulation,[],[f26936,f16150])).
fof(f26936,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | p(s(bool,sK1('T_0',X97,bool))) | s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X95))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(subsumption_resolution,[],[f26887,f4107])).
fof(f26887,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | p(s(bool,sK1('T_0',X97,bool))) | s(bool,happ(s(fun(bool,bool),X97),s(bool,X98))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X95))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X96)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(superposition,[],[f26463,f16151])).
fof(f26463,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK1('T_0',X0,bool))) | p(s(bool,X1)) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(resolution,[],[f24129,f2432])).
fof(f26933,plain,(
  ( ! [X88,X87,X86] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) | p(s(bool,X86)) | p(s(bool,sK1('T_0',X87,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(forward_demodulation,[],[f26932,f12328])).
fof(f26932,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | p(s(bool,sK1('T_0',X87,bool))) | s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(subsumption_resolution,[],[f26884,f4107])).
fof(f26884,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | p(s(bool,sK1('T_0',X87,bool))) | s(bool,happ(s(fun(bool,bool),X87),s(bool,X88))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X86)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(superposition,[],[f26463,f12329])).
fof(f26929,plain,(
  ( ! [X78,X76,X77] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) | p(s(bool,X76)) | p(s(bool,sK1('T_0',X77,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(forward_demodulation,[],[f26928,f11496])).
fof(f26928,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,sK1('T_0',X77,bool))) | s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26881,f4107])).
fof(f26881,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,sK1('T_0',X77,bool))) | s(bool,happ(s(fun(bool,bool),X77),s(bool,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75))))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f11475])).
fof(f26925,plain,(
  ( ! [X70,X71,X69] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) | p(s(bool,X69)) | p(s(bool,sK1('T_0',X70,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(forward_demodulation,[],[f26924,f15070])).
fof(f26924,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,sK1('T_0',X70,bool))) | s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26879,f4107])).
fof(f26879,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,sK1('T_0',X70,bool))) | s(bool,happ(s(fun(bool,bool),X70),s(bool,X71))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68))))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f15039])).
fof(f26923,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,sK1('T_0',X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26878,f4107])).
fof(f26878,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,sK1('T_0',X66,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f259])).
fof(f26922,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,sK1('T_0',X63,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,X64))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26877,f8638])).
fof(f26877,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X62)))) | p(s(bool,sK1('T_0',X63,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,X64))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X62))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f8674])).
fof(f24494,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,sK1('T_0',X25,bool))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) | $spl396),
  inference(cnf_transformation,[],[f24494_D])).
fof(f24494_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,sK1('T_0',X25,bool))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26)))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f26912,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1('T_0',X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26868,f4107])).
fof(f26868,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1('T_0',X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f4721])).
fof(f26911,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1('T_0',X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26867,f4107])).
fof(f26867,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1('T_0',X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f4721])).
fof(f26901,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26854,f4107])).
fof(f26854,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f4721])).
fof(f26900,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26853,f4107])).
fof(f26853,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f26463,f4721])).
fof(f26691,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26680,f4107])).
fof(f26680,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f26690,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f26679,f4107])).
fof(f26679,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f26643,plain,(
  ( ! [X13] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X13,'type.closSem.v'))))) ) | $spl438),
  inference(cnf_transformation,[],[f26643_D])).
fof(f26643_D,plain,(
  ( ! [X13] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X13,'type.closSem.v'))))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f26644,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),X15),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X15),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X15,'type.closSem.v'))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f26631,f4107])).
fof(f26631,plain,(
  ( ! [X15,X16] : (p(s(bool,X16)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),X15),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X15),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X15,'type.closSem.v')))) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f2475,f4721])).
fof(f2475,plain,(
  ( ! [X51] : (p(s(bool,happ(s(fun('type.closSem.v',bool),X51),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X51,'type.closSem.v'))))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),X51),s('type.closSem.v','const.closSem.Unit_0'))))) )),
  inference(superposition,[],[f95,f117])).
fof(f26642,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X13,'type.closSem.v'))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f26630,f4107])).
fof(f26630,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X13),s('type.closSem.v',sK1('const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')),X13,'type.closSem.v'))))) ) | $spl6),
  inference(superposition,[],[f2475,f4721])).
fof(f26462,plain,(
  ( ! [X332,X330] : (~p(s(bool,X332)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X330)) | s(bool,X332) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X330)))) ) | $spl436),
  inference(cnf_transformation,[],[f26462_D])).
fof(f26462_D,plain,(
  ( ! [X332,X330] : (~p(s(bool,X332)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X330)) | s(bool,X332) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X330)))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f26331,plain,(
  ( ! [X331,X332,X330] : (~p(s(bool,happ(s(fun(bool,bool),X331),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X332) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X330))) | p(s(bool,sK1('T_0',X331,bool))) | ~p(s(bool,X332)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X330))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f8674])).
fof(f26459,plain,(
  ( ! [X308,X310] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X308))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X310))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X308))) ) | $spl434),
  inference(cnf_transformation,[],[f26459_D])).
fof(f26459_D,plain,(
  ( ! [X308,X310] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X308))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X310))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X308))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f26458,plain,(
  ( ! [X308,X310,X309] : (s(bool,'T_0') = s(bool,X309) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X308))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X310))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X308))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl292)),
  inference(subsumption_resolution,[],[f26322,f259])).
fof(f26322,plain,(
  ( ! [X308,X310,X309] : (s(bool,X309) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X309) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X308))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X310))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X308))) ) | ($spl1 | $spl2 | $spl96 | $spl292)),
  inference(superposition,[],[f16156,f8674])).
fof(f26456,plain,(
  ( ! [X293,X295] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X293))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X295))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) | $spl432),
  inference(cnf_transformation,[],[f26456_D])).
fof(f26456_D,plain,(
  ( ! [X293,X295] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X293))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X295))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f26455,plain,(
  ( ! [X294,X293,X295] : (s(bool,'T_0') = s(bool,X294) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X293))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X295))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) | ($spl1 | $spl2 | $spl6 | $spl96)),
  inference(subsumption_resolution,[],[f26316,f259])).
fof(f26316,plain,(
  ( ! [X294,X293,X295] : (s(bool,X294) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X294) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X293))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X295))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X293))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f15349,f8674])).
fof(f26453,plain,(
  ( ! [X280,X278] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X278))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X280))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X278))) ) | $spl430),
  inference(cnf_transformation,[],[f26453_D])).
fof(f26453_D,plain,(
  ( ! [X280,X278] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X278))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X280))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X278))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f26452,plain,(
  ( ! [X280,X279,X278] : (s(bool,'T_0') = s(bool,X279) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X278))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X280))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X278))) ) | ($spl1 | $spl2 | $spl6 | $spl96 | $spl258)),
  inference(subsumption_resolution,[],[f26310,f259])).
fof(f26310,plain,(
  ( ! [X280,X279,X278] : (s(bool,X279) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X279) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X278))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X280))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X278))) ) | ($spl1 | $spl2 | $spl96 | $spl258)),
  inference(superposition,[],[f12334,f8674])).
fof(f26450,plain,(
  ( ! [X263,X265] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X263))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X265))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X263))) ) | $spl428),
  inference(cnf_transformation,[],[f26450_D])).
fof(f26450_D,plain,(
  ( ! [X263,X265] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X263))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X265))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X263))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f26449,plain,(
  ( ! [X263,X265,X264] : (s(bool,'T_0') = s(bool,X264) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X263))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X265))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X263))) ) | ($spl1 | $spl2 | $spl6 | $spl96)),
  inference(subsumption_resolution,[],[f26304,f259])).
fof(f26304,plain,(
  ( ! [X263,X265,X264] : (s(bool,X264) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X264) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X263))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X265))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X263))) ) | ($spl1 | $spl2 | $spl96)),
  inference(superposition,[],[f11751,f8674])).
fof(f10062,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X19))) | s(bool,'T_0') = s(bool,X17) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19))) ) | $spl204),
  inference(cnf_transformation,[],[f10062_D])).
fof(f10062_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X19))) | s(bool,'T_0') = s(bool,X17) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f26440,plain,(
  ( ! [X210,X208,X209] : (s(bool,X210) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X208))) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X209) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X208))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f26282,f259])).
fof(f26282,plain,(
  ( ! [X210,X208,X209] : (s(bool,X209) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X210) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X208))) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X209) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X208))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f8674])).
fof(f26431,plain,(
  ( ! [X175,X173,X174] : (s(bool,sK1(X174,X175,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X174)) | ~p(s(bool,happ(s(fun(bool,bool),X175),s(bool,X174)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X173))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f26267,f8638])).
fof(f26267,plain,(
  ( ! [X175,X173,X174] : (s(bool,sK1(X174,X175,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X173)))) | ~p(s(bool,X174)) | ~p(s(bool,happ(s(fun(bool,bool),X175),s(bool,X174)))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X173))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f8674])).
fof(f26421,plain,(
  ( ! [X121,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X121) = s(bool,sK1(X121,X120,bool)) | p(s(bool,X121)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X119))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26247,f8638])).
fof(f26247,plain,(
  ( ! [X121,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X121) = s(bool,sK1(X121,X120,bool)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X119)))) | p(s(bool,X121)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X119))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1730,f8674])).
fof(f26402,plain,(
  ( ! [X189] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X189))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X189))) ) | $spl424),
  inference(cnf_transformation,[],[f26402_D])).
fof(f26402_D,plain,(
  ( ! [X189] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X189))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X189))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f26156,plain,(
  ( ! [X302,X301] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X301))) | ~p(s(bool,happ(s(fun(bool,bool),X302),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1(sK1('T_0',X302,bool),X302,bool))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X301))) ) | ($spl1 | $spl2 | $spl344)),
  inference(superposition,[],[f21356,f8674])).
fof(f26152,plain,(
  ( ! [X292,X291] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X291))) | ~p(s(bool,happ(s(fun(bool,bool),X292),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | p(s(bool,sK1('T_0',X292,bool))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X291))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f8674])).
fof(f26410,plain,(
  ( ! [X290,X289] : (p(s(bool,sK1(sK1('T_0',X290,bool),X290,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X290),s(bool,'T_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X289))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f26151,f8638])).
fof(f26151,plain,(
  ( ! [X290,X289] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X289)))) | p(s(bool,sK1(sK1('T_0',X290,bool),X290,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X290),s(bool,'T_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X289))) ) | ($spl1 | $spl2 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f8674])).
fof(f26405,plain,(
  ( ! [X215,X213,X214] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X214),s('type.closSem.v',X214)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X215))),s('type.closSem.v',X214))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X213))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26131,f8638])).
fof(f26131,plain,(
  ( ! [X215,X213,X214] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X213)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X214),s('type.closSem.v',X214)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X215))),s('type.closSem.v',X214))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X213))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f9506,f8674])).
fof(f26404,plain,(
  ( ! [X212,X210,X211] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X211),s('type.closSem.v',X211)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X212))),s('type.closSem.v',X211))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X210))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26130,f8638])).
fof(f26130,plain,(
  ( ! [X212,X210,X211] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X210)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X211),s('type.closSem.v',X211)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X212))),s('type.closSem.v',X211))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X210))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f9393,f8674])).
fof(f10201,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | $spl210),
  inference(cnf_transformation,[],[f10201_D])).
fof(f10201_D,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f26403,plain,(
  ( ! [X204,X205] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X205),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X205)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X204))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26129,f8638])).
fof(f26129,plain,(
  ( ! [X204,X205] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X204)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X205),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X205)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X204))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f8770,f8674])).
fof(f8770,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X2),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X2))) )),
  inference(resolution,[],[f8750,f88])).
fof(f26124,plain,(
  ( ! [X191,X189,X190] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X189))) | p(s(bool,sK1(X190,X191,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | ~p(s(bool,X190)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X189))) ) | ($spl1 | $spl2 | $spl106)),
  inference(superposition,[],[f8009,f8674])).
fof(f26398,plain,(
  ( ! [X169,X170,X168] : (s(bool,X169) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X170) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X169) = s(bool,X170) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X168))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26116,f8638])).
fof(f26116,plain,(
  ( ! [X169,X170,X168] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X168)))) | s(bool,X169) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X170) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X169) = s(bool,X170) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X168))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f5987,f8674])).
fof(f26391,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X122),s('type.closSem.v',X122)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X122),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X123))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X121))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26095,f8638])).
fof(f26095,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X121)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X122),s('type.closSem.v',X122)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X122),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X123))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X121))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2527,f8674])).
fof(f26390,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,happ(s(fun(X118,bool),X119),s(X118,sK1(X120,X119,X118))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X118,bool),X119),s(X118,X120))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X117))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26094,f8638])).
fof(f26094,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X117)))) | p(s(bool,happ(s(fun(X118,bool),X119),s(X118,sK1(X120,X119,X118))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X118,bool),X119),s(X118,X120))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X117))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2432,f8674])).
fof(f26389,plain,(
  ( ! [X111,X112,X110] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v',X111)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X112))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X110))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26092,f8638])).
fof(f26092,plain,(
  ( ! [X111,X112,X110] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X110)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v',X111)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X112))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X110))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2382,f8674])).
fof(f26388,plain,(
  ( ! [X109,X107,X108] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X108))),s('type.closSem.v',X109))) | s('type.closSem.v',X109) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X108))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X107))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26091,f8638])).
fof(f26091,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X107)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X108))),s('type.closSem.v',X109))) | s('type.closSem.v',X109) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X108))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X107))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2217,f8674])).
fof(f26387,plain,(
  ( ! [X105,X106,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X105))),s('type.closSem.v',X106))) | s('type.closSem.v',X106) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X105))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X104))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26090,f8638])).
fof(f26090,plain,(
  ( ! [X105,X106,X104] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X104)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X105))),s('type.closSem.v',X106))) | s('type.closSem.v',X106) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X105))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X104))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f2079,f8674])).
fof(f26386,plain,(
  ( ! [X103,X101,X102] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X102),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))))) | s('type.closSem.v',X102) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X101))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26089,f8638])).
fof(f26089,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X101)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X102),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))))) | s('type.closSem.v',X102) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X101))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1911,f8674])).
fof(f26384,plain,(
  ( ! [X99,X100,X98] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X99),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))))) | s('type.closSem.v',X99) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X98))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f26088,f8638])).
fof(f26088,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X98)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X99),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))))) | s('type.closSem.v',X99) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X98))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1773,f8674])).
fof(f25998,plain,(
  ( ! [X92,X90,X91] : (s(bool,X92) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X90)) | s(bool,X91) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X91) = s(bool,X92)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f25997,f16150])).
fof(f25997,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | s(bool,X91) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X92) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X91) = s(bool,X92)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f25964,f4107])).
fof(f25964,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | s(bool,X91) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X92) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X91) = s(bool,X92) | s(bool,'T_0') = s(bool,X90)) ) | ($spl1 | $spl2 | $spl8 | $spl292)),
  inference(superposition,[],[f5987,f16151])).
fof(f25995,plain,(
  ( ! [X80,X81,X82] : (s(bool,X82) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X80)) | s(bool,X81) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X81) = s(bool,X82)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f25994,f12328])).
fof(f25994,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,X81) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X81) = s(bool,X82)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f25961,f4107])).
fof(f25961,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,X81) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X80)) ) | ($spl1 | $spl2 | $spl8 | $spl258)),
  inference(superposition,[],[f5987,f12329])).
fof(f25992,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X70)) | s(bool,X71) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X71) = s(bool,X72)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f25991,f15070])).
fof(f25991,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | s(bool,X71) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X72) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) | s(bool,X71) = s(bool,X72)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f25958,f4107])).
fof(f25958,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | s(bool,X71) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X72) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X70)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f5987,f15039])).
fof(f25989,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X63)) | s(bool,X64) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X64) = s(bool,X65)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f25988,f11496])).
fof(f25988,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,X64) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X65) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) | s(bool,X64) = s(bool,X65)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f25956,f4107])).
fof(f25956,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,X64) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X65) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) | s(bool,X64) = s(bool,X65) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f5987,f11475])).
fof(f25986,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | s(bool,X60) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X61) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X60) = s(bool,X61)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25955,f4107])).
fof(f25955,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | s(bool,X60) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X61) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f5987,f259])).
fof(f25976,plain,(
  ( ! [X26,X23,X25] : (s(bool,X25) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,X26) | s(bool,X23) = s(bool,X26)) ) | $spl416),
  inference(cnf_transformation,[],[f25976_D])).
fof(f25976_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X25) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,X26) | s(bool,X23) = s(bool,X26)) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f25977,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X29) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X27) = s(bool,X30) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25947,f4107])).
fof(f25947,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X29) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X27) = s(bool,X30) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f5987,f4721])).
fof(f25975,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X25) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X23) = s(bool,X26) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25946,f4107])).
fof(f25946,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X25) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X23) = s(bool,X26) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f5987,f4721])).
fof(f25808,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))) | p(s(bool,X114)) | p(s(bool,happ(s(fun(X115,bool),X116),s(X115,sK1(X117,X116,X115)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f25807,f16150])).
fof(f25807,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X114)) | p(s(bool,happ(s(fun(X115,bool),X116),s(X115,sK1(X117,X116,X115))))) | s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f25751,f4107])).
fof(f25751,plain,(
  ( ! [X116,X114,X117,X115,X113] : (p(s(bool,X114)) | p(s(bool,happ(s(fun(X115,bool),X116),s(X115,sK1(X117,X116,X115))))) | s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X114)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2432,f16151])).
fof(f25804,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X102,bool),X103),s(X102,X104))) | p(s(bool,X101)) | p(s(bool,happ(s(fun(X102,bool),X103),s(X102,sK1(X104,X103,X102)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f25803,f12328])).
fof(f25803,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,happ(s(fun(X102,bool),X103),s(X102,sK1(X104,X103,X102))))) | s(bool,happ(s(fun(X102,bool),X103),s(X102,X104))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f25748,f4107])).
fof(f25748,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,X101)) | p(s(bool,happ(s(fun(X102,bool),X103),s(X102,sK1(X104,X103,X102))))) | s(bool,happ(s(fun(X102,bool),X103),s(X102,X104))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X100))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X101)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2432,f12329])).
fof(f25800,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X89,bool),X90),s(X89,X91))) | p(s(bool,X88)) | p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK1(X91,X90,X89)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f25799,f15070])).
fof(f25799,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,X88)) | p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK1(X91,X90,X89))))) | s(bool,happ(s(fun(X89,bool),X90),s(X89,X91))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X87)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f25745,f4107])).
fof(f25745,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,X88)) | p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK1(X91,X90,X89))))) | s(bool,happ(s(fun(X89,bool),X90),s(X89,X91))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X87))))) | s(bool,'T_0') = s(bool,X88)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2432,f15039])).
fof(f25796,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X80,bool),X81),s(X80,X82))) | p(s(bool,X79)) | p(s(bool,happ(s(fun(X80,bool),X81),s(X80,sK1(X82,X81,X80)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f25795,f11496])).
fof(f25795,plain,(
  ( ! [X80,X78,X81,X79,X82] : (p(s(bool,X79)) | p(s(bool,happ(s(fun(X80,bool),X81),s(X80,sK1(X82,X81,X80))))) | s(bool,happ(s(fun(X80,bool),X81),s(X80,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f25743,f4107])).
fof(f25743,plain,(
  ( ! [X80,X78,X81,X79,X82] : (p(s(bool,X79)) | p(s(bool,happ(s(fun(X80,bool),X81),s(X80,sK1(X82,X81,X80))))) | s(bool,happ(s(fun(X80,bool),X81),s(X80,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78))))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2432,f11475])).
fof(f25793,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK1(X77,X76,X75))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,X77)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25742,f4107])).
fof(f25742,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | p(s(bool,happ(s(fun(X75,bool),X76),s(X75,sK1(X77,X76,X75))))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,X77))) | s(bool,'T_0') = s(bool,X74)) ) | $spl6),
  inference(superposition,[],[f2432,f259])).
fof(f25783,plain,(
  ( ! [X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK1(X33,X32,X31))))) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl412),
  inference(cnf_transformation,[],[f25783_D])).
fof(f25783_D,plain,(
  ( ! [X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK1(X33,X32,X31))))) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f25784,plain,(
  ( ! [X37,X35,X38,X36,X34] : (p(s(bool,X35)) | p(s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36))))) | s(bool,X34) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25734,f4107])).
fof(f25734,plain,(
  ( ! [X37,X35,X38,X36,X34] : (p(s(bool,X35)) | p(s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36))))) | s(bool,X34) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | $spl6),
  inference(superposition,[],[f2432,f4721])).
fof(f25782,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK1(X33,X32,X31))))) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25733,f4107])).
fof(f25733,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK1(X33,X32,X31))))) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | $spl6),
  inference(superposition,[],[f2432,f4721])).
fof(f25771,plain,(
  ( ! [X37,X38,X36,X40] : (p(s(bool,X40)) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36)))) | s(bool,X40) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) ) | $spl410),
  inference(cnf_transformation,[],[f25771_D])).
fof(f25771_D,plain,(
  ( ! [X37,X38,X36,X40] : (p(s(bool,X40)) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36)))) | s(bool,X40) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f25772,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) | s(bool,'T_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK1(X43,X42,X41))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25721,f4107])).
fof(f25721,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) | s(bool,'T_0') = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK1(X43,X42,X41)))) | s(bool,'T_0') = s(bool,X44)) ) | $spl6),
  inference(superposition,[],[f2432,f4721])).
fof(f25770,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f25720,f4107])).
fof(f25720,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X40) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK1(X38,X37,X36))))) ) | $spl6),
  inference(superposition,[],[f2432,f4721])).
fof(f25516,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl396)),
  inference(subsumption_resolution,[],[f25505,f4107])).
fof(f25505,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl396)),
  inference(superposition,[],[f24494,f4721])).
fof(f25515,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl396)),
  inference(subsumption_resolution,[],[f25504,f4107])).
fof(f25504,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl6 | $spl396)),
  inference(superposition,[],[f24494,f4721])).
fof(f25471,plain,(
  ( ! [X92,X90,X91] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) | p(s(bool,X90)) | s(bool,'T_0') = s(bool,sK1('T_0',X91,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292 | $spl394)),
  inference(forward_demodulation,[],[f25470,f16150])).
fof(f25470,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | s(bool,'T_0') = s(bool,sK1('T_0',X91,bool)) | s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292 | $spl394)),
  inference(subsumption_resolution,[],[f25435,f4107])).
fof(f25435,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | s(bool,'T_0') = s(bool,sK1('T_0',X91,bool)) | s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl8 | $spl292 | $spl394)),
  inference(superposition,[],[f24482,f16151])).
fof(f25467,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) | p(s(bool,X80)) | s(bool,'T_0') = s(bool,sK1('T_0',X81,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl394)),
  inference(forward_demodulation,[],[f25466,f12328])).
fof(f25466,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'T_0') = s(bool,sK1('T_0',X81,bool)) | s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl394)),
  inference(subsumption_resolution,[],[f25432,f4107])).
fof(f25432,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'T_0') = s(bool,sK1('T_0',X81,bool)) | s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl8 | $spl258 | $spl394)),
  inference(superposition,[],[f24482,f12329])).
fof(f25463,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) | p(s(bool,X70)) | s(bool,'T_0') = s(bool,sK1('T_0',X71,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(forward_demodulation,[],[f25462,f15070])).
fof(f25462,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | s(bool,'T_0') = s(bool,sK1('T_0',X71,bool)) | s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f25429,f4107])).
fof(f25429,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | s(bool,'T_0') = s(bool,sK1('T_0',X71,bool)) | s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl1 | $spl2 | $spl8 | $spl394)),
  inference(superposition,[],[f24482,f15039])).
fof(f25459,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) | p(s(bool,X63)) | s(bool,'T_0') = s(bool,sK1('T_0',X64,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(forward_demodulation,[],[f25458,f11496])).
fof(f25458,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,'T_0') = s(bool,sK1('T_0',X64,bool)) | s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f25427,f4107])).
fof(f25427,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,'T_0') = s(bool,sK1('T_0',X64,bool)) | s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl2 | $spl8 | $spl394)),
  inference(superposition,[],[f24482,f11475])).
fof(f25457,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | s(bool,'T_0') = s(bool,sK1('T_0',X60,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61)))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f25426,f4107])).
fof(f25426,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | s(bool,'T_0') = s(bool,sK1('T_0',X60,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f259])).
fof(f25448,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,sK1('T_0',X29,bool)) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f25418,f4107])).
fof(f25418,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,sK1('T_0',X29,bool)) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f4721])).
fof(f25446,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl394)),
  inference(subsumption_resolution,[],[f25417,f4107])).
fof(f25417,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,sK1('T_0',X25,bool)) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl394)),
  inference(superposition,[],[f24482,f4721])).
fof(f8887,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,happ(s(fun('type.closSem.v',bool),X19),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X19,'type.closSem.v'))) ) | $spl184),
  inference(cnf_transformation,[],[f8887_D])).
fof(f8887_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,happ(s(fun('type.closSem.v',bool),X19),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X19,'type.closSem.v'))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f25277,plain,(
  ( ! [X30,X31,X29] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X29,'type.closSem.v')) | ~p(s(bool,X30)) | s(bool,X31) = s(bool,happ(s(fun('type.closSem.v',bool),X29),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl6 | $spl8 | $spl402)),
  inference(subsumption_resolution,[],[f25252,f396])).
fof(f25252,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,X30) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X29,'type.closSem.v')) | ~p(s(bool,X30)) | s(bool,X31) = s(bool,happ(s(fun('type.closSem.v',bool),X29),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl6 | $spl402)),
  inference(superposition,[],[f25031,f4721])).
fof(f25276,plain,(
  ( ! [X21,X22,X20] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X20,'type.closSem.v')) | ~p(s(bool,X21)) | s(bool,X22) = s(bool,happ(s(fun('type.closSem.v',bool),X20),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl8 | $spl402)),
  inference(subsumption_resolution,[],[f25249,f396])).
fof(f25249,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X20,'type.closSem.v')) | ~p(s(bool,X21)) | s(bool,X22) = s(bool,happ(s(fun('type.closSem.v',bool),X20),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl402)),
  inference(superposition,[],[f25031,f4721])).
fof(f25202,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl398)),
  inference(subsumption_resolution,[],[f25190,f4107])).
fof(f25190,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl398)),
  inference(superposition,[],[f24505,f4721])).
fof(f25200,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl398)),
  inference(subsumption_resolution,[],[f25189,f4107])).
fof(f25189,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl6 | $spl398)),
  inference(superposition,[],[f24505,f4721])).
fof(f25154,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f25142,f4107])).
fof(f25142,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,sK1('T_0',X20,bool)) = s(bool,happ(s(fun(bool,bool),X20),s(bool,X22))) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f25152,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f25141,f4107])).
fof(f25141,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,sK1('T_0',X17,bool)) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X19))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24430,f4721])).
fof(f25031,plain,(
  ( ! [X252,X250] : (s(bool,X252) != s(bool,happ(s(fun('type.closSem.v',bool),X250),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X250,'type.closSem.v')) | ~p(s(bool,X252))) ) | $spl402),
  inference(cnf_transformation,[],[f25031_D])).
fof(f25031_D,plain,(
  ( ! [X252,X250] : (s(bool,X252) != s(bool,happ(s(fun('type.closSem.v',bool),X250),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X250,'type.closSem.v')) | ~p(s(bool,X252))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f24953,plain,(
  ( ! [X251,X252,X250] : (~p(s(bool,happ(s(fun(bool,bool),X251),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X252) != s(bool,happ(s(fun('type.closSem.v',bool),X250),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,sK1('T_0',X251,bool))) | ~p(s(bool,X252)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X250,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f5283])).
fof(f25002,plain,(
  ( ! [X130,X128,X129] : (s(bool,X130) = s(bool,happ(s(fun('type.closSem.v',bool),X128),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,X129) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X128,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f24904,f259])).
fof(f24904,plain,(
  ( ! [X130,X128,X129] : (s(bool,X129) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X130) = s(bool,happ(s(fun('type.closSem.v',bool),X128),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,X129) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X128,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f5283])).
fof(f24994,plain,(
  ( ! [X97,X98,X96] : (s(bool,sK1(X97,X98,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X97)) | ~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X96,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f24890,f4041])).
fof(f24890,plain,(
  ( ! [X97,X98,X96] : (s(bool,sK1(X97,X98,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,happ(s(fun('type.closSem.v',bool),X96),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,X97)) | ~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X96,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f5283])).
fof(f24984,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X48) = s(bool,sK1(X48,X47,bool)) | p(s(bool,X48)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X46,'type.closSem.v'))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f24871,f4041])).
fof(f24871,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X48) = s(bool,sK1(X48,X47,bool)) | p(s(bool,happ(s(fun('type.closSem.v',bool),X46),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,X48)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X46,'type.closSem.v'))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1730,f5283])).
fof(f24063,plain,(
  ( ! [X127,X128,X129] : (s(X127,X128) = s(X127,sK1(X128,X129,X127)) | s(bool,'T_0') != s(bool,happ(s(fun(X127,bool),X129),s(X127,X128)))) ) | $spl392),
  inference(cnf_transformation,[],[f24063_D])).
fof(f24063_D,plain,(
  ( ! [X127,X128,X129] : (s(X127,X128) = s(X127,sK1(X128,X129,X127)) | s(bool,'T_0') != s(bool,happ(s(fun(X127,bool),X129),s(X127,X128)))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f24006,plain,(
  ( ! [X127,X130,X128,X129] : (s(X127,X128) = s(X127,sK1(X128,X129,X127)) | s(bool,'T_0') != s(bool,happ(s(fun(X127,bool),X129),s(X127,X128))) | ~p(s(bool,X130))) ) | $spl388),
  inference(resolution,[],[f23980,f105])).
fof(f105,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ((p(s(bool,'T_0')) & p(s(bool,X9))) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.AND_CLAUSES|split|0')).
fof(f24007,plain,(
  ( ! [X134,X132,X133,X131] : (s(X131,X132) = s(X131,sK1(X132,X133,X131)) | s(bool,'T_0') != s(bool,happ(s(fun(X131,bool),X133),s(X131,X132))) | ~p(s(bool,X134))) ) | $spl388),
  inference(resolution,[],[f23980,f103])).
fof(f103,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9] : ((p(s(bool,X9)) & p(s(bool,'T_0'))) <=> p(s(bool,X9)))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.AND_CLAUSES|split|1')).
fof(f24519,plain,(
  ( ! [X92,X90,X91] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) | p(s(bool,X90)) | p(s(bool,sK1('T_0',X91,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(forward_demodulation,[],[f24518,f16150])).
fof(f24518,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,sK1('T_0',X91,bool))) | s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(subsumption_resolution,[],[f24469,f4107])).
fof(f24469,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,sK1('T_0',X91,bool))) | s(bool,happ(s(fun(bool,bool),X91),s(bool,X92))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl292 | $spl390)),
  inference(superposition,[],[f24208,f16151])).
fof(f24515,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) | p(s(bool,X80)) | p(s(bool,sK1('T_0',X81,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(forward_demodulation,[],[f24514,f12328])).
fof(f24514,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,sK1('T_0',X81,bool))) | s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(subsumption_resolution,[],[f24466,f4107])).
fof(f24466,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,sK1('T_0',X81,bool))) | s(bool,happ(s(fun(bool,bool),X81),s(bool,X82))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl258 | $spl390)),
  inference(superposition,[],[f24208,f12329])).
fof(f24511,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) | p(s(bool,X70)) | p(s(bool,sK1('T_0',X71,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(forward_demodulation,[],[f24510,f15070])).
fof(f24510,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,sK1('T_0',X71,bool))) | s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24463,f4107])).
fof(f24463,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,sK1('T_0',X71,bool))) | s(bool,happ(s(fun(bool,bool),X71),s(bool,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f15039])).
fof(f24507,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) | p(s(bool,X63)) | p(s(bool,sK1('T_0',X64,bool)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(forward_demodulation,[],[f24506,f11496])).
fof(f24506,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,sK1('T_0',X64,bool))) | s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62)))))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24461,f4107])).
fof(f24461,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,sK1('T_0',X64,bool))) | s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f11475])).
fof(f24504,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,sK1('T_0',X60,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61)))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24460,f4107])).
fof(f24460,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,sK1('T_0',X60,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f259])).
fof(f24495,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1('T_0',X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24452,f4107])).
fof(f24452,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,sK1('T_0',X29,bool))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f4721])).
fof(f24493,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1('T_0',X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24451,f4107])).
fof(f24451,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,sK1('T_0',X25,bool))) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f4721])).
fof(f24483,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24438,f4107])).
fof(f24438,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X30))) | s(bool,'T_0') = s(bool,sK1('T_0',X27,bool)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f4721])).
fof(f24481,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(subsumption_resolution,[],[f24437,f4107])).
fof(f24437,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,sK1('T_0',X23,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106 | $spl390)),
  inference(superposition,[],[f24208,f4721])).
fof(f24071,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,X69)) | s(X70,X71) = s(X70,sK1(X71,X72,X70)) | s(bool,'T_0') != s(bool,happ(s(fun(X70,bool),X72),s(X70,X71))) | s(bool,'T_0') = s(bool,X73)) ) | ($spl6 | $spl176 | $spl388)),
  inference(subsumption_resolution,[],[f24044,f8143])).
fof(f24044,plain,(
  ( ! [X70,X72,X71,X69,X73] : (~p(s(bool,X69)) | s(X70,X71) = s(X70,sK1(X71,X72,X70)) | s(bool,'T_0') != s(bool,happ(s(fun(X70,bool),X72),s(X70,X71))) | s(bool,X69) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | ($spl6 | $spl388)),
  inference(superposition,[],[f23980,f4721])).
fof(f24070,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,X64)) | s(X65,X66) = s(X65,sK1(X66,X67,X65)) | s(bool,'T_0') != s(bool,happ(s(fun(X65,bool),X67),s(X65,X66))) | s(bool,'T_0') = s(bool,X68)) ) | ($spl6 | $spl176 | $spl388)),
  inference(subsumption_resolution,[],[f24043,f8143])).
fof(f24043,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,X64)) | s(X65,X66) = s(X65,sK1(X66,X67,X65)) | s(bool,'T_0') != s(bool,happ(s(fun(X65,bool),X67),s(X65,X66))) | s(bool,X64) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) ) | ($spl6 | $spl388)),
  inference(superposition,[],[f23980,f4721])).
fof(f23980,plain,(
  ( ! [X516,X512,X514,X513] : (~p(s(bool,X516)) | s(X512,X514) = s(X512,sK1(X514,X513,X512)) | s(bool,X516) != s(bool,happ(s(fun(X512,bool),X513),s(X512,X514)))) ) | $spl388),
  inference(cnf_transformation,[],[f23980_D])).
fof(f23980_D,plain,(
  ( ! [X516,X512,X514,X513] : (~p(s(bool,X516)) | s(X512,X514) = s(X512,sK1(X514,X513,X512)) | s(bool,X516) != s(bool,happ(s(fun(X512,bool),X513),s(X512,X514)))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f23882,plain,(
  ( ! [X516,X512,X514,X513,X515] : (~p(s(bool,happ(s(fun(bool,bool),X515),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X516) != s(bool,happ(s(fun(X512,bool),X513),s(X512,X514))) | p(s(bool,sK1('T_0',X515,bool))) | ~p(s(bool,X516)) | s(X512,X514) = s(X512,sK1(X514,X513,X512))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f4403])).
fof(f8449,plain,(
  ( ! [X33,X31,X29,X32] : (s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X29) | s(X31,X33) = s(X31,sK1(X33,X32,X31))) ) | $spl180),
  inference(cnf_transformation,[],[f8449_D])).
fof(f8449_D,plain,(
  ( ! [X33,X31,X29,X32] : (s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X29) | s(X31,X33) = s(X31,sK1(X33,X32,X31))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f23954,plain,(
  ( ! [X292,X294,X296,X293,X295] : (s(bool,X296) = s(bool,happ(s(fun(X292,bool),X293),s(X292,X294))) | s(bool,'T_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X295) | s(X292,X294) = s(X292,sK1(X294,X293,X292))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f23833,f259])).
fof(f23833,plain,(
  ( ! [X292,X294,X296,X293,X295] : (s(bool,X295) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X296) = s(bool,happ(s(fun(X292,bool),X293),s(X292,X294))) | s(bool,'T_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X295) | s(X292,X294) = s(X292,sK1(X294,X293,X292))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f4403])).
fof(f23946,plain,(
  ( ! [X235,X233,X236,X234,X232] : (s(bool,sK1(X235,X236,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X235)) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X235)))) | s(X232,X234) = s(X232,sK1(X234,X233,X232))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f23819,f94])).
fof(f23819,plain,(
  ( ! [X235,X233,X236,X234,X232] : (s(bool,sK1(X235,X236,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,happ(s(fun(X232,bool),X233),s(X232,X234)))) | ~p(s(bool,X235)) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,X235)))) | s(X232,X234) = s(X232,sK1(X234,X233,X232))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f4403])).
fof(f23936,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,happ(s(fun(bool,bool),X153),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X154) = s(bool,sK1(X154,X153,bool)) | p(s(bool,X154)) | s(X150,X152) = s(X150,sK1(X152,X151,X150))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f23801,f94])).
fof(f23801,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,happ(s(fun(bool,bool),X153),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X154) = s(bool,sK1(X154,X153,bool)) | p(s(bool,happ(s(fun(X150,bool),X151),s(X150,X152)))) | p(s(bool,X154)) | s(X150,X152) = s(X150,sK1(X152,X151,X150))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1730,f4403])).
fof(f8994,plain,(
  ( ! [X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | s(bool,'T_0') = s(bool,X38) | s(bool,X35) = s(bool,X38) | p(s(bool,sK1('T_0',X36,bool)))) ) | $spl190),
  inference(cnf_transformation,[],[f8994_D])).
fof(f8994_D,plain,(
  ( ! [X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | s(bool,'T_0') = s(bool,X38) | s(bool,X35) = s(bool,X38) | p(s(bool,sK1('T_0',X36,bool)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f23508,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | p(s(bool,sK1('T_0',X56,bool))) | ~p(s(bool,X57)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f23469,f396])).
fof(f23469,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | s(bool,'T_0') != s(bool,X57) | p(s(bool,sK1('T_0',X56,bool))) | ~p(s(bool,X57)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f4721])).
fof(f23507,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X51)))) | p(s(bool,sK1('T_0',X52,bool))) | ~p(s(bool,X53)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f23468,f396])).
fof(f23468,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X51)))) | s(bool,'T_0') != s(bool,X53) | p(s(bool,sK1('T_0',X52,bool))) | ~p(s(bool,X53)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21098,f4721])).
fof(f670,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl22),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f22701,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | p(s(bool,sK1(sK1(X21,X20,bool),X20,bool)))) ) | $spl384),
  inference(cnf_transformation,[],[f22701_D])).
fof(f22701_D,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | p(s(bool,sK1(sK1(X21,X20,bool),X20,bool)))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f22702,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,sK1(sK1(X33,X32,bool),X32,bool))) | ~p(s(bool,X33)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f22660,f4239])).
fof(f22660,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,sK1(sK1(X33,X32,bool),X32,bool))) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f19729,f4721])).
fof(f19729,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | p(s(bool,sK1(sK1(X9,X11,bool),X11,bool))) | ~p(s(bool,X9)) | ~p(s(bool,X10))) ) | ($spl1 | $spl2 | $spl15)),
  inference(duplicate_literal_removal,[],[f19688])).
fof(f19688,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | p(s(bool,sK1(sK1(X9,X11,bool),X11,bool))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f19652,f91])).
fof(f19652,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | p(s(bool,sK1(sK1(X38,X39,bool),X39,bool))) | ~p(s(bool,X38))) ) | ($spl1 | $spl2 | $spl15)),
  inference(duplicate_literal_removal,[],[f19651])).
fof(f19651,plain,(
  ( ! [X39,X38] : (~p(s(bool,X38)) | p(s(bool,sK1(sK1(X38,X39,bool),X39,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38))))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f19531,f94])).
fof(f22700,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | p(s(bool,sK1(sK1(X21,X20,bool),X20,bool))) | ~p(s(bool,X21)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f22657,f4239])).
fof(f22657,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | p(s(bool,sK1(sK1(X21,X20,bool),X20,bool))) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f19729,f4721])).
fof(f20845,plain,(
  ( ! [X11] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | $spl336),
  inference(cnf_transformation,[],[f20845_D])).
fof(f20845_D,plain,(
  ( ! [X11] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f22624,plain,(
  ( ! [X101,X100] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106 | $spl292)),
  inference(forward_demodulation,[],[f22623,f16150])).
fof(f22623,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X102))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl6 | $spl8 | $spl96 | $spl106 | $spl292)),
  inference(subsumption_resolution,[],[f22577,f4721])).
fof(f22577,plain,(
  ( ! [X101,X102,X100] : (s(bool,X101) != s(bool,sK1('T_0',X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X102))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl8 | $spl96 | $spl106 | $spl292)),
  inference(superposition,[],[f16156,f20778])).
fof(f22595,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl382),
  inference(cnf_transformation,[],[f22595_D])).
fof(f22595_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f22581,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl8 | $spl96 | $spl292)),
  inference(duplicate_literal_removal,[],[f22532])).
fof(f22532,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl8 | $spl96 | $spl292)),
  inference(superposition,[],[f16156,f16151])).
fof(f22594,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl380),
  inference(cnf_transformation,[],[f22594_D])).
fof(f22594_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f22582,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(duplicate_literal_removal,[],[f22531])).
fof(f22531,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(superposition,[],[f16156,f15039])).
fof(f22593,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl378),
  inference(cnf_transformation,[],[f22593_D])).
fof(f22593_D,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f22583,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl8 | $spl96 | $spl258 | $spl292)),
  inference(duplicate_literal_removal,[],[f22530])).
fof(f22530,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl8 | $spl96 | $spl258 | $spl292)),
  inference(superposition,[],[f16156,f12329])).
fof(f22592,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl376),
  inference(cnf_transformation,[],[f22592_D])).
fof(f22592_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f22584,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(duplicate_literal_removal,[],[f22529])).
fof(f22529,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(superposition,[],[f16156,f11475])).
fof(f22456,plain,(
  ( ! [X101,X100] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106 | $spl258)),
  inference(forward_demodulation,[],[f22455,f12328])).
fof(f22455,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl6 | $spl8 | $spl96 | $spl106 | $spl258)),
  inference(subsumption_resolution,[],[f22409,f4721])).
fof(f22409,plain,(
  ( ! [X101,X102,X100] : (s(bool,X101) != s(bool,sK1('T_0',X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl8 | $spl96 | $spl106 | $spl258)),
  inference(superposition,[],[f12334,f20778])).
fof(f22427,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl374),
  inference(cnf_transformation,[],[f22427_D])).
fof(f22427_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f22413,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl8 | $spl96 | $spl258 | $spl292)),
  inference(duplicate_literal_removal,[],[f22364])).
fof(f22364,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl8 | $spl96 | $spl258 | $spl292)),
  inference(superposition,[],[f12334,f16151])).
fof(f22426,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl372),
  inference(cnf_transformation,[],[f22426_D])).
fof(f22426_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f22414,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(duplicate_literal_removal,[],[f22363])).
fof(f22363,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(superposition,[],[f12334,f15039])).
fof(f22425,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl370),
  inference(cnf_transformation,[],[f22425_D])).
fof(f22425_D,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f22415,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl8 | $spl96 | $spl258)),
  inference(duplicate_literal_removal,[],[f22362])).
fof(f22362,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl8 | $spl96 | $spl258)),
  inference(superposition,[],[f12334,f12329])).
fof(f22424,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl368),
  inference(cnf_transformation,[],[f22424_D])).
fof(f22424_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f22416,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(duplicate_literal_removal,[],[f22361])).
fof(f22361,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(superposition,[],[f12334,f11475])).
fof(f22090,plain,(
  ( ! [X101,X100] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106)),
  inference(forward_demodulation,[],[f22089,f15070])).
fof(f22089,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X102))))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106)),
  inference(subsumption_resolution,[],[f22043,f4721])).
fof(f22043,plain,(
  ( ! [X101,X102,X100] : (s(bool,X101) != s(bool,sK1('T_0',X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X102))))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl106)),
  inference(superposition,[],[f15349,f20778])).
fof(f22061,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50)))))) ) | $spl366),
  inference(cnf_transformation,[],[f22061_D])).
fof(f22061_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50)))))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f22047,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(duplicate_literal_removal,[],[f21998])).
fof(f21998,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(superposition,[],[f15349,f16151])).
fof(f22060,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))))) ) | $spl364),
  inference(cnf_transformation,[],[f22060_D])).
fof(f22060_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f22048,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(duplicate_literal_removal,[],[f21997])).
fof(f21997,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(superposition,[],[f15349,f15039])).
fof(f22059,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42)))))) ) | $spl362),
  inference(cnf_transformation,[],[f22059_D])).
fof(f22059_D,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42)))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f22049,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(duplicate_literal_removal,[],[f21996])).
fof(f21996,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(superposition,[],[f15349,f12329])).
fof(f22058,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38)))))) ) | $spl360),
  inference(cnf_transformation,[],[f22058_D])).
fof(f22058_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38)))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f22050,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(duplicate_literal_removal,[],[f21995])).
fof(f21995,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X38))))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(superposition,[],[f15349,f11475])).
fof(f21984,plain,(
  ( ! [X101,X100] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106)),
  inference(forward_demodulation,[],[f21983,f11496])).
fof(f21983,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl96 | $spl106)),
  inference(subsumption_resolution,[],[f21937,f4721])).
fof(f21937,plain,(
  ( ! [X101,X102,X100] : (s(bool,X101) != s(bool,sK1('T_0',X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,sK1('T_0',X100,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl106)),
  inference(superposition,[],[f11751,f20778])).
fof(f21955,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50)))))) ) | $spl358),
  inference(cnf_transformation,[],[f21955_D])).
fof(f21955_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50)))))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f21941,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(duplicate_literal_removal,[],[f21892])).
fof(f21892,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X49) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X49) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl292)),
  inference(superposition,[],[f11751,f16151])).
fof(f21954,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46)))))) ) | $spl356),
  inference(cnf_transformation,[],[f21954_D])).
fof(f21954_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46)))))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f21942,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(duplicate_literal_removal,[],[f21891])).
fof(f21891,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,X45) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(superposition,[],[f11751,f15039])).
fof(f21953,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42)))))) ) | $spl354),
  inference(cnf_transformation,[],[f21953_D])).
fof(f21953_D,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42)))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f21943,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(duplicate_literal_removal,[],[f21890])).
fof(f21890,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X41) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl8 | $spl96 | $spl258)),
  inference(superposition,[],[f11751,f12329])).
fof(f21952,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38)))))) ) | $spl352),
  inference(cnf_transformation,[],[f21952_D])).
fof(f21952_D,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38)))))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f21944,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38)))))) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(duplicate_literal_removal,[],[f21889])).
fof(f21889,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X37) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X38))))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8 | $spl96)),
  inference(superposition,[],[f11751,f11475])).
fof(f21710,plain,(
  ( ! [X19,X17] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X17,bool),X17,bool)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) ) | $spl350),
  inference(cnf_transformation,[],[f21710_D])).
fof(f21710_D,plain,(
  ( ! [X19,X17] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X17,bool),X17,bool)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f21711,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X20,bool),X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21665,f4107])).
fof(f21665,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X20,bool),X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f4721])).
fof(f21709,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X17,bool),X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21664,f4107])).
fof(f21664,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1(sK1('T_0',X17,bool),X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f4721])).
fof(f21703,plain,(
  ( ! [X66,X67] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) | p(s(bool,X66)) | p(s(bool,sK1(sK1('T_0',X67,bool),X67,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl292)),
  inference(forward_demodulation,[],[f21702,f16150])).
fof(f21702,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | p(s(bool,sK1(sK1('T_0',X67,bool),X67,bool))) | s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl292)),
  inference(subsumption_resolution,[],[f21652,f4107])).
fof(f21652,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | p(s(bool,sK1(sK1('T_0',X67,bool),X67,bool))) | s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X66)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl292)),
  inference(superposition,[],[f20737,f16151])).
fof(f21699,plain,(
  ( ! [X59,X60] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) | p(s(bool,X59)) | p(s(bool,sK1(sK1('T_0',X60,bool),X60,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl258)),
  inference(forward_demodulation,[],[f21698,f12328])).
fof(f21698,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,sK1(sK1('T_0',X60,bool),X60,bool))) | s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X58))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl258)),
  inference(subsumption_resolution,[],[f21649,f4107])).
fof(f21649,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,sK1(sK1('T_0',X60,bool),X60,bool))) | s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X58))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl258)),
  inference(superposition,[],[f20737,f12329])).
fof(f21695,plain,(
  ( ! [X52,X53] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) | p(s(bool,X52)) | p(s(bool,sK1(sK1('T_0',X53,bool),X53,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(forward_demodulation,[],[f21694,f15070])).
fof(f21694,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,sK1(sK1('T_0',X53,bool),X53,bool))) | s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X51)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21646,f4107])).
fof(f21646,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,sK1(sK1('T_0',X53,bool),X53,bool))) | s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X51))))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f15039])).
fof(f21691,plain,(
  ( ! [X47,X48] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) | p(s(bool,X47)) | p(s(bool,sK1(sK1('T_0',X48,bool),X48,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(forward_demodulation,[],[f21690,f11496])).
fof(f21690,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X47)) | p(s(bool,sK1(sK1('T_0',X48,bool),X48,bool))) | s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21644,f4107])).
fof(f21644,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X47)) | p(s(bool,sK1(sK1('T_0',X48,bool),X48,bool))) | s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f11475])).
fof(f21688,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,sK1(sK1('T_0',X45,bool),X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21643,f4107])).
fof(f21643,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,sK1(sK1('T_0',X45,bool),X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f259])).
fof(f21678,plain,(
  ( ! [X19,X17] : (p(s(bool,sK1(sK1('T_0',X19,bool),X19,bool))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) ) | $spl346),
  inference(cnf_transformation,[],[f21678_D])).
fof(f21678_D,plain,(
  ( ! [X19,X17] : (p(s(bool,sK1(sK1('T_0',X19,bool),X19,bool))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f21679,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,sK1(sK1('T_0',X22,bool),X22,bool))) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X22),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21635,f4107])).
fof(f21635,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,sK1(sK1('T_0',X22,bool),X22,bool))) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X22),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f4721])).
fof(f21677,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,sK1(sK1('T_0',X19,bool),X19,bool))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21634,f4107])).
fof(f21634,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,sK1(sK1('T_0',X19,bool),X19,bool))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(superposition,[],[f20737,f4721])).
fof(f21027,plain,(
  ( ! [X145,X143] : (s(bool,X145) = s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X143,bool)) | s(bool,'T_0') = s(bool,X145)) ) | $spl338),
  inference(cnf_transformation,[],[f21027_D])).
fof(f21027_D,plain,(
  ( ! [X145,X143] : (s(bool,X145) = s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X143,bool)) | s(bool,'T_0') = s(bool,X145)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f21346,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21331,f4107])).
fof(f21331,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,sK1('T_0',X20,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl8 | $spl106)),
  inference(superposition,[],[f20590,f4721])).
fof(f20590,plain,(
  ( ! [X10,X11] : (p(s(bool,sK1('T_0',X10,bool))) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f20508,f396])).
fof(f21345,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21330,f4107])).
fof(f21330,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool))) ) | ($spl6 | $spl8 | $spl106)),
  inference(superposition,[],[f20590,f4721])).
fof(f21342,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) | p(s(bool,sK1(sK1('T_0',X6,bool),X6,bool)))) ) | $spl344),
  inference(cnf_transformation,[],[f21342_D])).
fof(f21342_D,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) | p(s(bool,sK1(sK1('T_0',X6,bool),X6,bool)))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f21341,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,X5) | p(s(bool,sK1(sK1('T_0',X6,bool),X6,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl176)),
  inference(subsumption_resolution,[],[f21311,f8143])).
fof(f21311,plain,(
  ( ! [X6,X5] : (s(bool,X5) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | p(s(bool,sK1(sK1('T_0',X6,bool),X6,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(resolution,[],[f20590,f19531])).
fof(f21103,plain,(
  ( ! [X39] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X39,bool)))) ) | $spl342),
  inference(cnf_transformation,[],[f21103_D])).
fof(f21103_D,plain,(
  ( ! [X39] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X39,bool)))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f21092,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X42,bool))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(trivial_inequality_removal,[],[f21072])).
fof(f21072,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) != s(bool,X41) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X42,bool))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f4721])).
fof(f21093,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X39,bool))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(trivial_inequality_removal,[],[f21071])).
fof(f21071,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,X38) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,sK1('T_0',X39,bool))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f4721])).
fof(f21101,plain,(
  ( ! [X6,X8] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X6)))) | ~p(s(bool,X6)) | p(s(bool,sK1('T_0',X8,bool)))) ) | $spl340),
  inference(cnf_transformation,[],[f21101_D])).
fof(f21101_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X6)))) | ~p(s(bool,X6)) | p(s(bool,sK1('T_0',X8,bool)))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f21102,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X9)))) | p(s(bool,sK1('T_0',X11,bool))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21061,f396])).
fof(f21061,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X9)))) | p(s(bool,sK1('T_0',X11,bool))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f91])).
fof(f21100,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X6)))) | p(s(bool,sK1('T_0',X8,bool))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f21060,f396])).
fof(f21060,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X6)))) | p(s(bool,sK1('T_0',X8,bool))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f21004,f91])).
fof(f21026,plain,(
  ( ! [X144,X145,X143] : (s(bool,X145) = s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,sK1('T_0',X143,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20945,f4721])).
fof(f20945,plain,(
  ( ! [X144,X145,X143] : (s(bool,X144) != s(bool,sK1('T_0',X143,bool)) | s(bool,X145) = s(bool,happ(s(fun(bool,bool),X143),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,sK1('T_0',X143,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f6249,f20778])).
fof(f21025,plain,(
  ( ! [X140,X139] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X139),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,sK1('T_0',X139,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20943,f4721])).
fof(f20943,plain,(
  ( ! [X140,X139] : (s(bool,X140) != s(bool,sK1('T_0',X139,bool)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X139),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,sK1('T_0',X139,bool))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(superposition,[],[f6091,f20778])).
fof(f20846,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X13,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20834,f4107])).
fof(f20834,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X13,bool)) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2 | $spl6 | $spl106)),
  inference(superposition,[],[f20629,f4721])).
fof(f20629,plain,(
  ( ! [X127] : (p(s(bool,sK1('T_0',X127,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X127),s(bool,'T_0')))) ) | ($spl1 | $spl2 | $spl106)),
  inference(resolution,[],[f20508,f175])).
fof(f20844,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20833,f4107])).
fof(f20833,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl106)),
  inference(superposition,[],[f20629,f4721])).
fof(f20803,plain,(
  ( ! [X11] : (s(bool,sK1('T_0',X11,bool)) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | $spl334),
  inference(cnf_transformation,[],[f20803_D])).
fof(f20803_D,plain,(
  ( ! [X11] : (s(bool,sK1('T_0',X11,bool)) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f20804,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,sK1('T_0',X13,bool)) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X13,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20792,f4107])).
fof(f20792,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,sK1('T_0',X13,bool)) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X13,bool)) | s(bool,'T_0') = s(bool,X14)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f20688,f4721])).
fof(f20802,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,sK1('T_0',X11,bool)) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20791,f4107])).
fof(f20791,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,sK1('T_0',X11,bool)) = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,sK1('T_0',X11,bool))) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f20688,f4721])).
fof(f20753,plain,(
  ( ! [X45] : (p(s(bool,sK1('T_0',X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) | $spl332),
  inference(cnf_transformation,[],[f20753_D])).
fof(f20753_D,plain,(
  ( ! [X45] : (p(s(bool,sK1('T_0',X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f20767,plain,(
  ( ! [X66,X67] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) | p(s(bool,X66)) | p(s(bool,sK1('T_0',X67,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl292)),
  inference(forward_demodulation,[],[f20766,f16150])).
fof(f20766,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | p(s(bool,sK1('T_0',X67,bool))) | s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl292)),
  inference(subsumption_resolution,[],[f20727,f4107])).
fof(f20727,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | p(s(bool,sK1('T_0',X67,bool))) | s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X66)) ) | ($spl8 | $spl106 | $spl292)),
  inference(superposition,[],[f20508,f16151])).
fof(f20763,plain,(
  ( ! [X59,X60] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) | p(s(bool,X59)) | p(s(bool,sK1('T_0',X60,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl258)),
  inference(forward_demodulation,[],[f20762,f12328])).
fof(f20762,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,sK1('T_0',X60,bool))) | s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X58))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106 | $spl258)),
  inference(subsumption_resolution,[],[f20724,f4107])).
fof(f20724,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,sK1('T_0',X60,bool))) | s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X58))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl8 | $spl106 | $spl258)),
  inference(superposition,[],[f20508,f12329])).
fof(f20759,plain,(
  ( ! [X52,X53] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) | p(s(bool,X52)) | p(s(bool,sK1('T_0',X53,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(forward_demodulation,[],[f20758,f15070])).
fof(f20758,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,sK1('T_0',X53,bool))) | s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X51)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20721,f4107])).
fof(f20721,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,sK1('T_0',X53,bool))) | s(bool,happ(s(fun(bool,bool),X53),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X51))))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl1 | $spl2 | $spl8 | $spl106)),
  inference(superposition,[],[f20508,f15039])).
fof(f20755,plain,(
  ( ! [X47,X48] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) | p(s(bool,X47)) | p(s(bool,sK1('T_0',X48,bool)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(forward_demodulation,[],[f20754,f11496])).
fof(f20754,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X47)) | p(s(bool,sK1('T_0',X48,bool))) | s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20719,f4107])).
fof(f20719,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X47)) | p(s(bool,sK1('T_0',X48,bool))) | s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl106)),
  inference(superposition,[],[f20508,f11475])).
fof(f20752,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,sK1('T_0',X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20718,f4107])).
fof(f20718,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | p(s(bool,sK1('T_0',X45,bool))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f20508,f259])).
fof(f20742,plain,(
  ( ! [X19,X17] : (p(s(bool,sK1('T_0',X19,bool))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) ) | $spl330),
  inference(cnf_transformation,[],[f20742_D])).
fof(f20742_D,plain,(
  ( ! [X19,X17] : (p(s(bool,sK1('T_0',X19,bool))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f20743,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,sK1('T_0',X22,bool))) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X22),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20710,f4107])).
fof(f20710,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,sK1('T_0',X22,bool))) | s(bool,X20) = s(bool,happ(s(fun(bool,bool),X22),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f20508,f4721])).
fof(f20741,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,sK1('T_0',X19,bool))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl106)),
  inference(subsumption_resolution,[],[f20709,f4107])).
fof(f20709,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,sK1('T_0',X19,bool))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f20508,f4721])).
fof(f20560,plain,(
  ( ! [X95,X96] : (p(s(bool,sK1('T_0',X95,bool))) | p(s(bool,X96)) | s(bool,X96) = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0')))) ) | $spl328),
  inference(cnf_transformation,[],[f20560_D])).
fof(f20560_D,plain,(
  ( ! [X95,X96] : (p(s(bool,sK1('T_0',X95,bool))) | p(s(bool,X96)) | s(bool,X96) = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0')))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f20568,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,sK1('T_0',X56,bool))) | s(bool,X57) = s(bool,happ(s(fun(bool,bool),X56),s(bool,'T_0'))) | p(s(bool,X57)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl8 | $spl106)),
  inference(subsumption_resolution,[],[f20543,f4239])).
fof(f20543,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,sK1('T_0',X56,bool))) | s(bool,X57) = s(bool,happ(s(fun(bool,bool),X56),s(bool,'T_0'))) | p(s(bool,X57)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f8010,f4721])).
fof(f20567,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,sK1('T_0',X52,bool))) | s(bool,X53) = s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0'))) | p(s(bool,X53)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl6 | $spl8 | $spl106)),
  inference(subsumption_resolution,[],[f20542,f4239])).
fof(f20542,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | p(s(bool,sK1('T_0',X52,bool))) | s(bool,X53) = s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0'))) | p(s(bool,X53)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl6 | $spl106)),
  inference(superposition,[],[f8010,f4721])).
fof(f20511,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,sK1('T_0',X98,bool))) | s(bool,X99) = s(bool,happ(s(fun(bool,bool),X98),s(bool,'T_0'))) | p(s(bool,X99)) | ~p(s(bool,X100))) ) | $spl106),
  inference(resolution,[],[f8010,f103])).
fof(f20510,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,sK1('T_0',X95,bool))) | s(bool,X96) = s(bool,happ(s(fun(bool,bool),X95),s(bool,'T_0'))) | p(s(bool,X96)) | ~p(s(bool,X97))) ) | $spl106),
  inference(resolution,[],[f8010,f105])).
fof(f20184,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,sK1(X21,X20,bool))) ) | $spl326),
  inference(cnf_transformation,[],[f20184_D])).
fof(f20184_D,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,sK1(X21,X20,bool))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f20185,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK1(X33,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl8 | $spl176)),
  inference(subsumption_resolution,[],[f20145,f8143])).
fof(f20145,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK1(X33,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f7010,f4721])).
fof(f7010,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | s(bool,'T_0') = s(bool,sK1(X9,X11,bool)) | ~p(s(bool,X9)) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f6985])).
fof(f6985,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | s(bool,'T_0') = s(bool,sK1(X9,X11,bool)) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl8),
  inference(superposition,[],[f4284,f91])).
fof(f4284,plain,(
  ( ! [X28,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | s(bool,'T_0') = s(bool,sK1(X28,X29,bool)) | ~p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f396,f94])).
fof(f20183,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK1(X21,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl8 | $spl176)),
  inference(subsumption_resolution,[],[f20142,f8143])).
fof(f20142,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK1(X21,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f7010,f4721])).
fof(f19525,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,sK1(X0,X1,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) ) | $spl322),
  inference(cnf_transformation,[],[f19525_D])).
fof(f19525_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,sK1(X0,X1,bool))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f19526,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3))) ) | $spl324),
  inference(cnf_transformation,[],[f19526_D])).
fof(f19526_D,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f19527,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | p(s(bool,sK1(X4,X5,bool)))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19474,f4124])).
fof(f19474,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | p(s(bool,sK1(X4,X5,bool))) | p(s(bool,X6))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f88])).
fof(f19524,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,sK1(X0,X1,bool)))) ) | ($spl1 | $spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19473,f4124])).
fof(f19473,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,X2)) | p(s(bool,sK1(X0,X1,bool)))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4312,f88])).
fof(f19503,plain,(
  ( ! [X47,X50,X49] : (s(bool,sK1(X49,X50,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X47)) | ~p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X49))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f19502,f16150])).
fof(f19502,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,sK1(X49,X50,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X47)) | ~p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X49))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f19458,f4107])).
fof(f19458,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,sK1(X49,X50,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X47)) | ~p(s(bool,X49)) | ~p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X49)))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(superposition,[],[f4312,f16151])).
fof(f19501,plain,(
  ( ! [X45,X43,X46] : (s(bool,sK1(X45,X46,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X43)) | ~p(s(bool,X45)) | ~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f19500,f15070])).
fof(f19500,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,sK1(X45,X46,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | p(s(bool,X43)) | ~p(s(bool,X45)) | ~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f19457,f4107])).
fof(f19457,plain,(
  ( ! [X45,X43,X46,X44] : (s(bool,sK1(X45,X46,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))))) | p(s(bool,X43)) | ~p(s(bool,X45)) | ~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X45)))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f4312,f15039])).
fof(f19499,plain,(
  ( ! [X39,X41,X42] : (s(bool,sK1(X41,X42,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X39)) | ~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X41))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f19498,f12328])).
fof(f19498,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,sK1(X41,X42,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X39)) | ~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X41))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f19456,f4107])).
fof(f19456,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,sK1(X41,X42,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X39)) | ~p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X41)))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(superposition,[],[f4312,f12329])).
fof(f19497,plain,(
  ( ! [X37,X35,X38] : (s(bool,sK1(X37,X38,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X35)) | ~p(s(bool,X37)) | ~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f19496,f11496])).
fof(f19496,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,sK1(X37,X38,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | p(s(bool,X35)) | ~p(s(bool,X37)) | ~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f19455,f4107])).
fof(f19455,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,sK1(X37,X38,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))))) | p(s(bool,X35)) | ~p(s(bool,X37)) | ~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X37)))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f4312,f11475])).
fof(f19494,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,sK1(X25,X26,bool)) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | ~p(s(bool,X25))) ) | $spl320),
  inference(cnf_transformation,[],[f19494_D])).
fof(f19494_D,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,sK1(X25,X26,bool)) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | ~p(s(bool,X25))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f19495,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,sK1(X29,X30,bool)) | p(s(bool,X27)) | ~p(s(bool,X29)) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,X29)))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19453,f4107])).
fof(f19453,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,sK1(X29,X30,bool)) | p(s(bool,X27)) | ~p(s(bool,X29)) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,X29)))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4312,f4721])).
fof(f19493,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,sK1(X25,X26,bool)) | p(s(bool,X23)) | ~p(s(bool,X25)) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19452,f4107])).
fof(f19452,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,sK1(X25,X26,bool)) | p(s(bool,X23)) | ~p(s(bool,X25)) | ~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X25)))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4312,f4721])).
fof(f19491,plain,(
  ( ! [X17,X18,X16] : (s(bool,sK1(X17,X18,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19450,f4107])).
fof(f19450,plain,(
  ( ! [X17,X18,X16] : (s(bool,sK1(X17,X18,bool)) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X17)))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4312,f259])).
fof(f19180,plain,(
  ( ! [X92,X90,X91] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X92))),s('type.closSem.v',X91))) | p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f19179,f16150])).
fof(f19179,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X92))),s('type.closSem.v',X91)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f19101,f4107])).
fof(f19101,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X92))),s('type.closSem.v',X91))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f9506,f16151])).
fof(f19176,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))),s('type.closSem.v',X81))) | p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f19175,f12328])).
fof(f19175,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))),s('type.closSem.v',X81)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f19098,f4107])).
fof(f19098,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))),s('type.closSem.v',X81))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f9506,f12329])).
fof(f19172,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X72))),s('type.closSem.v',X71))) | p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f19171,f15070])).
fof(f19171,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X72))),s('type.closSem.v',X71)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f19095,f4107])).
fof(f19095,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X72))),s('type.closSem.v',X71))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f9506,f15039])).
fof(f19168,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v',X64))) | p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f19167,f11496])).
fof(f19167,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v',X64)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f19093,f4107])).
fof(f19093,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X65))),s('type.closSem.v',X64))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f9506,f11475])).
fof(f19165,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X61))),s('type.closSem.v',X60)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19092,f4107])).
fof(f19092,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X61))),s('type.closSem.v',X60))) | s(bool,'T_0') = s(bool,X59)) ) | $spl6),
  inference(superposition,[],[f9506,f259])).
fof(f19155,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))),s('type.closSem.v',X25)))) ) | $spl314),
  inference(cnf_transformation,[],[f19155_D])).
fof(f19155_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))),s('type.closSem.v',X25)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f19156,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))),s('type.closSem.v',X29))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19084,f4107])).
fof(f19084,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))),s('type.closSem.v',X29))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f9506,f4721])).
fof(f19154,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))),s('type.closSem.v',X25))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19083,f4107])).
fof(f19083,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))),s('type.closSem.v',X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f9506,f4721])).
fof(f19143,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))),s('type.closSem.v',X32)))) ) | $spl312),
  inference(cnf_transformation,[],[f19143_D])).
fof(f19143_D,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))),s('type.closSem.v',X32)))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f19144,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39))),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19071,f4107])).
fof(f19071,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39))),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f9506,f4721])).
fof(f19142,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))),s('type.closSem.v',X32))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f19070,f4107])).
fof(f19070,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))),s('type.closSem.v',X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | $spl6),
  inference(superposition,[],[f9506,f4721])).
fof(f18890,plain,(
  ( ! [X92,X90,X91] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X92))),s('type.closSem.v',X91))) | p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f18889,f16150])).
fof(f18889,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X92))),s('type.closSem.v',X91)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f18811,f4107])).
fof(f18811,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X91),s('type.closSem.v',X91)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X92))),s('type.closSem.v',X91))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f9393,f16151])).
fof(f18886,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))),s('type.closSem.v',X81))) | p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f18885,f12328])).
fof(f18885,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))),s('type.closSem.v',X81)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f18808,f4107])).
fof(f18808,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v',X81)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))),s('type.closSem.v',X81))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f9393,f12329])).
fof(f18882,plain,(
  ( ! [X70,X72,X71] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X72))),s('type.closSem.v',X71))) | p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f18881,f15070])).
fof(f18881,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X72))),s('type.closSem.v',X71)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f18805,f4107])).
fof(f18805,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X71),s('type.closSem.v',X71)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X69))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X72))),s('type.closSem.v',X71))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f9393,f15039])).
fof(f18878,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X65))),s('type.closSem.v',X64))) | p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f18877,f11496])).
fof(f18877,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X65))),s('type.closSem.v',X64)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f18803,f4107])).
fof(f18803,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X64),s('type.closSem.v',X64)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X65))),s('type.closSem.v',X64))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f9393,f11475])).
fof(f18875,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X61))),s('type.closSem.v',X60)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18802,f4107])).
fof(f18802,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X60),s('type.closSem.v',X60)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X61))),s('type.closSem.v',X60))) | s(bool,'T_0') = s(bool,X59)) ) | $spl6),
  inference(superposition,[],[f9393,f259])).
fof(f18865,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))),s('type.closSem.v',X25)))) ) | $spl308),
  inference(cnf_transformation,[],[f18865_D])).
fof(f18865_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))),s('type.closSem.v',X25)))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f18866,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))),s('type.closSem.v',X29))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18794,f4107])).
fof(f18794,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))),s('type.closSem.v',X29))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f9393,f4721])).
fof(f18864,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))),s('type.closSem.v',X25))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18793,f4107])).
fof(f18793,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))),s('type.closSem.v',X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f9393,f4721])).
fof(f18853,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X32)))) ) | $spl306),
  inference(cnf_transformation,[],[f18853_D])).
fof(f18853_D,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X32)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f18854,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18781,f4107])).
fof(f18781,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f9393,f4721])).
fof(f18852,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X32))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18780,f4107])).
fof(f18780,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | $spl6),
  inference(superposition,[],[f9393,f4721])).
fof(f18416,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',X18)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',sK1(X18,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X17)),'type.closSem.v'))))) ) | $spl304),
  inference(cnf_transformation,[],[f18416_D])).
fof(f18416_D,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',X18)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',sK1(X18,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X17)),'type.closSem.v'))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f18417,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X20),s('type.closSem.v',X21)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X20),s('type.closSem.v',sK1(X21,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X20)),'type.closSem.v'))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18400,f4107])).
fof(f18400,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X20),s('type.closSem.v',X21)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X20),s('type.closSem.v',sK1(X21,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X20)),'type.closSem.v')))) | s(bool,'T_0') = s(bool,X22)) ) | $spl6),
  inference(superposition,[],[f2980,f4721])).
fof(f2980,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v',sK1(X1,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0)),'type.closSem.v'))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v',X1))))) )),
  inference(forward_demodulation,[],[f2954,f92])).
fof(f2954,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v',X1)))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v',sK1(X1,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0)),'type.closSem.v')))))) )),
  inference(superposition,[],[f95,f92])).
fof(f18415,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',X18)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',sK1(X18,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X17)),'type.closSem.v'))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f18399,f4107])).
fof(f18399,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',X18)))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X17),s('type.closSem.v',sK1(X18,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X17)),'type.closSem.v'))))) ) | $spl6),
  inference(superposition,[],[f2980,f4721])).
fof(f17345,plain,(
  ( ! [X218,X219] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X219),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X218)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X219))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17344,f16150])).
fof(f17344,plain,(
  ( ! [X218,X219,X217] : (p(s(bool,X218)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X219),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X217))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X219))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17057,f4107])).
fof(f17057,plain,(
  ( ! [X218,X219,X217] : (p(s(bool,X218)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X219),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X217))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X219)) | s(bool,'T_0') = s(bool,X218)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f8770,f16151])).
fof(f10065,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X33))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | $spl206),
  inference(cnf_transformation,[],[f10065_D])).
fof(f10065_D,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X33))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f17343,plain,(
  ( ! [X216,X215] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X216))) | p(s(bool,X215)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X216))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17342,f16150])).
fof(f17342,plain,(
  ( ! [X216,X215,X214] : (p(s(bool,X215)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X216))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X214))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X216))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17056,f4107])).
fof(f17056,plain,(
  ( ! [X216,X215,X214] : (p(s(bool,X215)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X216))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X214))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X216)) | s(bool,'T_0') = s(bool,X215)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f8671,f16151])).
fof(f8671,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X2))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X2))) )),
  inference(resolution,[],[f8638,f88])).
fof(f8890,plain,(
  ( ! [X37] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X37),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X37,'type.closSem.v'))) ) | $spl186),
  inference(cnf_transformation,[],[f8890_D])).
fof(f8890_D,plain,(
  ( ! [X37] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X37),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X37,'type.closSem.v'))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f17336,plain,(
  ( ! [X175,X174] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X175),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X174)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X175,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17335,f16150])).
fof(f17335,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X174)) | s(bool,happ(s(fun('type.closSem.v',bool),X175),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X173))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X175,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17044,f4107])).
fof(f17044,plain,(
  ( ! [X175,X173,X174] : (p(s(bool,X174)) | s(bool,happ(s(fun('type.closSem.v',bool),X175),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X173))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X175,'type.closSem.v')) | s(bool,'T_0') = s(bool,X174)) ) | ($spl2 | $spl8 | $spl292)),
  inference(superposition,[],[f5284,f16151])).
fof(f5284,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X4) = s(bool,happ(s(fun('type.closSem.v',bool),X3),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X3,'type.closSem.v'))) ) | $spl2),
  inference(resolution,[],[f4041,f88])).
fof(f17324,plain,(
  ( ! [X121,X122,X120] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122))))) | p(s(bool,X120)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v',X121))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17323,f16150])).
fof(f17323,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,X120)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v',X121)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X119))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17027,f4107])).
fof(f17027,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,X120)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v',X121)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X119))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X121),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122))))) | s(bool,'T_0') = s(bool,X120)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2527,f16151])).
fof(f17322,plain,(
  ( ! [X118,X116,X117] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X118))))) | p(s(bool,X116)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v',X117))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17321,f16150])).
fof(f17321,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,X116)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v',X117)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X115))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X118)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17026,f4107])).
fof(f17026,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,X116)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v',X117)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X115))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X117),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X118))))) | s(bool,'T_0') = s(bool,X116)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2382,f16151])).
fof(f17320,plain,(
  ( ! [X114,X112,X113] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))),s('type.closSem.v',X114))) | p(s(bool,X112)) | s('type.closSem.v',X114) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17319,f16150])).
fof(f17319,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,X112)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X111))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))),s('type.closSem.v',X114))) | s('type.closSem.v',X114) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17025,f4107])).
fof(f17025,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,X112)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X111))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))),s('type.closSem.v',X114))) | s('type.closSem.v',X114) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))) | s(bool,'T_0') = s(bool,X112)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2217,f16151])).
fof(f17318,plain,(
  ( ! [X109,X110,X108] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))),s('type.closSem.v',X110))) | p(s(bool,X108)) | s('type.closSem.v',X110) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17317,f16150])).
fof(f17317,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,X108)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))),s('type.closSem.v',X110))) | s('type.closSem.v',X110) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17024,f4107])).
fof(f17024,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,X108)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))),s('type.closSem.v',X110))) | s('type.closSem.v',X110) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X109))) | s(bool,'T_0') = s(bool,X108)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f2079,f16151])).
fof(f17316,plain,(
  ( ! [X105,X106,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106))))) | p(s(bool,X104)) | s('type.closSem.v',X105) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17315,f16150])).
fof(f17315,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106))))) | s('type.closSem.v',X105) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17023,f4107])).
fof(f17023,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X103))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106))))) | s('type.closSem.v',X105) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X106))) | s(bool,'T_0') = s(bool,X104)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f1911,f16151])).
fof(f17314,plain,(
  ( ! [X101,X102,X100] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X101),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))))) | p(s(bool,X100)) | s('type.closSem.v',X101) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17313,f16150])).
fof(f17313,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X99))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X101),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))))) | s('type.closSem.v',X101) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17022,f4107])).
fof(f17022,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X99))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X101),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))))) | s('type.closSem.v',X101) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X102))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f1773,f16151])).
fof(f8452,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X63,bool),X64),s(X63,X65))) | s(X63,X65) = s(X63,sK1(X65,X64,X63))) ) | $spl182),
  inference(cnf_transformation,[],[f8452_D])).
fof(f8452_D,plain,(
  ( ! [X64,X65,X63] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X63,bool),X64),s(X63,X65))) | s(X63,X65) = s(X63,sK1(X65,X64,X63))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f17305,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) | p(s(bool,X75)) | s(X76,X78) = s(X76,sK1(X78,X77,X76))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f17304,f16150])).
fof(f17304,plain,(
  ( ! [X78,X76,X74,X77,X75] : (p(s(bool,X75)) | s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X74))),s('type.closSem.v','const.closSem.Unit_0'))) | s(X76,X78) = s(X76,sK1(X78,X77,X76))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f17016,f4107])).
fof(f17016,plain,(
  ( ! [X78,X76,X74,X77,X75] : (p(s(bool,X75)) | s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X74))),s('type.closSem.v','const.closSem.Unit_0'))) | s(X76,X78) = s(X76,sK1(X78,X77,X76)) | s(bool,'T_0') = s(bool,X75)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f1718,f16151])).
fof(f1718,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X7)) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X6),s(X4,X5))) | s(X4,X5) = s(X4,sK1(X5,X6,X4))) )),
  inference(resolution,[],[f94,f88])).
fof(f17158,plain,(
  ( ! [X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl302),
  inference(cnf_transformation,[],[f17158_D])).
fof(f17158_D,plain,(
  ( ! [X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f17166,plain,(
  ( ! [X182,X181] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X181)) ) | ($spl8 | $spl84 | $spl176 | $spl292)),
  inference(forward_demodulation,[],[f17083,f2028])).
fof(f17083,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))))) | s(bool,'T_0') = s(bool,X181)) ) | ($spl8 | $spl176 | $spl292)),
  inference(duplicate_literal_removal,[],[f16861])).
fof(f16861,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | ($spl8 | $spl176 | $spl292)),
  inference(superposition,[],[f9055,f16151])).
fof(f9055,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X15))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X15))))) | s(bool,'T_0') = s(bool,X16)) ) | $spl176),
  inference(resolution,[],[f1999,f8143])).
fof(f17165,plain,(
  ( ! [X178,X179] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X178)) ) | ($spl8 | $spl82 | $spl176 | $spl292)),
  inference(forward_demodulation,[],[f17084,f1863])).
fof(f17084,plain,(
  ( ! [X180,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))))) | s(bool,'T_0') = s(bool,X178)) ) | ($spl8 | $spl176 | $spl292)),
  inference(duplicate_literal_removal,[],[f16860])).
fof(f16860,plain,(
  ( ! [X180,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))))) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X178)) ) | ($spl8 | $spl176 | $spl292)),
  inference(superposition,[],[f9019,f16151])).
fof(f9019,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))))) | s(bool,'T_0') = s(bool,X16)) ) | $spl176),
  inference(resolution,[],[f1836,f8143])).
fof(f17157,plain,(
  ( ! [X165,X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl8 | $spl86 | $spl176 | $spl292)),
  inference(forward_demodulation,[],[f17085,f2717])).
fof(f17085,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.list.LIST_REL_3'(s(fun(X167,fun(X168,bool)),X169),s('type.list.list'(X167),'const.list.NIL_0'),s('type.list.list'(X168),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl8 | $spl86 | $spl176 | $spl292)),
  inference(duplicate_literal_removal,[],[f16856])).
fof(f16856,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.list.LIST_REL_3'(s(fun(X167,fun(X168,bool)),X169),s('type.list.list'(X167),'const.list.NIL_0'),s('type.list.list'(X168),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X165) | s(bool,'T_0') = s(bool,X165)) ) | ($spl8 | $spl86 | $spl176 | $spl292)),
  inference(superposition,[],[f8200,f16151])).
fof(f8200,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X44) != s(bool,'const.list.LIST_REL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.list.list'(X41),'const.list.NIL_0'),s('type.list.list'(X42),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl86 | $spl176)),
  inference(subsumption_resolution,[],[f8191,f124])).
fof(f8191,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X44) != s(bool,'const.list.LIST_REL_3'(s(fun(X41,fun(X42,bool)),X43),s('type.list.list'(X41),'const.list.NIL_0'),s('type.list.list'(X42),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl86 | $spl176)),
  inference(superposition,[],[f8143,f2717])).
fof(f17154,plain,(
  ( ! [X154,X152,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X154) | s(bool,X153) = s(bool,X154)) ) | $spl300),
  inference(cnf_transformation,[],[f17154_D])).
fof(f17154_D,plain,(
  ( ! [X154,X152,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X154) | s(bool,X153) = s(bool,X154)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f17086,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X151)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292)),
  inference(duplicate_literal_removal,[],[f16852])).
fof(f16852,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292)),
  inference(superposition,[],[f6249,f16151])).
fof(f17153,plain,(
  ( ! [X146,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl298),
  inference(cnf_transformation,[],[f17153_D])).
fof(f17153_D,plain,(
  ( ! [X146,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f17087,plain,(
  ( ! [X146,X144,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X144)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(duplicate_literal_removal,[],[f16850])).
fof(f16850,plain,(
  ( ! [X146,X144,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(superposition,[],[f6091,f16151])).
fof(f17126,plain,(
  ( ! [X43,X41,X44] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8 | $spl292)),
  inference(forward_demodulation,[],[f17125,f16150])).
fof(f17125,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl8 | $spl292)),
  inference(subsumption_resolution,[],[f16820,f396])).
fof(f16820,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X41)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl8 | $spl292)),
  inference(superposition,[],[f1730,f16151])).
fof(f16393,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) | $spl294),
  inference(cnf_transformation,[],[f16393_D])).
fof(f16393_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f16646,plain,(
  ( ! [X125,X126,X124] : (s(bool,X126) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X124))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16569,f259])).
fof(f16569,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X126) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X124))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl292)),
  inference(superposition,[],[f6249,f16150])).
fof(f16402,plain,(
  ( ! [X42] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl296),
  inference(cnf_transformation,[],[f16402_D])).
fof(f16402_D,plain,(
  ( ! [X42] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f16412,plain,(
  ( ! [X57,X56] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X56))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl292)),
  inference(forward_demodulation,[],[f16411,f12328])).
fof(f16411,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X56)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X55))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258 | $spl292)),
  inference(subsumption_resolution,[],[f16383,f4107])).
fof(f16383,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X56)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X55))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl8 | $spl258 | $spl292)),
  inference(superposition,[],[f16141,f12329])).
fof(f16141,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl292),
  inference(resolution,[],[f16131,f88])).
fof(f16408,plain,(
  ( ! [X50,X49] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X49))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f16407,f15070])).
fof(f16407,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X49)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16380,f4107])).
fof(f16380,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X49)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X48))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl2 | $spl8 | $spl292)),
  inference(superposition,[],[f16141,f15039])).
fof(f16404,plain,(
  ( ! [X45,X44] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X44))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(forward_demodulation,[],[f16403,f11496])).
fof(f16403,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X44)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16378,f4107])).
fof(f16378,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X44)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl1 | $spl2 | $spl8 | $spl292)),
  inference(superposition,[],[f16141,f11475])).
fof(f16401,plain,(
  ( ! [X41,X42] : (p(s(bool,X41)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16377,f4107])).
fof(f16377,plain,(
  ( ! [X41,X42] : (p(s(bool,X41)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl6 | $spl292)),
  inference(superposition,[],[f16141,f259])).
fof(f16394,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X22))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16370,f4107])).
fof(f16370,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X22))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl292)),
  inference(superposition,[],[f16141,f4721])).
fof(f16392,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl292)),
  inference(subsumption_resolution,[],[f16369,f4107])).
fof(f16369,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl292)),
  inference(superposition,[],[f16141,f4721])).
fof(f15999,plain,(
  ( ! [X212,X213] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X213),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X212)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X213))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15998,f15070])).
fof(f15998,plain,(
  ( ! [X212,X213,X211] : (p(s(bool,X212)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X213),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X211))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X213))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15720,f4107])).
fof(f15720,plain,(
  ( ! [X212,X213,X211] : (p(s(bool,X212)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X213),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X211))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X213)) | s(bool,'T_0') = s(bool,X212)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f8770,f15039])).
fof(f15997,plain,(
  ( ! [X210,X209] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X210))) | p(s(bool,X209)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X210))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15996,f15070])).
fof(f15996,plain,(
  ( ! [X210,X208,X209] : (p(s(bool,X209)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X210))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X208))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X210))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15719,f4107])).
fof(f15719,plain,(
  ( ! [X210,X208,X209] : (p(s(bool,X209)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X210))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X208))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X210)) | s(bool,'T_0') = s(bool,X209)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f8671,f15039])).
fof(f15990,plain,(
  ( ! [X169,X168] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X169),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X168)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X169,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15989,f15070])).
fof(f15989,plain,(
  ( ! [X167,X169,X168] : (p(s(bool,X168)) | s(bool,happ(s(fun('type.closSem.v',bool),X169),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X167))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X169,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15707,f4107])).
fof(f15707,plain,(
  ( ! [X167,X169,X168] : (p(s(bool,X168)) | s(bool,happ(s(fun('type.closSem.v',bool),X169),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X167))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X169,'type.closSem.v')) | s(bool,'T_0') = s(bool,X168)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f5284,f15039])).
fof(f15978,plain,(
  ( ! [X116,X114,X115] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116))))) | p(s(bool,X114)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15977,f15070])).
fof(f15977,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,X114)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15690,f4107])).
fof(f15690,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,X114)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X113))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116))))) | s(bool,'T_0') = s(bool,X114)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2527,f15039])).
fof(f15976,plain,(
  ( ! [X111,X112,X110] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X112))))) | p(s(bool,X110)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v',X111))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15975,f15070])).
fof(f15975,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,X110)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v',X111)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X109))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X112)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15689,f4107])).
fof(f15689,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,X110)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v',X111)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X109))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X112))))) | s(bool,'T_0') = s(bool,X110)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2382,f15039])).
fof(f15974,plain,(
  ( ! [X107,X108,X106] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))),s('type.closSem.v',X108))) | p(s(bool,X106)) | s('type.closSem.v',X108) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15973,f15070])).
fof(f15973,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,X106)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X105))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))),s('type.closSem.v',X108))) | s('type.closSem.v',X108) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15688,f4107])).
fof(f15688,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,X106)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X105))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))),s('type.closSem.v',X108))) | s('type.closSem.v',X108) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X107))) | s(bool,'T_0') = s(bool,X106)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2217,f15039])).
fof(f15972,plain,(
  ( ! [X103,X102,X104] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103))),s('type.closSem.v',X104))) | p(s(bool,X102)) | s('type.closSem.v',X104) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15971,f15070])).
fof(f15971,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X101))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103))),s('type.closSem.v',X104))) | s('type.closSem.v',X104) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15687,f4107])).
fof(f15687,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X101))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103))),s('type.closSem.v',X104))) | s('type.closSem.v',X104) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X103))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2079,f15039])).
fof(f15970,plain,(
  ( ! [X99,X100,X98] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X99),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100))))) | p(s(bool,X98)) | s('type.closSem.v',X99) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15969,f15070])).
fof(f15969,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,X98)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X97))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X99),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100))))) | s('type.closSem.v',X99) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15686,f4107])).
fof(f15686,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,X98)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X97))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X99),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100))))) | s('type.closSem.v',X99) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X100))) | s(bool,'T_0') = s(bool,X98)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1911,f15039])).
fof(f15968,plain,(
  ( ! [X94,X95,X96] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X95),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96))))) | p(s(bool,X94)) | s('type.closSem.v',X95) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15967,f15070])).
fof(f15967,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X95),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96))))) | s('type.closSem.v',X95) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15685,f4107])).
fof(f15685,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X95),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96))))) | s('type.closSem.v',X95) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X96))) | s(bool,'T_0') = s(bool,X94)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1773,f15039])).
fof(f15959,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))) | p(s(bool,X69)) | s(X70,X72) = s(X70,sK1(X72,X71,X70))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15958,f15070])).
fof(f15958,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X69)) | s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68))))) | s(X70,X72) = s(X70,sK1(X72,X71,X70))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15679,f4107])).
fof(f15679,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X69)) | s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X68))))) | s(X70,X72) = s(X70,sK1(X72,X71,X70)) | s(bool,'T_0') = s(bool,X69)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1718,f15039])).
fof(f15947,plain,(
  ( ! [X20] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X20)))) ) | $spl292),
  inference(cnf_transformation,[],[f15947_D])).
fof(f15947_D,plain,(
  ( ! [X20] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X20)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f15946,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X20)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15660,f4107])).
fof(f15660,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X20))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f112,f15039])).
fof(f15817,plain,(
  ( ! [X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166)))))) ) | $spl290),
  inference(cnf_transformation,[],[f15817_D])).
fof(f15817_D,plain,(
  ( ! [X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166)))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f15825,plain,(
  ( ! [X182,X181] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))))) | s(bool,'T_0') = s(bool,X181)) ) | ($spl1 | $spl2 | $spl8 | $spl84 | $spl176)),
  inference(forward_demodulation,[],[f15744,f2028])).
fof(f15744,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))))) | s(bool,'T_0') = s(bool,X181)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(duplicate_literal_removal,[],[f15534])).
fof(f15534,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X182))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X183))))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,X181)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(superposition,[],[f9055,f15039])).
fof(f15824,plain,(
  ( ! [X178,X179] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) | s(bool,'T_0') = s(bool,X178)) ) | ($spl1 | $spl2 | $spl8 | $spl82 | $spl176)),
  inference(forward_demodulation,[],[f15745,f1863])).
fof(f15745,plain,(
  ( ! [X180,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))))) | s(bool,'T_0') = s(bool,X178)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(duplicate_literal_removal,[],[f15533])).
fof(f15533,plain,(
  ( ! [X180,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X180))))) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X178)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(superposition,[],[f9019,f15039])).
fof(f15816,plain,(
  ( ! [X165,X166] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(forward_demodulation,[],[f15746,f2717])).
fof(f15746,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))))) != s(bool,'const.list.LIST_REL_3'(s(fun(X167,fun(X168,bool)),X169),s('type.list.list'(X167),'const.list.NIL_0'),s('type.list.list'(X168),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(duplicate_literal_removal,[],[f15529])).
fof(f15529,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X166))))) != s(bool,'const.list.LIST_REL_3'(s(fun(X167,fun(X168,bool)),X169),s('type.list.list'(X167),'const.list.NIL_0'),s('type.list.list'(X168),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X165) | s(bool,'T_0') = s(bool,X165)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(superposition,[],[f8200,f15039])).
fof(f15813,plain,(
  ( ! [X154,X152,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))))) | s(bool,'T_0') = s(bool,X154) | s(bool,X153) = s(bool,X154)) ) | $spl288),
  inference(cnf_transformation,[],[f15813_D])).
fof(f15813_D,plain,(
  ( ! [X154,X152,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))))) | s(bool,'T_0') = s(bool,X154) | s(bool,X153) = s(bool,X154)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f15747,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))))) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X151)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(duplicate_literal_removal,[],[f15525])).
fof(f15525,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,X153) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X152))))) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f15039])).
fof(f15812,plain,(
  ( ! [X146,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl286),
  inference(cnf_transformation,[],[f15812_D])).
fof(f15812_D,plain,(
  ( ! [X146,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f15748,plain,(
  ( ! [X146,X144,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X144)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(duplicate_literal_removal,[],[f15523])).
fof(f15523,plain,(
  ( ! [X146,X144,X145] : (s(bool,X146) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X145))))) | s(bool,X146) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f6091,f15039])).
fof(f15785,plain,(
  ( ! [X43,X41,X44] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f15784,f15070])).
fof(f15784,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42)))))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f15493,f396])).
fof(f15493,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42)))))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X41)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1730,f15039])).
fof(f15138,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl282),
  inference(cnf_transformation,[],[f15138_D])).
fof(f15138_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f15350,plain,(
  ( ! [X123,X124,X122] : (s(bool,X124) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122))))) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X123)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15278,f259])).
fof(f15278,plain,(
  ( ! [X123,X124,X122] : (s(bool,X123) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X124) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122))))) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X123)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f15070])).
fof(f15145,plain,(
  ( ! [X39] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39)))))) ) | $spl284),
  inference(cnf_transformation,[],[f15145_D])).
fof(f15145_D,plain,(
  ( ! [X39] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39)))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f15151,plain,(
  ( ! [X48,X49] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X49))))) | p(s(bool,X48))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f15150,f12328])).
fof(f15150,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X48)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X49))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X47))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f15128,f4107])).
fof(f15128,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X48)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X49))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X47))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X48)) ) | ($spl1 | $spl2 | $spl8 | $spl258)),
  inference(superposition,[],[f14827,f12329])).
fof(f15147,plain,(
  ( ! [X41,X42] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))))) | p(s(bool,X41))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15146,f11496])).
fof(f15146,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X41)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f15125,f4107])).
fof(f15125,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X41)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X42))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X40))))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f14827,f11475])).
fof(f15144,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f15124,f4107])).
fof(f15124,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39))))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f14827,f259])).
fof(f15139,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X22))))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f15118,f4107])).
fof(f15118,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X22))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f14827,f4721])).
fof(f15137,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f15117,f4107])).
fof(f15117,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f14827,f4721])).
fof(f15031,plain,(
  ( ! [X66,X67,X65] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X67))))) | p(s(bool,X65)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v',X66))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f15030,f12328])).
fof(f15030,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v',X66)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X64))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X67)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f14978,f4107])).
fof(f14978,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v',X66)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X64))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X66),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X67))))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2527,f12329])).
fof(f15027,plain,(
  ( ! [X57,X56,X55] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57))))) | p(s(bool,X55)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v',X56))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f15026,f11496])).
fof(f15026,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v',X56)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X54))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f14975,f4107])).
fof(f14975,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v',X56)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X54))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X56),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X57))))) | s(bool,'T_0') = s(bool,X55)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2527,f11475])).
fof(f15024,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v',X52)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X53)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14974,f4107])).
fof(f14974,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v',X52)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X53))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl6),
  inference(superposition,[],[f2527,f259])).
fof(f15018,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))))) ) | $spl278),
  inference(cnf_transformation,[],[f15018_D])).
fof(f15018_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f15019,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14968,f4107])).
fof(f14968,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f2527,f4721])).
fof(f15017,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14967,f4107])).
fof(f14967,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2527,f4721])).
fof(f15010,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35)))))) ) | $spl276),
  inference(cnf_transformation,[],[f15010_D])).
fof(f15010_D,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35)))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f15011,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14957,f4107])).
fof(f14957,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X39))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f2527,f4721])).
fof(f15009,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14956,f4107])).
fof(f14956,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X35))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | $spl6),
  inference(superposition,[],[f2527,f4721])).
fof(f14990,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10)))))) ) | $spl274),
  inference(cnf_transformation,[],[f14990_D])).
fof(f14990_D,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10)))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f15007,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f14947,f4107])).
fof(f14947,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2527,f259])).
fof(f14830,plain,(
  ( ! [X74,X75] : (p(s(bool,X74)) | s(bool,X74) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X75))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))) )),
  inference(resolution,[],[f2527,f8638])).
fof(f11379,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')))) | $spl238),
  inference(cnf_transformation,[],[f11379_D])).
fof(f11379_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f14811,plain,(
  ( ! [X14,X13] : (p(s(bool,X13)) | s(bool,X13) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X14))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))))) )),
  inference(resolution,[],[f2527,f3685])).
fof(f3685,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X0))))) )),
  inference(superposition,[],[f3494,f117])).
fof(f3494,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0))))) )),
  inference(forward_demodulation,[],[f121,f92])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK3(X2,X1,X0)))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f11377,plain,(
  s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')) | $spl236),
  inference(cnf_transformation,[],[f11377_D])).
fof(f11377_D,plain,(
  s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f14810,plain,(
  ( ! [X12,X11] : (p(s(bool,X11)) | s(bool,X11) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))) )),
  inference(resolution,[],[f2527,f8750])).
fof(f11374,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Unit_0')),'type.closSem.v')) | $spl232),
  inference(cnf_transformation,[],[f11374_D])).
fof(f11374_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Unit_0')),'type.closSem.v')) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f14809,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Unit_0')),'type.closSem.v'))) ) | $spl2),
  inference(resolution,[],[f2527,f5287])).
fof(f5287,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0)),'type.closSem.v'))) ) | $spl2),
  inference(superposition,[],[f4041,f92])).
fof(f14062,plain,(
  ( ! [X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))))))) ) | $spl272),
  inference(cnf_transformation,[],[f14062_D])).
fof(f14062_D,plain,(
  ( ! [X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f14063,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X16),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18)))))) | ~p(s(bool,X17)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X16))))) ) | ($spl8 | $spl256)),
  inference(subsumption_resolution,[],[f14050,f396])).
fof(f14050,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X16),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18)))))) | ~p(s(bool,X17)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X16))))) ) | $spl256),
  inference(superposition,[],[f12271,f91])).
fof(f12271,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X3))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X3),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X2))))))) ) | $spl256),
  inference(superposition,[],[f12085,f109])).
fof(f14061,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15)))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,X14))) ) | ($spl8 | $spl256)),
  inference(subsumption_resolution,[],[f14049,f396])).
fof(f14049,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15)))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,X14))) ) | $spl256),
  inference(superposition,[],[f12271,f91])).
fof(f13896,plain,(
  ( ! [X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))),s('type.closSem.v',X13))))) ) | $spl270),
  inference(cnf_transformation,[],[f13896_D])).
fof(f13896_D,plain,(
  ( ! [X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))),s('type.closSem.v',X13))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f13897,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))),s('type.closSem.v',X16)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X16))))) ) | ($spl8 | $spl256)),
  inference(subsumption_resolution,[],[f13884,f396])).
fof(f13884,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))),s('type.closSem.v',X16)))) | ~p(s(bool,X17)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X16))))) ) | $spl256),
  inference(superposition,[],[f12270,f91])).
fof(f12270,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X1))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0))),s('type.closSem.v',X1))))) ) | $spl256),
  inference(superposition,[],[f12085,f113])).
fof(f13895,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,X14))) ) | ($spl8 | $spl256)),
  inference(subsumption_resolution,[],[f13883,f396])).
fof(f13883,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X15))),s('type.closSem.v',X13)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X13)))) | ~p(s(bool,X14))) ) | $spl256),
  inference(superposition,[],[f12270,f91])).
fof(f13318,plain,(
  ( ! [X191,X190] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X191),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X190)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X191))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13317,f12328])).
fof(f13317,plain,(
  ( ! [X191,X189,X190] : (p(s(bool,X190)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X191),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X189))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X191))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13083,f4107])).
fof(f13083,plain,(
  ( ! [X191,X189,X190] : (p(s(bool,X190)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X191),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X189))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X191)) | s(bool,'T_0') = s(bool,X190)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f8770,f12329])).
fof(f13316,plain,(
  ( ! [X187,X188] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X188))) | p(s(bool,X187)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X188))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13315,f12328])).
fof(f13315,plain,(
  ( ! [X187,X188,X186] : (p(s(bool,X187)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X188))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X186))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X188))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13082,f4107])).
fof(f13082,plain,(
  ( ! [X187,X188,X186] : (p(s(bool,X187)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X188))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X186))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X188)) | s(bool,'T_0') = s(bool,X187)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f8671,f12329])).
fof(f13309,plain,(
  ( ! [X146,X147] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X147),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X146)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X147,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13308,f12328])).
fof(f13308,plain,(
  ( ! [X146,X147,X145] : (p(s(bool,X146)) | s(bool,happ(s(fun('type.closSem.v',bool),X147),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X147,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13070,f4107])).
fof(f13070,plain,(
  ( ! [X146,X147,X145] : (p(s(bool,X146)) | s(bool,happ(s(fun('type.closSem.v',bool),X147),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X145))),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X147,'type.closSem.v')) | s(bool,'T_0') = s(bool,X146)) ) | ($spl2 | $spl8 | $spl258)),
  inference(superposition,[],[f5284,f12329])).
fof(f13297,plain,(
  ( ! [X94,X92,X93] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X94))))) | p(s(bool,X92)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v',X93))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13296,f12328])).
fof(f13296,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v',X93)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X91))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X94)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13053,f4107])).
fof(f13053,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v',X93)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X91))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X93),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X94))))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2382,f12329])).
fof(f13295,plain,(
  ( ! [X90,X88,X89] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v',X90))) | p(s(bool,X88)) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13294,f12328])).
fof(f13294,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X87))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v',X90))) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13052,f4107])).
fof(f13052,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X87))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))),s('type.closSem.v',X90))) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X89))) | s(bool,'T_0') = s(bool,X88)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2217,f12329])).
fof(f13293,plain,(
  ( ! [X85,X86,X84] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v',X86))) | p(s(bool,X84)) | s('type.closSem.v',X86) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13292,f12328])).
fof(f13292,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X83))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v',X86))) | s('type.closSem.v',X86) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13051,f4107])).
fof(f13051,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X83))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))),s('type.closSem.v',X86))) | s('type.closSem.v',X86) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X85))) | s(bool,'T_0') = s(bool,X84)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f2079,f12329])).
fof(f13291,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))))) | p(s(bool,X80)) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13290,f12328])).
fof(f13290,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))))) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13050,f4107])).
fof(f13050,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))))) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X82))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f1911,f12329])).
fof(f13289,plain,(
  ( ! [X78,X76,X77] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X77),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78))))) | p(s(bool,X76)) | s('type.closSem.v',X77) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13288,f12328])).
fof(f13288,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X77),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78))))) | s('type.closSem.v',X77) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13049,f4107])).
fof(f13049,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X75))),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X77),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78))))) | s('type.closSem.v',X77) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X78))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f1773,f12329])).
fof(f13280,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X52,bool),X53),s(X52,X54))) | p(s(bool,X51)) | s(X52,X54) = s(X52,sK1(X54,X53,X52))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f13279,f12328])).
fof(f13279,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | s(bool,happ(s(fun(X52,bool),X53),s(X52,X54))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | s(X52,X54) = s(X52,sK1(X54,X53,X52))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f13043,f4107])).
fof(f13043,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | s(bool,happ(s(fun(X52,bool),X53),s(X52,X54))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X50))),s('type.closSem.v','const.closSem.Unit_0'))) | s(X52,X54) = s(X52,sK1(X54,X53,X52)) | s(bool,'T_0') = s(bool,X51)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f1718,f12329])).
fof(f13170,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl268),
  inference(cnf_transformation,[],[f13170_D])).
fof(f13170_D,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f13178,plain,(
  ( ! [X177,X178] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X177)) ) | ($spl8 | $spl84 | $spl176 | $spl258)),
  inference(forward_demodulation,[],[f13101,f2028])).
fof(f13101,plain,(
  ( ! [X177,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) | s(bool,'T_0') = s(bool,X177)) ) | ($spl8 | $spl176 | $spl258)),
  inference(duplicate_literal_removal,[],[f12932])).
fof(f12932,plain,(
  ( ! [X177,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) | s(bool,'T_0') = s(bool,X177) | s(bool,'T_0') = s(bool,X177)) ) | ($spl8 | $spl176 | $spl258)),
  inference(superposition,[],[f9055,f12329])).
fof(f13177,plain,(
  ( ! [X175,X174] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X174)) ) | ($spl8 | $spl82 | $spl176 | $spl258)),
  inference(forward_demodulation,[],[f13102,f1863])).
fof(f13102,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))))) | s(bool,'T_0') = s(bool,X174)) ) | ($spl8 | $spl176 | $spl258)),
  inference(duplicate_literal_removal,[],[f12931])).
fof(f12931,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))))) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl8 | $spl176 | $spl258)),
  inference(superposition,[],[f9019,f12329])).
fof(f13169,plain,(
  ( ! [X161,X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X161)) ) | ($spl8 | $spl86 | $spl176 | $spl258)),
  inference(forward_demodulation,[],[f13103,f2717])).
fof(f13103,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.list.LIST_REL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.list.list'(X163),'const.list.NIL_0'),s('type.list.list'(X164),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X161)) ) | ($spl8 | $spl86 | $spl176 | $spl258)),
  inference(duplicate_literal_removal,[],[f12927])).
fof(f12927,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v','const.closSem.Unit_0'))) != s(bool,'const.list.LIST_REL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.list.list'(X163),'const.list.NIL_0'),s('type.list.list'(X164),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,X161)) ) | ($spl8 | $spl86 | $spl176 | $spl258)),
  inference(superposition,[],[f8200,f12329])).
fof(f13166,plain,(
  ( ! [X149,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,X149) = s(bool,X150)) ) | $spl266),
  inference(cnf_transformation,[],[f13166_D])).
fof(f13166_D,plain,(
  ( ! [X149,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,X149) = s(bool,X150)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f13104,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X149) = s(bool,X150) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X147)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258)),
  inference(duplicate_literal_removal,[],[f12923])).
fof(f12923,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X149) = s(bool,X150) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,X147)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258)),
  inference(superposition,[],[f6249,f12329])).
fof(f13165,plain,(
  ( ! [X142,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl264),
  inference(cnf_transformation,[],[f13165_D])).
fof(f13165_D,plain,(
  ( ! [X142,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f13105,plain,(
  ( ! [X142,X140,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X140)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(duplicate_literal_removal,[],[f12921])).
fof(f12921,plain,(
  ( ! [X142,X140,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(superposition,[],[f6091,f12329])).
fof(f13140,plain,(
  ( ! [X43,X41,X44] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8 | $spl258)),
  inference(forward_demodulation,[],[f13139,f12328])).
fof(f13139,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl8 | $spl258)),
  inference(subsumption_resolution,[],[f12892,f396])).
fof(f12892,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42))),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X41)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl8 | $spl258)),
  inference(superposition,[],[f1730,f12329])).
fof(f12519,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) | $spl260),
  inference(cnf_transformation,[],[f12519_D])).
fof(f12519_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f12732,plain,(
  ( ! [X109,X107,X108] : (s(bool,X109) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X107))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X108)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f12665,f259])).
fof(f12665,plain,(
  ( ! [X109,X107,X108] : (s(bool,X108) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X109) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X107))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X108)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15 | $spl258)),
  inference(superposition,[],[f6249,f12328])).
fof(f12524,plain,(
  ( ! [X36] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl262),
  inference(cnf_transformation,[],[f12524_D])).
fof(f12524_D,plain,(
  ( ! [X36] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f12526,plain,(
  ( ! [X39,X38] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X38))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(forward_demodulation,[],[f12525,f11496])).
fof(f12525,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X38)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X37))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f12508,f4107])).
fof(f12508,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X38)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X37))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl2 | $spl8 | $spl258)),
  inference(superposition,[],[f12320,f11475])).
fof(f12320,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl258),
  inference(resolution,[],[f12309,f88])).
fof(f12523,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f12507,f4107])).
fof(f12507,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X36))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl6 | $spl258)),
  inference(superposition,[],[f12320,f259])).
fof(f12520,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f12502,f4107])).
fof(f12502,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl258)),
  inference(superposition,[],[f12320,f4721])).
fof(f12518,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl258)),
  inference(subsumption_resolution,[],[f12501,f4107])).
fof(f12501,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl258)),
  inference(superposition,[],[f12320,f4721])).
fof(f12230,plain,(
  ( ! [X194,X195] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X195),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X194)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X195))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12229,f11496])).
fof(f12229,plain,(
  ( ! [X194,X195,X193] : (p(s(bool,X194)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X195),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X193))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X195))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f12002,f4107])).
fof(f12002,plain,(
  ( ! [X194,X195,X193] : (p(s(bool,X194)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X195),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X193))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X195)) | s(bool,'T_0') = s(bool,X194)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f8770,f11475])).
fof(f12228,plain,(
  ( ! [X191,X192] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X192))) | p(s(bool,X191)) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X192))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12227,f11496])).
fof(f12227,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X192))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X190))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X192))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f12001,f4107])).
fof(f12001,plain,(
  ( ! [X191,X192,X190] : (p(s(bool,X191)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X192))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X190))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X192)) | s(bool,'T_0') = s(bool,X191)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f8671,f11475])).
fof(f12221,plain,(
  ( ! [X151,X150] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X151),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X150)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X151,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12220,f11496])).
fof(f12220,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X150)) | s(bool,happ(s(fun('type.closSem.v',bool),X151),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X149))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X151,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11989,f4107])).
fof(f11989,plain,(
  ( ! [X151,X149,X150] : (p(s(bool,X150)) | s(bool,happ(s(fun('type.closSem.v',bool),X151),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X149))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X151,'type.closSem.v')) | s(bool,'T_0') = s(bool,X150)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f5284,f11475])).
fof(f12209,plain,(
  ( ! [X97,X98,X96] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X98))))) | p(s(bool,X96)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v',X97))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12208,f11496])).
fof(f12208,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v',X97)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X95))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X98)))))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11972,f4107])).
fof(f11972,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,X96)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v',X97)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X95))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X97),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X98))))) | s(bool,'T_0') = s(bool,X96)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2382,f11475])).
fof(f12207,plain,(
  ( ! [X94,X92,X93] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))),s('type.closSem.v',X94))) | p(s(bool,X92)) | s('type.closSem.v',X94) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12206,f11496])).
fof(f12206,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X91))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))),s('type.closSem.v',X94))) | s('type.closSem.v',X94) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11971,f4107])).
fof(f11971,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X92)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X91))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))),s('type.closSem.v',X94))) | s('type.closSem.v',X94) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X93))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2217,f11475])).
fof(f12205,plain,(
  ( ! [X90,X88,X89] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89))),s('type.closSem.v',X90))) | p(s(bool,X88)) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12204,f11496])).
fof(f12204,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X87))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89))),s('type.closSem.v',X90))) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11970,f4107])).
fof(f11970,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X87))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89))),s('type.closSem.v',X90))) | s('type.closSem.v',X90) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X89))) | s(bool,'T_0') = s(bool,X88)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f2079,f11475])).
fof(f12203,plain,(
  ( ! [X85,X86,X84] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X85),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86))))) | p(s(bool,X84)) | s('type.closSem.v',X85) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12202,f11496])).
fof(f12202,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X83))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X85),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86))))) | s('type.closSem.v',X85) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11969,f4107])).
fof(f11969,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X83))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X85),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86))))) | s('type.closSem.v',X85) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X86))) | s(bool,'T_0') = s(bool,X84)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1911,f11475])).
fof(f12201,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))))) | p(s(bool,X80)) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12200,f11496])).
fof(f12200,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))))) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11968,f4107])).
fof(f11968,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X79))))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X81),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))))) | s('type.closSem.v',X81) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X82))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1773,f11475])).
fof(f12192,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X56,bool),X57),s(X56,X58))) | p(s(bool,X55)) | s(X56,X58) = s(X56,sK1(X58,X57,X56))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f12191,f11496])).
fof(f12191,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X55)) | s(bool,happ(s(fun(X56,bool),X57),s(X56,X58))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X54))))) | s(X56,X58) = s(X56,sK1(X58,X57,X56))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11962,f4107])).
fof(f11962,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X55)) | s(bool,happ(s(fun(X56,bool),X57),s(X56,X58))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X54))))) | s(X56,X58) = s(X56,sK1(X58,X57,X56)) | s(bool,'T_0') = s(bool,X55)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1718,f11475])).
fof(f12180,plain,(
  ( ! [X16] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16)))) ) | $spl258),
  inference(cnf_transformation,[],[f12180_D])).
fof(f12180_D,plain,(
  ( ! [X16] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16)))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f12179,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16)))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11947,f4107])).
fof(f11947,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f110,f11475])).
fof(f12085,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162)))))) ) | $spl256),
  inference(cnf_transformation,[],[f12085_D])).
fof(f12085_D,plain,(
  ( ! [X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162)))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f12093,plain,(
  ( ! [X177,X178] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))))) | s(bool,'T_0') = s(bool,X177)) ) | ($spl1 | $spl2 | $spl8 | $spl84 | $spl176)),
  inference(forward_demodulation,[],[f12018,f2028])).
fof(f12018,plain,(
  ( ! [X177,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) | s(bool,'T_0') = s(bool,X177)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(duplicate_literal_removal,[],[f11856])).
fof(f11856,plain,(
  ( ! [X177,X178,X179] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X178))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X179))))) | s(bool,'T_0') = s(bool,X177) | s(bool,'T_0') = s(bool,X177)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(superposition,[],[f9055,f11475])).
fof(f12092,plain,(
  ( ! [X175,X174] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))))) | s(bool,'T_0') = s(bool,X174)) ) | ($spl1 | $spl2 | $spl8 | $spl82 | $spl176)),
  inference(forward_demodulation,[],[f12019,f1863])).
fof(f12019,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))))) | s(bool,'T_0') = s(bool,X174)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(duplicate_literal_removal,[],[f11855])).
fof(f11855,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X175))))) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X176))))) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl1 | $spl2 | $spl8 | $spl176)),
  inference(superposition,[],[f9019,f11475])).
fof(f12084,plain,(
  ( ! [X161,X162] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))))) | s(bool,'T_0') = s(bool,X161)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(forward_demodulation,[],[f12020,f2717])).
fof(f12020,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))))) != s(bool,'const.list.LIST_REL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.list.list'(X163),'const.list.NIL_0'),s('type.list.list'(X164),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X161)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(duplicate_literal_removal,[],[f11851])).
fof(f11851,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))))) != s(bool,'const.list.LIST_REL_3'(s(fun(X163,fun(X164,bool)),X165),s('type.list.list'(X163),'const.list.NIL_0'),s('type.list.list'(X164),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,X161)) ) | ($spl1 | $spl2 | $spl8 | $spl86 | $spl176)),
  inference(superposition,[],[f8200,f11475])).
fof(f12081,plain,(
  ( ! [X149,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))))) | s(bool,'T_0') = s(bool,X150) | s(bool,X149) = s(bool,X150)) ) | $spl254),
  inference(cnf_transformation,[],[f12081_D])).
fof(f12081_D,plain,(
  ( ! [X149,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))))) | s(bool,'T_0') = s(bool,X150) | s(bool,X149) = s(bool,X150)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f12021,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))))) | s(bool,X149) = s(bool,X150) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X147)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(duplicate_literal_removal,[],[f11847])).
fof(f11847,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X148))))) | s(bool,X149) = s(bool,X150) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,X147)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f11475])).
fof(f12080,plain,(
  ( ! [X142,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl252),
  inference(cnf_transformation,[],[f12080_D])).
fof(f12080_D,plain,(
  ( ! [X142,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f12022,plain,(
  ( ! [X142,X140,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X140)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(duplicate_literal_removal,[],[f11845])).
fof(f11845,plain,(
  ( ! [X142,X140,X141] : (s(bool,X142) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X141))))) | s(bool,X142) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f6091,f11475])).
fof(f12055,plain,(
  ( ! [X43,X41,X44] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f12054,f11496])).
fof(f12054,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42)))))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f11816,f396])).
fof(f11816,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X42)))))))) | s(bool,X44) = s(bool,sK1(X44,X43,bool)) | p(s(bool,X41)) | p(s(bool,X44)) | s(bool,'T_0') = s(bool,X41)) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f1730,f11475])).
fof(f11553,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl248),
  inference(cnf_transformation,[],[f11553_D])).
fof(f11553_D,plain,(
  ( ! [X19,X17] : (s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f11752,plain,(
  ( ! [X111,X112,X110] : (s(bool,X112) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X110))))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X111)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f11690,f259])).
fof(f11690,plain,(
  ( ! [X111,X112,X110] : (s(bool,X111) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X112) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X110))))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X111)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6249,f11496])).
fof(f11556,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33)))))) ) | $spl250),
  inference(cnf_transformation,[],[f11556_D])).
fof(f11556_D,plain,(
  ( ! [X33] : (s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33)))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f11555,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11543,f4107])).
fof(f11543,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))))) | s(bool,'T_0') = s(bool,X32)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f11243,f259])).
fof(f11554,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11539,f4107])).
fof(f11539,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f11243,f4721])).
fof(f11552,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11538,f4107])).
fof(f11538,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6)),
  inference(superposition,[],[f11243,f4721])).
fof(f11405,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v',X44)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11363,f4107])).
fof(f11363,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v',X44)))) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f2382,f259])).
fof(f11403,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26)))))) ) | $spl244),
  inference(cnf_transformation,[],[f11403_D])).
fof(f11403_D,plain,(
  ( ! [X26,X23,X25] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f11404,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11359,f4107])).
fof(f11359,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v',X29)))) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f2382,f4721])).
fof(f11402,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11358,f4107])).
fof(f11358,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v',X25)))) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2382,f4721])).
fof(f11399,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35)))))) ) | $spl242),
  inference(cnf_transformation,[],[f11399_D])).
fof(f11399_D,plain,(
  ( ! [X35,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32))) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f11400,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11350,f4107])).
fof(f11350,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X39))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X36),s('type.closSem.v',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl6),
  inference(superposition,[],[f2382,f4721])).
fof(f11398,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11349,f4107])).
fof(f11349,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X34) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X32),s('type.closSem.v',X32)))) ) | $spl6),
  inference(superposition,[],[f2382,f4721])).
fof(f11375,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X4)))))) ) | $spl234),
  inference(cnf_transformation,[],[f11375_D])).
fof(f11375_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X4)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f11396,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X2)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f11340,f4107])).
fof(f11340,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2382,f259])).
fof(f11244,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,X36) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X37))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))) )),
  inference(resolution,[],[f2382,f8638])).
fof(f11234,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X8))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))))) )),
  inference(resolution,[],[f2382,f3685])).
fof(f11233,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X6))))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))) )),
  inference(resolution,[],[f2382,f8750])).
fof(f11232,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X4))))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Unit_0')),'type.closSem.v'))) ) | $spl2),
  inference(resolution,[],[f2382,f5287])).
fof(f10856,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(sK1(X27,X26,X25),X26,X25))))) ) | $spl230),
  inference(cnf_transformation,[],[f10856_D])).
fof(f10856_D,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(sK1(X27,X26,X25),X26,X25))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f10857,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X32)) | ~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31)))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,sK1(sK1(X31,X30,X29),X30,X29))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10849,f4107])).
fof(f10849,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X32)) | ~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31)))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,sK1(sK1(X31,X30,X29),X30,X29)))) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f2423,f4721])).
fof(f2423,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(sK1(X2,X1,X0),X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(resolution,[],[f95,f95])).
fof(f10855,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X28)) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(sK1(X27,X26,X25),X26,X25))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10848,f4107])).
fof(f10848,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X28)) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(sK1(X27,X26,X25),X26,X25))))) ) | $spl6),
  inference(superposition,[],[f2423,f4721])).
fof(f10773,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10761,f4107])).
fof(f10761,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X44))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f2217,f259])).
fof(f10771,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))),s('type.closSem.v',X26))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25)))) ) | $spl226),
  inference(cnf_transformation,[],[f10771_D])).
fof(f10771_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))),s('type.closSem.v',X26))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f10772,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v',X30))) | s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10757,f4107])).
fof(f10757,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v',X30))) | s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f2217,f4721])).
fof(f10770,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))),s('type.closSem.v',X26))) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10756,f4107])).
fof(f10756,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))),s('type.closSem.v',X26))) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2217,f4721])).
fof(f10665,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10653,f4107])).
fof(f10653,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))),s('type.closSem.v',X45))) | s('type.closSem.v',X45) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f2079,f259])).
fof(f10663,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))),s('type.closSem.v',X26))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25)))) ) | $spl222),
  inference(cnf_transformation,[],[f10663_D])).
fof(f10663_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))),s('type.closSem.v',X26))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f10664,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v',X30))) | s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10649,f4107])).
fof(f10649,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v',X30))) | s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f2079,f4721])).
fof(f10662,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))),s('type.closSem.v',X26))) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10648,f4107])).
fof(f10648,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))),s('type.closSem.v',X26))) | s('type.closSem.v',X26) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f2079,f4721])).
fof(f10557,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10545,f4107])).
fof(f10545,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X45))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f1911,f259])).
fof(f10555,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))) ) | $spl218),
  inference(cnf_transformation,[],[f10555_D])).
fof(f10555_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f10556,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))))) | s('type.closSem.v',X29) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10541,f4107])).
fof(f10541,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))))) | s('type.closSem.v',X29) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f1911,f4721])).
fof(f10554,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10540,f4107])).
fof(f10540,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))))) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1911,f4721])).
fof(f10449,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10437,f4107])).
fof(f10437,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X44),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))))) | s('type.closSem.v',X44) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X45))) | s(bool,'T_0') = s(bool,X43)) ) | $spl6),
  inference(superposition,[],[f1773,f259])).
fof(f10447,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26)))) ) | $spl214),
  inference(cnf_transformation,[],[f10447_D])).
fof(f10447_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s(bool,'T_0') = s(bool,X23) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f10448,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))))) | s('type.closSem.v',X29) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10433,f4107])).
fof(f10433,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))))) | s('type.closSem.v',X29) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f1773,f4721])).
fof(f10446,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10432,f4107])).
fof(f10432,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))))) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1773,f4721])).
fof(f10355,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X19)) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK1(X19,X20,bool))) ) | $spl212),
  inference(cnf_transformation,[],[f10355_D])).
fof(f10355_D,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X19)) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK1(X19,X20,bool))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f10356,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,X33) = s(bool,sK1(X31,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl176)),
  inference(subsumption_resolution,[],[f10332,f8143])).
fof(f10332,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,X33) = s(bool,sK1(X31,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f1756,f4721])).
fof(f1756,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | s(bool,X20) = s(bool,sK1(X18,X19,bool)) | ~p(s(bool,X20)) | ~p(s(bool,X18))) )),
  inference(superposition,[],[f91,f94])).
fof(f10354,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,X21) = s(bool,sK1(X19,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl176)),
  inference(subsumption_resolution,[],[f10329,f8143])).
fof(f10329,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,X21) = s(bool,sK1(X19,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl6),
  inference(superposition,[],[f1756,f4721])).
fof(f10200,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10188,f4107])).
fof(f10188,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X33),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33)) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f8770,f259])).
fof(f10199,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X22),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X22)) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10184,f4107])).
fof(f10184,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X22),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X22)) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f8770,f4721])).
fof(f10197,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X19),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19)) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10183,f4107])).
fof(f10183,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X19),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19)) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | $spl6),
  inference(superposition,[],[f8770,f4721])).
fof(f10064,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X33))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10052,f4107])).
fof(f10052,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X33))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X33)) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f8671,f259])).
fof(f10063,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X22))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X22)) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10048,f4107])).
fof(f10048,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X22))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X22)) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f8671,f4721])).
fof(f10061,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X19))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19)) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f10047,f4107])).
fof(f10047,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X19))) | s('type.list.list'('type.closSem.v'),'const.list.NIL_0') = s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X19)) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | $spl6),
  inference(superposition,[],[f8671,f4721])).
fof(f9493,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK1(X21,X20,bool))) ) | $spl202),
  inference(cnf_transformation,[],[f9493_D])).
fof(f9493_D,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,sK1(X21,X20,bool))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f9494,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,X33) = s(bool,sK1(X33,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl176)),
  inference(subsumption_resolution,[],[f9470,f8143])).
fof(f9470,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,X33) = s(bool,sK1(X33,X32,bool)) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f1732,f4721])).
fof(f1732,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | s(bool,X10) = s(bool,sK1(X10,X12,bool)) | ~p(s(bool,X10)) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f94,f91])).
fof(f9492,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,X21) = s(bool,sK1(X21,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,X22)) ) | ($spl6 | $spl176)),
  inference(subsumption_resolution,[],[f9467,f8143])).
fof(f9467,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,X21) = s(bool,sK1(X21,X20,bool)) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl6),
  inference(superposition,[],[f1732,f4721])).
fof(f9374,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X26)) | s(bool,X26) = s(bool,sK1(X26,X25,bool))) ) | $spl200),
  inference(cnf_transformation,[],[f9374_D])).
fof(f9374_D,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X26)) | s(bool,X26) = s(bool,sK1(X26,X25,bool))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f9375,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | s(bool,X30) = s(bool,sK1(X30,X29,bool)) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f9348,f396])).
fof(f9348,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X28)))) | s(bool,X30) = s(bool,sK1(X30,X29,bool)) | p(s(bool,X27)) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | $spl6),
  inference(superposition,[],[f1730,f4721])).
fof(f9373,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,X26) = s(bool,sK1(X26,X25,bool)) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f9347,f396])).
fof(f9347,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,X26) = s(bool,sK1(X26,X25,bool)) | p(s(bool,X23)) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f1730,f4721])).
fof(f9371,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X18) = s(bool,sK1(X18,X17,bool)) | p(s(bool,X18)) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f9345,f396])).
fof(f9345,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))))) | s(bool,X18) = s(bool,sK1(X18,X17,bool)) | p(s(bool,X16)) | p(s(bool,X18)) | s(bool,'T_0') = s(bool,X16)) ) | $spl6),
  inference(superposition,[],[f1730,f259])).
fof(f9001,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))),X47,bool)))) ) | $spl194),
  inference(cnf_transformation,[],[f9001_D])).
fof(f9001_D,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))),X47,bool)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f9279,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X31) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29)))),X30,bool)))) ) | ($spl8 | $spl84)),
  inference(subsumption_resolution,[],[f9278,f396])).
fof(f9278,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X31) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | p(s(bool,X31)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29)))),X30,bool)))) ) | $spl84),
  inference(forward_demodulation,[],[f9223,f2028])).
fof(f9223,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | s(bool,X31) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))))) | p(s(bool,X31)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29)))),X30,bool)))) ) | $spl84),
  inference(superposition,[],[f1720,f2028])).
fof(f1720,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X3)))) | s(bool,X3) = s(bool,X5) | p(s(bool,X5)) | p(s(bool,sK1(X3,X4,bool)))) )),
  inference(superposition,[],[f94,f88])).
fof(f8998,plain,(
  ( ! [X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43)))),X44,bool)))) ) | $spl192),
  inference(cnf_transformation,[],[f8998_D])).
fof(f8998_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43)))),X44,bool)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f9169,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X31) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29)))),X30,bool)))) ) | ($spl8 | $spl82)),
  inference(subsumption_resolution,[],[f9168,f396])).
fof(f9168,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X31) | ~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | p(s(bool,X31)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29)))),X30,bool)))) ) | $spl82),
  inference(forward_demodulation,[],[f9113,f1863])).
fof(f9113,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) | s(bool,X31) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))))) | p(s(bool,X31)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X29)))),X30,bool)))) ) | $spl82),
  inference(superposition,[],[f1720,f1863])).
fof(f2028,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10)))))) ) | $spl84),
  inference(cnf_transformation,[],[f2028_D])).
fof(f2028_D,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9077,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X13))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X13)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f9071,f4107])).
fof(f9071,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X13))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X13))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f1999,f4721])).
fof(f9076,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f9070,f4107])).
fof(f9070,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X11)))))) ) | $spl6),
  inference(superposition,[],[f1999,f4721])).
fof(f1863,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10)))))) ) | $spl82),
  inference(cnf_transformation,[],[f1863_D])).
fof(f1863_D,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9041,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X13))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X13)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f9035,f4107])).
fof(f9035,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X13))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X13))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f1836,f4721])).
fof(f9040,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X11))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X11)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f9034,f4107])).
fof(f9034,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X11))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X11)))))) ) | $spl6),
  inference(superposition,[],[f1836,f4721])).
fof(f9005,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | p(s(bool,sK1('const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0')),X59,bool)))) ) | $spl196),
  inference(cnf_transformation,[],[f9005_D])).
fof(f9005_D,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | p(s(bool,sK1('const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0')),X59,bool)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f9004,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X60) | ~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | p(s(bool,sK1('const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0')),X59,bool)))) ) | ($spl8 | $spl86)),
  inference(subsumption_resolution,[],[f9003,f396])).
fof(f9003,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,X60) | ~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | p(s(bool,X60)) | p(s(bool,sK1('const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0')),X59,bool)))) ) | $spl86),
  inference(forward_demodulation,[],[f8975,f2717])).
fof(f8975,plain,(
  ( ! [X59,X57,X60,X58,X56] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | s(bool,X60) = s(bool,'const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0'))) | p(s(bool,X60)) | p(s(bool,sK1('const.list.LIST_REL_3'(s(fun(X56,fun(X57,bool)),X58),s('type.list.list'(X56),'const.list.NIL_0'),s('type.list.list'(X57),'const.list.NIL_0')),X59,bool)))) ) | $spl86),
  inference(superposition,[],[f1720,f2717])).
fof(f9000,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X48) | ~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))),X47,bool)))) ) | ($spl8 | $spl84)),
  inference(subsumption_resolution,[],[f8999,f396])).
fof(f8999,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X48) | ~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) | p(s(bool,X48)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))),X47,bool)))) ) | $spl84),
  inference(forward_demodulation,[],[f8972,f2028])).
fof(f8972,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))) | s(bool,X48) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))))) | p(s(bool,X48)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X46)))),X47,bool)))) ) | $spl84),
  inference(superposition,[],[f1720,f2028])).
fof(f8997,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,X45) | ~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43)))),X44,bool)))) ) | ($spl8 | $spl82)),
  inference(subsumption_resolution,[],[f8996,f396])).
fof(f8996,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,X45) | ~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))) | p(s(bool,X45)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43)))),X44,bool)))) ) | $spl82),
  inference(forward_demodulation,[],[f8971,f1863])).
fof(f8971,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))) | s(bool,X45) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))))) | p(s(bool,X45)) | p(s(bool,sK1('const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X43)))),X44,bool)))) ) | $spl82),
  inference(superposition,[],[f1720,f1863])).
fof(f8995,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X40),s(bool,X39)))) | s(bool,'T_0') = s(bool,X41) | p(s(bool,sK1('T_0',X40,bool))) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f8970,f396])).
fof(f8970,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X40),s(bool,X39)))) | s(bool,'T_0') = s(bool,X41) | p(s(bool,X41)) | p(s(bool,sK1('T_0',X40,bool))) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | $spl6),
  inference(superposition,[],[f1720,f4721])).
fof(f8993,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | s(bool,'T_0') = s(bool,X37) | p(s(bool,sK1('T_0',X36,bool))) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f8969,f396])).
fof(f8969,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,happ(s(fun(bool,bool),X36),s(bool,X35)))) | s(bool,'T_0') = s(bool,X37) | p(s(bool,X37)) | p(s(bool,sK1('T_0',X36,bool))) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl6),
  inference(superposition,[],[f1720,f4721])).
fof(f8991,plain,(
  ( ! [X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X22) | p(s(bool,sK1(X19,X20,bool)))) ) | $spl188),
  inference(cnf_transformation,[],[f8991_D])).
fof(f8991_D,plain,(
  ( ! [X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X22) | p(s(bool,sK1(X19,X20,bool)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f8992,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,X33)) | p(s(bool,sK1(X31,X32,bool))) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8968,f4936])).
fof(f8968,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | s(bool,X31) = s(bool,X33) | p(s(bool,X33)) | p(s(bool,sK1(X31,X32,bool))) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f1720,f4721])).
fof(f8990,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | p(s(bool,X21)) | p(s(bool,sK1(X19,X20,bool))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8965,f4936])).
fof(f8965,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | s(bool,X19) = s(bool,X21) | p(s(bool,X21)) | p(s(bool,sK1(X19,X20,bool))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | $spl6),
  inference(superposition,[],[f1720,f4721])).
fof(f8889,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X37),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X37,'type.closSem.v'))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8879,f4107])).
fof(f8879,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun('type.closSem.v',bool),X37),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X37,'type.closSem.v')) | s(bool,'T_0') = s(bool,X36)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f5284,f259])).
fof(f8888,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,happ(s(fun('type.closSem.v',bool),X22),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X22,'type.closSem.v')) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8873,f4107])).
fof(f8873,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X20) = s(bool,happ(s(fun('type.closSem.v',bool),X22),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X22,'type.closSem.v')) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f5284,f4721])).
fof(f8886,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,happ(s(fun('type.closSem.v',bool),X19),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X19,'type.closSem.v')) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8872,f4107])).
fof(f8872,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X17) = s(bool,happ(s(fun('type.closSem.v',bool),X19),s('type.closSem.v','const.closSem.Unit_0'))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X19,'type.closSem.v')) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f5284,f4721])).
fof(f8451,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X63,bool),X64),s(X63,X65))) | s(X63,X65) = s(X63,sK1(X65,X64,X63))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8441,f4107])).
fof(f8441,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) = s(bool,happ(s(fun(X63,bool),X64),s(X63,X65))) | s(X63,X65) = s(X63,sK1(X65,X64,X63)) | s(bool,'T_0') = s(bool,X62)) ) | $spl6),
  inference(superposition,[],[f1718,f259])).
fof(f8450,plain,(
  ( ! [X37,X35,X38,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(X36,X38) = s(X36,sK1(X38,X37,X36)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8435,f4107])).
fof(f8435,plain,(
  ( ! [X37,X35,X38,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,happ(s(fun(X36,bool),X37),s(X36,X38))) | s(X36,X38) = s(X36,sK1(X38,X37,X36)) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | $spl6),
  inference(superposition,[],[f1718,f4721])).
fof(f8448,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK1(X33,X32,X31)) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f8434,f4107])).
fof(f8434,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | s(bool,X29) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(X31,X33) = s(X31,sK1(X33,X32,X31)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | $spl6),
  inference(superposition,[],[f1718,f4721])).
fof(f8331,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl178),
  inference(cnf_transformation,[],[f8331_D])).
fof(f8331_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f8332,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl176)),
  inference(subsumption_resolution,[],[f8305,f8143])).
fof(f8305,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f6257,f91])).
fof(f6257,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(resolution,[],[f4936,f175])).
fof(f8330,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl176)),
  inference(subsumption_resolution,[],[f8304,f8143])).
fof(f8304,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f6257,f91])).
fof(f397,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f397_D])).
fof(f397_D,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f8143,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl176),
  inference(cnf_transformation,[],[f8143_D])).
fof(f8143_D,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f8144,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,X11) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f8105,f4043])).
fof(f8105,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,X11) | s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f6091,f91])).
fof(f8142,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f8104,f4043])).
fof(f8104,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(superposition,[],[f6091,f91])).
fof(f7928,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35)) ) | $spl134),
  inference(cnf_transformation,[],[f7928_D])).
fof(f7928_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7839,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(equality_factoring,[],[f670])).
fof(f7843,plain,(
  ( ! [X136,X137] : (s(bool,'T_0') = s(bool,X136) | s(bool,X137) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f7795])).
fof(f7795,plain,(
  ( ! [X136,X137] : (s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,X136) | s(bool,X137) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f259,f670])).
fof(f7980,plain,(
  ( ! [X88,X85] : (~p(s(bool,X85)) | s(bool,X85) = s(bool,X88)) ) | $spl174),
  inference(cnf_transformation,[],[f7980_D])).
fof(f7980_D,plain,(
  ( ! [X88,X85] : (~p(s(bool,X85)) | s(bool,X85) = s(bool,X88)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7447,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | p(s(bool,sK1(X21,X20,bool)))) ) | $spl106),
  inference(cnf_transformation,[],[f7447_D])).
fof(f7447_D,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | ~p(s(bool,X21)) | p(s(bool,sK1(X21,X20,bool)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7775,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,'T_0')))) | p(s(bool,sK1(X87,X86,bool))) | ~p(s(bool,X87)) | ~p(s(bool,X85)) | s(bool,X85) = s(bool,X88)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(superposition,[],[f6960,f670])).
fof(f6960,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | p(s(bool,sK1(X6,X8,bool))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f6932,f396])).
fof(f6932,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | p(s(bool,sK1(X6,X8,bool))) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4118,f91])).
fof(f4118,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | p(s(bool,sK1(X10,X11,bool))) | s(bool,'T_0') != s(bool,X10)) ) | ($spl1 | $spl2 | $spl15)),
  inference(superposition,[],[f4107,f94])).
fof(f7877,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1))) ) | $spl110),
  inference(cnf_transformation,[],[f7877_D])).
fof(f7877_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7978,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X48)) | s(bool,X49) = s(bool,X50) | p(s(bool,X50)) | s(bool,X48) = s(bool,X51)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7765,f4239])).
fof(f7765,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,X49) | p(s(bool,X48)) | s(bool,X49) = s(bool,X50) | p(s(bool,X50)) | s(bool,X48) = s(bool,X51)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4300,f670])).
fof(f7973,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7748,f4107])).
fof(f7748,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7972,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X2)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7747,f4107])).
fof(f7747,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X2)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7970,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X101)) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X103) | s(bool,X101) = s(bool,X104)) ) | ($spl6 | $spl22 | $spl96)),
  inference(subsumption_resolution,[],[f7739,f4721])).
fof(f7739,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X101)) | s(bool,'T_0') = s(bool,X102) | s(bool,X102) != s(bool,X103) | s(bool,'T_0') = s(bool,X103) | s(bool,X101) = s(bool,X104)) ) | ($spl22 | $spl96)),
  inference(superposition,[],[f5453,f670])).
fof(f7969,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | s(bool,'T_0') = s(bool,X99) | s(bool,X98) = s(bool,X100)) ) | ($spl22 | $spl94)),
  inference(subsumption_resolution,[],[f7738,f670])).
fof(f7738,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | s(bool,X99) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X99) | s(bool,X98) = s(bool,X100)) ) | ($spl22 | $spl94)),
  inference(superposition,[],[f5451,f670])).
fof(f7967,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,X78)) | s(bool,X79) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X78) = s(bool,X80)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7732,f259])).
fof(f7732,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,X78)) | s(bool,X79) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') != s(bool,X79) | s(bool,X78) = s(bool,X80)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4394,f670])).
fof(f7966,plain,(
  ( ! [X74,X77,X75] : (s(bool,X74) != s(bool,X75) | s(bool,X74) = s(bool,X77) | p(s(bool,X75))) ) | $spl172),
  inference(cnf_transformation,[],[f7966_D])).
fof(f7966_D,plain,(
  ( ! [X74,X77,X75] : (s(bool,X74) != s(bool,X75) | s(bool,X74) = s(bool,X77) | p(s(bool,X75))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7965,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X74) != s(bool,X75) | p(s(bool,X75)) | p(s(bool,X76)) | s(bool,X74) = s(bool,X77)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7731,f4107])).
fof(f7731,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X74) != s(bool,X75) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X76) | p(s(bool,X76)) | s(bool,X74) = s(bool,X77)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4300,f670])).
fof(f7847,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,X66)) ) | ($spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f7728])).
fof(f7728,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,X64)) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,X64) = s(bool,X66)) ) | ($spl8 | $spl22)),
  inference(superposition,[],[f4239,f670])).
fof(f7960,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) = s(bool,X35)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7716,f4107])).
fof(f7716,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,X34) | p(s(bool,X34)) | s(bool,X33) = s(bool,X35)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7959,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) = s(bool,X32)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7715,f4107])).
fof(f7715,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X31)) | s(bool,X30) = s(bool,X32)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7706,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f105,f670])).
fof(f7705,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X2)) ) | $spl22),
  inference(superposition,[],[f103,f670])).
fof(f7923,plain,(
  ( ! [X261,X260,X258] : (p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X260)))) | ~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X258))))) ) | $spl132),
  inference(cnf_transformation,[],[f7923_D])).
fof(f7923_D,plain,(
  ( ! [X261,X260,X258] : (p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X260)))) | ~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X258))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7958,plain,(
  ( ! [X271,X273,X270,X272] : (p(s(bool,happ(s(fun(bool,bool),X272),s(bool,X271)))) | ~p(s(bool,happ(s(fun(bool,bool),X272),s(bool,X273)))) | s(bool,'T_0') = s(bool,X270)) ) | $spl22),
  inference(subsumption_resolution,[],[f7702,f670])).
fof(f7702,plain,(
  ( ! [X271,X273,X270,X272] : (p(s(bool,happ(s(fun(bool,bool),X272),s(bool,X271)))) | ~p(s(bool,happ(s(fun(bool,bool),X272),s(bool,X273)))) | s(bool,X270) != s(bool,X273) | s(bool,'T_0') = s(bool,X270)) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f7957,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,happ(s(fun(bool,bool),X268),s(bool,X267)))) | p(s(bool,happ(s(fun(bool,bool),X268),s(bool,X269)))) | s(bool,'T_0') = s(bool,X266)) ) | $spl22),
  inference(subsumption_resolution,[],[f7701,f670])).
fof(f7701,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,happ(s(fun(bool,bool),X268),s(bool,X267)))) | p(s(bool,happ(s(fun(bool,bool),X268),s(bool,X269)))) | s(bool,X266) != s(bool,X269) | s(bool,'T_0') = s(bool,X266)) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f713,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X104),s('type.list.list'('type.closSem.v'),X103)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X103,X104,X105)),s('type.list.list'('type.closSem.v'),X103)))) ) | $spl78),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X104),s('type.list.list'('type.closSem.v'),X103)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X103,X104,X105)),s('type.list.list'('type.closSem.v'),X103)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7696,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X251)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X250),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X249),s('type.list.list'('type.closSem.v'),X248)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X248,X249,X250)),s('type.list.list'('type.closSem.v'),X248)))) ) | $spl22),
  inference(superposition,[],[f4051,f670])).
fof(f709,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X96))),s('type.closSem.v',X99)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X96),s('type.list.list'('type.closSem.v'),X97))) | s('type.closSem.v',X99) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X97)))) ) | $spl74),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X96))),s('type.closSem.v',X99)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X96),s('type.list.list'('type.closSem.v'),X97))) | s('type.closSem.v',X99) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X97)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7695,plain,(
  ( ! [X243,X246,X244,X247,X245] : (~p(s(bool,X245)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X246),s('type.list.list'('type.closSem.v'),X243))),s('type.closSem.v',X247)))) | s('type.closSem.v',X247) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X246),s('type.list.list'('type.closSem.v'),X244))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X243),s('type.list.list'('type.closSem.v'),X244)))) ) | $spl22),
  inference(superposition,[],[f6970,f670])).
fof(f707,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X94),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X93)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X92),s('type.list.list'('type.closSem.v'),X93))) | s('type.closSem.v',X94) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X92)))) ) | $spl72),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X94),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X93)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X92),s('type.list.list'('type.closSem.v'),X93))) | s('type.closSem.v',X94) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X92)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7694,plain,(
  ( ! [X241,X239,X242,X240,X238] : (~p(s(bool,X240)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X241),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X242),s('type.list.list'('type.closSem.v'),X239)))))) | s('type.closSem.v',X241) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X242),s('type.list.list'('type.closSem.v'),X238))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X238),s('type.list.list'('type.closSem.v'),X239)))) ) | $spl22),
  inference(superposition,[],[f7256,f670])).
fof(f7256,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3)))) )),
  inference(forward_demodulation,[],[f123,f92])).
fof(f123,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X0))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X3)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f711,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X101),s('type.list.list'('type.closSem.v'),X100))),s('type.closSem.v',X102)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X100),s('type.list.list'('type.closSem.v'),sK2(X100,X101,X102))))) ) | $spl76),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X101),s('type.list.list'('type.closSem.v'),X100))),s('type.closSem.v',X102)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X100),s('type.list.list'('type.closSem.v'),sK2(X100,X101,X102))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7693,plain,(
  ( ! [X237,X235,X236,X234] : (p(s(bool,X237)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X234))),s('type.closSem.v',X236)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X234),s('type.list.list'('type.closSem.v'),sK2(X234,X235,X236))))) ) | $spl22),
  inference(superposition,[],[f3923,f670])).
fof(f704,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.list.list'(X85),'const.list.NIL_0') = s('type.list.list'(X85),X87) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X84,fun(X85,bool)),X86),s('type.list.list'(X84),'const.list.NIL_0'),s('type.list.list'(X85),X87)))) ) | $spl68),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.list.list'(X85),'const.list.NIL_0') = s('type.list.list'(X85),X87) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X84,fun(X85,bool)),X86),s('type.list.list'(X84),'const.list.NIL_0'),s('type.list.list'(X85),X87)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7692,plain,(
  ( ! [X231,X229,X233,X230,X232] : (~p(s(bool,X233)) | s('type.list.list'(X230),'const.list.NIL_0') = s('type.list.list'(X230),X232) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X229,fun(X230,bool)),X231),s('type.list.list'(X229),'const.list.NIL_0'),s('type.list.list'(X230),X232)))) ) | $spl22),
  inference(superposition,[],[f86,f670])).
fof(f705,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.list.list'(X89),'const.list.NIL_0') != s('type.list.list'(X89),X91) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X88,fun(X89,bool)),X90),s('type.list.list'(X88),'const.list.NIL_0'),s('type.list.list'(X89),X91)))) ) | $spl70),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.list.list'(X89),'const.list.NIL_0') != s('type.list.list'(X89),X91) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X88,fun(X89,bool)),X90),s('type.list.list'(X88),'const.list.NIL_0'),s('type.list.list'(X89),X91)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7691,plain,(
  ( ! [X227,X225,X228,X226,X224] : (p(s(bool,X228)) | s('type.list.list'(X225),'const.list.NIL_0') != s('type.list.list'(X225),X227) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X224,fun(X225,bool)),X226),s('type.list.list'(X224),'const.list.NIL_0'),s('type.list.list'(X225),X227)))) ) | $spl22),
  inference(superposition,[],[f87,f670])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),X3)))) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) )),
  inference(cnf_transformation,[],[f61])).
fof(f702,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.list.list'(X76),'const.list.NIL_0') = s('type.list.list'(X76),X79) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X76,fun(X77,bool)),X78),s('type.list.list'(X76),X79),s('type.list.list'(X77),'const.list.NIL_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.list.list'(X76),'const.list.NIL_0') = s('type.list.list'(X76),X79) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X76,fun(X77,bool)),X78),s('type.list.list'(X76),X79),s('type.list.list'(X77),'const.list.NIL_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7690,plain,(
  ( ! [X222,X220,X219,X223,X221] : (~p(s(bool,X223)) | s('type.list.list'(X219),'const.list.NIL_0') = s('type.list.list'(X219),X222) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X219,fun(X220,bool)),X221),s('type.list.list'(X219),X222),s('type.list.list'(X220),'const.list.NIL_0')))) ) | $spl22),
  inference(superposition,[],[f84,f670])).
fof(f703,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.list.list'(X80),'const.list.NIL_0') != s('type.list.list'(X80),X83) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X80,fun(X81,bool)),X82),s('type.list.list'(X80),X83),s('type.list.list'(X81),'const.list.NIL_0')))) ) | $spl66),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.list.list'(X80),'const.list.NIL_0') != s('type.list.list'(X80),X83) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X80,fun(X81,bool)),X82),s('type.list.list'(X80),X83),s('type.list.list'(X81),'const.list.NIL_0')))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7689,plain,(
  ( ! [X218,X216,X215,X217,X214] : (p(s(bool,X218)) | s('type.list.list'(X214),'const.list.NIL_0') != s('type.list.list'(X214),X217) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X214,fun(X215,bool)),X216),s('type.list.list'(X214),X217),s('type.list.list'(X215),'const.list.NIL_0')))) ) | $spl22),
  inference(superposition,[],[f85,f670])).
fof(f85,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) )),
  inference(cnf_transformation,[],[f60])).
fof(f697,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun(X59,bool),X60),s(X59,X61)))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,sK1(X61,X60,X59))))) ) | $spl54),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun(X59,bool),X60),s(X59,X61)))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,sK1(X61,X60,X59))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7687,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X209)) | ~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,X208)))) | s(bool,'T_0') = s(bool,happ(s(fun(X206,bool),X207),s(X206,sK1(X208,X207,X206))))) ) | $spl22),
  inference(superposition,[],[f95,f670])).
fof(f7956,plain,(
  ( ! [X204] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X204,'type.closSem.v')) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X204),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl170),
  inference(cnf_transformation,[],[f7956_D])).
fof(f7956_D,plain,(
  ( ! [X204] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X204,'type.closSem.v')) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X204),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7686,plain,(
  ( ! [X204,X205] : (~p(s(bool,X205)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK1(sK0,X204,'type.closSem.v')) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X204),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl2 | $spl22)),
  inference(superposition,[],[f4041,f670])).
fof(f7955,plain,(
  ( ! [X202] : (p(s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v',sK1(sK0,X202,'type.closSem.v'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl168),
  inference(cnf_transformation,[],[f7955_D])).
fof(f7955_D,plain,(
  ( ! [X202] : (p(s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v',sK1(sK0,X202,'type.closSem.v'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7685,plain,(
  ( ! [X202,X203] : (~p(s(bool,X203)) | p(s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v',sK1(sK0,X202,'type.closSem.v'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),X202),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl2 | $spl22)),
  inference(superposition,[],[f4018,f670])).
fof(f4018,plain,(
  ( ! [X38] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),X38),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun('type.closSem.v',bool),X38),s('type.closSem.v',sK1(sK0,X38,'type.closSem.v')))))) ) | $spl2),
  inference(superposition,[],[f95,f128])).
fof(f700,plain,(
  ( ! [X70,X72,X71] : (s(X70,X72) = s(X70,sK1(X72,X71,X70)) | s(bool,'T_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,X72)))) ) | $spl60),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  ( ! [X70,X72,X71] : (s(X70,X72) = s(X70,sK1(X72,X71,X70)) | s(bool,'T_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,X72)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7684,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,X201)) | s(X198,X200) = s(X198,sK1(X200,X199,X198)) | s(bool,'T_0') = s(bool,happ(s(fun(X198,bool),X199),s(X198,X200)))) ) | $spl22),
  inference(superposition,[],[f94,f670])).
fof(f701,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,happ(s(fun(X73,bool),X74),s(X73,sK1(X75,X74,X73))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) ) | $spl62),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,happ(s(fun(X73,bool),X74),s(X73,sK1(X75,X74,X73))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7683,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,X197)) | p(s(bool,happ(s(fun(X194,bool),X195),s(X194,sK1(X196,X195,X194))))) | s(bool,'T_0') = s(bool,happ(s(fun(X194,bool),X195),s(X194,X196)))) ) | $spl22),
  inference(superposition,[],[f95,f670])).
fof(f699,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,X69)))) | s(bool,'T_0') = s(bool,happ(s(fun(X66,bool),X67),s(X66,X68))) | s(X66,X68) != s(X66,X69)) ) | $spl58),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,X69)))) | s(bool,'T_0') = s(bool,happ(s(fun(X66,bool),X67),s(X66,X68))) | s(X66,X68) != s(X66,X69)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7682,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X192)) | ~p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X193)))) | s(X189,X191) != s(X189,X193) | s(bool,'T_0') = s(bool,happ(s(fun(X189,bool),X190),s(X189,X191)))) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f698,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,happ(s(fun(X62,bool),X63),s(X62,X65)))) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s(X62,X64) != s(X62,X65)) ) | $spl56),
  inference(cnf_transformation,[],[f698_D])).
fof(f698_D,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,happ(s(fun(X62,bool),X63),s(X62,X65)))) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s(X62,X64) != s(X62,X65)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7681,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,X187)) | p(s(bool,happ(s(fun(X184,bool),X185),s(X184,X188)))) | s(X184,X186) != s(X184,X188) | s(bool,'T_0') = s(bool,happ(s(fun(X184,bool),X185),s(X184,X186)))) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f7954,plain,(
  ( ! [X182,X181] : (p(s(bool,sK1(X182,X181,bool))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X181),s(bool,X182))) | s(bool,'T_0') != s(bool,X182)) ) | $spl166),
  inference(cnf_transformation,[],[f7954_D])).
fof(f7954_D,plain,(
  ( ! [X182,X181] : (p(s(bool,sK1(X182,X181,bool))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X181),s(bool,X182))) | s(bool,'T_0') != s(bool,X182)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7680,plain,(
  ( ! [X182,X183,X181] : (~p(s(bool,X183)) | p(s(bool,sK1(X182,X181,bool))) | s(bool,'T_0') != s(bool,X182) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X181),s(bool,X182)))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4118,f670])).
fof(f7953,plain,(
  ( ! [X178,X179] : (s(bool,'T_0') = s(bool,sK1(X179,X178,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X178),s(bool,X179))) | ~p(s(bool,X179))) ) | $spl164),
  inference(cnf_transformation,[],[f7953_D])).
fof(f7953_D,plain,(
  ( ! [X178,X179] : (s(bool,'T_0') = s(bool,sK1(X179,X178,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X178),s(bool,X179))) | ~p(s(bool,X179))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f7679,plain,(
  ( ! [X180,X178,X179] : (~p(s(bool,X180)) | s(bool,'T_0') = s(bool,sK1(X179,X178,bool)) | ~p(s(bool,X179)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X178),s(bool,X179)))) ) | ($spl8 | $spl22)),
  inference(superposition,[],[f4284,f670])).
fof(f7952,plain,(
  ( ! [X177,X175,X174] : (p(s(bool,sK1(X177,X174,bool))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X174),s(bool,X175))) | ~p(s(bool,X175)) | ~p(s(bool,X177))) ) | $spl162),
  inference(cnf_transformation,[],[f7952_D])).
fof(f7952_D,plain,(
  ( ! [X177,X175,X174] : (p(s(bool,sK1(X177,X174,bool))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X174),s(bool,X175))) | ~p(s(bool,X175)) | ~p(s(bool,X177))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7678,plain,(
  ( ! [X177,X175,X176,X174] : (~p(s(bool,X176)) | p(s(bool,sK1(X177,X174,bool))) | ~p(s(bool,X177)) | ~p(s(bool,X175)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X174),s(bool,X175)))) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(superposition,[],[f6960,f670])).
fof(f7951,plain,(
  ( ! [X169,X168] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X168))),s('type.closSem.v',X169)))) ) | $spl160),
  inference(cnf_transformation,[],[f7951_D])).
fof(f7951_D,plain,(
  ( ! [X169,X168] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X168))),s('type.closSem.v',X169)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f7950,plain,(
  ( ! [X169,X170,X168] : (~p(s(bool,X170)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X168))),s('type.closSem.v',X169)))) ) | ($spl22 | $spl40)),
  inference(subsumption_resolution,[],[f7676,f684])).
fof(f7676,plain,(
  ( ! [X169,X170,X168] : (~p(s(bool,X170)) | s('type.closSem.v',X169) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X168))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X168))),s('type.closSem.v',X169)))) ) | $spl22),
  inference(superposition,[],[f113,f670])).
fof(f7949,plain,(
  ( ! [X163,X162] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X163),s('type.closSem.v',X163))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v',X163)))) ) | $spl158),
  inference(cnf_transformation,[],[f7949_D])).
fof(f7949_D,plain,(
  ( ! [X163,X162] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X163),s('type.closSem.v',X163))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v',X163)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f7674,plain,(
  ( ! [X163,X164,X162] : (~p(s(bool,X164)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X163),s('type.closSem.v',X163))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X162))),s('type.closSem.v',X163)))) ) | ($spl22 | $spl82)),
  inference(superposition,[],[f5251,f670])).
fof(f5251,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X12))),s('type.closSem.v',X13)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v',X13)))) ) | $spl82),
  inference(superposition,[],[f1863,f113])).
fof(f7948,plain,(
  ( ! [X158,X157] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X157))),s('type.closSem.v',X158)))) ) | $spl156),
  inference(cnf_transformation,[],[f7948_D])).
fof(f7948_D,plain,(
  ( ! [X158,X157] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X157))),s('type.closSem.v',X158)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f7947,plain,(
  ( ! [X158,X159,X157] : (~p(s(bool,X159)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X157))),s('type.closSem.v',X158)))) ) | ($spl22 | $spl36)),
  inference(subsumption_resolution,[],[f7672,f682])).
fof(f7672,plain,(
  ( ! [X158,X159,X157] : (~p(s(bool,X159)) | s('type.closSem.v',X158) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X157))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X157))),s('type.closSem.v',X158)))) ) | $spl22),
  inference(superposition,[],[f115,f670])).
fof(f7946,plain,(
  ( ! [X152,X151] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X152),s('type.closSem.v',X152))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X151))),s('type.closSem.v',X152)))) ) | $spl154),
  inference(cnf_transformation,[],[f7946_D])).
fof(f7946_D,plain,(
  ( ! [X152,X151] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X152),s('type.closSem.v',X152))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X151))),s('type.closSem.v',X152)))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7670,plain,(
  ( ! [X152,X151,X153] : (~p(s(bool,X153)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X152),s('type.closSem.v',X152))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X151))),s('type.closSem.v',X152)))) ) | ($spl22 | $spl84)),
  inference(superposition,[],[f5602,f670])).
fof(f5602,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12))),s('type.closSem.v',X13)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X13),s('type.closSem.v',X13)))) ) | $spl84),
  inference(superposition,[],[f2028,f115])).
fof(f686,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.v',X41))) | s('type.closSem.v',X41) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),sK2(X40,X39,X41))))) ) | $spl42),
  inference(cnf_transformation,[],[f686_D])).
fof(f686_D,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.v',X41))) | s('type.closSem.v',X41) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),sK2(X40,X39,X41))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7668,plain,(
  ( ! [X146,X147,X145,X148] : (~p(s(bool,X148)) | s('type.closSem.v',X147) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X145),s('type.list.list'('type.closSem.v'),sK2(X146,X145,X147)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X145),s('type.list.list'('type.closSem.v'),X146))),s('type.closSem.v',X147)))) ) | $spl22),
  inference(superposition,[],[f3268,f670])).
fof(f3268,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0))))) )),
  inference(forward_demodulation,[],[f118,f92])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),sK2(X2,X1,X0)))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.closSem.v'),X2))))),s('type.closSem.v',X0))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f7943,plain,(
  ( ! [X139] : (s('type.closSem.v',X139) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X139)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X139)))) ) | $spl152),
  inference(cnf_transformation,[],[f7943_D])).
fof(f7943_D,plain,(
  ( ! [X139] : (s('type.closSem.v',X139) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X139)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X139)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f7664,plain,(
  ( ! [X140,X139] : (~p(s(bool,X140)) | s('type.closSem.v',X139) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X139)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X139)))) ) | $spl22),
  inference(superposition,[],[f3429,f670])).
fof(f3429,plain,(
  ( ! [X0] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X0))))) )),
  inference(superposition,[],[f3268,f117])).
fof(f7942,plain,(
  ( ! [X136,X137] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X136),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X137)))))) ) | $spl150),
  inference(cnf_transformation,[],[f7942_D])).
fof(f7942_D,plain,(
  ( ! [X136,X137] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X136),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X137)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f7941,plain,(
  ( ! [X138,X136,X137] : (~p(s(bool,X138)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X136),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X137)))))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f7663,f678])).
fof(f7663,plain,(
  ( ! [X138,X136,X137] : (~p(s(bool,X138)) | s('type.closSem.v',X136) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X137))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X136),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X137)))))) ) | $spl22),
  inference(superposition,[],[f109,f670])).
fof(f7940,plain,(
  ( ! [X130,X131] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v',X130)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X131)))))) ) | $spl148),
  inference(cnf_transformation,[],[f7940_D])).
fof(f7940_D,plain,(
  ( ! [X130,X131] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v',X130)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X131)))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f7661,plain,(
  ( ! [X132,X130,X131] : (~p(s(bool,X132)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v',X130)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X130),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X131)))))) ) | $spl22),
  inference(superposition,[],[f1850,f670])).
fof(f7939,plain,(
  ( ! [X127,X128] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v',X127))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X128)))))) ) | $spl146),
  inference(cnf_transformation,[],[f7939_D])).
fof(f7939_D,plain,(
  ( ! [X127,X128] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v',X127))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X128)))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7660,plain,(
  ( ! [X127,X128,X129] : (~p(s(bool,X129)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v',X127))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X127),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X128)))))) ) | ($spl22 | $spl82)),
  inference(superposition,[],[f4581,f670])).
fof(f4581,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X0)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1)))) ) | $spl82),
  inference(superposition,[],[f1863,f109])).
fof(f7938,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X124),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X125)))))) ) | $spl144),
  inference(cnf_transformation,[],[f7938_D])).
fof(f7938_D,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X124),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X125)))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7937,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X124),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X125)))))) ) | ($spl22 | $spl32)),
  inference(subsumption_resolution,[],[f7659,f680])).
fof(f7659,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,X126)) | s('type.closSem.v',X124) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X125))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X124),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X125)))))) ) | $spl22),
  inference(superposition,[],[f111,f670])).
fof(f7936,plain,(
  ( ! [X118,X119] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v',X118)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X119)))))) ) | $spl142),
  inference(cnf_transformation,[],[f7936_D])).
fof(f7936_D,plain,(
  ( ! [X118,X119] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v',X118)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X119)))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7657,plain,(
  ( ! [X118,X120,X119] : (~p(s(bool,X120)) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v',X118)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X118),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X119)))))) ) | $spl22),
  inference(superposition,[],[f2013,f670])).
fof(f7935,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116)))))) ) | $spl140),
  inference(cnf_transformation,[],[f7935_D])).
fof(f7935_D,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116)))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7656,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,X117)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v',X115))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X115),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X116)))))) ) | ($spl22 | $spl84)),
  inference(superposition,[],[f4626,f670])).
fof(f4626,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X0)))))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X1),s('type.closSem.v',X1)))) ) | $spl84),
  inference(superposition,[],[f2028,f111])).
fof(f692,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),X51))))) | s('type.closSem.v',X49) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),sK3(X51,X50,X49))))) ) | $spl48),
  inference(cnf_transformation,[],[f692_D])).
fof(f692_D,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),X51))))) | s('type.closSem.v',X49) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),sK3(X51,X50,X49))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7655,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X114)) | s('type.closSem.v',X111) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X112),s('type.list.list'('type.closSem.v'),sK3(X113,X112,X111)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X111),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X112),s('type.list.list'('type.closSem.v'),X113)))))) ) | $spl22),
  inference(superposition,[],[f3494,f670])).
fof(f7934,plain,(
  ( ! [X109] : (s('type.closSem.v',X109) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X109)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X109),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl138),
  inference(cnf_transformation,[],[f7934_D])).
fof(f7934_D,plain,(
  ( ! [X109] : (s('type.closSem.v',X109) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X109)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X109),s('type.closSem.v','const.closSem.Unit_0')))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7654,plain,(
  ( ! [X109,X110] : (~p(s(bool,X110)) | s('type.closSem.v',X109) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK3('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',X109)))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X109),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl22),
  inference(superposition,[],[f3685,f670])).
fof(f7932,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X57)) | s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7635,f670])).
fof(f7635,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X57)) | s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X56) != s(bool,X58) | s(bool,'T_0') = s(bool,X56)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4394,f670])).
fof(f7931,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,X51)) | s(bool,X50) = s(bool,X51)) ) | $spl136),
  inference(cnf_transformation,[],[f7931_D])).
fof(f7931_D,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,X51)) | s(bool,X50) = s(bool,X51)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f7930,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7633,f4107])).
fof(f7633,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51)) | s(bool,'T_0') = s(bool,X48)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4300,f670])).
fof(f7929,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7630,f4239])).
fof(f7630,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41) | s(bool,'T_0') = s(bool,X39)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4124,f670])).
fof(f7927,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | p(s(bool,X34))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7628,f4107])).
fof(f7628,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | p(s(bool,X34)) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4107,f670])).
fof(f7852,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X30)) ) | ($spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f7626])).
fof(f7626,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl8 | $spl22)),
  inference(superposition,[],[f396,f670])).
fof(f7853,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f7625])).
fof(f7625,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f259,f670])).
fof(f7926,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X5))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7616,f4107])).
fof(f7616,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7925,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7615,f4107])).
fof(f7615,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7924,plain,(
  ( ! [X263,X265,X262,X264] : (p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X262)))) | ~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X264)))) | s(bool,'T_0') = s(bool,X263)) ) | $spl22),
  inference(subsumption_resolution,[],[f7610,f670])).
fof(f7610,plain,(
  ( ! [X263,X265,X262,X264] : (s(bool,X263) != s(bool,X264) | p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X262)))) | ~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,X264)))) | s(bool,'T_0') = s(bool,X263)) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f7922,plain,(
  ( ! [X261,X259,X260,X258] : (p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X260)))) | ~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X258)))) | s(bool,'T_0') = s(bool,X259)) ) | $spl22),
  inference(subsumption_resolution,[],[f7609,f670])).
fof(f7609,plain,(
  ( ! [X261,X259,X260,X258] : (s(bool,X259) != s(bool,X260) | p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X260)))) | ~p(s(bool,happ(s(fun(bool,bool),X261),s(bool,X258)))) | s(bool,'T_0') = s(bool,X259)) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f7921,plain,(
  ( ! [X249,X250,X248] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X250),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X249),s('type.list.list'('type.closSem.v'),X248))))))) ) | $spl130),
  inference(cnf_transformation,[],[f7921_D])).
fof(f7921_D,plain,(
  ( ! [X249,X250,X248] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X250),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X249),s('type.list.list'('type.closSem.v'),X248))))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7920,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X251)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X250),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X249),s('type.list.list'('type.closSem.v'),X248))))))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7606,f4107])).
fof(f7606,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X251)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X250),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X249),s('type.list.list'('type.closSem.v'),X248)))))) | s(bool,'T_0') = s(bool,X251)) ) | $spl22),
  inference(superposition,[],[f4051,f670])).
fof(f7919,plain,(
  ( ! [X235,X236,X234] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X234))),s('type.closSem.v',X236))))) ) | $spl128),
  inference(cnf_transformation,[],[f7919_D])).
fof(f7919_D,plain,(
  ( ! [X235,X236,X234] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X234))),s('type.closSem.v',X236))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7918,plain,(
  ( ! [X237,X235,X236,X234] : (p(s(bool,X237)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X234))),s('type.closSem.v',X236))))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7603,f4107])).
fof(f7603,plain,(
  ( ! [X237,X235,X236,X234] : (p(s(bool,X237)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X235),s('type.list.list'('type.closSem.v'),X234))),s('type.closSem.v',X236)))) | s(bool,'T_0') = s(bool,X237)) ) | $spl22),
  inference(superposition,[],[f3923,f670])).
fof(f7916,plain,(
  ( ! [X217,X214] : (s('type.list.list'(X214),'const.list.NIL_0') != s('type.list.list'(X214),X217)) ) | $spl126),
  inference(cnf_transformation,[],[f7916_D])).
fof(f7916_D,plain,(
  ( ! [X217,X214] : (s('type.list.list'(X214),'const.list.NIL_0') != s('type.list.list'(X214),X217)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7917,plain,(
  ( ! [X227,X225,X228] : (p(s(bool,X228)) | s('type.list.list'(X225),'const.list.NIL_0') != s('type.list.list'(X225),X227)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7601,f4107])).
fof(f7601,plain,(
  ( ! [X227,X225,X228] : (p(s(bool,X228)) | s('type.list.list'(X225),'const.list.NIL_0') != s('type.list.list'(X225),X227) | s(bool,'T_0') = s(bool,X228)) ) | $spl22),
  inference(superposition,[],[f87,f670])).
fof(f7915,plain,(
  ( ! [X218,X217,X214] : (p(s(bool,X218)) | s('type.list.list'(X214),'const.list.NIL_0') != s('type.list.list'(X214),X217)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7599,f4107])).
fof(f7599,plain,(
  ( ! [X218,X217,X214] : (p(s(bool,X218)) | s('type.list.list'(X214),'const.list.NIL_0') != s('type.list.list'(X214),X217) | s(bool,'T_0') = s(bool,X218)) ) | $spl22),
  inference(superposition,[],[f85,f670])).
fof(f7914,plain,(
  ( ! [X208,X206,X207] : (~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))))) ) | $spl124),
  inference(cnf_transformation,[],[f7914_D])).
fof(f7914_D,plain,(
  ( ! [X208,X206,X207] : (~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f7913,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X209)) | ~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7597,f4107])).
fof(f7597,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X209)) | ~p(s(bool,happ(s(fun(X206,bool),X207),s(X206,X208)))) | s(bool,'T_0') = s(bool,X209)) ) | $spl22),
  inference(superposition,[],[f95,f670])).
fof(f7912,plain,(
  ( ! [X191,X189,X190,X193] : (~p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X193)))) | s(X189,X191) != s(X189,X193)) ) | $spl122),
  inference(cnf_transformation,[],[f7912_D])).
fof(f7912_D,plain,(
  ( ! [X191,X189,X190,X193] : (~p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X193)))) | s(X189,X191) != s(X189,X193)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7911,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X192)) | ~p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X193)))) | s(X189,X191) != s(X189,X193)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7592,f4107])).
fof(f7592,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,X192)) | ~p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X193)))) | s(X189,X191) != s(X189,X193) | s(bool,'T_0') = s(bool,X192)) ) | $spl22),
  inference(superposition,[],[f93,f670])).
fof(f7908,plain,(
  ( ! [X134,X133] : (s('type.closSem.v',X133) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X134)))) ) | $spl120),
  inference(cnf_transformation,[],[f7908_D])).
fof(f7908_D,plain,(
  ( ! [X134,X133] : (s('type.closSem.v',X133) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X134)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7910,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X167)) | s('type.closSem.v',X166) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X165)))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7585,f4107])).
fof(f7585,plain,(
  ( ! [X167,X165,X166] : (p(s(bool,X167)) | s('type.closSem.v',X166) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X165))) | s(bool,'T_0') = s(bool,X167)) ) | $spl22),
  inference(superposition,[],[f114,f670])).
fof(f114,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f7906,plain,(
  ( ! [X121,X122] : (s('type.closSem.v',X121) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122)))) ) | $spl118),
  inference(cnf_transformation,[],[f7906_D])).
fof(f7906_D,plain,(
  ( ! [X121,X122] : (s('type.closSem.v',X121) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f7909,plain,(
  ( ! [X156,X154,X155] : (p(s(bool,X156)) | s('type.closSem.v',X155) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X154)))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7581,f4107])).
fof(f7581,plain,(
  ( ! [X156,X154,X155] : (p(s(bool,X156)) | s('type.closSem.v',X155) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X154))) | s(bool,'T_0') = s(bool,X156)) ) | $spl22),
  inference(superposition,[],[f116,f670])).
fof(f116,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f7907,plain,(
  ( ! [X134,X135,X133] : (p(s(bool,X135)) | s('type.closSem.v',X133) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X134)))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7572,f4107])).
fof(f7572,plain,(
  ( ! [X134,X135,X133] : (p(s(bool,X135)) | s('type.closSem.v',X133) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X134))) | s(bool,'T_0') = s(bool,X135)) ) | $spl22),
  inference(superposition,[],[f110,f670])).
fof(f7905,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,X123)) | s('type.closSem.v',X121) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122)))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7568,f4107])).
fof(f7568,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,X123)) | s('type.closSem.v',X121) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X122))) | s(bool,'T_0') = s(bool,X123)) ) | $spl22),
  inference(superposition,[],[f112,f670])).
fof(f7904,plain,(
  ( ! [X101,X102] : (p(s(bool,X102)) | s(bool,'T_0') = s(bool,X101)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7561,f4239])).
fof(f7561,plain,(
  ( ! [X101,X102] : (s(bool,X101) != s(bool,X102) | p(s(bool,X102)) | s(bool,'T_0') = s(bool,X101)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4107,f670])).
fof(f7900,plain,(
  ( ! [X95,X96] : (p(s(bool,X95)) | ~p(s(bool,X96))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7557,f4107])).
fof(f7557,plain,(
  ( ! [X95,X96] : (p(s(bool,X95)) | ~p(s(bool,X96)) | s(bool,'T_0') = s(bool,X95)) ) | $spl22),
  inference(superposition,[],[f105,f670])).
fof(f7899,plain,(
  ( ! [X94,X93] : (p(s(bool,X93)) | ~p(s(bool,X94))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7556,f4107])).
fof(f7556,plain,(
  ( ! [X94,X93] : (p(s(bool,X93)) | ~p(s(bool,X94)) | s(bool,'T_0') = s(bool,X93)) ) | $spl22),
  inference(superposition,[],[f103,f670])).
fof(f7898,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X90)))) | s(bool,'T_0') = s(bool,X90) | ~p(s(bool,X92)) | p(s(bool,sK1(X92,X91,bool)))) ) | $spl116),
  inference(cnf_transformation,[],[f7898_D])).
fof(f7898_D,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X90)))) | s(bool,'T_0') = s(bool,X90) | ~p(s(bool,X92)) | p(s(bool,sK1(X92,X91,bool)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f7555,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,X90)))) | p(s(bool,sK1(X92,X91,bool))) | ~p(s(bool,X92)) | ~p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(superposition,[],[f6960,f670])).
fof(f7875,plain,(
  ( ! [X22,X20] : (s(bool,X20) = s(bool,X22)) ) | $spl108),
  inference(cnf_transformation,[],[f7875_D])).
fof(f7875_D,plain,(
  ( ! [X22,X20] : (s(bool,X20) = s(bool,X22)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7897,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,'T_0') = s(bool,X87) | s(bool,X85) = s(bool,X88)) ) | ($spl1 | $spl2 | $spl15 | $spl22 | $spl96)),
  inference(subsumption_resolution,[],[f7896,f670])).
fof(f7896,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,'T_0') = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X85) = s(bool,X88)) ) | ($spl1 | $spl2 | $spl15 | $spl22 | $spl96)),
  inference(subsumption_resolution,[],[f7554,f4107])).
fof(f7554,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | s(bool,'T_0') = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X86)) ) | ($spl22 | $spl96)),
  inference(superposition,[],[f5453,f670])).
fof(f7895,plain,(
  ( ! [X84,X82] : (s(bool,X84) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,X84)) ) | $spl114),
  inference(cnf_transformation,[],[f7895_D])).
fof(f7895_D,plain,(
  ( ! [X84,X82] : (s(bool,X84) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,X84)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f7894,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X83)) | s(bool,X84) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,X84)) ) | ($spl1 | $spl2 | $spl15 | $spl22 | $spl94)),
  inference(subsumption_resolution,[],[f7553,f4107])).
fof(f7553,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X83)) | s(bool,X84) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl22 | $spl94)),
  inference(superposition,[],[f5451,f670])).
fof(f7892,plain,(
  ( ! [X58,X56] : (s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X56) != s(bool,X58)) ) | $spl112),
  inference(cnf_transformation,[],[f7892_D])).
fof(f7892_D,plain,(
  ( ! [X58,X56] : (s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X56) != s(bool,X58)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f7891,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X57)) | s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X56) != s(bool,X58)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7545,f4107])).
fof(f7545,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X57)) | s(bool,X58) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X56) != s(bool,X58) | s(bool,'T_0') = s(bool,X57)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4394,f670])).
fof(f7890,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,X54)) | s(bool,X52) = s(bool,X55) | p(s(bool,X55)) | s(bool,'T_0') = s(bool,X53)) ) | ($spl1 | $spl2 | $spl8 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7544,f4239])).
fof(f7544,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X53) != s(bool,X54) | p(s(bool,X54)) | s(bool,X52) = s(bool,X55) | p(s(bool,X55)) | s(bool,'T_0') = s(bool,X53)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4300,f670])).
fof(f7889,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51)) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7543,f670])).
fof(f7543,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51)) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(superposition,[],[f4300,f670])).
fof(f7887,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | s(bool,X32) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7537,f4107])).
fof(f7537,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | s(bool,X32) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl1 | $spl2 | $spl22)),
  inference(superposition,[],[f4043,f670])).
fof(f7884,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X11) | p(s(bool,X11))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7528,f4107])).
fof(f7528,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s(bool,'T_0') = s(bool,X10)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7883,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,X8) | p(s(bool,X8))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7527,f4107])).
fof(f7527,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X7)) ) | $spl22),
  inference(superposition,[],[f88,f670])).
fof(f7859,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f7522])).
fof(f7522,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f670,f259])).
fof(f7862,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X19)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f7519])).
fof(f7519,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl22),
  inference(superposition,[],[f670,f670])).
fof(f7882,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7514,f4107])).
fof(f7514,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl22),
  inference(superposition,[],[f670,f88])).
fof(f7881,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7513,f4107])).
fof(f7513,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl22),
  inference(superposition,[],[f670,f88])).
fof(f7867,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X14)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f7483])).
fof(f7483,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl22),
  inference(superposition,[],[f670,f670])).
fof(f7868,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl22)),
  inference(duplicate_literal_removal,[],[f7482])).
fof(f7482,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f670,f259])).
fof(f7878,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7479,f4107])).
fof(f7479,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl22),
  inference(superposition,[],[f670,f88])).
fof(f7876,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl15 | $spl22)),
  inference(subsumption_resolution,[],[f7478,f4107])).
fof(f7478,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl22),
  inference(superposition,[],[f670,f88])).
fof(f7874,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f7464])).
fof(f7464,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22)) ) | $spl22),
  inference(superposition,[],[f670,f670])).
fof(f7448,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,sK1(X33,X32,bool))) | ~p(s(bool,X33)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f7424,f4239])).
fof(f7424,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,sK1(X33,X32,bool))) | ~p(s(bool,X33)) | ~p(s(bool,X31)) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6960,f4721])).
fof(f7446,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | p(s(bool,sK1(X21,X20,bool))) | ~p(s(bool,X21)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f7421,f4239])).
fof(f7421,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0')))) | p(s(bool,sK1(X21,X20,bool))) | ~p(s(bool,X21)) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) ) | ($spl1 | $spl2 | $spl6 | $spl8 | $spl15)),
  inference(superposition,[],[f6960,f4721])).
fof(f6918,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X30)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),X27)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X27,X28,X29)),s('type.list.list'('type.closSem.v'),X27)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6912,f4107])).
fof(f6912,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X30)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X29),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),X27)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X27,X28,X29)),s('type.list.list'('type.closSem.v'),X27))) | s(bool,'T_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f4051,f4721])).
fof(f6917,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X24),s('type.list.list'('type.closSem.v'),X23)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X23,X24,X25)),s('type.list.list'('type.closSem.v'),X23)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6911,f4107])).
fof(f6911,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X24),s('type.list.list'('type.closSem.v'),X23)))))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X23,X24,X25)),s('type.list.list'('type.closSem.v'),X23)))) ) | $spl6),
  inference(superposition,[],[f4051,f4721])).
fof(f6795,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X30)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),X27))),s('type.closSem.v',X29)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X27),s('type.list.list'('type.closSem.v'),sK2(X27,X28,X29))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6789,f4107])).
fof(f6789,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X30)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X28),s('type.list.list'('type.closSem.v'),X27))),s('type.closSem.v',X29)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X27),s('type.list.list'('type.closSem.v'),sK2(X27,X28,X29)))) | s(bool,'T_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f3923,f4721])).
fof(f6794,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X24),s('type.list.list'('type.closSem.v'),X23))),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X23),s('type.list.list'('type.closSem.v'),sK2(X23,X24,X25))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6788,f4107])).
fof(f6788,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X26)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X24),s('type.list.list'('type.closSem.v'),X23))),s('type.closSem.v',X25)))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X23),s('type.list.list'('type.closSem.v'),sK2(X23,X24,X25))))) ) | $spl6),
  inference(superposition,[],[f3923,f4721])).
fof(f6127,plain,(
  ( ! [X19,X17] : (s(bool,X19) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) != s(bool,X19)) ) | $spl100),
  inference(cnf_transformation,[],[f6127_D])).
fof(f6127_D,plain,(
  ( ! [X19,X17] : (s(bool,X19) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) != s(bool,X19)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6708,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) != s(bool,X51) | s(bool,X51) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl96)),
  inference(subsumption_resolution,[],[f6699,f4107])).
fof(f6699,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) != s(bool,X51) | s(bool,X51) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl6 | $spl96)),
  inference(superposition,[],[f5453,f259])).
fof(f6286,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X26)) ) | $spl104),
  inference(cnf_transformation,[],[f6286_D])).
fof(f6286_D,plain,(
  ( ! [X26,X23,X25] : (s(bool,X23) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X26)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f6707,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) != s(bool,X30) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl96)),
  inference(subsumption_resolution,[],[f6693,f4107])).
fof(f6693,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) != s(bool,X30) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl6 | $spl96)),
  inference(superposition,[],[f5453,f4721])).
fof(f6706,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl96)),
  inference(subsumption_resolution,[],[f6692,f4107])).
fof(f6692,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) != s(bool,X26) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl96)),
  inference(superposition,[],[f5453,f4721])).
fof(f6583,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X40)) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X41)))) | s(X37,X39) != s(X37,X41) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6455,f4107])).
fof(f6455,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X40)) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X41)))) | s(X37,X39) != s(X37,X41) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,X39))) | s(bool,'T_0') = s(bool,X40)) ) | $spl6),
  inference(superposition,[],[f93,f4721])).
fof(f6582,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X35)) | ~p(s(bool,happ(s(fun(X32,bool),X33),s(X32,X36)))) | s(X32,X34) != s(X32,X36) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),X33),s(X32,X34)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6454,f4107])).
fof(f6454,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X35)) | ~p(s(bool,happ(s(fun(X32,bool),X33),s(X32,X36)))) | s(X32,X34) != s(X32,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,happ(s(fun(X32,bool),X33),s(X32,X34)))) ) | $spl6),
  inference(superposition,[],[f93,f4721])).
fof(f6388,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X41)) | s('type.list.list'(X38),'const.list.NIL_0') != s('type.list.list'(X38),X40) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X37,fun(X38,bool)),X39),s('type.list.list'(X37),'const.list.NIL_0'),s('type.list.list'(X38),X40)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6361,f4107])).
fof(f6361,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X41)) | s('type.list.list'(X38),'const.list.NIL_0') != s('type.list.list'(X38),X40) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X37,fun(X38,bool)),X39),s('type.list.list'(X37),'const.list.NIL_0'),s('type.list.list'(X38),X40))) | s(bool,'T_0') = s(bool,X41)) ) | $spl6),
  inference(superposition,[],[f87,f4721])).
fof(f6387,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X36)) | s('type.list.list'(X33),'const.list.NIL_0') != s('type.list.list'(X33),X35) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X32,fun(X33,bool)),X34),s('type.list.list'(X32),'const.list.NIL_0'),s('type.list.list'(X33),X35)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6360,f4107])).
fof(f6360,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X36)) | s('type.list.list'(X33),'const.list.NIL_0') != s('type.list.list'(X33),X35) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X32,fun(X33,bool)),X34),s('type.list.list'(X32),'const.list.NIL_0'),s('type.list.list'(X33),X35)))) ) | $spl6),
  inference(superposition,[],[f87,f4721])).
fof(f6170,plain,(
  ( ! [X19,X17] : (s(bool,X19) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,X19)) ) | $spl102),
  inference(cnf_transformation,[],[f6170_D])).
fof(f6170_D,plain,(
  ( ! [X19,X17] : (s(bool,X19) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,X19)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6288,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6278,f4107])).
fof(f6278,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X49)) | s(bool,X50) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4936,f259])).
fof(f6287,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X27)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6272,f4107])).
fof(f6272,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4936,f4721])).
fof(f6285,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6271,f4107])).
fof(f6271,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4936,f4721])).
fof(f6171,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl94)),
  inference(subsumption_resolution,[],[f6156,f4107])).
fof(f6156,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl6 | $spl94)),
  inference(superposition,[],[f5451,f4721])).
fof(f6169,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15 | $spl94)),
  inference(subsumption_resolution,[],[f6155,f4107])).
fof(f6155,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl94)),
  inference(superposition,[],[f5451,f4721])).
fof(f6128,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X20) != s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6113,f4107])).
fof(f6113,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | s(bool,X22) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X20) != s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4394,f4721])).
fof(f6126,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) != s(bool,X19) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f6112,f4107])).
fof(f6112,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | s(bool,X19) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) != s(bool,X19) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4394,f4721])).
fof(f5982,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X41)) | s('type.list.list'(X37),'const.list.NIL_0') != s('type.list.list'(X37),X40) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X37,fun(X38,bool)),X39),s('type.list.list'(X37),X40),s('type.list.list'(X38),'const.list.NIL_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5964,f4107])).
fof(f5964,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,X41)) | s('type.list.list'(X37),'const.list.NIL_0') != s('type.list.list'(X37),X40) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X37,fun(X38,bool)),X39),s('type.list.list'(X37),X40),s('type.list.list'(X38),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X41)) ) | $spl6),
  inference(superposition,[],[f85,f4721])).
fof(f5981,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X36)) | s('type.list.list'(X32),'const.list.NIL_0') != s('type.list.list'(X32),X35) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X32,fun(X33,bool)),X34),s('type.list.list'(X32),X35),s('type.list.list'(X33),'const.list.NIL_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5963,f4107])).
fof(f5963,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X36)) | s('type.list.list'(X32),'const.list.NIL_0') != s('type.list.list'(X32),X35) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X32,fun(X33,bool)),X34),s('type.list.list'(X32),X35),s('type.list.list'(X33),'const.list.NIL_0')))) ) | $spl6),
  inference(superposition,[],[f85,f4721])).
fof(f5843,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X32)) | ~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31)))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,sK1(X31,X30,X29))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5794,f4107])).
fof(f5794,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X32)) | ~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31)))) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X30),s(X29,sK1(X31,X30,X29)))) | s(bool,'T_0') = s(bool,X32)) ) | $spl6),
  inference(superposition,[],[f95,f4721])).
fof(f5842,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X28)) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(X27,X26,X25))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5793,f4107])).
fof(f5793,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X28)) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,sK1(X27,X26,X25))))) ) | $spl6),
  inference(superposition,[],[f95,f4721])).
fof(f682,plain,(
  ( ! [X31,X32] : (s('type.closSem.v',X32) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))),s('type.closSem.v',X32)))) ) | $spl36),
  inference(cnf_transformation,[],[f682_D])).
fof(f682_D,plain,(
  ( ! [X31,X32] : (s('type.closSem.v',X32) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))),s('type.closSem.v',X32)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5702,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.closSem.v',X22) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X21))),s('type.closSem.v',X22)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5673,f4107])).
fof(f5673,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.closSem.v',X22) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X21))),s('type.closSem.v',X22))) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f116,f4721])).
fof(f5701,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X19) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X18))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X18))),s('type.closSem.v',X19)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5672,f4107])).
fof(f5672,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X19) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X18))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X18))),s('type.closSem.v',X19)))) ) | $spl6),
  inference(superposition,[],[f116,f4721])).
fof(f5481,plain,(
  $false | ($spl1 | $spl2 | $spl80)),
  inference(resolution,[],[f1303,f175])).
fof(f5468,plain,(
  ( ! [X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X38) | s(bool,X35) = s(bool,X38) | p(s(bool,X36))) ) | $spl98),
  inference(cnf_transformation,[],[f5468_D])).
fof(f5468_D,plain,(
  ( ! [X35,X38,X36] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X38) | s(bool,X35) = s(bool,X38) | p(s(bool,X36))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5469,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5440,f4107])).
fof(f5440,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | p(s(bool,X40)) | s(bool,'T_0') = s(bool,X41) | p(s(bool,X41)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4300,f4721])).
fof(f5467,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X36)) | p(s(bool,X37)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5439,f4107])).
fof(f5439,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X35) != s(bool,X36) | p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | p(s(bool,X37)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4300,f4721])).
fof(f5453,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X25) = s(bool,X26)) ) | $spl96),
  inference(cnf_transformation,[],[f5453_D])).
fof(f5453_D,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,X24) != s(bool,X25) | s(bool,X25) = s(bool,X26)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5454,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,X29) | p(s(bool,X27)) | s(bool,X29) = s(bool,X30) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5420,f4107])).
fof(f5420,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,X29) | p(s(bool,X27)) | s(bool,X29) = s(bool,X30) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X28)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4300,f4721])).
fof(f5452,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,X25) | p(s(bool,X23)) | s(bool,X25) = s(bool,X26) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5419,f4107])).
fof(f5419,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X24) != s(bool,X25) | p(s(bool,X23)) | s(bool,X25) = s(bool,X26) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4300,f4721])).
fof(f5451,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) = s(bool,X18)) ) | $spl94),
  inference(cnf_transformation,[],[f5451_D])).
fof(f5451_D,plain,(
  ( ! [X17,X18] : (p(s(bool,X18)) | s(bool,X17) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X17) = s(bool,X18)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5450,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X16)) | s(bool,X17) = s(bool,X18) | p(s(bool,X18))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5417,f4107])).
fof(f5417,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X16)) | s(bool,X17) = s(bool,X18) | p(s(bool,X18)) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(superposition,[],[f4300,f259])).
fof(f684,plain,(
  ( ! [X35,X36] : (s('type.closSem.v',X36) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X36)))) ) | $spl40),
  inference(cnf_transformation,[],[f684_D])).
fof(f684_D,plain,(
  ( ! [X35,X36] : (s('type.closSem.v',X36) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X36)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5409,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.closSem.v',X22) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X21))),s('type.closSem.v',X22)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5383,f4107])).
fof(f5383,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X23)) | s('type.closSem.v',X22) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X21))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X21))),s('type.closSem.v',X22))) | s(bool,'T_0') = s(bool,X23)) ) | $spl6),
  inference(superposition,[],[f114,f4721])).
fof(f5408,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X19) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))),s('type.closSem.v',X19)))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f5382,f4107])).
fof(f5382,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X19) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X18))),s('type.closSem.v',X19)))) ) | $spl6),
  inference(superposition,[],[f114,f4721])).
fof(f2717,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.NIL_0')))) ) | $spl86),
  inference(cnf_transformation,[],[f2717_D])).
fof(f2717_D,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.NIL_0')))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5146,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X84)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X81,fun(X82,bool)),X83),s('type.list.list'(X81),'const.list.NIL_0'),s('type.list.list'(X82),'const.list.NIL_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4928,f4107])).
fof(f4928,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X84)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X81,fun(X82,bool)),X83),s('type.list.list'(X81),'const.list.NIL_0'),s('type.list.list'(X82),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f2685,f4721])).
fof(f2685,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(equality_resolution,[],[f85])).
fof(f680,plain,(
  ( ! [X28,X27] : (s('type.closSem.v',X27) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X27),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28)))))) ) | $spl32),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  ( ! [X28,X27] : (s('type.closSem.v',X27) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X27),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5145,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s('type.closSem.v',X58) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X58),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4918,f4107])).
fof(f4918,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s('type.closSem.v',X58) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X58),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59))))) | s(bool,'T_0') = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f112,f4721])).
fof(f678,plain,(
  ( ! [X24,X23] : (s('type.closSem.v',X23) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X23),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24)))))) ) | $spl28),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  ( ! [X24,X23] : (s('type.closSem.v',X23) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X23),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5144,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X54)) | s('type.closSem.v',X52) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4916,f4107])).
fof(f4916,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X54)) | s('type.closSem.v',X52) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53))))) | s(bool,'T_0') = s(bool,X54)) ) | $spl6),
  inference(superposition,[],[f110,f4721])).
fof(f5139,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X84)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X81,fun(X82,bool)),X83),s('type.list.list'(X81),'const.list.NIL_0'),s('type.list.list'(X82),'const.list.NIL_0')))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4891,f4107])).
fof(f4891,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,X84)) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X81,fun(X82,bool)),X83),s('type.list.list'(X81),'const.list.NIL_0'),s('type.list.list'(X82),'const.list.NIL_0')))) ) | $spl6),
  inference(superposition,[],[f2685,f4721])).
fof(f5138,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s('type.closSem.v',X58) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X58),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4881,f4107])).
fof(f4881,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,X60)) | s('type.closSem.v',X58) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59))) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X58),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X59)))))) ) | $spl6),
  inference(superposition,[],[f112,f4721])).
fof(f5137,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X54)) | s('type.closSem.v',X52) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53)))))) ) | ($spl1 | $spl2 | $spl6 | $spl15)),
  inference(subsumption_resolution,[],[f4879,f4107])).
fof(f4879,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X54)) | s('type.closSem.v',X52) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X53)))))) ) | $spl6),
  inference(superposition,[],[f110,f4721])).
fof(f4223,plain,(
  $false | ($spl1 | $spl2 | $spl10)),
  inference(resolution,[],[f4193,f108])).
fof(f4193,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl1 | $spl2 | $spl10)),
  inference(superposition,[],[f175,f397])).
fof(f4224,plain,(
  $false | ($spl1 | $spl2 | $spl10)),
  inference(resolution,[],[f4193,f107])).
fof(f107,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.bool.TRUTH')).
fof(f4227,plain,(
  $false | ($spl1 | $spl2 | $spl10)),
  inference(resolution,[],[f4193,f124])).
fof(f4217,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')))) | $spl92),
  inference(cnf_transformation,[],[f4217_D])).
fof(f4217_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0')))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4194,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Unit_0'))))) ) | $spl10),
  inference(superposition,[],[f3429,f397])).
fof(f4075,plain,(
  $false | ($spl1 | $spl2 | $spl14)),
  inference(subsumption_resolution,[],[f4053,f108])).
fof(f4053,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl14)),
  inference(backward_demodulation,[],[f4052,f175])).
fof(f4052,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ($spl2 | $spl14)),
  inference(forward_demodulation,[],[f468,f128])).
fof(f669,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl21),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f4035,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) | $spl90),
  inference(cnf_transformation,[],[f4035_D])).
fof(f4035_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4033,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl2),
  inference(forward_demodulation,[],[f4012,f128])).
fof(f4012,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',sK0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) | $spl2),
  inference(superposition,[],[f3429,f128])).
fof(f3996,plain,(
  $false | ($spl1 | $spl2 | $spl12)),
  inference(subsumption_resolution,[],[f3995,f108])).
fof(f3995,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl12)),
  inference(forward_demodulation,[],[f175,f427])).
fof(f3994,plain,(
  $false | ($spl12 | $spl21)),
  inference(subsumption_resolution,[],[f3993,f108])).
fof(f3993,plain,(
  ~p(s(bool,'T_0')) | ($spl12 | $spl21)),
  inference(forward_demodulation,[],[f669,f427])).
fof(f3982,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3981,f108])).
fof(f3981,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3980,f441])).
fof(f441,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) | $spl0),
  inference(resolution,[],[f431,f108])).
fof(f431,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) ) | $spl0),
  inference(resolution,[],[f126,f91])).
fof(f3980,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3955,f117])).
fof(f3955,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))),s('type.closSem.v',sK0)))) | ($spl0 | $spl3)),
  inference(resolution,[],[f3923,f3615])).
fof(f3615,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun('type.closSem.v',bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3614,f129])).
fof(f3614,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK0) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun('type.closSem.v',bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3565,f117])).
fof(f3565,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',sK0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun('type.closSem.v',bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))))) ) | $spl0),
  inference(superposition,[],[f3547,f86])).
fof(f3547,plain,(
  s('type.closSem.v',sK0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) | $spl0),
  inference(subsumption_resolution,[],[f3533,f124])).
fof(f3533,plain,(
  ~p(s(bool,'T_0')) | s('type.closSem.v',sK0) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0',sK0)))) | $spl0),
  inference(superposition,[],[f3429,f441])).
fof(f427,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | $spl12),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3770,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl0 | $spl20)),
  inference(subsumption_resolution,[],[f3765,f474])).
fof(f474,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(backward_demodulation,[],[f441,f431])).
fof(f3765,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl0 | $spl20)),
  inference(superposition,[],[f668,f546])).
fof(f546,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(resolution,[],[f513,f474])).
fof(f513,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(resolution,[],[f474,f88])).
fof(f3769,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl0 | $spl20)),
  inference(subsumption_resolution,[],[f3764,f474])).
fof(f3764,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl0 | $spl20)),
  inference(superposition,[],[f668,f546])).
fof(f259,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3750,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | ($spl0 | $spl21)),
  inference(subsumption_resolution,[],[f3737,f474])).
fof(f3737,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | ($spl0 | $spl21)),
  inference(superposition,[],[f3551,f546])).
fof(f3551,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl21),
  inference(resolution,[],[f669,f88])).
fof(f3749,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl0 | $spl21)),
  inference(subsumption_resolution,[],[f3736,f474])).
fof(f3736,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl0 | $spl21)),
  inference(superposition,[],[f3551,f546])).
fof(f3546,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) | $spl88),
  inference(cnf_transformation,[],[f3546_D])).
fof(f3546_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3544,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))))),
  inference(forward_demodulation,[],[f3527,f117])).
fof(f3527,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),sK2('const.list.NIL_0','const.dec_to_exh.tuple_tag_0','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0'))))))),
  inference(superposition,[],[f3429,f117])).
fof(f3244,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(X33,bool),X34),s(X33,X37)))) | s(X33,X35) != s(X33,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35)))) ) | $spl0),
  inference(subsumption_resolution,[],[f3132,f474])).
fof(f3132,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X36)) | ~p(s(bool,happ(s(fun(X33,bool),X34),s(X33,X37)))) | s(X33,X35) != s(X33,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,X35)))) ) | $spl0),
  inference(superposition,[],[f93,f546])).
fof(f3243,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X32)))) | s(X28,X30) != s(X28,X32) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(subsumption_resolution,[],[f3131,f474])).
fof(f3131,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X31)) | ~p(s(bool,happ(s(fun(X28,bool),X29),s(X28,X32)))) | s(X28,X30) != s(X28,X32) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(bool,'T_0') = s(bool,X31)) ) | $spl0),
  inference(superposition,[],[f93,f546])).
fof(f2884,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s('type.list.list'(X34),'const.list.NIL_0') != s('type.list.list'(X34),X36) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X33,fun(X34,bool)),X35),s('type.list.list'(X33),'const.list.NIL_0'),s('type.list.list'(X34),X36)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2858,f474])).
fof(f2858,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X37)) | s('type.list.list'(X34),'const.list.NIL_0') != s('type.list.list'(X34),X36) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X33,fun(X34,bool)),X35),s('type.list.list'(X33),'const.list.NIL_0'),s('type.list.list'(X34),X36)))) ) | $spl0),
  inference(superposition,[],[f87,f546])).
fof(f2883,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.list.list'(X29),'const.list.NIL_0') != s('type.list.list'(X29),X31) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X28,fun(X29,bool)),X30),s('type.list.list'(X28),'const.list.NIL_0'),s('type.list.list'(X29),X31))) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f2857,f474])).
fof(f2857,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X32)) | s('type.list.list'(X29),'const.list.NIL_0') != s('type.list.list'(X29),X31) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X28,fun(X29,bool)),X30),s('type.list.list'(X28),'const.list.NIL_0'),s('type.list.list'(X29),X31))) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f87,f546])).
fof(f2718,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(X25,bool)),X26),s('type.list.list'(X24),'const.list.NIL_0'),s('type.list.list'(X25),'const.list.NIL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f2712,f474])).
fof(f2712,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X27)) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(X25,bool)),X26),s('type.list.list'(X24),'const.list.NIL_0'),s('type.list.list'(X25),'const.list.NIL_0')))) ) | $spl0),
  inference(superposition,[],[f2685,f546])).
fof(f2716,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl0),
  inference(subsumption_resolution,[],[f2711,f474])).
fof(f2711,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X23)) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X20,fun(X21,bool)),X22),s('type.list.list'(X20),'const.list.NIL_0'),s('type.list.list'(X21),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f2685,f546])).
fof(f2693,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.list.list'(X30),'const.list.NIL_0') != s('type.list.list'(X30),X33) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X30,fun(X31,bool)),X32),s('type.list.list'(X30),X33),s('type.list.list'(X31),'const.list.NIL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f2678,f474])).
fof(f2678,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X34)) | s('type.list.list'(X30),'const.list.NIL_0') != s('type.list.list'(X30),X33) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X30,fun(X31,bool)),X32),s('type.list.list'(X30),X33),s('type.list.list'(X31),'const.list.NIL_0')))) ) | $spl0),
  inference(superposition,[],[f85,f546])).
fof(f2692,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s('type.list.list'(X25),'const.list.NIL_0') != s('type.list.list'(X25),X28) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X25,fun(X26,bool)),X27),s('type.list.list'(X25),X28),s('type.list.list'(X26),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl0),
  inference(subsumption_resolution,[],[f2677,f474])).
fof(f2677,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,X29)) | s('type.list.list'(X25),'const.list.NIL_0') != s('type.list.list'(X25),X28) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X25,fun(X26,bool)),X27),s('type.list.list'(X25),X28),s('type.list.list'(X26),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f85,f546])).
fof(f2494,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK1(X26,X25,X24))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2454,f474])).
fof(f2454,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,X27)) | ~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,sK1(X26,X25,X24))))) ) | $spl0),
  inference(superposition,[],[f95,f546])).
fof(f2493,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,sK1(X22,X21,X20)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl0),
  inference(subsumption_resolution,[],[f2453,f474])).
fof(f2453,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,X23)) | ~p(s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,sK1(X22,X21,X20)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl0),
  inference(superposition,[],[f95,f546])).
fof(f2373,plain,(
  ( ! [X21,X19,X20] : (s('type.closSem.v',X20) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v',X20)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2348,f474])).
fof(f2348,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X21)) | s('type.closSem.v',X20) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))),s('type.closSem.v',X20)))) ) | $spl0),
  inference(superposition,[],[f116,f546])).
fof(f2372,plain,(
  ( ! [X17,X18,X16] : (s('type.closSem.v',X17) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))),s('type.closSem.v',X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl0),
  inference(subsumption_resolution,[],[f2347,f474])).
fof(f2347,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X18)) | s('type.closSem.v',X17) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))),s('type.closSem.v',X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl0),
  inference(superposition,[],[f116,f546])).
fof(f2211,plain,(
  ( ! [X21,X19,X20] : (s('type.closSem.v',X20) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v',X20)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2189,f474])).
fof(f2189,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X21)) | s('type.closSem.v',X20) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))),s('type.closSem.v',X20)))) ) | $spl0),
  inference(superposition,[],[f114,f546])).
fof(f2210,plain,(
  ( ! [X17,X18,X16] : (s('type.closSem.v',X17) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))),s('type.closSem.v',X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl0),
  inference(subsumption_resolution,[],[f2188,f474])).
fof(f2188,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X18)) | s('type.closSem.v',X17) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))),s('type.closSem.v',X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl0),
  inference(superposition,[],[f114,f546])).
fof(f2029,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2022,f474])).
fof(f2022,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X12)))))) ) | $spl0),
  inference(superposition,[],[f1999,f546])).
fof(f2027,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f2021,f474])).
fof(f2021,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f1999,f546])).
fof(f2004,plain,(
  ( ! [X19,X20,X18] : (s('type.closSem.v',X18) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X18),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1987,f474])).
fof(f1987,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X18) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X18),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X19)))))) ) | $spl0),
  inference(superposition,[],[f112,f546])).
fof(f2003,plain,(
  ( ! [X17,X15,X16] : (s('type.closSem.v',X15) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X15),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl0),
  inference(subsumption_resolution,[],[f1986,f474])).
fof(f1986,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.closSem.v',X15) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X15),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f112,f546])).
fof(f1864,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X12))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1858,f474])).
fof(f1858,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X12))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X12)))))) ) | $spl0),
  inference(superposition,[],[f1836,f546])).
fof(f1862,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f1857,f474])).
fof(f1857,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f1836,f546])).
fof(f1841,plain,(
  ( ! [X19,X20,X18] : (s('type.closSem.v',X18) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X18),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1827,f474])).
fof(f1827,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | s('type.closSem.v',X18) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X18),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X19)))))) ) | $spl0),
  inference(superposition,[],[f110,f546])).
fof(f1840,plain,(
  ( ! [X17,X15,X16] : (s('type.closSem.v',X15) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X15),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl0),
  inference(subsumption_resolution,[],[f1826,f474])).
fof(f1826,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.closSem.v',X15) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X15),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl0),
  inference(superposition,[],[f110,f546])).
fof(f1716,plain,(
  $false | ($spl3 | $spl80)),
  inference(subsumption_resolution,[],[f1713,f1685])).
fof(f1685,plain,(
  ( ! [X2,X1] : (s('type.closSem.v',X1) = s('type.closSem.v',X2)) ) | $spl80),
  inference(superposition,[],[f1679,f1679])).
fof(f1679,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) ) | $spl80),
  inference(subsumption_resolution,[],[f109,f1303])).
fof(f1713,plain,(
  ( ! [X10] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X10)) ) | ($spl3 | $spl80)),
  inference(superposition,[],[f129,f1685])).
fof(f1715,plain,(
  $false | ($spl3 | $spl80)),
  inference(subsumption_resolution,[],[f1706,f1685])).
fof(f1706,plain,(
  ( ! [X10] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X10)) ) | ($spl3 | $spl80)),
  inference(superposition,[],[f129,f1685])).
fof(f1461,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f1459,f1435])).
fof(f1435,plain,(
  ( ! [X2,X1] : (s('type.closSem.v',X1) = s('type.closSem.v',X2)) ) | $spl16),
  inference(superposition,[],[f1286,f1286])).
fof(f1286,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f1198,f124])).
fof(f1198,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X1)))) ) | $spl16),
  inference(backward_demodulation,[],[f664,f109])).
fof(f1459,plain,(
  ( ! [X10] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X10)) ) | ($spl3 | $spl16)),
  inference(superposition,[],[f129,f1435])).
fof(f1460,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f1453,f1435])).
fof(f1453,plain,(
  ( ! [X10] : (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X10)) ) | ($spl3 | $spl16)),
  inference(superposition,[],[f129,f1435])).
fof(f1238,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3))) ) | $spl16),
  inference(superposition,[],[f105,f664])).
fof(f1237,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl16),
  inference(superposition,[],[f103,f664])).
fof(f1145,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X152),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150)))) ) | $spl0),
  inference(forward_demodulation,[],[f1144,f92])).
fof(f1144,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X152))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150)))) ) | $spl0),
  inference(subsumption_resolution,[],[f881,f474])).
fof(f881,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,X153)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X152))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150)))) ) | $spl0),
  inference(superposition,[],[f122,f546])).
fof(f1143,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148))))) ) | $spl0),
  inference(forward_demodulation,[],[f1142,f92])).
fof(f1142,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148))))) ) | $spl0),
  inference(subsumption_resolution,[],[f880,f474])).
fof(f880,plain,(
  ( ! [X146,X149,X147,X148] : (p(s(bool,X149)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148))))) ) | $spl0),
  inference(superposition,[],[f119,f546])).
fof(f1141,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s('type.list.list'(X132),'const.list.NIL_0') != s('type.list.list'(X132),X134) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X131,fun(X132,bool)),X133),s('type.list.list'(X131),'const.list.NIL_0'),s('type.list.list'(X132),X134)))) ) | $spl0),
  inference(subsumption_resolution,[],[f877,f474])).
fof(f877,plain,(
  ( ! [X134,X132,X135,X133,X131] : (p(s(bool,X135)) | s('type.list.list'(X132),'const.list.NIL_0') != s('type.list.list'(X132),X134) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X131,fun(X132,bool)),X133),s('type.list.list'(X131),'const.list.NIL_0'),s('type.list.list'(X132),X134)))) ) | $spl0),
  inference(superposition,[],[f87,f546])).
fof(f1140,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s('type.list.list'(X121),'const.list.NIL_0') != s('type.list.list'(X121),X124) | s(bool,'T_0') = s(bool,X125) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X121,fun(X122,bool)),X123),s('type.list.list'(X121),X124),s('type.list.list'(X122),'const.list.NIL_0')))) ) | $spl0),
  inference(subsumption_resolution,[],[f875,f474])).
fof(f875,plain,(
  ( ! [X125,X123,X121,X124,X122] : (p(s(bool,X125)) | s('type.list.list'(X121),'const.list.NIL_0') != s('type.list.list'(X121),X124) | s(bool,'T_0') = s(bool,X125) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X121,fun(X122,bool)),X123),s('type.list.list'(X121),X124),s('type.list.list'(X122),'const.list.NIL_0')))) ) | $spl0),
  inference(superposition,[],[f85,f546])).
fof(f1139,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,happ(s(fun(X103,bool),X104),s(X103,X107)))) | s(X103,X105) != s(X103,X107) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,happ(s(fun(X103,bool),X104),s(X103,X105)))) ) | $spl0),
  inference(subsumption_resolution,[],[f871,f474])).
fof(f871,plain,(
  ( ! [X103,X107,X105,X106,X104] : (p(s(bool,X106)) | ~p(s(bool,happ(s(fun(X103,bool),X104),s(X103,X107)))) | s(X103,X105) != s(X103,X107) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,happ(s(fun(X103,bool),X104),s(X103,X105)))) ) | $spl0),
  inference(superposition,[],[f93,f546])).
fof(f1138,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK1(X96,X95,X94))))) ) | $spl0),
  inference(subsumption_resolution,[],[f869,f474])).
fof(f869,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X97)) | ~p(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK1(X96,X95,X94))))) ) | $spl0),
  inference(superposition,[],[f95,f546])).
fof(f696,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X55),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X57))))) | s('type.closSem.v',X55) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X58))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X58),s('type.list.list'('type.closSem.v'),X57))))) ) | $spl52),
  inference(cnf_transformation,[],[f696_D])).
fof(f696_D,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X55),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X57))))) | s('type.closSem.v',X55) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X58))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X58),s('type.list.list'('type.closSem.v'),X57))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1137,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X89),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(forward_demodulation,[],[f1136,f92])).
fof(f1136,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X89))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f868,f474])).
fof(f868,plain,(
  ( ! [X92,X90,X93,X91,X89] : (p(s(bool,X92)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X89))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91)))))) ) | $spl0),
  inference(superposition,[],[f123,f546])).
fof(f690,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X46))),s('type.closSem.v',X47))) | s('type.closSem.v',X47) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X48))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X46),s('type.list.list'('type.closSem.v'),X48))))) ) | $spl46),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X46))),s('type.closSem.v',X47))) | s('type.closSem.v',X47) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X48))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X46),s('type.list.list'('type.closSem.v'),X48))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1135,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))),s('type.closSem.v',X78))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(forward_demodulation,[],[f1134,f92])).
fof(f1134,plain,(
  ( ! [X80,X78,X76,X79,X77] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))))),s('type.closSem.v',X78)))) ) | $spl0),
  inference(subsumption_resolution,[],[f865,f474])).
fof(f865,plain,(
  ( ! [X80,X78,X76,X79,X77] : (p(s(bool,X79)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))))),s('type.closSem.v',X78)))) ) | $spl0),
  inference(superposition,[],[f120,f546])).
fof(f1133,plain,(
  ( ! [X64,X62,X63] : (s('type.closSem.v',X63) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))),s('type.closSem.v',X63)))) ) | $spl0),
  inference(subsumption_resolution,[],[f861,f474])).
fof(f861,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | s('type.closSem.v',X63) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))),s('type.closSem.v',X63)))) ) | $spl0),
  inference(superposition,[],[f114,f546])).
fof(f1132,plain,(
  ( ! [X57,X58,X56] : (s('type.closSem.v',X57) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))),s('type.closSem.v',X57)))) ) | $spl0),
  inference(subsumption_resolution,[],[f859,f474])).
fof(f859,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X58)) | s('type.closSem.v',X57) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))),s('type.closSem.v',X57)))) ) | $spl0),
  inference(superposition,[],[f116,f546])).
fof(f1131,plain,(
  ( ! [X50,X51,X49] : (s('type.closSem.v',X49) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f856,f474])).
fof(f856,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X51)) | s('type.closSem.v',X49) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50)))))) ) | $spl0),
  inference(superposition,[],[f112,f546])).
fof(f1130,plain,(
  ( ! [X45,X43,X44] : (s('type.closSem.v',X43) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X43),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f854,f474])).
fof(f854,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X45)) | s('type.closSem.v',X43) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X43),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44)))))) ) | $spl0),
  inference(superposition,[],[f110,f546])).
fof(f1127,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X152),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150))) | s(bool,'T_0') = s(bool,X153)) ) | $spl0),
  inference(forward_demodulation,[],[f1126,f92])).
fof(f1126,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X152))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150))) | s(bool,'T_0') = s(bool,X153)) ) | $spl0),
  inference(subsumption_resolution,[],[f827,f474])).
fof(f827,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,X153)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X152))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X151),s('type.list.list'('type.closSem.v'),X150)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X150,X151,X152)),s('type.list.list'('type.closSem.v'),X150))) | s(bool,'T_0') = s(bool,X153)) ) | $spl0),
  inference(superposition,[],[f122,f546])).
fof(f1125,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148)))) | s(bool,'T_0') = s(bool,X149)) ) | $spl0),
  inference(forward_demodulation,[],[f1124,f92])).
fof(f1124,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148)))) | s(bool,'T_0') = s(bool,X149)) ) | $spl0),
  inference(subsumption_resolution,[],[f826,f474])).
fof(f826,plain,(
  ( ! [X146,X149,X147,X148] : (p(s(bool,X149)) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X147),s('type.list.list'('type.closSem.v'),X146))))),s('type.closSem.v',X148)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X146),s('type.list.list'('type.closSem.v'),sK2(X146,X147,X148)))) | s(bool,'T_0') = s(bool,X149)) ) | $spl0),
  inference(superposition,[],[f119,f546])).
fof(f1123,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s('type.list.list'(X132),'const.list.NIL_0') != s('type.list.list'(X132),X134) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X131,fun(X132,bool)),X133),s('type.list.list'(X131),'const.list.NIL_0'),s('type.list.list'(X132),X134))) | s(bool,'T_0') = s(bool,X135)) ) | $spl0),
  inference(subsumption_resolution,[],[f823,f474])).
fof(f823,plain,(
  ( ! [X134,X132,X135,X133,X131] : (p(s(bool,X135)) | s('type.list.list'(X132),'const.list.NIL_0') != s('type.list.list'(X132),X134) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X131,fun(X132,bool)),X133),s('type.list.list'(X131),'const.list.NIL_0'),s('type.list.list'(X132),X134))) | s(bool,'T_0') = s(bool,X135)) ) | $spl0),
  inference(superposition,[],[f87,f546])).
fof(f1122,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s('type.list.list'(X121),'const.list.NIL_0') != s('type.list.list'(X121),X124) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X121,fun(X122,bool)),X123),s('type.list.list'(X121),X124),s('type.list.list'(X122),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(subsumption_resolution,[],[f821,f474])).
fof(f821,plain,(
  ( ! [X125,X123,X121,X124,X122] : (p(s(bool,X125)) | s('type.list.list'(X121),'const.list.NIL_0') != s('type.list.list'(X121),X124) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X121,fun(X122,bool)),X123),s('type.list.list'(X121),X124),s('type.list.list'(X122),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X125)) ) | $spl0),
  inference(superposition,[],[f85,f546])).
fof(f1121,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,happ(s(fun(X103,bool),X104),s(X103,X107)))) | s(X103,X105) != s(X103,X107) | s(bool,'T_0') = s(bool,happ(s(fun(X103,bool),X104),s(X103,X105))) | s(bool,'T_0') = s(bool,X106)) ) | $spl0),
  inference(subsumption_resolution,[],[f817,f474])).
fof(f817,plain,(
  ( ! [X103,X107,X105,X106,X104] : (p(s(bool,X106)) | ~p(s(bool,happ(s(fun(X103,bool),X104),s(X103,X107)))) | s(X103,X105) != s(X103,X107) | s(bool,'T_0') = s(bool,happ(s(fun(X103,bool),X104),s(X103,X105))) | s(bool,'T_0') = s(bool,X106)) ) | $spl0),
  inference(superposition,[],[f93,f546])).
fof(f1120,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK1(X96,X95,X94)))) | s(bool,'T_0') = s(bool,X97)) ) | $spl0),
  inference(subsumption_resolution,[],[f815,f474])).
fof(f815,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X97)) | ~p(s(bool,happ(s(fun(X94,bool),X95),s(X94,X96)))) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),X95),s(X94,sK1(X96,X95,X94)))) | s(bool,'T_0') = s(bool,X97)) ) | $spl0),
  inference(superposition,[],[f95,f546])).
fof(f1119,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X89),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(forward_demodulation,[],[f1118,f92])).
fof(f1118,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X89))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(subsumption_resolution,[],[f814,f474])).
fof(f814,plain,(
  ( ! [X92,X90,X93,X91,X89] : (p(s(bool,X92)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X93),s('type.list.list'('type.closSem.v'),X91)))) | s('type.closSem.v',X89) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X93))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X89))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X90),s('type.list.list'('type.closSem.v'),X91))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl0),
  inference(superposition,[],[f123,f546])).
fof(f1117,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))),s('type.closSem.v',X78))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(forward_demodulation,[],[f1116,f92])).
fof(f1116,plain,(
  ( ! [X80,X78,X76,X79,X77] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))))),s('type.closSem.v',X78))) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(subsumption_resolution,[],[f811,f474])).
fof(f811,plain,(
  ( ! [X80,X78,X76,X79,X77] : (p(s(bool,X79)) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X77),s('type.list.list'('type.closSem.v'),X80)))) | s('type.closSem.v',X78) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X80))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X76),s('type.list.list'('type.closSem.v'),X77))))),s('type.closSem.v',X78))) | s(bool,'T_0') = s(bool,X79)) ) | $spl0),
  inference(superposition,[],[f120,f546])).
fof(f1115,plain,(
  ( ! [X64,X62,X63] : (s('type.closSem.v',X63) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))),s('type.closSem.v',X63))) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(subsumption_resolution,[],[f807,f474])).
fof(f807,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | s('type.closSem.v',X63) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X62))),s('type.closSem.v',X63))) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(superposition,[],[f114,f546])).
fof(f1114,plain,(
  ( ! [X57,X58,X56] : (s('type.closSem.v',X57) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))),s('type.closSem.v',X57))) | s(bool,'T_0') = s(bool,X58)) ) | $spl0),
  inference(subsumption_resolution,[],[f805,f474])).
fof(f805,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X58)) | s('type.closSem.v',X57) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X56))),s('type.closSem.v',X57))) | s(bool,'T_0') = s(bool,X58)) ) | $spl0),
  inference(superposition,[],[f116,f546])).
fof(f1113,plain,(
  ( ! [X50,X51,X49] : (s('type.closSem.v',X49) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f802,f474])).
fof(f802,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,X51)) | s('type.closSem.v',X49) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X50))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl0),
  inference(superposition,[],[f112,f546])).
fof(f1112,plain,(
  ( ! [X45,X43,X44] : (s('type.closSem.v',X43) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X43),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))))) | s(bool,'T_0') = s(bool,X45)) ) | $spl0),
  inference(subsumption_resolution,[],[f800,f474])).
fof(f800,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X45)) | s('type.closSem.v',X43) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X43),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X44))))) | s(bool,'T_0') = s(bool,X45)) ) | $spl0),
  inference(superposition,[],[f110,f546])).
fof(f428,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | $spl13),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f660,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | $spl6),
  inference(equality_factoring,[],[f259])).
fof(f717,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f647,f513])).
fof(f647,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X10)) ) | $spl6),
  inference(superposition,[],[f91,f259])).
fof(f716,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f646,f513])).
fof(f646,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,X8)) ) | $spl6),
  inference(superposition,[],[f91,f259])).
fof(f665,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f665_D])).
fof(f665_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f715,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f645,f474])).
fof(f645,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | s(bool,X7) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f88,f259])).
fof(f714,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f644,f474])).
fof(f644,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X5) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f88,f259])).
fof(f668,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl20),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f712,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X105),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X104),s('type.list.list'('type.closSem.v'),X103)))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X103,X104,X105)),s('type.list.list'('type.closSem.v'),X103)))) ) | $spl6),
  inference(forward_demodulation,[],[f633,f92])).
fof(f633,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X105))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X104),s('type.list.list'('type.closSem.v'),X103)))))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X103,X104,X105)),s('type.list.list'('type.closSem.v'),X103)))) ) | $spl6),
  inference(superposition,[],[f122,f259])).
fof(f710,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X101),s('type.list.list'('type.closSem.v'),X100))),s('type.closSem.v',X102)))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X100),s('type.list.list'('type.closSem.v'),sK2(X100,X101,X102))))) ) | $spl6),
  inference(forward_demodulation,[],[f632,f92])).
fof(f632,plain,(
  ( ! [X101,X102,X100] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X101),s('type.list.list'('type.closSem.v'),X100))))),s('type.closSem.v',X102)))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X100),s('type.list.list'('type.closSem.v'),sK2(X100,X101,X102))))) ) | $spl6),
  inference(superposition,[],[f119,f259])).
fof(f708,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X96))),s('type.closSem.v',X99)))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X99) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X97))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X96),s('type.list.list'('type.closSem.v'),X97)))) ) | $spl6),
  inference(forward_demodulation,[],[f631,f92])).
fof(f631,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X96))))),s('type.closSem.v',X99)))) | s('type.closSem.v',X99) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X98),s('type.list.list'('type.closSem.v'),X97))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X96),s('type.list.list'('type.closSem.v'),X97)))) ) | $spl6),
  inference(superposition,[],[f120,f259])).
fof(f706,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X94),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X93)))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X94) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X92))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X92),s('type.list.list'('type.closSem.v'),X93)))) ) | $spl6),
  inference(forward_demodulation,[],[f630,f92])).
fof(f630,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X94))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X93)))))) | s('type.closSem.v',X94) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X95),s('type.list.list'('type.closSem.v'),X92))) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X92),s('type.list.list'('type.closSem.v'),X93)))) ) | $spl6),
  inference(superposition,[],[f123,f259])).
fof(f629,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.list.list'(X89),'const.list.NIL_0') != s('type.list.list'(X89),X91) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X88,fun(X89,bool)),X90),s('type.list.list'(X88),'const.list.NIL_0'),s('type.list.list'(X89),X91)))) ) | $spl6),
  inference(superposition,[],[f87,f259])).
fof(f628,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.list.list'(X85),'const.list.NIL_0') = s('type.list.list'(X85),X87) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X84,fun(X85,bool)),X86),s('type.list.list'(X84),'const.list.NIL_0'),s('type.list.list'(X85),X87)))) ) | $spl6),
  inference(superposition,[],[f86,f259])).
fof(f627,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.list.list'(X80),'const.list.NIL_0') != s('type.list.list'(X80),X83) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X80,fun(X81,bool)),X82),s('type.list.list'(X80),X83),s('type.list.list'(X81),'const.list.NIL_0')))) ) | $spl6),
  inference(superposition,[],[f85,f259])).
fof(f626,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.list.list'(X76),'const.list.NIL_0') = s('type.list.list'(X76),X79) | s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X76,fun(X77,bool)),X78),s('type.list.list'(X76),X79),s('type.list.list'(X77),'const.list.NIL_0')))) ) | $spl6),
  inference(superposition,[],[f84,f259])).
fof(f625,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun(X73,bool),X74),s(X73,sK1(X75,X74,X73))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75)))) ) | $spl6),
  inference(superposition,[],[f95,f259])).
fof(f624,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(X70,X72) = s(X70,sK1(X72,X71,X70)) | s(bool,'T_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,X72)))) ) | $spl6),
  inference(superposition,[],[f94,f259])).
fof(f623,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,X69)))) | s(X66,X68) != s(X66,X69) | s(bool,'T_0') = s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))) ) | $spl6),
  inference(superposition,[],[f93,f259])).
fof(f622,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,happ(s(fun(X62,bool),X63),s(X62,X65)))) | s(X62,X64) != s(X62,X65) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),X63),s(X62,X64)))) ) | $spl6),
  inference(superposition,[],[f93,f259])).
fof(f621,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,happ(s(fun(X59,bool),X60),s(X59,X61)))) | s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),X60),s(X59,sK1(X61,X60,X59))))) ) | $spl6),
  inference(superposition,[],[f95,f259])).
fof(f695,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X55),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X57))))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X58),s('type.list.list'('type.closSem.v'),X57)))) | s('type.closSem.v',X55) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X58)))) ) | $spl6),
  inference(forward_demodulation,[],[f620,f92])).
fof(f620,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X58),s('type.list.list'('type.closSem.v'),X57)))) | s('type.closSem.v',X55) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X55))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X56),s('type.list.list'('type.closSem.v'),X57)))))) ) | $spl6),
  inference(superposition,[],[f123,f259])).
fof(f694,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X53),s('type.list.list'('type.closSem.v'),X54))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X54,X53,X52)),s('type.list.list'('type.closSem.v'),X54))))) ) | $spl50),
  inference(cnf_transformation,[],[f694_D])).
fof(f694_D,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X53),s('type.list.list'('type.closSem.v'),X54))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X54,X53,X52)),s('type.list.list'('type.closSem.v'),X54))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f693,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X52),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X53),s('type.list.list'('type.closSem.v'),X54))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X54,X53,X52)),s('type.list.list'('type.closSem.v'),X54))))) ) | $spl6),
  inference(forward_demodulation,[],[f619,f92])).
fof(f619,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),sK3(X54,X53,X52)),s('type.list.list'('type.closSem.v'),X54)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X52))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X53),s('type.list.list'('type.closSem.v'),X54)))))) ) | $spl6),
  inference(superposition,[],[f122,f259])).
fof(f691,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X49),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),X51))))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X49) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),sK3(X51,X50,X49))))) ) | $spl6),
  inference(forward_demodulation,[],[f618,f92])).
fof(f618,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X49) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),sK3(X51,X50,X49)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v',X49))),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X50),s('type.list.list'('type.closSem.v'),X51)))))) ) | $spl6),
  inference(superposition,[],[f121,f259])).
fof(f689,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X46))),s('type.closSem.v',X47))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X46),s('type.list.list'('type.closSem.v'),X48)))) | s('type.closSem.v',X47) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X48)))) ) | $spl6),
  inference(forward_demodulation,[],[f617,f92])).
fof(f617,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X46),s('type.list.list'('type.closSem.v'),X48)))) | s('type.closSem.v',X47) != s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X45),s('type.list.list'('type.closSem.v'),X46))))),s('type.closSem.v',X47)))) ) | $spl6),
  inference(superposition,[],[f120,f259])).
fof(f688,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X42),s('type.list.list'('type.closSem.v'),X43))),s('type.closSem.v',X44))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X43),s('type.list.list'('type.closSem.v'),sK2(X43,X42,X44)))))) ) | $spl44),
  inference(cnf_transformation,[],[f688_D])).
fof(f688_D,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X42),s('type.list.list'('type.closSem.v'),X43))),s('type.closSem.v',X44))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X43),s('type.list.list'('type.closSem.v'),sK2(X43,X42,X44)))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f687,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X42),s('type.list.list'('type.closSem.v'),X43))),s('type.closSem.v',X44))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X43),s('type.list.list'('type.closSem.v'),sK2(X43,X42,X44)))))) ) | $spl6),
  inference(forward_demodulation,[],[f616,f92])).
fof(f616,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.list.list'('type.closSem.v'),X43),s('type.list.list'('type.closSem.v'),sK2(X43,X42,X44))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X42),s('type.list.list'('type.closSem.v'),X43))))),s('type.closSem.v',X44)))) ) | $spl6),
  inference(superposition,[],[f119,f259])).
fof(f685,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),X40))),s('type.closSem.v',X41))) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X41) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),sK2(X40,X39,X41))))) ) | $spl6),
  inference(forward_demodulation,[],[f615,f92])).
fof(f615,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X41) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),sK2(X40,X39,X41)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_numberProof.v_rel_0'),s('type.closSem.v','const.closSem.Block_2'(s('type.num.num',X39),s('type.list.list'('type.closSem.v'),X40))))),s('type.closSem.v',X41)))) ) | $spl6),
  inference(superposition,[],[f118,f259])).
fof(f613,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X36) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X35))),s('type.closSem.v',X36)))) ) | $spl6),
  inference(superposition,[],[f114,f259])).
fof(f683,plain,(
  ( ! [X33,X34] : (s('type.closSem.v',X34) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))),s('type.closSem.v',X34)))) ) | $spl38),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  ( ! [X33,X34] : (s('type.closSem.v',X34) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))),s('type.closSem.v',X34)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f612,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X34) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X33))),s('type.closSem.v',X34)))) ) | $spl6),
  inference(superposition,[],[f113,f259])).
fof(f611,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X32) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X31))),s('type.closSem.v',X32)))) ) | $spl6),
  inference(superposition,[],[f116,f259])).
fof(f681,plain,(
  ( ! [X30,X29] : (s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v',X30)))) ) | $spl34),
  inference(cnf_transformation,[],[f681_D])).
fof(f681_D,plain,(
  ( ! [X30,X29] : (s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v',X30)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f610,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X30) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X29))),s('type.closSem.v',X30)))) ) | $spl6),
  inference(superposition,[],[f115,f259])).
fof(f608,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X27) != s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X27),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X28)))))) ) | $spl6),
  inference(superposition,[],[f112,f259])).
fof(f679,plain,(
  ( ! [X26,X25] : (s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))))) ) | $spl30),
  inference(cnf_transformation,[],[f679_D])).
fof(f679_D,plain,(
  ( ! [X26,X25] : (s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f607,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X25) = s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X25),s('type.closSem.v','const.closSem.Number_1'(s('type.integer.int',X26)))))) ) | $spl6),
  inference(superposition,[],[f111,f259])).
fof(f606,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X23) != s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X23),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X24)))))) ) | $spl6),
  inference(superposition,[],[f110,f259])).
fof(f677,plain,(
  ( ! [X21,X22] : (s('type.closSem.v',X21) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X21),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22)))))) ) | $spl26),
  inference(cnf_transformation,[],[f677_D])).
fof(f677_D,plain,(
  ( ! [X21,X22] : (s('type.closSem.v',X21) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X21),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f605,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v',X21) = s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22))) | s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X21),s('type.closSem.v','const.closSem.RefPtr_1'(s('type.num.num',X22)))))) ) | $spl6),
  inference(superposition,[],[f109,f259])).
fof(f676,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl24),
  inference(cnf_transformation,[],[f676_D])).
fof(f676_D,plain,(
  ( ! [X19,X20] : (s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f604,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f513,f259])).
fof(f662,plain,(
  ( ! [X18] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl0 | $spl6)),
  inference(duplicate_literal_removal,[],[f603])).
fof(f603,plain,(
  ( ! [X18] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f474,f259])).
fof(f675,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f600,f513])).
fof(f600,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X14) = s(bool,X15) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f91,f259])).
fof(f674,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f599,f513])).
fof(f599,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s(bool,X12) = s(bool,X13) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f91,f259])).
fof(f673,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f594,f474])).
fof(f594,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f88,f259])).
fof(f672,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f593,f474])).
fof(f593,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f88,f259])).
fof(f671,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f592,f513])).
fof(f592,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ~p(s(bool,X9))) ) | $spl6),
  inference(superposition,[],[f259,f91])).
fof(f667,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f591,f513])).
fof(f591,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X7)) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))))) ) | $spl6),
  inference(superposition,[],[f259,f91])).
fof(f666,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X4))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f581,f474])).
fof(f581,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f259,f88])).
fof(f663,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X2))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f580,f474])).
fof(f580,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl6),
  inference(superposition,[],[f259,f88])).
fof(f468,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) | $spl14),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f444,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) | ~p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f431,f103])).
fof(f443,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))) | ~p(s(bool,X4))) ) | $spl0),
  inference(resolution,[],[f431,f105])).
fof(f392,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2)),
  inference(equality_factoring,[],[f233])).
fof(f233,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f186,f108])).
fof(f186,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f176,f91])).
fof(f176,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f175,f88])).
fof(f396,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f394,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,X21) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2)),
  inference(duplicate_literal_removal,[],[f385])).
fof(f385,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,X21) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,X21) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f186,f233])).
fof(f424,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X16) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f381,f176])).
fof(f381,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | s(bool,X16) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f91,f233])).
fof(f423,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,X14) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f380,f176])).
fof(f380,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,X14) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X13)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f91,f233])).
fof(f404,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X10)) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f327,f176])).
fof(f327,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X11)) | ~p(s(bool,X10)) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f91,f233])).
fof(f403,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f326,f176])).
fof(f326,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,X8)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f91,f233])).
fof(f398,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X7) | ~p(s(bool,X7))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f313,f176])).
fof(f313,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,X8))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f233,f91])).
fof(f395,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X5))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f312,f176])).
fof(f312,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | s(bool,'T_0') = s(bool,X5) | ~p(s(bool,X6)) | ~p(s(bool,X5))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f233,f91])).
fof(f283,plain,(
  $false | $spl4),
  inference(resolution,[],[f258,f108])).
fof(f284,plain,(
  $false | $spl4),
  inference(resolution,[],[f258,f107])).
fof(f287,plain,(
  $false | $spl4),
  inference(resolution,[],[f258,f124])).
fof(f236,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X12) | s(bool,X12) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X13))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f186,f103])).
fof(f235,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X10) | s(bool,X10) = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ~p(s(bool,X11))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f186,f105])).
fof(f126,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | $spl0),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f128,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK0) | $spl2),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK0) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f82,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK0) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK0)) & (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f58])).
fof(f58,plain,(
  ? [X0] : ((p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0)) & (~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ? [X0] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) <~> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X6)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X6)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X6))),
  file('thm.clos_numberProof.v_rel_Unit|split|1','thm.clos_numberProof.v_rel_Unit|split|1')).
fof(f127,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f129,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK0) | $spl3),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK0) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f83,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK0) | ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK0))))),
  inference(cnf_transformation,[],[f59])).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 455
% Passive clauses: 11293
% Generated clauses: 30353
% Final active clauses: 312
% Final passive clauses: 8497
% Input formulas: 33
% Initial clauses: 42
% 
% Selected by SInE selection: 19
% SInE iterations: 1
% 
% Duplicate literals: 940
% Trivial inequalities: 219
% Fw subsumption resolutions: 2753
% Fw demodulations: 3435
% Bw demodulations: 62
% 
% Simple tautologies: 940
% Equational tautologies: 77
% Forward subsumptions: 10572
% 
% Binary resolution: 6411
% Factoring: 11
% Forward superposition: 10928
% Backward superposition: 4781
% Self superposition: 205
% Equality factoring: 70
% Equality resolution: 110
% 
% Split clauses: 664
% Split components: 1342
% SAT solver clauses: 686
% SAT solver unit clauses: 3
% SAT solver binary clauses: 15
% 
% Sat splits: 664
% Sat splitting refutations: 19
% 
% Memory used [KB]: 14711
% Time elapsed: 1.333 s
% ------------------------------
% ------------------------------
% Success in time 13.071 s
