Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% (104821)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1002
% Passive clauses: 6363
% Generated clauses: 19021
% Final active clauses: 996
% Final passive clauses: 4949
% Input formulas: 69
% Initial clauses: 420
% 
% Duplicate literals: 255
% Trivial inequalities: 21
% Fw subsumption resolutions: 1745
% Fw demodulations: 1635
% Bw demodulations: 4
% 
% Simple tautologies: 290
% Equational tautologies: 45
% Forward subsumptions: 8025
% 
% Binary resolution: 7412
% Factoring: 1261
% Forward superposition: 1983
% Backward superposition: 4018
% Self superposition: 39
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 651
% Split components: 1304
% SAT solver clauses: 898
% SAT solver unit clauses: 11
% SAT solver binary clauses: 58
% 
% Sat splits: 651
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 898
% TWLsolver calls for satisfiability: 445
% 
% Memory used [KB]: 11385
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.dec_to_exhProof
% (104855)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 467
% Passive clauses: 2940
% Generated clauses: 9864
% Final active clauses: 461
% Final passive clauses: 2466
% Input formulas: 69
% Initial clauses: 420
% 
% Duplicate literals: 102
% Trivial inequalities: 16
% Fw subsumption resolutions: 486
% Fw demodulations: 361
% Bw demodulations: 11
% Condensations: 117
% Global subsumptions: 18
% 
% Simple tautologies: 232
% Equational tautologies: 12
% Forward subsumptions: 2394
% 
% Binary resolution: 2947
% Factoring: 329
% Forward superposition: 1327
% Backward superposition: 731
% Self superposition: 39
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 3576
% SAT solver unit clauses: 401
% SAT solver binary clauses: 133
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 70
% 
% Sat splits: 16
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3504
% TWLsolver calls for satisfiability: 21230
% 
% Memory used [KB]: 21236
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.dec_to_exhProof
% (104888)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1481
% Passive clauses: 9106
% Generated clauses: 160125
% Final active clauses: 1255
% Final passive clauses: 3083
% Input formulas: 69
% Initial clauses: 420
% Discarded non-redundant clauses: 13150
% 
% Duplicate literals: 260
% Trivial inequalities: 44
% Fw subsumption resolutions: 3008
% Fw demodulations: 75578
% Bw demodulations: 1936
% Condensations: 258
% Global subsumptions: 587
% 
% Simple tautologies: 127
% Equational tautologies: 103
% Forward subsumptions: 41135
% Fw demodulations to eq. taut.: 16752
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 71631
% Factoring: 1
% Forward superposition: 5455
% Backward superposition: 4112
% Self superposition: 95
% Equality factoring: 26
% Equality resolution: 42
% 
% Split clauses: 53
% Split components: 106
% SAT solver clauses: 12285
% SAT solver unit clauses: 3478
% SAT solver binary clauses: 4033
% 
% Sat splits: 53
% Sat splitting refutations: 2484
% 
% Memory used [KB]: 119998
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f8725,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f824,f826,f828,f830,f831,f905,f932,f1043,f1036,f1373,f1374,f1375,f1466,f1465,f1464,f1473,f1475,f1546,f1545,f1551,f1474,f1552,f2115,f2116,f2616,f2618,f2619,f2620,f2617,f2814,f2815,f2824,f2829,f2884,f2885,f3324,f3325,f4078,f4077,f4098,f4097,f4211,f4199,f4198,f4197,f4400,f4401,f4402,f4403,f4404,f4405,f4721,f4723,f4724,f4725,f4726,f4727,f4722,f4728,f4729,f4716,f4715,f4738,f4737,f4822,f4824,f4825,f4826,f4827,f4828,f4823,f4830,f4832,f4833,f4831,f5072,f5074,f5075,f5076,f5077,f5078,f5073,f5079,f5080,f5067,f5066,f5086,f5085,f5484,f5486,f5487,f5488,f5485,f5561,f5563,f5564,f5565,f5566,f5567,f5562,f5568,f5569,f5570,f5572,f5573,f5571,f5887,f5889,f5890,f5891,f1471,f5888,f5964,f5966,f5967,f5968,f5969,f5970,f5965,f5971,f5972,f5973,f5975,f5976,f5974,f5959,f5958,f5979,f5978,f6487,f6489,f6490,f6491,f6492,f6493,f6488,f6494,f6495,f6482,f6481,f6500,f6499,f6781,f6783,f6784,f6785,f6786,f6787,f6782,f6788,f6789,f6776,f6775,f6794,f6793,f7133,f7134,f7296,f7308,f7655,f7702,f7714,f7704,f7664,f7968,f7969,f8064,f8065,f8044,f8043,f8042,f8070,f8066,f8278,f8279,f8286,f8287,f8513,f8517,f8659,f8523,f8527,f8665,f8529,f8658,f8533,f8666,f8660,f8538,f8540,f8542,f933,f8544,f936,f8545,f934,f8619,f937,f8667,f8715,f8716,f8724])).
fof(f8724,plain,(
  $false | $spl92),
  inference(subsumption_resolution,[],[f8723,f610])).
fof(f610,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4',aHL_FALSITY)).
fof(f8723,plain,(
  p(s(bool,'F_0')) | $spl92),
  inference(forward_demodulation,[],[f8717,f894])).
fof(f894,plain,(
  s(bool,'F_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))),
  inference(resolution,[],[f412,f573])).
fof(f573,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X16))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_F_EQ_F')).
fof(f412,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))))),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34,sK35])],[f146])).
fof(f146,plain,(
  ? [X0,X1] : ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X1))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7))))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.dec_to_exhProof.v_rel_rules|split|4')).
fof(f8717,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) | $spl92),
  inference(resolution,[],[f8659,f673])).
fof(f673,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ~sP16(sK74(X2,X1,X0),sK75(X2,X1,X0),sK76(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ! [X12,X13,X14] : (sP16(X12,X13,X14) | (sP14(X13,sK68(X14,X13,X12,X2,X1,X0),X12,sK70(X14,X13,X12,X2,X1,X0),sK69(X14,X13,X12,X2,X1,X0)) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X13),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK68(X14,X13,X12,X2,X1,X0)))),s('type.list.list'('type.conSem.v'),sK69(X14,X13,X12,X2,X1,X0)))),s('type.list.list'('type.exhSem.v'),sK70(X14,X13,X12,X2,X1,X0)))))) | (sP15(X14,sK71(X14,X13,X12,X2,X1,X0),X12,sK73(X14,X13,X12,X2,X1,X0),sK72(X14,X13,X12,X2,X1,X0)) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X14),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK71(X14,X13,X12,X2,X1,X0)))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),sK72(X14,X13,X12,X2,X1,X0)))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),sK73(X14,X13,X12,X2,X1,X0)))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2)))))) & ((~sP16(sK74(X2,X1,X0),sK75(X2,X1,X0),sK76(X2,X1,X0)) & ! [X9,X10,X11] : (~sP14(sK75(X2,X1,X0),X9,sK74(X2,X1,X0),X11,X10) | p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),sK75(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9))),s('type.list.list'('type.conSem.v'),X10))),s('type.list.list'('type.exhSem.v'),X11))))) & ! [X6,X7,X8] : (~sP15(sK76(X2,X1,X0),X6,sK74(X2,X1,X0),X8,X7) | p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),sK76(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK68,sK69,sK70,sK71,sK72,sK73,sK74,sK75,sK76])],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ! [X12,X13,X14] : (sP16(X12,X13,X14) | ? [X18,X19,X20] : (sP14(X13,X18,X12,X20,X19) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X13),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X18))),s('type.list.list'('type.conSem.v'),X19))),s('type.list.list'('type.exhSem.v'),X20))))) | ? [X15,X16,X17] : (sP15(X14,X15,X12,X17,X16) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X14),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X15))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X16))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X17))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2)))))) & (? [X3,X4,X5] : (~sP16(X3,X4,X5) & ! [X9,X10,X11] : (~sP14(X4,X9,X3,X11,X10) | p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X9))),s('type.list.list'('type.conSem.v'),X10))),s('type.list.list'('type.exhSem.v'),X11))))) & ! [X6,X7,X8] : (~sP15(X5,X6,X3,X8,X7) | p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(rectify,[],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ! [X3,X4,X5] : (sP16(X3,X4,X5) | ? [X14,X15,X16] : (sP14(X4,X14,X3,X16,X15) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) | ? [X6,X7,X8] : (sP15(X5,X6,X3,X8,X7) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2)))))) & (? [X3,X4,X5] : (~sP16(X3,X4,X5) & ! [X14,X15,X16] : (~sP14(X4,X14,X3,X16,X15) | p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) & ! [X6,X7,X8] : (~sP15(X5,X6,X3,X8,X7) | p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2))))))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> ! [X3,X4,X5] : (sP16(X3,X4,X5) | ? [X14,X15,X16] : (sP14(X4,X14,X3,X16,X15) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) | ? [X6,X7,X8] : (sP15(X5,X6,X3,X8,X7) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))))),
  inference(definition_folding,[],[f182,f210,f209,f208,f207,f206,f205,f204])).
fof(f204,plain,(
  ! [X5,X21,X23,X22] : (sP10(X5,X21,X23,X22) <=> ? [X32,X33,X34,X35,X36] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32),s('type.list.list'('type.string.char'),X33),s('type.conLang.exp',X34))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35),s('type.list.list'('type.string.char'),X33),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.conLang.exp',X34))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f205,plain,(
  ! [X5,X21,X23,X22] : (sP11(X5,X21,X23,X22) <=> ? [X27,X28,X29,X30,X31] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30),s('type.list.list'('type.string.char'),X29))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30))),s('type.list.list'('type.string.char'),X29))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f206,plain,(
  ! [X4,X21,X23,X22] : (sP12(X4,X21,X23,X22) <=> ? [X37,X38,X39] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37),s('type.list.list'('type.conSem.v'),X38))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37))),s('type.list.list'('type.exhSem.v'),X39))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X38))),s('type.list.list'('type.exhSem.v'),X39))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f207,plain,(
  ! [X4,X21,X23,X22] : (sP13(X4,X21,X23,X22) <=> ? [X24,X25] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X24))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X25))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X24))),s('type.list.list'('type.exhSem.v'),X25))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f208,plain,(
  ! [X4,X14,X3,X16,X15] : (sP14(X4,X14,X3,X16,X15) <=> ((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X15) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X16)) | ? [X17,X18,X19,X20] : (s('type.list.list'('type.conSem.v'),X15) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X17),s('type.list.list'('type.conSem.v'),X19))) & s('type.list.list'('type.exhSem.v'),X16) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X18),s('type.list.list'('type.exhSem.v'),X20))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.conSem.v',X17))),s('type.exhSem.v',X18)))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X19))),s('type.list.list'('type.exhSem.v'),X20)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f209,plain,(
  ! [X5,X6,X3,X8,X7] : (sP15(X5,X6,X3,X8,X7) <=> ((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)) | ? [X9,X10,X11,X12,X13] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.conSem.v',X10))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.exhSem.v',X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.conSem.v',X10))),s('type.exhSem.v',X11)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f210,plain,(
  ! [X3,X4,X5] : (sP16(X3,X4,X5) <=> ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP12(X4,X21,X23,X22) | sP10(X5,X21,X23,X22) | sP11(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP13(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f182,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> ! [X3,X4,X5] : (? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | ? [X37,X38,X39] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37),s('type.list.list'('type.conSem.v'),X38))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37))),s('type.list.list'('type.exhSem.v'),X39))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X38))),s('type.list.list'('type.exhSem.v'),X39))))) | ? [X32,X33,X34,X35,X36] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32),s('type.list.list'('type.string.char'),X33),s('type.conLang.exp',X34))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35),s('type.list.list'('type.string.char'),X33),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.conLang.exp',X34))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35))))) | ? [X27,X28,X29,X30,X31] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30),s('type.list.list'('type.string.char'),X29))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30))),s('type.list.list'('type.string.char'),X29))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27))))) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | ? [X24,X25] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X24))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X25))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X24))),s('type.list.list'('type.exhSem.v'),X25)))))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) | ? [X14,X15,X16] : (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X15) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X16)) | ? [X17,X18,X19,X20] : (s('type.list.list'('type.conSem.v'),X15) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X17),s('type.list.list'('type.conSem.v'),X19))) & s('type.list.list'('type.exhSem.v'),X16) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X18),s('type.list.list'('type.exhSem.v'),X20))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.conSem.v',X17))),s('type.exhSem.v',X18)))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X19))),s('type.list.list'('type.exhSem.v'),X20)))))) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) | ? [X6,X7,X8] : (((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)) | ? [X9,X10,X11,X12,X13] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.conSem.v',X10))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.exhSem.v',X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.conSem.v',X10))),s('type.exhSem.v',X11)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13)))))) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> ! [X3,X4,X5] : ((? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | ? [X37,X38,X39] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37),s('type.list.list'('type.conSem.v'),X38))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37))),s('type.list.list'('type.exhSem.v'),X39))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X38))),s('type.list.list'('type.exhSem.v'),X39))))) | ? [X32,X33,X34,X35,X36] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32),s('type.list.list'('type.string.char'),X33),s('type.conLang.exp',X34))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35),s('type.list.list'('type.string.char'),X33),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.conLang.exp',X34))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35))))) | ? [X27,X28,X29,X30,X31] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30),s('type.list.list'('type.string.char'),X29))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30))),s('type.list.list'('type.string.char'),X29))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27))))) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | ? [X24,X25] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X24))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X25))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X24))),s('type.list.list'('type.exhSem.v'),X25)))))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) | ? [X14,X15,X16] : (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X15) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X16)) | ? [X17,X18,X19,X20] : (s('type.list.list'('type.conSem.v'),X15) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X17),s('type.list.list'('type.conSem.v'),X19))) & s('type.list.list'('type.exhSem.v'),X16) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X18),s('type.list.list'('type.exhSem.v'),X20))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.conSem.v',X17))),s('type.exhSem.v',X18)))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X19))),s('type.list.list'('type.exhSem.v'),X20)))))) & ~p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) | ? [X6,X7,X8] : (((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)) | ? [X9,X10,X11,X12,X13] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.conSem.v',X10))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.exhSem.v',X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.conSem.v',X10))),s('type.exhSem.v',X11)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13)))))) & ~p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) <=> ! [X3,X4,X5] : ((! [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | ? [X37,X38,X39] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37),s('type.list.list'('type.conSem.v'),X38))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X37))),s('type.list.list'('type.exhSem.v'),X39))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X38))),s('type.list.list'('type.exhSem.v'),X39))))) | ? [X32,X33,X34,X35,X36] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32),s('type.list.list'('type.string.char'),X33),s('type.conLang.exp',X34))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35),s('type.list.list'('type.string.char'),X33),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.conLang.exp',X34))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X36),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X32))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X35))))) | ? [X27,X28,X29,X30,X31] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30),s('type.list.list'('type.string.char'),X29))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X30))),s('type.list.list'('type.string.char'),X29))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X31),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X28))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X27))))) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | ? [X24,X25] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X24))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X25))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.list.list'('type.conSem.v'),X24))),s('type.list.list'('type.exhSem.v'),X25)))))) => p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) & ! [X14,X15,X16] : (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X15) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X16)) | ? [X17,X18,X19,X20] : (s('type.list.list'('type.conSem.v'),X15) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X17),s('type.list.list'('type.conSem.v'),X19))) & s('type.list.list'('type.exhSem.v'),X16) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X18),s('type.list.list'('type.exhSem.v'),X20))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.conSem.v',X17))),s('type.exhSem.v',X18)))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X19))),s('type.list.list'('type.exhSem.v'),X20)))))) => p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X4),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X14))),s('type.list.list'('type.conSem.v'),X15))),s('type.list.list'('type.exhSem.v'),X16))))) & ! [X6,X7,X8] : (((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)) | ? [X9,X10,X11,X12,X13] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.conSem.v',X10))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X9),s('type.exhSem.v',X11))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.conSem.v',X10))),s('type.exhSem.v',X11)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X12))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X13)))))) => p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X5),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X7))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X8)))))) => p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X5,X8,X9] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X5),s('type.conSem.v',X8),s('type.exhSem.v',X9)))) <=> ! [X10,X11,X12] : ((! [X13,X14,X15] : ((? [X7] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X16,X17,X18] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Conv_2'(s('type.num.num','const.dec_to_exhProof.get_tag_1'(s('type.option.option'('type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X11),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13))),s('type.list.list'('type.conSem.v'),X17))),s('type.list.list'('type.exhSem.v'),X18))))) | ? [X19,X20,X21,X22,X23] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19),s('type.list.list'('type.string.char'),X20),s('type.conLang.exp',X21))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Closure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22),s('type.list.list'('type.string.char'),X20),s('type.exhLang.exp','const.dec_to_exh.compile_exp_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X23),s('type.conLang.exp',X21))))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X23),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22))))) | ? [X22,X19,X20,X24,X23] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X24),s('type.list.list'('type.string.char'),X20))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Recclosure_3'(s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),'const.dec_to_exh.compile_funs_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X23),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.conLang.exp'))),X24))),s('type.list.list'('type.string.char'),X20))) & p(s(bool,'const.finite_map.SUBMAP_2'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X23),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22))))) | ? [X7] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X7))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X7)))) | ? [X17,X18] : (s('type.conSem.v',X14) = s('type.conSem.v','const.conSem.Vectorv_1'(s('type.list.list'('type.conSem.v'),X17))) & s('type.exhSem.v',X15) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X11),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13))),s('type.list.list'('type.conSem.v'),X17))),s('type.list.list'('type.exhSem.v'),X18)))))) => p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X13))),s('type.conSem.v',X14))),s('type.exhSem.v',X15))))) & ! [X25,X26,X27] : (((s('type.list.list'('type.conSem.v'),'const.list.NIL_0') = s('type.list.list'('type.conSem.v'),X26) & s('type.list.list'('type.exhSem.v'),'const.list.NIL_0') = s('type.list.list'('type.exhSem.v'),X27)) | ? [X28,X29,X17,X18] : (s('type.list.list'('type.conSem.v'),X26) = s('type.list.list'('type.conSem.v'),'const.list.CONS_2'(s('type.conSem.v',X28),s('type.list.list'('type.conSem.v'),X17))) & s('type.list.list'('type.exhSem.v'),X27) = s('type.list.list'('type.exhSem.v'),'const.list.CONS_2'(s('type.exhSem.v',X29),s('type.list.list'('type.exhSem.v'),X18))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))),s('type.conSem.v',X28))),s('type.exhSem.v',X29)))) & p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X11),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))),s('type.list.list'('type.conSem.v'),X17))),s('type.list.list'('type.exhSem.v'),X18)))))) => p(s(bool,happ(s(fun('type.list.list'('type.exhSem.v'),bool),happ(s(fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.conSem.v'),fun('type.list.list'('type.exhSem.v'),bool))),X11),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X25))),s('type.list.list'('type.conSem.v'),X26))),s('type.list.list'('type.exhSem.v'),X27))))) & ! [X30,X31,X32] : (((s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X31) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.NIL_0') = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X32)) | ? [X20,X28,X29,X19,X22] : (s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X31) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X20),s('type.conSem.v',X28))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19))) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X32) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),'const.list.CONS_2'(s('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v'),'const.pair.,_2'(s('type.list.list'('type.string.char'),X20),s('type.exhSem.v',X29))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22))) & p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30))),s('type.conSem.v',X28))),s('type.exhSem.v',X29)))) & p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X19))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X22)))))) => p(s(bool,happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),bool))),X12),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X30))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.conSem.v')),X31))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),X32)))))) => p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X10),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X5))),s('type.conSem.v',X8))),s('type.exhSem.v',X9))))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.dec_to_exhProof.v_rel2_def')).
fof(f8716,plain,(
  $false | $spl92),
  inference(resolution,[],[f8659,f820])).
fof(f820,plain,(
  ~sP16(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),sK75('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),sK76('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34))),
  inference(resolution,[],[f412,f673])).
fof(f8715,plain,(
  $false | $spl94),
  inference(trivial_inequality_removal,[],[f8714])).
fof(f8714,plain,(
  s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) | $spl94),
  inference(equality_resolution,[],[f8660])).
fof(f8667,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | s(bool,X52) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) | $spl100),
  inference(cnf_transformation,[],[f8667_D])).
fof(f8667_D,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | s(bool,X52) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f937,plain,(
  ( ! [X24] : (p(s(bool,X24))) ) | $spl8),
  inference(cnf_transformation,[],[f937_D])).
fof(f937_D,plain,(
  ( ! [X24] : (p(s(bool,X24))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f8619,plain,(
  ( ! [X184,X183] : (p(s(bool,X183)) | ~p(s(bool,X184)) | s(bool,X184) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f823,f814])).
fof(f814,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X20,X128,X43,X129] : ((s(bool,X20) = s(bool,X128) & (p(s(bool,X128)) => s(bool,X43) = s(bool,X129))) => ((p(s(bool,X20)) => p(s(bool,X43))) <=> (p(s(bool,X128)) => p(s(bool,X129)))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CONG')).
fof(f823,plain,(
  ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))))),
  inference(resolution,[],[f412,f676])).
fof(f676,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v',X1),s('type.exhSem.v',X2)))) | ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74(X2,X1,X0)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0))),s('type.conSem.v',X1))),s('type.exhSem.v',X2))))) )),
  inference(cnf_transformation,[],[f344])).
fof(f934,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | ~p(s(bool,X12))) ) | $spl4),
  inference(cnf_transformation,[],[f934_D])).
fof(f934_D,plain,(
  ( ! [X14,X12] : (s(bool,X12) != s(bool,X14) | ~p(s(bool,X12))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f8545,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | s(bool,X73) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | s(bool,X72) != s(bool,X74)) )),
  inference(resolution,[],[f823,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f936,plain,(
  ( ! [X26] : (~p(s(bool,X26))) ) | $spl6),
  inference(cnf_transformation,[],[f936_D])).
fof(f936_D,plain,(
  ( ! [X26] : (~p(s(bool,X26))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f8544,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f823,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f933,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | p(s(bool,X6))) ) | $spl2),
  inference(cnf_transformation,[],[f933_D])).
fof(f933_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | p(s(bool,X6))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f8542,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X63) != s(bool,X65) | s(bool,X64) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f823,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f8540,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | s(bool,X57) != s(bool,X59)) )),
  inference(resolution,[],[f823,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f8538,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34)),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34))),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f823,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f8660,plain,(
  ( ! [X11] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11)))) ) | $spl94),
  inference(cnf_transformation,[],[f8660_D])).
fof(f8660_D,plain,(
  ( ! [X11] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8666,plain,(
  ( ! [X37,X36] : (sP30(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X36,X37)) ) | $spl98),
  inference(cnf_transformation,[],[f8666_D])).
fof(f8666_D,plain,(
  ( ! [X37,X36] : (sP30(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X36,X37)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8533,plain,(
  ( ! [X47,X45,X46] : (sP30(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X45,X46) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X47))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X47)))) )),
  inference(resolution,[],[f823,f744])).
fof(f744,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP30(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6)))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1,X2] : ((~sP30(X0,X1,X2) | (((s('type.conSem.v',sK116(X2,X1,X0)) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK118(X2,X1,X0)))) & s('type.exhSem.v',sK117(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK118(X2,X1,X0))))) | sP26(X1,sK115(X2,X1,X0),sK117(X2,X1,X0),sK116(X2,X1,X0)) | sP24(X2,sK115(X2,X1,X0),sK117(X2,X1,X0),sK116(X2,X1,X0)) | sP25(X2,sK115(X2,X1,X0),sK117(X2,X1,X0),sK116(X2,X1,X0)) | (s('type.conSem.v',sK116(X2,X1,X0)) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK119(X2,X1,X0)))) & s('type.exhSem.v',sK117(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK119(X2,X1,X0))))) | sP27(X1,sK115(X2,X1,X0),sK117(X2,X1,X0),sK116(X2,X1,X0))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK115(X2,X1,X0)))),s('type.conSem.v',sK116(X2,X1,X0)))),s('type.exhSem.v',sK117(X2,X1,X0))))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP26(X1,X3,X5,X4) & ~sP24(X2,X3,X5,X4) & ~sP25(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP27(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP30(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK115,sK116,sK117,sK118,sK119])],[f372])).
fof(f372,plain,(
  ! [X0,X1,X2] : ((~sP30(X0,X1,X2) | ? [X8,X9,X10] : ((? [X12] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X12))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X12)))) | sP26(X1,X8,X10,X9) | sP24(X2,X8,X10,X9) | sP25(X2,X8,X10,X9) | ? [X11] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11)))) | sP27(X1,X8,X10,X9)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8))),s('type.conSem.v',X9))),s('type.exhSem.v',X10)))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP26(X1,X3,X5,X4) & ~sP24(X2,X3,X5,X4) & ~sP25(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP27(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP30(X0,X1,X2)))),
  inference(rectify,[],[f371])).
fof(f371,plain,(
  ! [X3,X4,X5] : ((~sP30(X3,X4,X5) | ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP26(X4,X21,X23,X22) | sP24(X5,X21,X23,X22) | sP25(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP27(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23)))))) & (! [X21,X22,X23] : ((! [X40] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) & ~sP26(X4,X21,X23,X22) & ~sP24(X5,X21,X23,X22) & ~sP25(X5,X21,X23,X22) & ! [X26] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) & ~sP27(X4,X21,X23,X22)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) | sP30(X3,X4,X5)))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X3,X4,X5] : (sP30(X3,X4,X5) <=> ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP26(X4,X21,X23,X22) | sP24(X5,X21,X23,X22) | sP25(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP27(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP30])])).
fof(f8658,plain,(
  ( ! [X2] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2)))) ) | $spl90),
  inference(cnf_transformation,[],[f8658_D])).
fof(f8658_D,plain,(
  ( ! [X2] : (s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8529,plain,(
  ( ! [X37,X38,X36] : (sP30(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X36,X37) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X38))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X38)))) )),
  inference(resolution,[],[f823,f740])).
fof(f740,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (sP30(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7)))) )),
  inference(cnf_transformation,[],[f373])).
fof(f8665,plain,(
  ( ! [X23,X22] : (sP23(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X22,X23)) ) | $spl96),
  inference(cnf_transformation,[],[f8665_D])).
fof(f8665_D,plain,(
  ( ! [X23,X22] : (sP23(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X22,X23)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8527,plain,(
  ( ! [X33,X31,X32] : (sP23(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X31,X32) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X33))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X33)))) )),
  inference(resolution,[],[f823,f687])).
fof(f687,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP23(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6)))) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | (((s('type.conSem.v',sK78(X2,X1,X0)) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK80(X2,X1,X0)))) & s('type.exhSem.v',sK79(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK80(X2,X1,X0))))) | sP19(X1,sK77(X2,X1,X0),sK79(X2,X1,X0),sK78(X2,X1,X0)) | sP17(X2,sK77(X2,X1,X0),sK79(X2,X1,X0),sK78(X2,X1,X0)) | sP18(X2,sK77(X2,X1,X0),sK79(X2,X1,X0),sK78(X2,X1,X0)) | (s('type.conSem.v',sK78(X2,X1,X0)) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK81(X2,X1,X0)))) & s('type.exhSem.v',sK79(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK81(X2,X1,X0))))) | sP20(X1,sK77(X2,X1,X0),sK79(X2,X1,X0),sK78(X2,X1,X0))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK77(X2,X1,X0)))),s('type.conSem.v',sK78(X2,X1,X0)))),s('type.exhSem.v',sK79(X2,X1,X0))))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP19(X1,X3,X5,X4) & ~sP17(X2,X3,X5,X4) & ~sP18(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP20(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP23(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK77,sK78,sK79,sK80,sK81])],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | ? [X8,X9,X10] : ((? [X12] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X12))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X12)))) | sP19(X1,X8,X10,X9) | sP17(X2,X8,X10,X9) | sP18(X2,X8,X10,X9) | ? [X11] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11)))) | sP20(X1,X8,X10,X9)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8))),s('type.conSem.v',X9))),s('type.exhSem.v',X10)))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP19(X1,X3,X5,X4) & ~sP17(X2,X3,X5,X4) & ~sP18(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP20(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP23(X0,X1,X2)))),
  inference(rectify,[],[f345])).
fof(f345,plain,(
  ! [X3,X4,X5] : ((~sP23(X3,X4,X5) | ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP19(X4,X21,X23,X22) | sP17(X5,X21,X23,X22) | sP18(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP20(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23)))))) & (! [X21,X22,X23] : ((! [X40] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) & ~sP19(X4,X21,X23,X22) & ~sP17(X5,X21,X23,X22) & ~sP18(X5,X21,X23,X22) & ! [X26] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) & ~sP20(X4,X21,X23,X22)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) | sP23(X3,X4,X5)))),
  inference(nnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X3,X4,X5] : (sP23(X3,X4,X5) <=> ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP19(X4,X21,X23,X22) | sP17(X5,X21,X23,X22) | sP18(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP20(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f8523,plain,(
  ( ! [X24,X23,X22] : (sP23(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X22,X23) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X24))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X24)))) )),
  inference(resolution,[],[f823,f683])).
fof(f683,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (sP23(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7)))) )),
  inference(cnf_transformation,[],[f347])).
fof(f8659,plain,(
  ( ! [X0,X1] : (sP16(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X0,X1)) ) | $spl92),
  inference(cnf_transformation,[],[f8659_D])).
fof(f8659_D,plain,(
  ( ! [X0,X1] : (sP16(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X0,X1)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8517,plain,(
  ( ! [X10,X11,X9] : (sP16(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X9,X10) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11)))) )),
  inference(resolution,[],[f823,f629])).
fof(f629,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP16(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6)))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | (((s('type.conSem.v',sK40(X2,X1,X0)) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK42(X2,X1,X0)))) & s('type.exhSem.v',sK41(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK42(X2,X1,X0))))) | sP12(X1,sK39(X2,X1,X0),sK41(X2,X1,X0),sK40(X2,X1,X0)) | sP10(X2,sK39(X2,X1,X0),sK41(X2,X1,X0),sK40(X2,X1,X0)) | sP11(X2,sK39(X2,X1,X0),sK41(X2,X1,X0),sK40(X2,X1,X0)) | (s('type.conSem.v',sK40(X2,X1,X0)) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK43(X2,X1,X0)))) & s('type.exhSem.v',sK41(X2,X1,X0)) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK43(X2,X1,X0))))) | sP13(X1,sK39(X2,X1,X0),sK41(X2,X1,X0),sK40(X2,X1,X0))) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK39(X2,X1,X0)))),s('type.conSem.v',sK40(X2,X1,X0)))),s('type.exhSem.v',sK41(X2,X1,X0))))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP12(X1,X3,X5,X4) & ~sP10(X2,X3,X5,X4) & ~sP11(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP13(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP16(X0,X1,X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK39,sK40,sK41,sK42,sK43])],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | ? [X8,X9,X10] : ((? [X12] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X12))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X12)))) | sP12(X1,X8,X10,X9) | sP10(X2,X8,X10,X9) | sP11(X2,X8,X10,X9) | ? [X11] : (s('type.conSem.v',X9) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X11))) & s('type.exhSem.v',X10) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X11)))) | sP13(X1,X8,X10,X9)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8))),s('type.conSem.v',X9))),s('type.exhSem.v',X10)))))) & (! [X3,X4,X5] : ((! [X7] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7)))) & ~sP12(X1,X3,X5,X4) & ~sP10(X2,X3,X5,X4) & ~sP11(X2,X3,X5,X4) & ! [X6] : (s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X6))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X6)))) & ~sP13(X1,X3,X5,X4)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5))))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f319])).
fof(f319,plain,(
  ! [X3,X4,X5] : ((~sP16(X3,X4,X5) | ? [X21,X22,X23] : ((? [X40] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) | sP12(X4,X21,X23,X22) | sP10(X5,X21,X23,X22) | sP11(X5,X21,X23,X22) | ? [X26] : (s('type.conSem.v',X22) = s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) & s('type.exhSem.v',X23) = s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) | sP13(X4,X21,X23,X22)) & ~p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23)))))) & (! [X21,X22,X23] : ((! [X40] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X40))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X40)))) & ~sP12(X4,X21,X23,X22) & ~sP10(X5,X21,X23,X22) & ~sP11(X5,X21,X23,X22) & ! [X26] : (s('type.conSem.v',X22) != s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',X26))) | s('type.exhSem.v',X23) != s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',X26)))) & ~sP13(X4,X21,X23,X22)) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X3),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X21))),s('type.conSem.v',X22))),s('type.exhSem.v',X23))))) | sP16(X3,X4,X5)))),
  inference(nnf_transformation,[],[f210])).
fof(f8513,plain,(
  ( ! [X2,X0,X1] : (sP16(sK74('const.exhSem.Loc_1'(s('type.num.num',sK35)),'const.conSem.Loc_1'(s('type.num.num',sK35)),sK34),X0,X1) | s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X2))) | s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X2)))) )),
  inference(resolution,[],[f823,f625])).
fof(f625,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (sP16(X0,X1,X2) | p(s(bool,happ(s(fun('type.exhSem.v',bool),happ(s(fun('type.conSem.v',fun('type.exhSem.v',bool)),happ(s(fun('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),fun('type.conSem.v',fun('type.exhSem.v',bool))),X0),s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X3))),s('type.conSem.v',X4))),s('type.exhSem.v',X5)))) | s('type.exhSem.v',X5) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))) | s('type.conSem.v',X4) != s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7)))) )),
  inference(cnf_transformation,[],[f321])).
fof(f8287,plain,(
  ( ! [X219,X217] : (~sP6(X219,X217,'F_0') | ~p(s(bool,X219)) | s(bool,'F_0') = s(bool,X217)) ) | $spl88),
  inference(cnf_transformation,[],[f8287_D])).
fof(f8287_D,plain,(
  ( ! [X219,X217] : (~sP6(X219,X217,'F_0') | ~p(s(bool,X219)) | s(bool,'F_0') = s(bool,X217)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8286,plain,(
  ( ! [X218,X219,X217] : (s(bool,'F_0') = s(bool,X217) | p(s(bool,X218)) | ~p(s(bool,X219)) | ~sP6(X219,X217,'F_0')) ) | $spl14),
  inference(subsumption_resolution,[],[f8285,f573])).
fof(f8285,plain,(
  ( ! [X218,X219,X217] : (s(bool,'F_0') = s(bool,X217) | p(s(bool,X218)) | ~p(s(bool,X219)) | ~p(s(bool,X217)) | ~sP6(X219,X217,'F_0')) ) | $spl14),
  inference(subsumption_resolution,[],[f8262,f573])).
fof(f8262,plain,(
  ( ! [X218,X219,X217] : (s(bool,'F_0') = s(bool,X217) | p(s(bool,X218)) | ~p(s(bool,X219)) | s(bool,'F_0') != s(bool,X218) | ~p(s(bool,X217)) | ~sP6(X219,X217,'F_0')) ) | $spl14),
  inference(resolution,[],[f2655,f2227])).
fof(f2227,plain,(
  ( ! [X2,X3] : (sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f2226,f602])).
fof(f602,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f2226,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f2223,f2201])).
fof(f2201,plain,(
  ( ! [X4,X5] : (sP4(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2196,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2196,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP4(X4,'F_0',X5)) )),
  inference(resolution,[],[f2088,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f2088,plain,(
  ( ! [X8,X7] : (sP7(X8,'F_0',X7) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) )),
  inference(superposition,[],[f2043,f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP7(X2,X1,X0)) & (~sP7(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP7(X2,X1,X0))),
  inference(definition_folding,[],[f80,f198,f197,f196,f195])).
fof(f196,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X123,X124,X62] : ((p(s(bool,X62)) <=> s(bool,X123) = s(bool,X124)) <=> ((p(s(bool,X62)) | p(s(bool,X124)) | p(s(bool,X123))) & (p(s(bool,X62)) | ~p(s(bool,X123)) | ~p(s(bool,X124))) & (p(s(bool,X124)) | ~p(s(bool,X123)) | ~p(s(bool,X62))) & (p(s(bool,X123)) | ~p(s(bool,X124)) | ~p(s(bool,X62)))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.sat.sat_2|49|_prove_ltb')).
fof(f2043,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(superposition,[],[f902,f894])).
fof(f902,plain,(
  s(bool,'T_0') != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))),
  inference(resolution,[],[f412,f601])).
fof(f601,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f315])).
fof(f2223,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP7(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP6(X3,X2,'F_0') | ~sP4(X2,'F_0',X3)) )),
  inference(resolution,[],[f2202,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f2202,plain,(
  ( ! [X6,X7] : (sP5('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2197,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f257])).
fof(f257,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f196])).
fof(f2197,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP5('F_0',X6,X7)) )),
  inference(resolution,[],[f2088,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP5(X1,X2,X0) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f2655,plain,(
  ( ! [X10,X8,X11,X9] : (~sP7(X11,X9,X8) | s(bool,'F_0') = s(bool,X8) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl14),
  inference(superposition,[],[f2608,f492])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f2608,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | s(bool,'F_0') = s(bool,X47) | p(s(bool,X46))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f2591])).
fof(f2591,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X46) != s(bool,X47) | s(bool,'F_0') = s(bool,X47) | p(s(bool,X46))) ) | $spl14),
  inference(superposition,[],[f2130,f600])).
fof(f600,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.EQ_CLAUSES|split|2')).
fof(f2130,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f2116,f813])).
fof(f813,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4',aHL_BOOL_CASES)).
fof(f8279,plain,(
  ( ! [X6,X4,X7] : (~sP6(X6,X4,X7) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl86),
  inference(cnf_transformation,[],[f8279_D])).
fof(f8279_D,plain,(
  ( ! [X6,X4,X7] : (~sP6(X6,X4,X7) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X7)) | ~p(s(bool,X6))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8278,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,X4,X7)) ) | ($spl14 | $spl46)),
  inference(subsumption_resolution,[],[f8277,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f258])).
fof(f8277,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,X4,X7) | ~sP5(X7,X4,X6)) ) | ($spl14 | $spl46)),
  inference(subsumption_resolution,[],[f8276,f554])).
fof(f554,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X49,X50] : (s(bool,X49) = s(bool,X50) <=> ((p(s(bool,X49)) & p(s(bool,X50))) | (~p(s(bool,X49)) & ~p(s(bool,X50)))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.EQ_EXPAND')).
fof(f8276,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X7)) | ~sP6(X6,X4,X7) | ~sP5(X7,X4,X6)) ) | ($spl14 | $spl46)),
  inference(subsumption_resolution,[],[f8207,f5748])).
fof(f5748,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl46),
  inference(subsumption_resolution,[],[f5747,f487])).
fof(f5747,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | sP4(X0,X1,X2)) ) | $spl46),
  inference(subsumption_resolution,[],[f5727,f488])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f260])).
fof(f5727,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP4(X0,X1,X2)) ) | $spl46),
  inference(resolution,[],[f5569,f471])).
fof(f8207,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X7)) | ~sP6(X6,X4,X7) | ~sP5(X7,X4,X6) | ~sP4(X4,X7,X6)) ) | $spl14),
  inference(resolution,[],[f2655,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f8066,plain,(
  ( ! [X66,X67,X65] : (sP7(X67,X66,X65) | p(s(bool,X67)) | s(bool,'F_0') != s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) | $spl84),
  inference(cnf_transformation,[],[f8066_D])).
fof(f8066_D,plain,(
  ( ! [X66,X67,X65] : (sP7(X67,X66,X65) | p(s(bool,X67)) | s(bool,'F_0') != s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8070,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'F_0') != s(bool,X93) | s(bool,'F_0') = s(bool,X94) | p(s(bool,X92)) | sP7(X95,X94,X93) | p(s(bool,X95))) ) | $spl14),
  inference(subsumption_resolution,[],[f8028,f601])).
fof(f8028,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'F_0') != s(bool,X93) | s(bool,'F_0') = s(bool,X94) | p(s(bool,X92)) | sP7(X95,X94,X93) | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl14),
  inference(superposition,[],[f2654,f813])).
fof(f2654,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | s(bool,'F_0') = s(bool,X4) | p(s(bool,X6)) | sP7(X7,X4,X5) | p(s(bool,X7))) ) | $spl14),
  inference(superposition,[],[f2608,f490])).
fof(f8042,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X78) | p(s(bool,X76)) | sP7(X79,X78,X77) | p(s(bool,X79))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f8024])).
fof(f8024,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X78) | p(s(bool,X76)) | sP7(X79,X78,X77) | p(s(bool,X79)) | p(s(bool,X76))) ) | $spl14),
  inference(superposition,[],[f2654,f600])).
fof(f8043,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X74) | p(s(bool,X72)) | sP7(X75,X74,X73) | p(s(bool,X75))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f8023])).
fof(f8023,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X74) | p(s(bool,X72)) | sP7(X75,X74,X73) | p(s(bool,X75)) | p(s(bool,X72))) ) | $spl14),
  inference(superposition,[],[f2654,f598])).
fof(f598,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.EQ_CLAUSES|split|3')).
fof(f8044,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,'F_0') = s(bool,X66) | p(s(bool,X64)) | sP7(X67,X66,X65) | p(s(bool,X67))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f8021])).
fof(f8021,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,'F_0') = s(bool,X66) | p(s(bool,X64)) | sP7(X67,X66,X65) | p(s(bool,X67)) | p(s(bool,X64))) ) | $spl14),
  inference(superposition,[],[f2654,f573])).
fof(f8065,plain,(
  ( ! [X130,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X130)) | sP7(X130,X129,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) | $spl82),
  inference(cnf_transformation,[],[f8065_D])).
fof(f8065_D,plain,(
  ( ! [X130,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X130)) | sP7(X130,X129,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8064,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X128)) | sP7(X130,X129,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | p(s(bool,X130))) ) | $spl14),
  inference(subsumption_resolution,[],[f8005,f573])).
fof(f8005,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') != s(bool,X128) | s(bool,'F_0') = s(bool,X129) | p(s(bool,X128)) | sP7(X130,X129,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | p(s(bool,X130))) ) | $spl14),
  inference(superposition,[],[f2654,f894])).
fof(f7969,plain,(
  ( ! [X130,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X130)) | sP7(X130,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))),X129)) ) | $spl80),
  inference(cnf_transformation,[],[f7969_D])).
fof(f7969_D,plain,(
  ( ! [X130,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X130)) | sP7(X130,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))),X129)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7968,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') = s(bool,X129) | p(s(bool,X128)) | sP7(X130,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))),X129) | p(s(bool,X130))) ) | $spl14),
  inference(subsumption_resolution,[],[f7909,f573])).
fof(f7909,plain,(
  ( ! [X130,X128,X129] : (s(bool,'F_0') != s(bool,X128) | s(bool,'F_0') = s(bool,X129) | p(s(bool,X128)) | sP7(X130,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))),X129) | p(s(bool,X130))) ) | $spl14),
  inference(superposition,[],[f2653,f894])).
fof(f2653,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | sP7(X3,X1,X0) | p(s(bool,X3))) ) | $spl14),
  inference(superposition,[],[f2608,f490])).
fof(f7664,plain,(
  ( ! [X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,X24)) ) | $spl76),
  inference(cnf_transformation,[],[f7664_D])).
fof(f7664_D,plain,(
  ( ! [X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7704,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X28) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,X31) | s(bool,X29) != s(bool,X31)) ) | $spl14),
  inference(resolution,[],[f7696,f6903])).
fof(f6903,plain,(
  ( ! [X17,X15,X16] : (~sP4(X17,X15,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6886])).
fof(f6886,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP4(X17,X15,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl14),
  inference(resolution,[],[f6873,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP7(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f261])).
fof(f6873,plain,(
  ( ! [X6,X8,X7] : (sP7(X7,X6,X8) | s(bool,'T_0') = s(bool,X6) | p(s(bool,X7)) | ~sP4(X8,X6,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f6860,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f6860,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | sP7(X7,X6,X8) | p(s(bool,X7)) | ~sP6(X7,X8,X6) | ~sP4(X8,X6,X7)) ) | $spl14),
  inference(resolution,[],[f6833,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f6833,plain,(
  ( ! [X4,X5,X3] : (sP5(X3,X4,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl14),
  inference(subsumption_resolution,[],[f6832,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f258])).
fof(f6832,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | sP5(X3,X4,X5)) ) | $spl14),
  inference(subsumption_resolution,[],[f6808,f483])).
fof(f6808,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | p(s(bool,X5)) | sP5(X3,X4,X5)) ) | $spl14),
  inference(resolution,[],[f6773,f472])).
fof(f6773,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X1,X0) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) ) | $spl14),
  inference(equality_resolution,[],[f2462])).
fof(f2462,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,X4) | sP7(X7,X4,X5) | p(s(bool,X7))) ) | $spl14),
  inference(superposition,[],[f2433,f490])).
fof(f2433,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X55) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,X54)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f2416])).
fof(f2416,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X54) != s(bool,X55) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,X54)) ) | $spl14),
  inference(superposition,[],[f2131,f813])).
fof(f2131,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | ~p(s(bool,X2))) ) | $spl14),
  inference(superposition,[],[f2116,f604])).
fof(f604,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.EQ_CLAUSES|split|1')).
fof(f7696,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X2,X1,X3) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f7695,f487])).
fof(f7695,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | sP4(X2,X1,X3)) ) | $spl14),
  inference(subsumption_resolution,[],[f7666,f488])).
fof(f7666,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | sP4(X2,X1,X3)) ) | $spl14),
  inference(resolution,[],[f2622,f471])).
fof(f2622,plain,(
  ( ! [X6,X4,X7,X5] : (sP7(X7,X4,X5) | s(bool,'F_0') = s(bool,X6) | p(s(bool,X4)) | s(bool,X5) != s(bool,X6) | p(s(bool,X7))) ) | $spl14),
  inference(superposition,[],[f2608,f490])).
fof(f7714,plain,(
  ( ! [X23,X21,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') = s(bool,X20)) ) | $spl78),
  inference(cnf_transformation,[],[f7714_D])).
fof(f7714_D,plain,(
  ( ! [X23,X21,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X23) | s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') = s(bool,X20)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7702,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'F_0') = s(bool,X20) | p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23) | s(bool,X21) != s(bool,X23)) ) | $spl14),
  inference(resolution,[],[f7696,f3554])).
fof(f3554,plain,(
  ( ! [X4,X5,X3] : (~sP4(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3542])).
fof(f3542,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP4(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl14),
  inference(resolution,[],[f3486,f493])).
fof(f3486,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP4(X11,X9,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f3483,f481])).
fof(f3483,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP7(X10,X9,X11) | p(s(bool,X10)) | ~sP6(X10,X11,X9) | ~sP4(X11,X9,X10)) ) | $spl14),
  inference(resolution,[],[f3479,f477])).
fof(f3479,plain,(
  ( ! [X10,X11,X9] : (sP5(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f3478,f484])).
fof(f3478,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP5(X10,X9,X11)) ) | $spl14),
  inference(subsumption_resolution,[],[f3466,f483])).
fof(f3466,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP5(X10,X9,X11)) ) | $spl14),
  inference(resolution,[],[f3295,f472])).
fof(f3295,plain,(
  ( ! [X6,X8,X7] : (sP7(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl14),
  inference(resolution,[],[f3284,f481])).
fof(f3284,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | sP7(X5,X4,X3) | ~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f3283,f602])).
fof(f3283,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3) | ~p(s(bool,X3)) | ~sP6(X5,X3,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f3278,f3264])).
fof(f3264,plain,(
  ( ! [X6,X8,X7] : (sP4(X7,X6,X8) | s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f3250,f487])).
fof(f3250,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP4(X7,X6,X8)) ) | $spl14),
  inference(resolution,[],[f2149,f471])).
fof(f2149,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X3,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl14),
  inference(superposition,[],[f2116,f490])).
fof(f3278,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3) | ~p(s(bool,X3)) | ~sP6(X5,X3,X4) | ~sP4(X3,X4,X5)) ) | $spl14),
  inference(resolution,[],[f3265,f475])).
fof(f3265,plain,(
  ( ! [X10,X11,X9] : (sP5(X9,X10,X11) | s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl14),
  inference(subsumption_resolution,[],[f3251,f483])).
fof(f3251,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') != s(bool,X10) | p(s(bool,X11)) | sP5(X9,X10,X11)) ) | $spl14),
  inference(resolution,[],[f2149,f472])).
fof(f7655,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X26)) | s(bool,'T_0') = s(bool,X27) | s(bool,X25) != s(bool,X27)) ) | $spl14),
  inference(resolution,[],[f7648,f6600])).
fof(f6600,plain,(
  ( ! [X17,X15,X16] : (~sP4(X15,X17,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6583])).
fof(f6583,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP4(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl14),
  inference(resolution,[],[f6577,f493])).
fof(f6577,plain,(
  ( ! [X6,X8,X7] : (sP7(X7,X8,X6) | s(bool,'T_0') = s(bool,X6) | p(s(bool,X7)) | ~sP4(X6,X8,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f6565,f481])).
fof(f6565,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | sP7(X7,X8,X6) | p(s(bool,X7)) | ~sP6(X7,X6,X8) | ~sP4(X6,X8,X7)) ) | $spl14),
  inference(resolution,[],[f6539,f477])).
fof(f6539,plain,(
  ( ! [X4,X5,X3] : (sP5(X4,X3,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl14),
  inference(subsumption_resolution,[],[f6538,f485])).
fof(f6538,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | sP5(X4,X3,X5)) ) | $spl14),
  inference(subsumption_resolution,[],[f6514,f483])).
fof(f6514,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | p(s(bool,X5)) | sP5(X4,X3,X5)) ) | $spl14),
  inference(resolution,[],[f6479,f472])).
fof(f6479,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) ) | $spl14),
  inference(equality_resolution,[],[f2461])).
fof(f2461,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0) | sP7(X3,X1,X0) | p(s(bool,X3))) ) | $spl14),
  inference(superposition,[],[f2433,f490])).
fof(f7648,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X1,X2,X3) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f7647,f487])).
fof(f7647,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | sP4(X1,X2,X3)) ) | $spl14),
  inference(subsumption_resolution,[],[f7617,f489])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f260])).
fof(f7617,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | p(s(bool,X3)) | sP4(X1,X2,X3)) ) | $spl14),
  inference(resolution,[],[f2621,f471])).
fof(f2621,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X3,X1,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X0)) | s(bool,X1) != s(bool,X2) | p(s(bool,X3))) ) | $spl14),
  inference(superposition,[],[f2608,f490])).
fof(f7308,plain,(
  ( ! [X23,X21,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X23) | s(bool,'T_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X20)) ) | $spl74),
  inference(cnf_transformation,[],[f7308_D])).
fof(f7308_D,plain,(
  ( ! [X23,X21,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) != s(bool,X23) | s(bool,'T_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7296,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X20) | p(s(bool,X22)) | s(bool,'T_0') != s(bool,X23) | s(bool,X21) != s(bool,X23)) ) | $spl14),
  inference(resolution,[],[f7290,f3540])).
fof(f3540,plain,(
  ( ! [X4,X5,X3] : (~sP5(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3526])).
fof(f3526,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP5(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl14),
  inference(resolution,[],[f3335,f493])).
fof(f3335,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X11,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10)) | ~sP5(X11,X9,X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f3331,f481])).
fof(f3331,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP7(X10,X11,X9) | p(s(bool,X10)) | ~sP6(X10,X9,X11) | ~sP5(X11,X9,X10)) ) | $spl14),
  inference(resolution,[],[f3327,f477])).
fof(f3327,plain,(
  ( ! [X6,X8,X7] : (sP4(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f3326,f487])).
fof(f3326,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP4(X7,X6,X8)) ) | $spl14),
  inference(subsumption_resolution,[],[f3312,f488])).
fof(f3312,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP4(X7,X6,X8)) ) | $spl14),
  inference(resolution,[],[f3289,f471])).
fof(f3289,plain,(
  ( ! [X6,X8,X7] : (sP7(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl14),
  inference(resolution,[],[f3276,f481])).
fof(f3276,plain,(
  ( ! [X6,X8,X7] : (~sP6(X8,X6,X7) | sP7(X8,X7,X6) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f3275,f485])).
fof(f3275,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP7(X8,X7,X6) | p(s(bool,X7)) | ~sP6(X8,X6,X7) | ~sP5(X7,X6,X8)) ) | $spl14),
  inference(subsumption_resolution,[],[f3268,f573])).
fof(f3268,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | sP7(X8,X7,X6) | p(s(bool,X7)) | ~sP6(X8,X6,X7) | ~sP5(X7,X6,X8)) ) | $spl14),
  inference(resolution,[],[f3264,f476])).
fof(f7290,plain,(
  ( ! [X6,X4,X7,X5] : (sP5(X6,X4,X7) | s(bool,X5) != s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl14),
  inference(subsumption_resolution,[],[f7289,f484])).
fof(f7289,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) != s(bool,X6) | sP5(X6,X4,X7)) ) | $spl14),
  inference(subsumption_resolution,[],[f7261,f483])).
fof(f7261,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) != s(bool,X6) | p(s(bool,X7)) | sP5(X6,X4,X7)) ) | $spl14),
  inference(resolution,[],[f2493,f472])).
fof(f2493,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X3,X1,X0) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) != s(bool,X2) | p(s(bool,X3))) ) | $spl14),
  inference(superposition,[],[f2433,f490])).
fof(f7134,plain,(
  ( ! [X2,X3,X1] : (~sP6(X2,X3,X1) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X3))) ) | $spl72),
  inference(cnf_transformation,[],[f7134_D])).
fof(f7134_D,plain,(
  ( ! [X2,X3,X1] : (~sP6(X2,X3,X1) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X3))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7133,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP6(X2,X3,X1)) ) | $spl14),
  inference(subsumption_resolution,[],[f7132,f489])).
fof(f7132,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP6(X2,X3,X1) | ~sP4(X3,X1,X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f7131,f483])).
fof(f7131,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP6(X2,X3,X1) | ~sP5(X1,X3,X2) | ~sP4(X3,X1,X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f7072,f553])).
fof(f553,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f285])).
fof(f7072,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X3)) | ~sP6(X2,X3,X1) | ~sP5(X1,X3,X2) | ~sP4(X3,X1,X2)) ) | $spl14),
  inference(resolution,[],[f2464,f475])).
fof(f2464,plain,(
  ( ! [X14,X12,X15,X13] : (~sP7(X15,X12,X13) | ~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X12) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl14),
  inference(superposition,[],[f2433,f492])).
fof(f6793,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X82,X81) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,'T_0') = s(bool,X82)) ) | $spl70),
  inference(cnf_transformation,[],[f6793_D])).
fof(f6793_D,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X82,X81) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,'T_0') = s(bool,X82)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f6794,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X90,X89) | p(s(bool,X91))) ) | $spl14),
  inference(subsumption_resolution,[],[f6759,f597])).
fof(f597,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f313])).
fof(f6759,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X90,X89) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2462,f813])).
fof(f6775,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X86,X85) | p(s(bool,X87))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6758])).
fof(f6758,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X86,X85) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2462,f604])).
fof(f6776,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X82,X81) | p(s(bool,X83))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6757])).
fof(f6757,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X82,X81) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2462,f602])).
fof(f6789,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X90,X88) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,'T_0') = s(bool,X90)) ) | $spl68),
  inference(cnf_transformation,[],[f6789_D])).
fof(f6789_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X90,X88) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,'T_0') = s(bool,X90)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f6788,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X90,X88) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f6727,f602])).
fof(f6727,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X90,X88) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2462,f813])).
fof(f6782,plain,(
  ( ! [X28,X24,X27] : (s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X27,X24)) ) | $spl66),
  inference(cnf_transformation,[],[f6782_D])).
fof(f6782_D,plain,(
  ( ! [X28,X24,X27] : (s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X27,X24)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f6787,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X86,X84) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f6726,f602])).
fof(f6726,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X85)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X86,X84) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2462,f604])).
fof(f6786,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X82,X80) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f6725,f602])).
fof(f6725,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X82,X80) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2462,f602])).
fof(f6785,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X62) | sP7(X63,X62,X59) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f6720,f553])).
fof(f6720,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X62) | sP7(X63,X62,X59) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f2462,f558])).
fof(f558,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X49,X50] : ((p(s(bool,X49)) => p(s(bool,X50))) => ((p(s(bool,X50)) => p(s(bool,X49))) => s(bool,X49) = s(bool,X50)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f6784,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X56)) | s(bool,'T_0') = s(bool,X57) | sP7(X58,X57,X54) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f6719,f553])).
fof(f6719,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | ~p(s(bool,X56)) | s(bool,'T_0') = s(bool,X57) | sP7(X58,X57,X54) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f2462,f558])).
fof(f6783,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | sP7(X33,X32,X29) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f6714,f553])).
fof(f6714,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | sP7(X33,X32,X29) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f2462,f553])).
fof(f6781,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X26)) | s(bool,'T_0') = s(bool,X27) | sP7(X28,X27,X24) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f6713,f553])).
fof(f6713,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X26)) | s(bool,'T_0') = s(bool,X27) | sP7(X28,X27,X24) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f2462,f553])).
fof(f6499,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X81,X82) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,'T_0') = s(bool,X82)) ) | $spl64),
  inference(cnf_transformation,[],[f6499_D])).
fof(f6499_D,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X81,X82) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,'T_0') = s(bool,X82)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f6500,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X89,X90) | p(s(bool,X91))) ) | $spl14),
  inference(subsumption_resolution,[],[f6465,f597])).
fof(f6465,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X89,X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2461,f813])).
fof(f6481,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X85,X86) | p(s(bool,X87))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6464])).
fof(f6464,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X85,X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2461,f604])).
fof(f6482,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X81,X82) | p(s(bool,X83))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f6463])).
fof(f6463,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X81,X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2461,f602])).
fof(f6495,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X88,X90) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,'T_0') = s(bool,X90)) ) | $spl62),
  inference(cnf_transformation,[],[f6495_D])).
fof(f6495_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X88,X90) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,'T_0') = s(bool,X90)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f6494,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X88,X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f6433,f602])).
fof(f6433,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X89)) | s(bool,'T_0') = s(bool,X90) | sP7(X91,X88,X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2461,f813])).
fof(f6488,plain,(
  ( ! [X28,X24,X27] : (s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X24,X27)) ) | $spl60),
  inference(cnf_transformation,[],[f6488_D])).
fof(f6488_D,plain,(
  ( ! [X28,X24,X27] : (s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X24,X27)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f6493,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X84,X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f6432,f602])).
fof(f6432,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X85)) | s(bool,'T_0') = s(bool,X86) | sP7(X87,X84,X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2461,f604])).
fof(f6492,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X80,X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f6431,f602])).
fof(f6431,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X82) | sP7(X83,X80,X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2461,f602])).
fof(f6491,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X62) | sP7(X63,X59,X62) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f6426,f553])).
fof(f6426,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X62) | sP7(X63,X59,X62) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f2461,f558])).
fof(f6490,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X56)) | s(bool,'T_0') = s(bool,X57) | sP7(X58,X54,X57) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f6425,f553])).
fof(f6425,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | ~p(s(bool,X56)) | s(bool,'T_0') = s(bool,X57) | sP7(X58,X54,X57) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f2461,f558])).
fof(f6489,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | sP7(X33,X29,X32) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f6420,f553])).
fof(f6420,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | sP7(X33,X29,X32) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f2461,f553])).
fof(f6487,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X26)) | s(bool,'T_0') = s(bool,X27) | sP7(X28,X24,X27) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f6419,f553])).
fof(f6419,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X26)) | s(bool,'T_0') = s(bool,X27) | sP7(X28,X24,X27) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f2461,f553])).
fof(f5978,plain,(
  ( ! [X83,X81,X82] : (sP7(X82,X81,X83) | p(s(bool,X83)) | p(s(bool,X82)) | s(bool,'T_0') != s(bool,X81)) ) | $spl58),
  inference(cnf_transformation,[],[f5978_D])).
fof(f5978_D,plain,(
  ( ! [X83,X81,X82] : (sP7(X82,X81,X83) | p(s(bool,X83)) | p(s(bool,X82)) | s(bool,'T_0') != s(bool,X81)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5979,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | sP7(X90,X89,X91) | p(s(bool,X90)) | p(s(bool,X91))) ) | $spl14),
  inference(subsumption_resolution,[],[f5942,f597])).
fof(f5942,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X88)) | sP7(X90,X89,X91) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f5876,f813])).
fof(f5876,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X77) != s(bool,X78) | ~p(s(bool,X78)) | sP7(X79,X77,X76) | p(s(bool,X79)) | p(s(bool,X76))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f5854])).
fof(f5854,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X77) != s(bool,X78) | ~p(s(bool,X78)) | sP7(X79,X77,X76) | p(s(bool,X79)) | p(s(bool,X76))) ) | $spl14),
  inference(superposition,[],[f2398,f600])).
fof(f2398,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6)) | sP7(X7,X4,X5) | p(s(bool,X7))) ) | $spl14),
  inference(superposition,[],[f2131,f490])).
fof(f5958,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | sP7(X86,X85,X87) | p(s(bool,X86)) | p(s(bool,X87))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f5941])).
fof(f5941,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X84)) | sP7(X86,X85,X87) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f5876,f604])).
fof(f5959,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | sP7(X82,X81,X83) | p(s(bool,X82)) | p(s(bool,X83))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f5940])).
fof(f5940,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X80)) | sP7(X82,X81,X83) | p(s(bool,X82)) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f5876,f602])).
fof(f5974,plain,(
  ( ! [X99,X98] : (sP7(X98,'T_0',X99) | p(s(bool,X99)) | p(s(bool,X98))) ) | $spl56),
  inference(cnf_transformation,[],[f5974_D])).
fof(f5974_D,plain,(
  ( ! [X99,X98] : (sP7(X98,'T_0',X99) | p(s(bool,X99)) | p(s(bool,X98))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5976,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X105)) | sP7(X106,'T_0',X107) | p(s(bool,X106)) | p(s(bool,X107)) | ~p(s(bool,X104))) ) | $spl14),
  inference(subsumption_resolution,[],[f5914,f553])).
fof(f5914,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X105)) | sP7(X106,'T_0',X107) | p(s(bool,X106)) | p(s(bool,X107)) | ~p(s(bool,X104))) ) | $spl14),
  inference(superposition,[],[f5876,f602])).
fof(f5975,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X101)) | sP7(X102,'T_0',X103) | p(s(bool,X102)) | p(s(bool,X103)) | ~p(s(bool,X100))) ) | $spl14),
  inference(subsumption_resolution,[],[f5913,f553])).
fof(f5913,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X101)) | sP7(X102,'T_0',X103) | p(s(bool,X102)) | p(s(bool,X103)) | ~p(s(bool,X100))) ) | $spl14),
  inference(superposition,[],[f5876,f604])).
fof(f5973,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X97)) | sP7(X98,'T_0',X99) | p(s(bool,X98)) | p(s(bool,X99))) ) | $spl14),
  inference(subsumption_resolution,[],[f5912,f2131])).
fof(f5912,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X97)) | sP7(X98,'T_0',X99) | p(s(bool,X98)) | p(s(bool,X99)) | s(bool,'F_0') = s(bool,X96)) ) | $spl14),
  inference(superposition,[],[f5876,f813])).
fof(f5972,plain,(
  ( ! [X90,X88,X91] : (sP7(X90,X88,X91) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | p(s(bool,X90))) ) | $spl54),
  inference(cnf_transformation,[],[f5972_D])).
fof(f5972_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X90,X88,X91) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | p(s(bool,X90))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5971,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | sP7(X90,X88,X91) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f5910,f602])).
fof(f5910,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X89)) | sP7(X90,X88,X91) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f5876,f813])).
fof(f5965,plain,(
  ( ! [X28,X24,X27] : (sP7(X27,X24,X28) | ~p(s(bool,X24)) | p(s(bool,X28)) | p(s(bool,X27))) ) | $spl52),
  inference(cnf_transformation,[],[f5965_D])).
fof(f5965_D,plain,(
  ( ! [X28,X24,X27] : (sP7(X27,X24,X28) | ~p(s(bool,X24)) | p(s(bool,X28)) | p(s(bool,X27))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5970,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP7(X86,X84,X87) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f5909,f602])).
fof(f5909,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X85)) | sP7(X86,X84,X87) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f5876,f604])).
fof(f5969,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP7(X82,X80,X83) | p(s(bool,X82)) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f5908,f602])).
fof(f5908,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X81)) | sP7(X82,X80,X83) | p(s(bool,X82)) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f5876,f602])).
fof(f5968,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X61)) | sP7(X62,X59,X63) | p(s(bool,X62)) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f5903,f553])).
fof(f5903,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X61)) | sP7(X62,X59,X63) | p(s(bool,X62)) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f5876,f558])).
fof(f5967,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X56)) | sP7(X57,X54,X58) | p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f5902,f553])).
fof(f5902,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | ~p(s(bool,X56)) | sP7(X57,X54,X58) | p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f5876,f558])).
fof(f5966,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X31)) | sP7(X32,X29,X33) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f5897,f553])).
fof(f5897,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | sP7(X32,X29,X33) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f5876,f553])).
fof(f5964,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X26)) | sP7(X27,X24,X28) | p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f5896,f553])).
fof(f5896,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X26)) | sP7(X27,X24,X28) | p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f5876,f553])).
fof(f5888,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X36) != s(bool,X37) | p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X36,X34) | ~p(s(bool,X37))) ) | $spl50),
  inference(cnf_transformation,[],[f5888_D])).
fof(f5888_D,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X36) != s(bool,X37) | p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X36,X34) | ~p(s(bool,X37))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1471,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) | $spl10),
  inference(cnf_transformation,[],[f1471_D])).
fof(f1471_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5891,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X52)) | sP7(X53,X51,X49) | p(s(bool,X53)) | p(s(bool,X49))) ) | $spl14),
  inference(subsumption_resolution,[],[f5848,f597])).
fof(f5848,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X52)) | sP7(X53,X51,X49) | p(s(bool,X53)) | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl14),
  inference(superposition,[],[f2398,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f162])).
fof(f5890,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | sP7(X48,X46,X44) | p(s(bool,X48)) | p(s(bool,X44))) ) | $spl14),
  inference(subsumption_resolution,[],[f5847,f597])).
fof(f5847,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | sP7(X48,X46,X44) | p(s(bool,X48)) | p(s(bool,X45)) | p(s(bool,X44))) ) | $spl14),
  inference(superposition,[],[f2398,f555])).
fof(f5889,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42)) | sP7(X43,X41,X39) | p(s(bool,X43)) | p(s(bool,X39))) ) | $spl14),
  inference(subsumption_resolution,[],[f5846,f597])).
fof(f5846,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42)) | sP7(X43,X41,X39) | p(s(bool,X43)) | p(s(bool,X39)) | p(s(bool,X40))) ) | $spl14),
  inference(superposition,[],[f2398,f554])).
fof(f5887,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | sP7(X38,X36,X34) | p(s(bool,X38)) | p(s(bool,X34))) ) | $spl14),
  inference(subsumption_resolution,[],[f5845,f597])).
fof(f5845,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | sP7(X38,X36,X34) | p(s(bool,X38)) | p(s(bool,X35)) | p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f2398,f554])).
fof(f5571,plain,(
  ( ! [X99,X98] : (sP7(X98,X99,'T_0') | p(s(bool,X99)) | p(s(bool,X98))) ) | $spl48),
  inference(cnf_transformation,[],[f5571_D])).
fof(f5571_D,plain,(
  ( ! [X99,X98] : (sP7(X98,X99,'T_0') | p(s(bool,X99)) | p(s(bool,X98))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5573,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X105)) | sP7(X106,X107,'T_0') | p(s(bool,X106)) | p(s(bool,X107)) | ~p(s(bool,X104))) ) | $spl14),
  inference(subsumption_resolution,[],[f5511,f553])).
fof(f5511,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X105)) | sP7(X106,X107,'T_0') | p(s(bool,X106)) | p(s(bool,X107)) | ~p(s(bool,X104))) ) | $spl14),
  inference(superposition,[],[f5473,f602])).
fof(f5473,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X77) != s(bool,X78) | ~p(s(bool,X78)) | sP7(X79,X76,X77) | p(s(bool,X79)) | p(s(bool,X76))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f5451])).
fof(f5451,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X77) != s(bool,X78) | ~p(s(bool,X78)) | sP7(X79,X76,X77) | p(s(bool,X79)) | p(s(bool,X76))) ) | $spl14),
  inference(superposition,[],[f2397,f600])).
fof(f2397,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2)) | sP7(X3,X1,X0) | p(s(bool,X3))) ) | $spl14),
  inference(superposition,[],[f2131,f490])).
fof(f5572,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X101)) | sP7(X102,X103,'T_0') | p(s(bool,X102)) | p(s(bool,X103)) | ~p(s(bool,X100))) ) | $spl14),
  inference(subsumption_resolution,[],[f5510,f553])).
fof(f5510,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X101)) | sP7(X102,X103,'T_0') | p(s(bool,X102)) | p(s(bool,X103)) | ~p(s(bool,X100))) ) | $spl14),
  inference(superposition,[],[f5473,f604])).
fof(f5570,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X97)) | sP7(X98,X99,'T_0') | p(s(bool,X98)) | p(s(bool,X99))) ) | $spl14),
  inference(subsumption_resolution,[],[f5509,f2131])).
fof(f5509,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X97)) | sP7(X98,X99,'T_0') | p(s(bool,X98)) | p(s(bool,X99)) | s(bool,'F_0') = s(bool,X96)) ) | $spl14),
  inference(superposition,[],[f5473,f813])).
fof(f5569,plain,(
  ( ! [X90,X88,X91] : (sP7(X90,X91,X88) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | p(s(bool,X90))) ) | $spl46),
  inference(cnf_transformation,[],[f5569_D])).
fof(f5569_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X90,X91,X88) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | p(s(bool,X90))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5568,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | sP7(X90,X91,X88) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f5507,f602])).
fof(f5507,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | ~p(s(bool,X89)) | sP7(X90,X91,X88) | p(s(bool,X90)) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f5473,f813])).
fof(f5562,plain,(
  ( ! [X28,X24,X27] : (sP7(X27,X28,X24) | ~p(s(bool,X24)) | p(s(bool,X28)) | p(s(bool,X27))) ) | $spl44),
  inference(cnf_transformation,[],[f5562_D])).
fof(f5562_D,plain,(
  ( ! [X28,X24,X27] : (sP7(X27,X28,X24) | ~p(s(bool,X24)) | p(s(bool,X28)) | p(s(bool,X27))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5567,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | sP7(X86,X87,X84) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f5506,f602])).
fof(f5506,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | ~p(s(bool,X85)) | sP7(X86,X87,X84) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f5473,f604])).
fof(f5566,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP7(X82,X83,X80) | p(s(bool,X82)) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f5505,f602])).
fof(f5505,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | ~p(s(bool,X81)) | sP7(X82,X83,X80) | p(s(bool,X82)) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f5473,f602])).
fof(f5565,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,X61)) | sP7(X62,X63,X59) | p(s(bool,X62)) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f5500,f553])).
fof(f5500,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X61)) | sP7(X62,X63,X59) | p(s(bool,X62)) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f5473,f558])).
fof(f5564,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,X56)) | sP7(X57,X58,X54) | p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f5499,f553])).
fof(f5499,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | ~p(s(bool,X56)) | sP7(X57,X58,X54) | p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f5473,f558])).
fof(f5563,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X31)) | sP7(X32,X33,X29) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f5494,f553])).
fof(f5494,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | sP7(X32,X33,X29) | p(s(bool,X32)) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f5473,f553])).
fof(f5561,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X26)) | sP7(X27,X28,X24) | p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f5493,f553])).
fof(f5493,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X26)) | sP7(X27,X28,X24) | p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f5473,f553])).
fof(f5485,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X36) != s(bool,X37) | p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X34,X36) | ~p(s(bool,X37))) ) | $spl42),
  inference(cnf_transformation,[],[f5485_D])).
fof(f5485_D,plain,(
  ( ! [X37,X38,X36,X34] : (s(bool,X36) != s(bool,X37) | p(s(bool,X34)) | p(s(bool,X38)) | sP7(X38,X34,X36) | ~p(s(bool,X37))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5488,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X52)) | sP7(X53,X49,X51) | p(s(bool,X53)) | p(s(bool,X49))) ) | $spl14),
  inference(subsumption_resolution,[],[f5445,f597])).
fof(f5445,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | ~p(s(bool,X52)) | sP7(X53,X49,X51) | p(s(bool,X53)) | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl14),
  inference(superposition,[],[f2397,f555])).
fof(f5487,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | sP7(X48,X44,X46) | p(s(bool,X48)) | p(s(bool,X44))) ) | $spl14),
  inference(subsumption_resolution,[],[f5444,f597])).
fof(f5444,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | sP7(X48,X44,X46) | p(s(bool,X48)) | p(s(bool,X45)) | p(s(bool,X44))) ) | $spl14),
  inference(superposition,[],[f2397,f555])).
fof(f5486,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42)) | sP7(X43,X39,X41) | p(s(bool,X43)) | p(s(bool,X39))) ) | $spl14),
  inference(subsumption_resolution,[],[f5443,f597])).
fof(f5443,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42)) | sP7(X43,X39,X41) | p(s(bool,X43)) | p(s(bool,X39)) | p(s(bool,X40))) ) | $spl14),
  inference(superposition,[],[f2397,f554])).
fof(f5484,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | sP7(X38,X34,X36) | p(s(bool,X38)) | p(s(bool,X34))) ) | $spl14),
  inference(subsumption_resolution,[],[f5442,f597])).
fof(f5442,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | sP7(X38,X34,X36) | p(s(bool,X38)) | p(s(bool,X35)) | p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f2397,f554])).
fof(f5085,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,'F_0',X82) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82)) ) | $spl40),
  inference(cnf_transformation,[],[f5085_D])).
fof(f5085_D,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,'F_0',X82) | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5086,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | ~p(s(bool,X88)) | sP7(X91,'F_0',X90) | p(s(bool,X91))) ) | $spl14),
  inference(subsumption_resolution,[],[f5050,f597])).
fof(f5050,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | ~p(s(bool,X88)) | sP7(X91,'F_0',X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2378,f813])).
fof(f2378,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X19) != s(bool,X20) | ~p(s(bool,X21)) | sP7(X22,'F_0',X19) | p(s(bool,X22))) ) | $spl14),
  inference(superposition,[],[f2131,f490])).
fof(f5066,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | ~p(s(bool,X84)) | sP7(X87,'F_0',X86) | p(s(bool,X87))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f5049])).
fof(f5049,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | ~p(s(bool,X84)) | sP7(X87,'F_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2378,f604])).
fof(f5067,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82) | ~p(s(bool,X80)) | sP7(X83,'F_0',X82) | p(s(bool,X83))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f5048])).
fof(f5048,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82) | ~p(s(bool,X80)) | sP7(X83,'F_0',X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2378,f602])).
fof(f5080,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,'F_0',X90) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,X88) != s(bool,X90)) ) | $spl38),
  inference(cnf_transformation,[],[f5080_D])).
fof(f5080_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,'F_0',X90) | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,X88) != s(bool,X90)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5079,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X90) | ~p(s(bool,X89)) | sP7(X91,'F_0',X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f5018,f602])).
fof(f5018,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X88) != s(bool,X90) | ~p(s(bool,X89)) | sP7(X91,'F_0',X90) | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2378,f813])).
fof(f5073,plain,(
  ( ! [X28,X24,X27] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,'F_0',X27)) ) | $spl36),
  inference(cnf_transformation,[],[f5073_D])).
fof(f5073_D,plain,(
  ( ! [X28,X24,X27] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,'F_0',X27)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5078,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X86) | ~p(s(bool,X85)) | sP7(X87,'F_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f5017,f602])).
fof(f5017,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X84) != s(bool,X86) | ~p(s(bool,X85)) | sP7(X87,'F_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2378,f604])).
fof(f5077,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X82) | ~p(s(bool,X81)) | sP7(X83,'F_0',X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f5016,f602])).
fof(f5016,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X80) != s(bool,X82) | ~p(s(bool,X81)) | sP7(X83,'F_0',X82) | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2378,f602])).
fof(f5076,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X59) != s(bool,X62) | ~p(s(bool,X61)) | sP7(X63,'F_0',X62) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f5011,f553])).
fof(f5011,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | s(bool,X59) != s(bool,X62) | ~p(s(bool,X61)) | sP7(X63,'F_0',X62) | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f2378,f558])).
fof(f5075,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X54) != s(bool,X57) | ~p(s(bool,X56)) | sP7(X58,'F_0',X57) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f5010,f553])).
fof(f5010,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | s(bool,X54) != s(bool,X57) | ~p(s(bool,X56)) | sP7(X58,'F_0',X57) | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f2378,f558])).
fof(f5074,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X29) != s(bool,X32) | ~p(s(bool,X31)) | sP7(X33,'F_0',X32) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f5005,f553])).
fof(f5005,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X29) != s(bool,X32) | ~p(s(bool,X31)) | sP7(X33,'F_0',X32) | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f2378,f553])).
fof(f5072,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X26)) | sP7(X28,'F_0',X27) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f5004,f553])).
fof(f5004,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27) | ~p(s(bool,X26)) | sP7(X28,'F_0',X27) | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f2378,f553])).
fof(f4831,plain,(
  ( ! [X78] : (sP7(X78,'T_0','F_0') | p(s(bool,X78))) ) | $spl34),
  inference(cnf_transformation,[],[f4831_D])).
fof(f4831_D,plain,(
  ( ! [X78] : (sP7(X78,'T_0','F_0') | p(s(bool,X78))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4833,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X83)) | sP7(X84,'T_0','F_0') | p(s(bool,X84)) | ~p(s(bool,X82))) ) | $spl14),
  inference(subsumption_resolution,[],[f4772,f553])).
fof(f4772,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) != s(bool,X83) | ~p(s(bool,X83)) | sP7(X84,'T_0','F_0') | p(s(bool,X84)) | ~p(s(bool,X82))) ) | $spl14),
  inference(superposition,[],[f4713,f602])).
fof(f4713,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | sP7(X2,X0,'F_0') | p(s(bool,X2))) ) | $spl14),
  inference(equality_resolution,[],[f2377])).
fof(f2377,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X15) != s(bool,X16) | ~p(s(bool,X17)) | sP7(X18,X15,'F_0') | p(s(bool,X18))) ) | $spl14),
  inference(superposition,[],[f2131,f490])).
fof(f4832,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,X80)) | sP7(X81,'T_0','F_0') | p(s(bool,X81)) | ~p(s(bool,X79))) ) | $spl14),
  inference(subsumption_resolution,[],[f4771,f553])).
fof(f4771,plain,(
  ( ! [X80,X81,X79] : (s(bool,X79) != s(bool,X80) | ~p(s(bool,X80)) | sP7(X81,'T_0','F_0') | p(s(bool,X81)) | ~p(s(bool,X79))) ) | $spl14),
  inference(superposition,[],[f4713,f604])).
fof(f4830,plain,(
  ( ! [X78,X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | sP7(X78,'T_0','F_0') | p(s(bool,X78))) ) | $spl14),
  inference(subsumption_resolution,[],[f4770,f2131])).
fof(f4770,plain,(
  ( ! [X78,X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | sP7(X78,'T_0','F_0') | p(s(bool,X78)) | s(bool,'F_0') = s(bool,X76)) ) | $spl14),
  inference(superposition,[],[f4713,f813])).
fof(f4823,plain,(
  ( ! [X23,X20] : (sP7(X23,X20,'F_0') | ~p(s(bool,X20)) | p(s(bool,X23))) ) | $spl32),
  inference(cnf_transformation,[],[f4823_D])).
fof(f4823_D,plain,(
  ( ! [X23,X20] : (sP7(X23,X20,'F_0') | ~p(s(bool,X20)) | p(s(bool,X23))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4828,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X68)) | sP7(X69,X67,'F_0') | p(s(bool,X69)) | ~p(s(bool,X67))) ) | $spl14),
  inference(subsumption_resolution,[],[f4767,f602])).
fof(f4767,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,X68) | ~p(s(bool,X68)) | sP7(X69,X67,'F_0') | p(s(bool,X69)) | ~p(s(bool,X67))) ) | $spl14),
  inference(superposition,[],[f4713,f604])).
fof(f4827,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,X65)) | sP7(X66,X64,'F_0') | p(s(bool,X66)) | ~p(s(bool,X64))) ) | $spl14),
  inference(subsumption_resolution,[],[f4766,f602])).
fof(f4766,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | ~p(s(bool,X65)) | sP7(X66,X64,'F_0') | p(s(bool,X66)) | ~p(s(bool,X64))) ) | $spl14),
  inference(superposition,[],[f4713,f602])).
fof(f4826,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X50)) | sP7(X51,X48,'F_0') | p(s(bool,X51)) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl14),
  inference(subsumption_resolution,[],[f4761,f553])).
fof(f4761,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,X50) | ~p(s(bool,X50)) | sP7(X51,X48,'F_0') | p(s(bool,X51)) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl14),
  inference(superposition,[],[f4713,f558])).
fof(f4825,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X46)) | sP7(X47,X44,'F_0') | p(s(bool,X47)) | ~p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl14),
  inference(subsumption_resolution,[],[f4760,f553])).
fof(f4760,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X45) != s(bool,X46) | ~p(s(bool,X46)) | sP7(X47,X44,'F_0') | p(s(bool,X47)) | ~p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl14),
  inference(superposition,[],[f4713,f558])).
fof(f4824,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X26)) | sP7(X27,X24,'F_0') | p(s(bool,X27)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl14),
  inference(subsumption_resolution,[],[f4755,f553])).
fof(f4755,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X26)) | sP7(X27,X24,'F_0') | p(s(bool,X27)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl14),
  inference(superposition,[],[f4713,f553])).
fof(f4822,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,X22)) | sP7(X23,X20,'F_0') | p(s(bool,X23)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | $spl14),
  inference(subsumption_resolution,[],[f4754,f553])).
fof(f4754,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X21) != s(bool,X22) | ~p(s(bool,X22)) | sP7(X23,X20,'F_0') | p(s(bool,X23)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) ) | $spl14),
  inference(superposition,[],[f4713,f553])).
fof(f4737,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82)) ) | $spl30),
  inference(cnf_transformation,[],[f4737_D])).
fof(f4737_D,plain,(
  ( ! [X83,X81,X82] : (sP7(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4738,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | ~p(s(bool,X88)) | sP7(X91,X90,'F_0') | p(s(bool,X91))) ) | $spl14),
  inference(subsumption_resolution,[],[f4699,f597])).
fof(f4699,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X89) != s(bool,X90) | ~p(s(bool,X88)) | sP7(X91,X90,'F_0') | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2377,f813])).
fof(f4715,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | ~p(s(bool,X84)) | sP7(X87,X86,'F_0') | p(s(bool,X87))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4698])).
fof(f4698,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) != s(bool,X86) | ~p(s(bool,X84)) | sP7(X87,X86,'F_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2377,f604])).
fof(f4716,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82) | ~p(s(bool,X80)) | sP7(X83,X82,'F_0') | p(s(bool,X83))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4697])).
fof(f4697,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,X82) | ~p(s(bool,X80)) | sP7(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2377,f602])).
fof(f4729,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X90,'F_0') | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,X88) != s(bool,X90)) ) | $spl28),
  inference(cnf_transformation,[],[f4729_D])).
fof(f4729_D,plain,(
  ( ! [X90,X88,X91] : (sP7(X91,X90,'F_0') | s(bool,'F_0') = s(bool,X88) | p(s(bool,X91)) | s(bool,X88) != s(bool,X90)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4728,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X90) | ~p(s(bool,X89)) | sP7(X91,X90,'F_0') | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(subsumption_resolution,[],[f4667,f602])).
fof(f4667,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,X88) != s(bool,X90) | ~p(s(bool,X89)) | sP7(X91,X90,'F_0') | p(s(bool,X91)) | s(bool,'F_0') = s(bool,X88)) ) | $spl14),
  inference(superposition,[],[f2377,f813])).
fof(f4722,plain,(
  ( ! [X28,X24,X27] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X27,'F_0')) ) | $spl26),
  inference(cnf_transformation,[],[f4722_D])).
fof(f4722_D,plain,(
  ( ! [X28,X24,X27] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X24)) | p(s(bool,X28)) | sP7(X28,X27,'F_0')) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4727,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X86) | ~p(s(bool,X85)) | sP7(X87,X86,'F_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(subsumption_resolution,[],[f4666,f602])).
fof(f4666,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X85) | s(bool,X84) != s(bool,X86) | ~p(s(bool,X85)) | sP7(X87,X86,'F_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl14),
  inference(superposition,[],[f2377,f604])).
fof(f4726,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X82) | ~p(s(bool,X81)) | sP7(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(subsumption_resolution,[],[f4665,f602])).
fof(f4665,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X80) != s(bool,X82) | ~p(s(bool,X81)) | sP7(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X80))) ) | $spl14),
  inference(superposition,[],[f2377,f602])).
fof(f4725,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X59) != s(bool,X62) | ~p(s(bool,X61)) | sP7(X63,X62,'F_0') | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(subsumption_resolution,[],[f4660,f553])).
fof(f4660,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) != s(bool,X61) | s(bool,X59) != s(bool,X62) | ~p(s(bool,X61)) | sP7(X63,X62,'F_0') | p(s(bool,X63)) | ~p(s(bool,X60)) | ~p(s(bool,X59))) ) | $spl14),
  inference(superposition,[],[f2377,f558])).
fof(f4724,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X54) != s(bool,X57) | ~p(s(bool,X56)) | sP7(X58,X57,'F_0') | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(subsumption_resolution,[],[f4659,f553])).
fof(f4659,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | s(bool,X54) != s(bool,X57) | ~p(s(bool,X56)) | sP7(X58,X57,'F_0') | p(s(bool,X58)) | ~p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f2377,f558])).
fof(f4723,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X29) != s(bool,X32) | ~p(s(bool,X31)) | sP7(X33,X32,'F_0') | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(subsumption_resolution,[],[f4654,f553])).
fof(f4654,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | s(bool,X29) != s(bool,X32) | ~p(s(bool,X31)) | sP7(X33,X32,'F_0') | p(s(bool,X33)) | ~p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f2377,f553])).
fof(f4721,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) != s(bool,X27) | ~p(s(bool,X26)) | sP7(X28,X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(subsumption_resolution,[],[f4653,f553])).
fof(f4653,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27) | ~p(s(bool,X26)) | sP7(X28,X27,'F_0') | p(s(bool,X28)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) ) | $spl14),
  inference(superposition,[],[f2377,f553])).
fof(f4405,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4196,f936])).
fof(f4196,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(resolution,[],[f936,f613])).
fof(f613,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f66])).
fof(f66,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.NOT_CLAUSES|split|1')).
fof(f4404,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4195,f936])).
fof(f4195,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(resolution,[],[f936,f611])).
fof(f611,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f65])).
fof(f65,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.NOT_CLAUSES|split|2')).
fof(f4403,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4192,f936])).
fof(f4192,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | $spl6),
  inference(resolution,[],[f936,f585])).
fof(f585,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,X16))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CLAUSES|split|3')).
fof(f4402,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4191,f936])).
fof(f4191,plain,(
  ( ! [X214] : (p(s(bool,X214))) ) | $spl6),
  inference(resolution,[],[f936,f582])).
fof(f582,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CLAUSES|split|1')).
fof(f4401,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4189,f936])).
fof(f4189,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(resolution,[],[f936,f579])).
fof(f579,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f301])).
fof(f301,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X16] : ((p(s(bool,'F_0')) => p(s(bool,X16))) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CLAUSES|split|2')).
fof(f4400,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f4188,f936])).
fof(f4188,plain,(
  ( ! [X212] : (p(s(bool,X212))) ) | $spl6),
  inference(resolution,[],[f936,f576])).
fof(f576,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X16] : ((p(s(bool,'T_0')) => p(s(bool,X16))) <=> p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CLAUSES|split|0')).
fof(f4197,plain,(
  $false | $spl6),
  inference(resolution,[],[f936,f615])).
fof(f615,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.TRUTH')).
fof(f4198,plain,(
  $false | $spl6),
  inference(resolution,[],[f936,f616])).
fof(f616,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4',aHL_TRUTH)).
fof(f4199,plain,(
  $false | $spl6),
  inference(resolution,[],[f936,f817])).
fof(f817,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f566])).
fof(f566,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X45] : ! [X20] : (s(X45,X20) = s(X45,X20) <=> p(s(bool,'T_0')))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.REFL_CLAUSE')).
fof(f4211,plain,(
  $false | $spl6),
  inference(resolution,[],[f936,f677])).
fof(f677,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X6),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))))))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.dec_to_exhProof.v_rel_rules|split|0')).
fof(f4097,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) != s(bool,X66)) ) | $spl24),
  inference(cnf_transformation,[],[f4097_D])).
fof(f4097_D,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) != s(bool,X66)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4098,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') != s(bool,X71) | s(bool,X71) != s(bool,X72) | ~p(s(bool,X70)) | s(bool,'T_0') = s(bool,X72)) ) | $spl14),
  inference(subsumption_resolution,[],[f4061,f597])).
fof(f4061,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') != s(bool,X71) | s(bool,X71) != s(bool,X72) | ~p(s(bool,X70)) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) ) | $spl14),
  inference(superposition,[],[f2372,f813])).
fof(f2372,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f2131,f813])).
fof(f4077,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,X68) | s(bool,X68) != s(bool,X69) | ~p(s(bool,X67)) | s(bool,'T_0') = s(bool,X69)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4060])).
fof(f4060,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,X68) | s(bool,X68) != s(bool,X69) | ~p(s(bool,X67)) | s(bool,'T_0') = s(bool,X69) | ~p(s(bool,X67))) ) | $spl14),
  inference(superposition,[],[f2372,f604])).
fof(f4078,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,X65) != s(bool,X66) | ~p(s(bool,X64)) | s(bool,'T_0') = s(bool,X66)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4059])).
fof(f4059,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,X65) != s(bool,X66) | ~p(s(bool,X64)) | s(bool,'T_0') = s(bool,X66) | ~p(s(bool,X64))) ) | $spl14),
  inference(superposition,[],[f2372,f602])).
fof(f3325,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | p(s(bool,X3))) ) | $spl22),
  inference(cnf_transformation,[],[f3325_D])).
fof(f3325_D,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | p(s(bool,X3))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3324,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3311])).
fof(f3311,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl14),
  inference(resolution,[],[f3289,f493])).
fof(f2885,plain,(
  $false | $spl10),
  inference(trivial_inequality_removal,[],[f2880])).
fof(f2880,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl10),
  inference(superposition,[],[f1471,f894])).
fof(f2884,plain,(
  $false | $spl10),
  inference(equality_resolution,[],[f1471])).
fof(f2829,plain,(
  ( ! [X15] : (sP7(X15,'F_0','T_0') | p(s(bool,X15))) ) | $spl20),
  inference(cnf_transformation,[],[f2829_D])).
fof(f2829_D,plain,(
  ( ! [X15] : (sP7(X15,'F_0','T_0') | p(s(bool,X15))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2824,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | p(s(bool,X15)) | sP7(X15,'F_0','T_0')) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2823])).
fof(f2823,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | s(bool,'F_0') != s(bool,X14) | p(s(bool,X15)) | sP7(X15,'F_0','T_0')) ) | $spl14),
  inference(resolution,[],[f2816,f2250])).
fof(f2250,plain,(
  ( ! [X6,X7] : (~sP4(X7,'F_0',X6) | p(s(bool,X6)) | sP7(X6,'F_0',X7)) )),
  inference(subsumption_resolution,[],[f2248,f481])).
fof(f2248,plain,(
  ( ! [X6,X7] : (sP7(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X6,X7,'F_0') | ~sP4(X7,'F_0',X6)) )),
  inference(resolution,[],[f2244,f477])).
fof(f2244,plain,(
  ( ! [X4,X5] : (sP5('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f2240,f483])).
fof(f2240,plain,(
  ( ! [X4,X5] : (sP5('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f2237,f481])).
fof(f2237,plain,(
  ( ! [X6,X7] : (~sP6(X7,X6,'F_0') | sP5('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f2233,f484])).
fof(f2233,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP6(X7,X6,'F_0') | sP5('F_0',X6,X7)) )),
  inference(resolution,[],[f2227,f472])).
fof(f2816,plain,(
  ( ! [X6,X8,X7] : (sP4('T_0',X7,X8) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f2806,f487])).
fof(f2806,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | sP4('T_0',X7,X8)) ) | $spl14),
  inference(resolution,[],[f2133,f471])).
fof(f2133,plain,(
  ( ! [X6,X8,X7] : (sP7(X8,X6,'T_0') | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | p(s(bool,X8))) ) | $spl14),
  inference(superposition,[],[f2116,f490])).
fof(f2815,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) | $spl18),
  inference(cnf_transformation,[],[f2815_D])).
fof(f2815_D,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2814,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2805])).
fof(f2805,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl14),
  inference(resolution,[],[f2133,f493])).
fof(f2617,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,X24)) ) | $spl16),
  inference(cnf_transformation,[],[f2617_D])).
fof(f2617_D,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2620,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl14),
  inference(subsumption_resolution,[],[f2585,f573])).
fof(f2585,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl14),
  inference(superposition,[],[f2130,f555])).
fof(f2619,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl14),
  inference(subsumption_resolution,[],[f2584,f573])).
fof(f2584,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl14),
  inference(superposition,[],[f2130,f555])).
fof(f2618,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl14),
  inference(subsumption_resolution,[],[f2583,f573])).
fof(f2583,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl14),
  inference(superposition,[],[f2130,f554])).
fof(f2616,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X23)) | p(s(bool,X22))) ) | $spl14),
  inference(subsumption_resolution,[],[f2582,f573])).
fof(f2582,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X23)) | p(s(bool,X22))) ) | $spl14),
  inference(superposition,[],[f2130,f554])).
fof(f2116,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f2116_D])).
fof(f2116_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2115,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2110])).
fof(f2110,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f2087,f493])).
fof(f2087,plain,(
  ( ! [X6,X5] : (sP7(X6,X5,'F_0') | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) )),
  inference(superposition,[],[f2043,f490])).
fof(f1552,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | ~p(s(bool,X54))) ) | $spl4),
  inference(subsumption_resolution,[],[f1500,f934])).
fof(f1500,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | ~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X54)) ) | $spl4),
  inference(superposition,[],[f934,f813])).
fof(f1474,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,X53)) ) | $spl12),
  inference(cnf_transformation,[],[f1474_D])).
fof(f1474_D,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,X53)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1551,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52))) ) | $spl4),
  inference(subsumption_resolution,[],[f1499,f934])).
fof(f1499,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | s(bool,'F_0') = s(bool,X52)) ) | $spl4),
  inference(superposition,[],[f934,f813])).
fof(f1545,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1498])).
fof(f1498,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50)) | ~p(s(bool,X50))) ) | $spl4),
  inference(superposition,[],[f934,f604])).
fof(f1546,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1497])).
fof(f1497,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48)) | ~p(s(bool,X48))) ) | $spl4),
  inference(superposition,[],[f934,f602])).
fof(f1475,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X54))) ) | $spl2),
  inference(subsumption_resolution,[],[f1420,f933])).
fof(f1420,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X54)) | s(bool,'T_0') = s(bool,X54)) ) | $spl2),
  inference(superposition,[],[f933,f813])).
fof(f1473,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X52))) ) | $spl2),
  inference(subsumption_resolution,[],[f1419,f933])).
fof(f1419,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X52)) ) | $spl2),
  inference(superposition,[],[f933,f813])).
fof(f1464,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1416])).
fof(f1416,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46)) | p(s(bool,X46))) ) | $spl2),
  inference(superposition,[],[f933,f600])).
fof(f1465,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1415])).
fof(f1415,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | p(s(bool,X44))) ) | $spl2),
  inference(superposition,[],[f933,f598])).
fof(f1466,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1413])).
fof(f1413,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl2),
  inference(superposition,[],[f933,f573])).
fof(f1375,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1038,f937])).
fof(f1038,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f937,f614])).
fof(f614,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f318])).
fof(f1374,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1037,f937])).
fof(f1037,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f937,f612])).
fof(f612,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f317])).
fof(f1373,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1032,f937])).
fof(f1032,plain,(
  ( ! [X215] : (~p(s(bool,X215))) ) | $spl8),
  inference(resolution,[],[f937,f571])).
fof(f571,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X16)))),
  file('thm.dec_to_exhProof.v_rel_rules|split|4','thm.bool.IMP_CLAUSES|split|4')).
fof(f1036,plain,(
  $false | $spl8),
  inference(resolution,[],[f937,f610])).
fof(f1043,plain,(
  $false | $spl8),
  inference(resolution,[],[f937,f412])).
fof(f932,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) | $spl0),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f905,plain,(
  ( ! [X138,X139] : (p(s(bool,X138)) | ~p(s(bool,X139)) | s(bool,X139) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f412,f814])).
fof(f831,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | s(bool,X27) != s(bool,X29)) )),
  inference(resolution,[],[f412,f436])).
fof(f830,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f412,f435])).
fof(f828,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f412,f434])).
fof(f826,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35))))) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f412,f432])).
fof(f824,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK34),s('type.conSem.v','const.conSem.Loc_1'(s('type.num.num',sK35))),s('type.exhSem.v','const.exhSem.Loc_1'(s('type.num.num',sK35)))))) )),
  inference(resolution,[],[f412,f430])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 567
% Passive clauses: 973
% Generated clauses: 7877
% Final active clauses: 560
% Final passive clauses: 600
% Input formulas: 69
% Initial clauses: 402
% 
% Selected by SInE selection: 56
% SInE iterations: 1
% 
% Duplicate literals: 575
% Trivial inequalities: 136
% Fw subsumption resolutions: 1156
% Fw demodulations: 223
% Bw demodulations: 1
% 
% Simple tautologies: 1757
% Equational tautologies: 86
% Forward subsumptions: 2806
% 
% Binary resolution: 3667
% Forward superposition: 1838
% Backward superposition: 229
% Equality resolution: 29
% 
% Split clauses: 131
% Split components: 298
% SAT solver clauses: 186
% SAT solver unit clauses: 20
% SAT solver binary clauses: 17
% 
% Sat splits: 131
% Sat splitting refutations: 20
% 
% Memory used [KB]: 21620
% Time elapsed: 0.371 s
% ------------------------------
% ------------------------------
% Success in time 12.113 s
