Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (63515)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 162
% Passive clauses: 3764
% Generated clauses: 19768
% Final active clauses: 144
% Final passive clauses: 3344
% Input formulas: 68
% Initial clauses: 100
% 
% Duplicate literals: 519
% Trivial inequalities: 303
% Fw subsumption resolutions: 1338
% Fw demodulations: 2908
% Bw demodulations: 16
% 
% Simple tautologies: 239
% Equational tautologies: 45
% Forward subsumptions: 10259
% 
% Binary resolution: 344
% Factoring: 1
% Forward superposition: 6633
% Backward superposition: 7252
% Self superposition: 181
% Equality factoring: 70
% Equality resolution: 34
% 
% Split clauses: 214
% Split components: 430
% SAT solver clauses: 555
% SAT solver unit clauses: 9
% SAT solver binary clauses: 120
% 
% Sat splits: 214
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 554
% TWLsolver calls for satisfiability: 143
% 
% Memory used [KB]: 21748
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (63542)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 191
% Passive clauses: 1547
% Generated clauses: 19802
% Final active clauses: 105
% Final passive clauses: 522
% Input formulas: 68
% Initial clauses: 100
% 
% Duplicate literals: 146
% Trivial inequalities: 178
% Fw subsumption resolutions: 1016
% Fw demodulations: 2989
% Bw demodulations: 708
% Condensations: 318
% Global subsumptions: 8
% 
% Simple tautologies: 73
% Equational tautologies: 15
% Forward subsumptions: 11192
% Bw demodulations to eq. taut.: 145
% 
% Binary resolution: 366
% Factoring: 1
% Forward superposition: 5709
% Backward superposition: 5875
% Self superposition: 172
% Equality factoring: 14
% Equality resolution: 23
% 
% Split clauses: 215
% Split components: 431
% SAT solver clauses: 2331
% SAT solver unit clauses: 150
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 20
% SAT solver learnt literals: 8
% 
% Sat splits: 215
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 2295
% TWLsolver calls for satisfiability: 416
% 
% Memory used [KB]: 17142
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f10238,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f212,f10171])).
fof(f10171,plain,(
  ( ! [X389] : (~p(s(bool,X389))) )),
  inference(forward_demodulation,[],[f10170,f271])).
fof(f271,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X16] : ! [X19,X31,X32,X33,X34] : s(X16,X31) = s(X16,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X16),X19),s(X16,X31),s(fun('type.num.num',X16),X32),s(fun(bool,X16),X33),s(X16,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|4','thm.bvp_to_wordProps.tag_case_def|split|1')).
fof(f10170,plain,(
  ( ! [X389,X391,X388,X390,X392] : (~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',bool),X388),s(bool,X389),s(fun('type.num.num',bool),X390),s(fun(bool,bool),X391),s(bool,X392))))) )),
  inference(subsumption_resolution,[],[f10052,f10004])).
fof(f10004,plain,(
  ( ! [X6,X8,X9] : (s(X6,X8) = s(X6,happ(s(fun('type.num.num',X6),X9),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f5321,f271])).
fof(f5321,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X1),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(superposition,[],[f274,f211])).
fof(f211,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f168])).
fof(f168,plain,(
  ? [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|4','thm.bvp_to_wordProps.tag_distinct|split|4')).
fof(f274,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X16] : ! [X6,X19,X31,X32,X33,X34] : s(X16,happ(s(fun('type.num.num',X16),X32),s('type.num.num',X6))) = s(X16,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6))),s(fun('type.num.num',X16),X19),s(X16,X31),s(fun('type.num.num',X16),X32),s(fun(bool,X16),X33),s(X16,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|4','thm.bvp_to_wordProps.tag_case_def|split|2')).
fof(f10052,plain,(
  ( ! [X389,X391,X388,X390,X392] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X390),s('type.num.num',sK0))) | ~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',bool),X388),s(bool,X389),s(fun('type.num.num',bool),X390),s(fun(bool,bool),X391),s(bool,X392))))) )),
  inference(superposition,[],[f236,f5321])).
fof(f236,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X29] : (s(bool,'F_0') = s(bool,X29) <=> ~p(s(bool,X29)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|4','thm.bool.EQ_CLAUSES|split|2')).
fof(f212,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_distinct|split|4',aHL_TRUTH)).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 144
% Passive clauses: 1237
% Generated clauses: 10029
% Final active clauses: 117
% Final passive clauses: 867
% Input formulas: 68
% Initial clauses: 100
% 
% Duplicate literals: 6
% Trivial inequalities: 22
% Fw subsumption resolutions: 163
% Fw demodulations: 1931
% Bw demodulations: 233
% Condensations: 31
% Global subsumptions: 1
% 
% Simple tautologies: 23
% Equational tautologies: 67
% Forward subsumptions: 6331
% Fw demodulations to eq. taut.: 152
% Bw demodulations to eq. taut.: 18
% 
% Unit resulting resolution: 977
% Factoring: 1
% Forward superposition: 2842
% Backward superposition: 3658
% Self superposition: 112
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 22
% Split components: 44
% SAT solver clauses: 1287
% SAT solver unit clauses: 358
% SAT solver binary clauses: 693
% 
% Sat splits: 22
% Sat splitting refutations: 3
% 
% Memory used [KB]: 23666
% Time elapsed: 0.541 s
% ------------------------------
% ------------------------------
% Success in time 1.9 s
