Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ImplicitQueue
% (1189)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 3228
% Generated clauses: 15101
% Final active clauses: 104
% Final passive clauses: 2884
% Input formulas: 31
% Initial clauses: 82
% 
% Duplicate literals: 693
% Trivial inequalities: 247
% Fw subsumption resolutions: 2603
% Fw demodulations: 780
% Bw demodulations: 4
% 
% Simple tautologies: 373
% Equational tautologies: 88
% Forward subsumptions: 5288
% 
% Binary resolution: 571
% Factoring: 4
% Forward superposition: 3928
% Backward superposition: 5556
% Self superposition: 262
% Equality factoring: 220
% Equality resolution: 29
% 
% Split clauses: 562
% Split components: 1169
% SAT solver clauses: 1490
% SAT solver unit clauses: 46
% SAT solver binary clauses: 460
% 
% Sat splits: 562
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1478
% TWLsolver calls for satisfiability: 176
% 
% Memory used [KB]: 7547
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ImplicitQueue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ImplicitQueue
% SZS output start Proof for thm.ImplicitQueue
fof(f10174,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f10154,f127,f226,f169,f249,f517,f525,f132,f230,f512,f514,f164,f243,f1062,f1069,f1020,f1025,f160,f242,f1044,f261,f264,f125,f224,f7487,f7489,f129,f228,f1010,f1019,f162,f240,f2036,f2040,f277,f279,f136,f232,f258,f260,f126,f222,f154,f238,f2033,f2034,f168,f248,f723,f725,f161,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))) ) | $spl24),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))))) & (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))))) & ((~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X13] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),X7),s('type.ImplicitQueue.queue'(X6),X13)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X6),X13)))) & s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X6),X13)))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.queue_inv_def')).
fof(f725,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))) | $spl158),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f723,plain,(
  s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))) | ($spl4 | $spl24 | $spl60)),
  inference(subsumption_resolution,[],[f714,f226])).
fof(f714,plain,(
  ~p(s(bool,'T_0')) | s('type.list.list'(sK0),sK2) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))) | ($spl24 | $spl60)),
  inference(superposition,[],[f241,f264])).
fof(f248,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) | s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2)))))) ) | $spl38),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) | s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f168,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2))))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) | s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) => s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) => s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),X2))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X9,X14,X12] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X12),s('type.ImplicitQueue.queue'(X6),X9)))) => s('type.list.list'(X6),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X6),X9),s('type.ImplicitQueue.times'(X6),X14))))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X6),X9))),s('type.list.list'(X6),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X6),X14))))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.ImplicitQueue_|49||49|2_prove_ltb')).
fof(f2034,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),X0))))) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),X0)))))) ) | $spl354),
  inference(cnf_transformation,[],[f2034_D])).
fof(f2034_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),X0))))) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),X0)))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f2033,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),X0))))) = s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),X0)))))) ) | ($spl4 | $spl38 | $spl158 | $spl238)),
  inference(forward_demodulation,[],[f2032,f725])).
fof(f2032,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),X0))))) = s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),X0)))))) ) | ($spl4 | $spl38 | $spl238)),
  inference(subsumption_resolution,[],[f2024,f226])).
fof(f2024,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),X0))))) = s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),sK1))),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),X0)))))) ) | ($spl38 | $spl238)),
  inference(superposition,[],[f248,f1069])).
fof(f238,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1)))))) ) | $spl18),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f154,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9] : s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X9)))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.exps_def|split|0')).
fof(f222,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | $spl1),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f126,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))))),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1)))) & ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f76])).
fof(f76,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X2),s('type.ImplicitQueue.queue'(X0),X1)))) & ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X3)))))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X2),s('type.ImplicitQueue.queue'(X0),X1)))) => p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X3)))))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X2),s('type.ImplicitQueue.queue'(X0),X1)))) => p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X3)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),X8),s('type.ImplicitQueue.queue'(X6),X7)))) => p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X6),X7),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X9)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),X8),s('type.ImplicitQueue.queue'(X6),X7)))) => p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.NIL_0'))))),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X6),X7),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X9)))))))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb')).
fof(f260,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | $spl59),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f258,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | ($spl1 | $spl18)),
  inference(forward_demodulation,[],[f222,f238])).
fof(f232,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f136,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f279,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) | $spl62),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f277,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK2),s('type.list.list'(sK0),'const.ImplicitQueue.exps_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) | ($spl14 | $spl59)),
  inference(resolution,[],[f260,f232])).
fof(f2040,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) | $spl358),
  inference(cnf_transformation,[],[f2040_D])).
fof(f2040_D,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f2036,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3))))))) | ($spl62 | $spl354)),
  inference(backward_demodulation,[],[f2034,f279])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2))))) ) | $spl22),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f162,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f1019,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.ImplicitQueue.queue'(X0),X1)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X1))))) ) | $spl218),
  inference(cnf_transformation,[],[f1019_D])).
fof(f1019_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.ImplicitQueue.queue'(X0),X1)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X1))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f1010,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X1))),s('type.ImplicitQueue.queue'(X0),X1)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X1))))) ) | $spl22),
  inference(equality_resolution,[],[f240])).
fof(f228,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f129,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb',aHL_FALSITY)).
fof(f7489,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | $spl711),
  inference(cnf_transformation,[],[f7489_D])).
fof(f7489_D,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) <=> ~$spl711),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl711])])).
fof(f7487,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | ($spl7 | $spl218 | $spl358)),
  inference(subsumption_resolution,[],[f7469,f228])).
fof(f7469,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(sK0),sK1),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))))) | ($spl218 | $spl358)),
  inference(superposition,[],[f1019,f2040])).
fof(f224,plain,(
  p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1)))) | $spl2),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f125,plain,(
  p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1))))),
  inference(cnf_transformation,[],[f89])).
fof(f264,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1))) | $spl60),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f261,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),sK2),s('type.ImplicitQueue.queue'(sK0),sK1))) | ($spl2 | $spl10)),
  inference(resolution,[],[f230,f224])).
fof(f1044,plain,(
  $false | ($spl4 | $spl60 | $spl220)),
  inference(subsumption_resolution,[],[f1039,f226])).
fof(f1039,plain,(
  ~p(s(bool,'T_0')) | ($spl60 | $spl220)),
  inference(superposition,[],[f1025,f264])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))) ) | $spl26),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f160,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f1025,plain,(
  ( ! [X0] : (~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),X0),s('type.ImplicitQueue.queue'(sK0),sK1))))) ) | $spl220),
  inference(cnf_transformation,[],[f1025_D])).
fof(f1025_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),X0),s('type.ImplicitQueue.queue'(sK0),sK1))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f1020,plain,(
  ( ! [X0] : (~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),X0),s('type.ImplicitQueue.queue'(sK0),sK1))))) ) | ($spl26 | $spl215)),
  inference(resolution,[],[f1017,f242])).
fof(f1017,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) | $spl215),
  inference(cnf_transformation,[],[f1017_D])).
fof(f1017_D,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f1069,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1))) | $spl238),
  inference(cnf_transformation,[],[f1069_D])).
fof(f1069_D,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f1062,plain,(
  s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1))) | ($spl10 | $spl214)),
  inference(resolution,[],[f1016,f230])).
fof(f1016,plain,(
  p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) | $spl214),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2))))))) ) | $spl28),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2))))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f164,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2))))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2)))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X1),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X2)))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X12,X9] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X12),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X9)))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X12))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.depth_def|split|0')).
fof(f514,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1))))))) ) | $spl104),
  inference(cnf_transformation,[],[f514_D])).
fof(f514_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f512,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X1))))))) ) | $spl28),
  inference(equality_resolution,[],[f243])).
fof(f230,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f132,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f525,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X7)))))) ) | $spl108),
  inference(cnf_transformation,[],[f525_D])).
fof(f525_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X7)))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f517,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X7)))))) ) | ($spl10 | $spl104)),
  inference(resolution,[],[f514,f230])).
fof(f249,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))) | ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1))))) ) | $spl40),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))) | ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f169,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))) | ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) | ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2)))) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) | ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2))))) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) & p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2))))) => p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),X1)))) & p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X3),s('type.ImplicitQueue.times'(X0),X2))))) => p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X3),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X0),X1),s('type.ImplicitQueue.times'(X0),X2)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X14,X12] : ((p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X12),s('type.ImplicitQueue.queue'(X6),X7)))) & p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num',X12),s('type.ImplicitQueue.times'(X6),X14))))) => p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X12),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.snoc_2'(s('type.ImplicitQueue.queue'(X6),X7),s('type.ImplicitQueue.times'(X6),X14)))))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb','thm.ImplicitQueue.ImplicitQueue_|49||49|3_prove_ltb')).
fof(f226,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f127,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||49|4_prove_ltb',aHL_TRUTH)).
fof(f10154,plain,(
  $false | ($spl4 | $spl40 | $spl108 | $spl238 | $spl711)),
  inference(subsumption_resolution,[],[f10153,f226])).
fof(f10153,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl40 | $spl108 | $spl238 | $spl711)),
  inference(forward_demodulation,[],[f10152,f1069])).
fof(f10152,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) | ($spl4 | $spl40 | $spl108 | $spl711)),
  inference(subsumption_resolution,[],[f10151,f226])).
fof(f10151,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) | ($spl40 | $spl108 | $spl711)),
  inference(forward_demodulation,[],[f10128,f525])).
fof(f10128,plain,(
  ~p(s(bool,'const.ImplicitQueue.depth_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.Once_1'(s(sK0,sK3)))))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(sK0),sK1)))) | ($spl40 | $spl711)),
  inference(resolution,[],[f7489,f249])).
% SZS output end Proof for thm.ImplicitQueue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 129
% Passive clauses: 462
% Generated clauses: 10032
% Final active clauses: 104
% Final passive clauses: 276
% Input formulas: 31
% Initial clauses: 82
% 
% Duplicate literals: 298
% Trivial inequalities: 255
% Fw subsumption resolutions: 1086
% Fw demodulations: 1444
% Bw demodulations: 69
% Condensations: 512
% Global subsumptions: 4
% 
% Simple tautologies: 164
% Equational tautologies: 27
% Forward subsumptions: 5067
% 
% Binary resolution: 167
% Factoring: 1
% Forward superposition: 2210
% Backward superposition: 3251
% Self superposition: 80
% Equality factoring: 47
% Equality resolution: 36
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 752
% SAT solver unit clauses: 66
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 9
% 
% Sat splits: 40
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 735
% TWLsolver calls for satisfiability: 409
% 
% Memory used [KB]: 7675
% Time elapsed: 0.315 s
% ------------------------------
% ------------------------------
% Success in time 0.914 s
