Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (145923)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 880
% Passive clauses: 7360
% Generated clauses: 34254
% Final active clauses: 762
% Final passive clauses: 6037
% Input formulas: 517
% Initial clauses: 782
% 
% Duplicate literals: 162
% Trivial inequalities: 46
% Fw subsumption resolutions: 1504
% Fw demodulations: 3311
% Bw demodulations: 115
% 
% Simple tautologies: 236
% Equational tautologies: 26
% Forward subsumptions: 21634
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 5136
% Factoring: 588
% Forward superposition: 11515
% Backward superposition: 10934
% Self superposition: 53
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 266
% Split components: 534
% SAT solver clauses: 409
% SAT solver unit clauses: 16
% SAT solver binary clauses: 71
% 
% Sat splits: 266
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 409
% TWLsolver calls for satisfiability: 324
% 
% Memory used [KB]: 18933
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (145943)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 485
% Passive clauses: 3109
% Generated clauses: 23115
% Final active clauses: 393
% Final passive clauses: 2367
% Input formulas: 517
% Initial clauses: 782
% 
% Duplicate literals: 79
% Trivial inequalities: 36
% Fw subsumption resolutions: 1136
% Fw demodulations: 2510
% Bw demodulations: 280
% Condensations: 100
% Global subsumptions: 24
% 
% Simple tautologies: 195
% Equational tautologies: 8
% Forward subsumptions: 11974
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 2000
% Factoring: 231
% Forward superposition: 6958
% Backward superposition: 5338
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 4457
% SAT solver unit clauses: 1208
% SAT solver binary clauses: 162
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 23
% Sat splitting refutations: 140
% 
% TWLsolver clauses: 4361
% TWLsolver calls for satisfiability: 6257
% 
% Memory used [KB]: 26353
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% (145964)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1568
% Passive clauses: 12933
% Generated clauses: 194130
% Final active clauses: 1253
% Final passive clauses: 4650
% Input formulas: 517
% Initial clauses: 782
% Discarded non-redundant clauses: 34437
% 
% Duplicate literals: 17
% Trivial inequalities: 32
% Fw subsumption resolutions: 1035
% Fw demodulations: 58643
% Bw demodulations: 2829
% Condensations: 30
% Global subsumptions: 3
% 
% Simple tautologies: 94
% Equational tautologies: 42
% Forward subsumptions: 104388
% Fw demodulations to eq. taut.: 787
% Bw demodulations to eq. taut.: 24
% 
% Unit resulting resolution: 40721
% Factoring: 1
% Forward superposition: 44876
% Backward superposition: 42760
% Self superposition: 132
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 13929
% SAT solver unit clauses: 3403
% SAT solver binary clauses: 8294
% 
% Sat splits: 17
% Sat splitting refutations: 957
% 
% Memory used [KB]: 180252
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f10666,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f813,f1112,f1110,f1113,f1115,f1116,f1117,f1120,f1195,f1219,f1237,f1252,f1270,f1272,f1273,f1275,f1277,f1354,f1406,f1361,f1250,f1362,f1412,f1414,f1476,f1478,f1479,f1480,f1483,f1558,f1649,f1563,f1564,f1565,f1566,f1657,f1568,f1569,f1570,f1571,f1572,f1644,f1643,f1642,f1655,f1674,f1676,f1677,f1678,f1681,f1756,f1847,f1761,f1762,f1763,f1764,f1853,f1766,f1767,f1768,f1769,f1770,f1842,f1841,f1840,f1852,f1973,f1967,f2064,f2065,f2066,f2067,f2068,f2069,f2070,f2071,f2072,f2075,f2076,f2976,f2978,f2979,f2980,f2983,f3058,f3160,f3063,f3064,f3065,f3066,f3165,f3068,f3069,f3070,f3071,f3072,f3146,f3145,f3144,f3164,f3177,f3194,f3193,f3214,f3216,f3219,f3222,f3224,f3228,f3237,f3239,f3240,f3242,f3244,f3321,f3422,f3327,f3329,f3330,f3331,f3332,f3333,f3334,f3335,f3336,f3342,f3343,f3416,f3415,f3428,f3613,f3612,f3611,f3618,f3841,f3840,f3846,f3992,f3990,f3989,f3972,f3971,f3970,f4087,f4088,f4089,f4090,f4091,f4092,f4097,f4098,f4099,f4100,f4101,f4102,f4103,f4104,f4197,f4199,f4200,f4201,f4204,f4279,f4383,f4285,f4286,f1653,f4287,f4288,f1656,f4388,f4290,f4291,f4292,f4293,f4294,f4369,f4368,f4367,f4387,f4415,f4417,f4418,f4420,f4422,f4499,f4601,f4506,f4508,f4509,f4510,f4511,f4512,f4513,f4514,f4515,f4521,f4522,f3434,f4595,f4594,f4607,f5175,f5177,f5210,f5211,f5176,f5245,f5247,f5281,f5282,f5246,f5596,f5598,f5599,f5600,f5603,f5678,f5784,f5817,f5819,f5820,f5821,f5824,f5899,f6004,f6428,f6781,f6430,f7121,f7462,f7122,f6429,f7465,f7467,f7468,f7470,f7472,f7549,f7653,f7558,f7560,f7561,f7562,f7563,f7564,f7565,f7566,f7567,f7573,f7574,f7662,f7647,f7646,f7658,f7671,f7673,f7672,f7701,f7703,f7704,f7705,f7708,f7783,f7878,f7994,f7996,f8110,f8112,f8113,f8115,f8117,f8194,f8299,f8219,f8220,f8304,f8344,f8345,f7995,f8346,f8348,f8349,f8351,f8353,f8430,f8535,f8440,f8442,f8443,f8444,f8445,f8446,f8447,f8448,f8449,f8455,f8456,f8544,f8529,f8528,f8540,f8557,f8559,f8560,f8562,f8564,f8641,f8705,f8707,f8706,f8710,f8712,f8713,f8715,f8717,f8794,f8856,f8858,f8859,f8860,f8863,f8938,f9035,f9047,f9049,f9050,f9051,f9054,f9129,f9178,f9236,f9238,f9239,f9241,f9243,f9320,f8847,f9420,f9422,f9423,f9425,f9427,f9504,f8695,f9589,f9591,f9588,f9593,f9592,f9599,f9607,f9609,f9637,f9639,f9663,f9664,f9638,f9662,f9666,f9665,f9673,f9675,f9676,f9678,f9680,f9757,f9851,f9853,f9949,f9950,f10090,f10088,f9951,f9953,f9954,f9955,f9958,f10033,f10091,f10098,f10100,f10101,f10103,f10105,f10182,f10288,f10208,f10209,f10293,f3432,f10323,f10324,f9590,f10348,f10349,f9852,f10356,f10358,f10359,f1220,f10361,f1222,f10363,f1221,f10440,f1223,f9816,f10654,f10656,f10665])).
fof(f10665,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f10664,f1101])).
fof(f1101,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb',aHL_TRUTH)).
fof(f10664,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f1112,f1193])).
fof(f1193,plain,(
  s(bool,'T_0') = s(bool,sK14)),
  inference(resolution,[],[f811,f1083])).
fof(f1083,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f807])).
fof(f807,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X30] : (s(bool,'T_0') = s(bool,X30) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f811,plain,(
  p(s(bool,sK14))),
  inference(cnf_transformation,[],[f702])).
fof(f702,plain,(
  p(s(bool,sK14)) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) & (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))) | ~p(s(bool,sK14)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f657])).
fof(f657,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,X4)) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3)))))) & (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))) | ~p(s(bool,X4))))),
  inference(flattening,[],[f656])).
fof(f656,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,X4)) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3)))))) & (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))) | ~p(s(bool,X4)))))),
  inference(ennf_transformation,[],[f520])).
fof(f520,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,X4)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3)))))) | (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))) & p(s(bool,X4)))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,X4)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3)))))) | (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X2),s('type.stackSem.state'(X0,X1),X3))))) & p(s(bool,X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,X10)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X9))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X9)))))) | (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X9))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X9))))) & p(s(bool,X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,X10)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X9))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X9)))))) | (s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X9))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X9))))) & p(s(bool,X10)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.stackSem.stackSem_5|48|7_prove_ltb')).
fof(f10656,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10655,f877])).
fof(f877,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2))))))) )),
  inference(cnf_transformation,[],[f568])).
fof(f568,plain,(
  ! [X0,X1,X2,X3] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2))))))),
  inference(flattening,[],[f567])).
fof(f567,plain,(
  ! [X0,X1] : ! [X2,X3] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X8] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X11))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X11))))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.stackSem.stackSem_5|48|5_prove_ltb|split|0')).
fof(f10655,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f10513,f1110])).
fof(f10513,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))),
  inference(resolution,[],[f1269,f1019])).
fof(f1019,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f768])).
fof(f768,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f767])).
fof(f767,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f444])).
fof(f444,axiom,(
  ! [X53,X16] : (s('type.num.num',X16) = s('type.num.num',X53) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X16)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X53))))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.arithmetic.EQ_LESS_EQ')).
fof(f1269,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))))))),
  inference(resolution,[],[f812,f816])).
fof(f816,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f704])).
fof(f704,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f523])).
fof(f523,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f509])).
fof(f509,axiom,(
  ! [X53,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num',X16)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X53)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f812,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))),
  inference(cnf_transformation,[],[f702])).
fof(f10654,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10653,f877])).
fof(f10653,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f10512,f1110])).
fof(f10512,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))),
  inference(resolution,[],[f1269,f1019])).
fof(f9816,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl124),
  inference(cnf_transformation,[],[f9816_D])).
fof(f9816_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1223,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl12),
  inference(cnf_transformation,[],[f1223_D])).
fof(f1223_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10440,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f1102])).
fof(f1102,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f889])).
fof(f889,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f711,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f710])).
fof(f710,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f661])).
fof(f661,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f660])).
fof(f660,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X31,X49,X34,X50] : ((s(bool,X31) = s(bool,X49) & (p(s(bool,X49)) => s(bool,X34) = s(bool,X50))) => ((p(s(bool,X31)) => p(s(bool,X34))) <=> (p(s(bool,X49)) => p(s(bool,X50)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CONG')).
fof(f1221,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl8),
  inference(cnf_transformation,[],[f1221_D])).
fof(f1221_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f10363,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X22) != s(bool,X24)) ) | $spl15),
  inference(resolution,[],[f1250,f892])).
fof(f892,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f1222,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl10),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f10361,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f891])).
fof(f891,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f1220,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl6),
  inference(cnf_transformation,[],[f1220_D])).
fof(f1220_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f10359,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f890])).
fof(f890,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f10358,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X7) != s(bool,X9)) ) | $spl15),
  inference(resolution,[],[f1250,f888])).
fof(f888,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f10356,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f886])).
fof(f886,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f711])).
fof(f9852,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X1,X0)) ) | $spl126),
  inference(cnf_transformation,[],[f9852_D])).
fof(f9852_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X1,X0)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f10349,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8(X4,'T_0') | p(s(bool,X3)) | p(s(bool,X5))) ) | $spl102),
  inference(subsumption_resolution,[],[f10326,f1055])).
fof(f1055,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f790])).
fof(f790,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f789])).
fof(f789,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f366])).
fof(f366,axiom,(
  ! [X30] : ((p(s(bool,X30)) | p(s(bool,'F_0'))) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.OR_CLAUSES|split|3')).
fof(f10326,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8(X4,'T_0') | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X5))) ) | $spl102),
  inference(resolution,[],[f10192,f996])).
fof(f996,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f759])).
fof(f759,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f758])).
fof(f758,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f10192,plain,(
  ( ! [X422,X424,X423] : (sP9(X422,X423) | sP8(X422,X424) | sP8(X424,'T_0')) ) | $spl102),
  inference(resolution,[],[f4764,f8706])).
fof(f4764,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP9(X0,X2) | sP8(X0,X1)) )),
  inference(resolution,[],[f4258,f1166])).
fof(f1166,plain,(
  ( ! [X113] : (~sP8(sK14,X113) | ~p(s(bool,X113))) )),
  inference(resolution,[],[f811,f981])).
fof(f981,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X0,X1)) )),
  inference(cnf_transformation,[],[f754])).
fof(f754,plain,(
  ! [X0,X1] : ((~sP8(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1)))),
  inference(rectify,[],[f753])).
fof(f753,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP8(X1,X0)))),
  inference(flattening,[],[f752])).
fof(f752,plain,(
  ! [X1,X0] : ((~sP8(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP8(X1,X0)))),
  inference(nnf_transformation,[],[f698])).
fof(f698,plain,(
  ! [X1,X0] : (sP8(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f4258,plain,(
  ( ! [X243,X244,X245] : (sP8(sK14,X243) | sP8(X244,X243) | sP9(X244,X245)) )),
  inference(resolution,[],[f1531,f998])).
fof(f998,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f759])).
fof(f1531,plain,(
  ( ! [X171,X170] : (sP8(sK14,X170) | sP8(X171,X170) | p(s(bool,X171))) )),
  inference(resolution,[],[f1167,f985])).
fof(f985,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f754])).
fof(f1167,plain,(
  ( ! [X114] : (p(s(bool,X114)) | sP8(sK14,X114)) )),
  inference(resolution,[],[f811,f982])).
fof(f982,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f754])).
fof(f10348,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8(X1,'T_0') | p(s(bool,X0)) | p(s(bool,X2))) ) | $spl102),
  inference(subsumption_resolution,[],[f10325,f1055])).
fof(f10325,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8(X1,'T_0') | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X2))) ) | $spl102),
  inference(resolution,[],[f10192,f1000])).
fof(f1000,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(flattening,[],[f760])).
fof(f760,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP9(X0,X1)))),
  inference(nnf_transformation,[],[f701])).
fof(f701,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f672,f700])).
fof(f672,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f671])).
fof(f671,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f465])).
fof(f465,axiom,(
  ! [X7,X6] : ((~(p(s(bool,X6)) | p(s(bool,X7))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X7)) => p(s(bool,'F_0')))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.sat.sat_9_prove_ltb')).
fof(f9590,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | sP8('T_0',X1)) ) | $spl110),
  inference(cnf_transformation,[],[f9590_D])).
fof(f9590_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1) | p(s(bool,X0)) | sP8('T_0',X1)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f10324,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X4) | p(s(bool,X3)) | p(s(bool,X5))) ) | $spl84),
  inference(subsumption_resolution,[],[f10301,f1055])).
fof(f10301,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X4) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X5))) ) | $spl84),
  inference(resolution,[],[f10191,f996])).
fof(f10191,plain,(
  ( ! [X420,X421,X419] : (sP9(X419,X420) | sP8(X419,X421) | sP8('T_0',X421)) ) | $spl84),
  inference(resolution,[],[f4764,f7672])).
fof(f10323,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X1) | p(s(bool,X0)) | p(s(bool,X2))) ) | $spl84),
  inference(subsumption_resolution,[],[f10300,f1055])).
fof(f10300,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X1) | p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X2))) ) | $spl84),
  inference(resolution,[],[f10191,f1000])).
fof(f3432,plain,(
  ( ! [X237,X236] : (p(s(bool,X236)) | ~sP9(X236,X237) | p(s(bool,X237))) ) | $spl48),
  inference(cnf_transformation,[],[f3432_D])).
fof(f3432_D,plain,(
  ( ! [X237,X236] : (p(s(bool,X236)) | ~sP9(X236,X237) | p(s(bool,X237))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10293,plain,(
  ( ! [X465,X464] : (sP9(X464,X465) | sP8(X464,'F_0')) ) | $spl136),
  inference(cnf_transformation,[],[f10293_D])).
fof(f10293_D,plain,(
  ( ! [X465,X464] : (sP9(X464,X465) | sP8(X464,'F_0')) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f10209,plain,(
  ( ! [X469,X471,X468,X470] : (sP9(X468,X469) | sP8(X468,'F_0') | p(s(bool,X470)) | p(s(bool,X471)) | ~sP9(X470,X471)) )),
  inference(resolution,[],[f4764,f1000])).
fof(f10208,plain,(
  ( ! [X466,X465,X467,X464] : (sP9(X464,X465) | sP8(X464,'F_0') | p(s(bool,X466)) | p(s(bool,X467)) | ~sP9(X466,X467)) )),
  inference(resolution,[],[f4764,f996])).
fof(f10288,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X0,X2)) ) | $spl134),
  inference(cnf_transformation,[],[f10288_D])).
fof(f10288_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X0,X2)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f10182,plain,(
  ( ! [X383,X385,X382,X384,X386] : (sP9(X382,X383) | sP8(X382,X384) | p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,X384) != s(bool,X386)) )),
  inference(resolution,[],[f4764,f1102])).
fof(f10105,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP9(X42,X43) | sP8(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f4764,f892])).
fof(f10103,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP9(X30,X31) | sP8(X30,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f4764,f891])).
fof(f10101,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP9(X18,X19) | sP8(X18,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4764,f890])).
fof(f10100,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP9(X12,X13) | sP8(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f4764,f888])).
fof(f10098,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1) | sP8(X0,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f4764,f886])).
fof(f10091,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl132),
  inference(cnf_transformation,[],[f10091_D])).
fof(f10091_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f10033,plain,(
  ( ! [X132,X133] : (p(s(bool,X132)) | p(s(bool,X133)) | s(bool,X133) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl14),
  inference(resolution,[],[f1249,f1102])).
fof(f1249,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) | $spl14),
  inference(cnf_transformation,[],[f1249_D])).
fof(f1249_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f9958,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X24) != s(bool,X25)) ) | $spl14),
  inference(resolution,[],[f1249,f892])).
fof(f9955,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl14),
  inference(resolution,[],[f1249,f891])).
fof(f9954,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl14),
  inference(resolution,[],[f1249,f890])).
fof(f9953,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X9) != s(bool,X10)) ) | $spl14),
  inference(resolution,[],[f1249,f888])).
fof(f9951,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl14),
  inference(resolution,[],[f1249,f886])).
fof(f10088,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12)))))) | $spl129),
  inference(cnf_transformation,[],[f10088_D])).
fof(f10088_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12)))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f10090,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | $spl130),
  inference(cnf_transformation,[],[f10090_D])).
fof(f10090_D,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f9950,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12)))))) | $spl14),
  inference(resolution,[],[f1249,f1019])).
fof(f9949,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) = s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12)))))) | $spl14),
  inference(resolution,[],[f1249,f1019])).
fof(f9853,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl102),
  inference(subsumption_resolution,[],[f9838,f1055])).
fof(f9838,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl102),
  inference(resolution,[],[f8944,f996])).
fof(f8944,plain,(
  ( ! [X227,X228,X226] : (sP8(X226,'T_0') | sP9(X227,X228) | sP8(X228,X226)) ) | $spl102),
  inference(resolution,[],[f8706,f4735])).
fof(f4735,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP9(X2,X0) | sP8(X0,X1)) )),
  inference(resolution,[],[f4257,f1166])).
fof(f4257,plain,(
  ( ! [X241,X242,X240] : (sP8(sK14,X240) | sP8(X241,X240) | sP9(X242,X241)) )),
  inference(resolution,[],[f1531,f997])).
fof(f997,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f759])).
fof(f9851,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl102),
  inference(subsumption_resolution,[],[f9837,f1055])).
fof(f9837,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl102),
  inference(resolution,[],[f8944,f1000])).
fof(f9757,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f1102])).
fof(f9680,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X22) != s(bool,X24)) ) | $spl15),
  inference(resolution,[],[f1250,f892])).
fof(f9678,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f891])).
fof(f9676,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f890])).
fof(f9675,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))) | s(bool,X7) != s(bool,X9)) ) | $spl15),
  inference(resolution,[],[f1250,f888])).
fof(f9673,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) ) | $spl15),
  inference(resolution,[],[f1250,f886])).
fof(f9665,plain,(
  ( ! [X1] : (sP9('T_0',X1)) ) | $spl120),
  inference(cnf_transformation,[],[f9665_D])).
fof(f9665_D,plain,(
  ( ! [X1] : (sP9('T_0',X1)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f9666,plain,(
  ( ! [X0] : (sP8(X0,'T_0')) ) | $spl122),
  inference(cnf_transformation,[],[f9666_D])).
fof(f9666_D,plain,(
  ( ! [X0] : (sP8(X0,'T_0')) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f9662,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | sP9('T_0',X1)) ) | $spl102),
  inference(factoring,[],[f8942])).
fof(f8942,plain,(
  ( ! [X222,X223,X224] : (sP8(X222,'T_0') | sP9(X223,X224) | sP8(X222,X223)) ) | $spl102),
  inference(resolution,[],[f8706,f3671])).
fof(f3671,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP9(X1,X2) | sP8(X0,X1)) )),
  inference(resolution,[],[f3037,f1166])).
fof(f3037,plain,(
  ( ! [X243,X244,X245] : (sP8(sK14,X243) | sP8(X243,X244) | sP9(X244,X245)) )),
  inference(resolution,[],[f1530,f998])).
fof(f1530,plain,(
  ( ! [X169,X168] : (sP8(sK14,X168) | sP8(X168,X169) | p(s(bool,X169))) )),
  inference(resolution,[],[f1167,f982])).
fof(f9638,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X0,X1)) ) | $spl118),
  inference(cnf_transformation,[],[f9638_D])).
fof(f9638_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | sP8(X0,X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f9664,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl102),
  inference(subsumption_resolution,[],[f9650,f1055])).
fof(f9650,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl102),
  inference(resolution,[],[f8942,f996])).
fof(f9663,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl102),
  inference(subsumption_resolution,[],[f9649,f1055])).
fof(f9649,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl102),
  inference(resolution,[],[f8942,f1000])).
fof(f9639,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl102),
  inference(subsumption_resolution,[],[f9624,f1055])).
fof(f9624,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,'T_0') | sP8(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl102),
  inference(resolution,[],[f8941,f996])).
fof(f8941,plain,(
  ( ! [X220,X219,X221] : (sP8(X219,'T_0') | sP9(X220,X221) | sP8(X219,X221)) ) | $spl102),
  inference(resolution,[],[f8706,f3642])).
fof(f3642,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP9(X2,X1) | sP8(X0,X1)) )),
  inference(resolution,[],[f3036,f1166])).
fof(f3036,plain,(
  ( ! [X241,X242,X240] : (sP8(sK14,X240) | sP8(X240,X241) | sP9(X242,X241)) )),
  inference(resolution,[],[f1530,f997])).
fof(f9637,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl102),
  inference(subsumption_resolution,[],[f9623,f1055])).
fof(f9623,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'T_0') | sP8(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl102),
  inference(resolution,[],[f8941,f1000])).
fof(f9609,plain,(
  $false | $spl112),
  inference(subsumption_resolution,[],[f9602,f1101])).
fof(f9602,plain,(
  ~p(s(bool,'T_0')) | $spl112),
  inference(resolution,[],[f9592,f1165])).
fof(f1165,plain,(
  ( ! [X112] : (~sP8(X112,sK14) | ~p(s(bool,X112))) )),
  inference(resolution,[],[f811,f981])).
fof(f9607,plain,(
  sP8(sK14,'T_0') | $spl116),
  inference(cnf_transformation,[],[f9607_D])).
fof(f9607_D,plain,(
  sP8(sK14,'T_0') <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f9599,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | sP8(sK14,'T_0')) ) | $spl112),
  inference(resolution,[],[f9592,f1529])).
fof(f1529,plain,(
  ( ! [X167,X166] : (~sP8(X166,X167) | ~p(s(bool,X167)) | sP8(sK14,X166)) )),
  inference(resolution,[],[f1167,f981])).
fof(f9592,plain,(
  ( ! [X1] : (sP8('T_0',X1)) ) | $spl112),
  inference(cnf_transformation,[],[f9592_D])).
fof(f9592_D,plain,(
  ( ! [X1] : (sP8('T_0',X1)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f9593,plain,(
  ( ! [X0] : (sP9(X0,'T_0')) ) | $spl114),
  inference(cnf_transformation,[],[f9593_D])).
fof(f9593_D,plain,(
  ( ! [X0] : (sP9(X0,'T_0')) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f9588,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0') | sP8('T_0',X1)) ) | $spl84),
  inference(factoring,[],[f8439])).
fof(f8439,plain,(
  ( ! [X420,X421,X419] : (sP9(X419,X420) | sP8(X420,X421) | sP8('T_0',X421)) ) | $spl84),
  inference(resolution,[],[f4735,f7672])).
fof(f9591,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X4) | p(s(bool,X5)) | p(s(bool,X3))) ) | $spl84),
  inference(subsumption_resolution,[],[f9567,f1055])).
fof(f9567,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X3))) ) | $spl84),
  inference(resolution,[],[f8439,f996])).
fof(f9589,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X1) | p(s(bool,X2)) | p(s(bool,X0))) ) | $spl84),
  inference(subsumption_resolution,[],[f9566,f1055])).
fof(f9566,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X0))) ) | $spl84),
  inference(resolution,[],[f8439,f1000])).
fof(f8695,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl100),
  inference(cnf_transformation,[],[f8695_D])).
fof(f8695_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f9504,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f1102])).
fof(f9427,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl23),
  inference(resolution,[],[f1414,f892])).
fof(f9425,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f891])).
fof(f9423,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f890])).
fof(f9422,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl23),
  inference(resolution,[],[f1414,f888])).
fof(f9420,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f886])).
fof(f8847,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl104),
  inference(cnf_transformation,[],[f8847_D])).
fof(f8847_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f9320,plain,(
  ( ! [X130,X131] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X131) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f1102])).
fof(f9243,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1412,f892])).
fof(f9241,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f891])).
fof(f9239,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f890])).
fof(f9238,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1412,f888])).
fof(f9236,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f886])).
fof(f9178,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl108),
  inference(cnf_transformation,[],[f9178_D])).
fof(f9178_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f9129,plain,(
  ( ! [X130,X131] : (p(s(bool,X130)) | p(s(bool,X131)) | s(bool,X131) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl20),
  inference(resolution,[],[f1411,f1102])).
fof(f1411,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | $spl20),
  inference(cnf_transformation,[],[f1411_D])).
fof(f1411_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9054,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X22) != s(bool,X23)) ) | $spl20),
  inference(resolution,[],[f1411,f892])).
fof(f9051,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl20),
  inference(resolution,[],[f1411,f891])).
fof(f9050,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl20),
  inference(resolution,[],[f1411,f890])).
fof(f9049,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X7) != s(bool,X8)) ) | $spl20),
  inference(resolution,[],[f1411,f888])).
fof(f9047,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl20),
  inference(resolution,[],[f1411,f886])).
fof(f9035,plain,(
  ( ! [X2,X0] : (sP8(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl106),
  inference(cnf_transformation,[],[f9035_D])).
fof(f9035_D,plain,(
  ( ! [X2,X0] : (sP8(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8938,plain,(
  ( ! [X212,X213,X214] : (sP8(X212,'T_0') | p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X212) != s(bool,X214)) ) | $spl102),
  inference(resolution,[],[f8706,f1102])).
fof(f8863,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl102),
  inference(resolution,[],[f8706,f892])).
fof(f8860,plain,(
  ( ! [X19,X17,X18,X16] : (sP8(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl102),
  inference(resolution,[],[f8706,f891])).
fof(f8859,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl102),
  inference(resolution,[],[f8706,f890])).
fof(f8858,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl102),
  inference(resolution,[],[f8706,f888])).
fof(f8856,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl102),
  inference(resolution,[],[f8706,f886])).
fof(f8794,plain,(
  ( ! [X130,X131] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X131) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f1102])).
fof(f8717,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1412,f892])).
fof(f8715,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f891])).
fof(f8713,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f890])).
fof(f8712,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1412,f888])).
fof(f8710,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl21),
  inference(resolution,[],[f1412,f886])).
fof(f8706,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8(X0,'T_0')) ) | $spl102),
  inference(cnf_transformation,[],[f8706_D])).
fof(f8706_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8(X0,'T_0')) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f8707,plain,(
  ( ! [X2,X3] : (sP8(X2,'T_0') | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f8704,f1055])).
fof(f8704,plain,(
  ( ! [X2,X3] : (sP8(X2,'T_0') | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(resolution,[],[f8452,f996])).
fof(f8452,plain,(
  ( ! [X455,X456] : (sP9(X455,X456) | sP8(X456,'T_0')) )),
  inference(resolution,[],[f4735,f1100])).
fof(f1100,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.TRUTH')).
fof(f8705,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f8703,f1055])).
fof(f8703,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(resolution,[],[f8452,f1000])).
fof(f8641,plain,(
  ( ! [X132,X131] : (p(s(bool,X131)) | ~p(s(bool,X132)) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f1102])).
fof(f8564,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl23),
  inference(resolution,[],[f1414,f892])).
fof(f8562,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f891])).
fof(f8560,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f890])).
fof(f8559,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl23),
  inference(resolution,[],[f1414,f888])).
fof(f8557,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl23),
  inference(resolution,[],[f1414,f886])).
fof(f8540,plain,(
  ( ! [X422,X423] : (sP9(X422,X423) | sP8(X423,'T_0')) ) | $spl96),
  inference(cnf_transformation,[],[f8540_D])).
fof(f8540_D,plain,(
  ( ! [X422,X423] : (sP9(X422,X423) | sP8(X423,'T_0')) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8528,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | sP9(X74,X75) | sP8(X75,'T_0')) )),
  inference(duplicate_literal_removal,[],[f8496])).
fof(f8496,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | sP9(X74,X75) | sP8(X75,'T_0') | ~p(s(bool,X73))) )),
  inference(superposition,[],[f4735,f1083])).
fof(f8529,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | sP9(X71,X72) | sP8(X72,'T_0')) )),
  inference(duplicate_literal_removal,[],[f8495])).
fof(f8495,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | sP9(X71,X72) | sP8(X72,'T_0') | ~p(s(bool,X70))) )),
  inference(superposition,[],[f4735,f1085])).
fof(f1085,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f808])).
fof(f808,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X30] : (s(bool,'T_0') = s(bool,X30) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f8544,plain,(
  ( ! [X461,X462] : (sP9(X461,X462) | sP8(X462,'F_0')) ) | $spl98),
  inference(cnf_transformation,[],[f8544_D])).
fof(f8544_D,plain,(
  ( ! [X461,X462] : (sP9(X461,X462) | sP8(X462,'F_0')) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8456,plain,(
  ( ! [X466,X465,X467,X468] : (sP9(X465,X466) | sP8(X466,'F_0') | p(s(bool,X467)) | p(s(bool,X468)) | ~sP9(X467,X468)) )),
  inference(resolution,[],[f4735,f1000])).
fof(f8455,plain,(
  ( ! [X461,X463,X462,X464] : (sP9(X461,X462) | sP8(X462,'F_0') | p(s(bool,X463)) | p(s(bool,X464)) | ~sP9(X463,X464)) )),
  inference(resolution,[],[f4735,f996])).
fof(f8449,plain,(
  ( ! [X449,X448,X450] : (sP9(X448,X449) | sP8(X449,'T_0') | ~p(s(bool,X450))) )),
  inference(resolution,[],[f4735,f1076])).
fof(f1076,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f804])).
fof(f804,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f803])).
fof(f803,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f335])).
fof(f335,axiom,(
  ! [X30] : ((p(s(bool,'T_0')) & p(s(bool,X30))) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.AND_CLAUSES|split|0')).
fof(f8448,plain,(
  ( ! [X445,X447,X446] : (sP9(X445,X446) | sP8(X446,'T_0') | ~p(s(bool,X447))) )),
  inference(resolution,[],[f4735,f1074])).
fof(f1074,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f801])).
fof(f801,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f336])).
fof(f336,axiom,(
  ! [X30] : ((p(s(bool,X30)) & p(s(bool,'T_0'))) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.AND_CLAUSES|split|1')).
fof(f8447,plain,(
  ( ! [X443,X444,X442] : (sP9(X442,X443) | sP8(X443,'T_0') | ~p(s(bool,X444))) )),
  inference(resolution,[],[f4735,f1052])).
fof(f1052,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f788])).
fof(f788,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f787])).
fof(f787,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f361])).
fof(f361,axiom,(
  ! [X30] : ((p(s(bool,'T_0')) | p(s(bool,X30))) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.OR_CLAUSES|split|0')).
fof(f8446,plain,(
  ( ! [X441,X440,X439] : (sP9(X439,X440) | sP8(X440,'T_0') | ~p(s(bool,X441))) )),
  inference(resolution,[],[f4735,f1048])).
fof(f1048,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f786])).
fof(f786,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f785])).
fof(f785,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f628])).
fof(f628,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f362])).
fof(f362,axiom,(
  ! [X30] : ((p(s(bool,X30)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.OR_CLAUSES|split|1')).
fof(f8445,plain,(
  ( ! [X436,X438,X437] : (sP9(X436,X437) | sP8(X437,'T_0') | ~p(s(bool,X438))) )),
  inference(resolution,[],[f4735,f1046])).
fof(f1046,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f784])).
fof(f784,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f783])).
fof(f783,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f681])).
fof(f681,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f329])).
fof(f329,axiom,(
  ! [X30] : ((p(s(bool,X30)) => p(s(bool,X30))) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CLAUSES|split|3')).
fof(f8444,plain,(
  ( ! [X434,X433,X435] : (sP9(X433,X434) | sP8(X434,'T_0') | p(s(bool,X435))) )),
  inference(resolution,[],[f4735,f1045])).
fof(f1045,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f784])).
fof(f8443,plain,(
  ( ! [X430,X432,X431] : (sP9(X430,X431) | sP8(X431,'T_0') | p(s(bool,X432))) )),
  inference(resolution,[],[f4735,f1042])).
fof(f1042,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f782])).
fof(f782,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f781])).
fof(f781,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f680])).
fof(f680,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f326])).
fof(f326,axiom,(
  ! [X30] : ((p(s(bool,'T_0')) => p(s(bool,X30))) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CLAUSES|split|0')).
fof(f8442,plain,(
  ( ! [X428,X429,X427] : (sP9(X427,X428) | sP8(X428,'T_0') | ~p(s(bool,X429))) )),
  inference(resolution,[],[f4735,f1040])).
fof(f1040,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f780])).
fof(f780,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f779])).
fof(f779,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f679])).
fof(f679,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f328])).
fof(f328,axiom,(
  ! [X30] : ((p(s(bool,'F_0')) => p(s(bool,X30))) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f8440,plain,(
  ( ! [X422,X424,X423] : (sP9(X422,X423) | sP8(X423,'T_0') | p(s(bool,X424))) )),
  inference(resolution,[],[f4735,f1036])).
fof(f1036,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f778])).
fof(f778,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f777])).
fof(f777,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f678])).
fof(f678,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f327])).
fof(f327,axiom,(
  ! [X30] : ((p(s(bool,X30)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CLAUSES|split|1')).
fof(f8535,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X1,X2)) ) | $spl94),
  inference(cnf_transformation,[],[f8535_D])).
fof(f8535_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X1,X2)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8430,plain,(
  ( ! [X383,X385,X382,X384,X386] : (sP9(X382,X383) | sP8(X383,X384) | p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,X384) != s(bool,X386)) )),
  inference(resolution,[],[f4735,f1102])).
fof(f8353,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP9(X42,X43) | sP8(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f4735,f892])).
fof(f8351,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP9(X30,X31) | sP8(X31,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f4735,f891])).
fof(f8349,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP9(X18,X19) | sP8(X19,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4735,f890])).
fof(f8348,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP9(X12,X13) | sP8(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f4735,f888])).
fof(f8346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1) | sP8(X1,X2) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f4735,f886])).
fof(f7995,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | p(s(bool,X1)) | sP8(X0,X1)) ) | $spl88),
  inference(cnf_transformation,[],[f7995_D])).
fof(f7995_D,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | p(s(bool,X1)) | sP8(X0,X1)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8345,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl84),
  inference(subsumption_resolution,[],[f8322,f1055])).
fof(f8322,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,X4) | sP8('T_0',X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl84),
  inference(resolution,[],[f8203,f996])).
fof(f8203,plain,(
  ( ! [X420,X421,X419] : (sP9(X419,X420) | sP8(X421,X419) | sP8('T_0',X421)) ) | $spl84),
  inference(resolution,[],[f3671,f7672])).
fof(f8344,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl84),
  inference(subsumption_resolution,[],[f8321,f1055])).
fof(f8321,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1) | sP8('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl84),
  inference(resolution,[],[f8203,f1000])).
fof(f8304,plain,(
  ( ! [X461,X462] : (sP9(X461,X462) | sP8('F_0',X461)) ) | $spl92),
  inference(cnf_transformation,[],[f8304_D])).
fof(f8304_D,plain,(
  ( ! [X461,X462] : (sP9(X461,X462) | sP8('F_0',X461)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8220,plain,(
  ( ! [X466,X465,X467,X468] : (sP9(X465,X466) | sP8('F_0',X465) | p(s(bool,X467)) | p(s(bool,X468)) | ~sP9(X467,X468)) )),
  inference(resolution,[],[f3671,f1000])).
fof(f8219,plain,(
  ( ! [X461,X463,X462,X464] : (sP9(X461,X462) | sP8('F_0',X461) | p(s(bool,X463)) | p(s(bool,X464)) | ~sP9(X463,X464)) )),
  inference(resolution,[],[f3671,f996])).
fof(f8299,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X2,X0)) ) | $spl90),
  inference(cnf_transformation,[],[f8299_D])).
fof(f8299_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X2,X0)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8194,plain,(
  ( ! [X383,X385,X382,X384,X386] : (sP9(X382,X383) | sP8(X384,X382) | p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,X384) != s(bool,X386)) )),
  inference(resolution,[],[f3671,f1102])).
fof(f8117,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP9(X42,X43) | sP8(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3671,f892])).
fof(f8115,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP9(X30,X31) | sP8(X32,X30) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f3671,f891])).
fof(f8113,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP9(X18,X19) | sP8(X20,X18) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3671,f890])).
fof(f8112,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP9(X12,X13) | sP8(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3671,f888])).
fof(f8110,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1) | sP8(X2,X0) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3671,f886])).
fof(f7996,plain,(
  ( ! [X4,X5,X3] : (sP8('T_0',X3) | sP8(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl84),
  inference(subsumption_resolution,[],[f7981,f1055])).
fof(f7981,plain,(
  ( ! [X4,X5,X3] : (sP8('T_0',X3) | sP8(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl84),
  inference(resolution,[],[f7786,f996])).
fof(f7786,plain,(
  ( ! [X220,X219,X221] : (sP8('T_0',X219) | sP9(X220,X221) | sP8(X219,X221)) ) | $spl84),
  inference(resolution,[],[f7672,f3642])).
fof(f7994,plain,(
  ( ! [X2,X0,X1] : (sP8('T_0',X0) | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl84),
  inference(subsumption_resolution,[],[f7980,f1055])).
fof(f7980,plain,(
  ( ! [X2,X0,X1] : (sP8('T_0',X0) | sP8(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl84),
  inference(resolution,[],[f7786,f1000])).
fof(f7878,plain,(
  ( ! [X2,X0] : (sP8('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl86),
  inference(cnf_transformation,[],[f7878_D])).
fof(f7878_D,plain,(
  ( ! [X2,X0] : (sP8('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7783,plain,(
  ( ! [X212,X213,X214] : (sP8('T_0',X212) | p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X212) != s(bool,X214)) ) | $spl84),
  inference(resolution,[],[f7672,f1102])).
fof(f7708,plain,(
  ( ! [X30,X28,X31,X29] : (sP8('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl84),
  inference(resolution,[],[f7672,f892])).
fof(f7705,plain,(
  ( ! [X19,X17,X18,X16] : (sP8('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl84),
  inference(resolution,[],[f7672,f891])).
fof(f7704,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl84),
  inference(resolution,[],[f7672,f890])).
fof(f7703,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl84),
  inference(resolution,[],[f7672,f888])).
fof(f7701,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl84),
  inference(resolution,[],[f7672,f886])).
fof(f7672,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8('T_0',X0)) ) | $spl84),
  inference(cnf_transformation,[],[f7672_D])).
fof(f7672_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP8('T_0',X0)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7673,plain,(
  ( ! [X2,X3] : (sP8('T_0',X2) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f7670,f1055])).
fof(f7670,plain,(
  ( ! [X2,X3] : (sP8('T_0',X2) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(resolution,[],[f7570,f996])).
fof(f7570,plain,(
  ( ! [X453,X452] : (sP9(X452,X453) | sP8('T_0',X453)) )),
  inference(resolution,[],[f3642,f1100])).
fof(f7671,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f7669,f1055])).
fof(f7669,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(resolution,[],[f7570,f1000])).
fof(f7658,plain,(
  ( ! [X420,X419] : (sP9(X419,X420) | sP8('T_0',X420)) ) | $spl80),
  inference(cnf_transformation,[],[f7658_D])).
fof(f7658_D,plain,(
  ( ! [X420,X419] : (sP9(X419,X420) | sP8('T_0',X420)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7646,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | sP9(X74,X75) | sP8('T_0',X75)) )),
  inference(duplicate_literal_removal,[],[f7614])).
fof(f7614,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | sP9(X74,X75) | sP8('T_0',X75) | ~p(s(bool,X73))) )),
  inference(superposition,[],[f3642,f1083])).
fof(f7647,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | sP9(X71,X72) | sP8('T_0',X72)) )),
  inference(duplicate_literal_removal,[],[f7613])).
fof(f7613,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | sP9(X71,X72) | sP8('T_0',X72) | ~p(s(bool,X70))) )),
  inference(superposition,[],[f3642,f1085])).
fof(f7662,plain,(
  ( ! [X459,X458] : (sP9(X458,X459) | sP8('F_0',X459)) ) | $spl82),
  inference(cnf_transformation,[],[f7662_D])).
fof(f7662_D,plain,(
  ( ! [X459,X458] : (sP9(X458,X459) | sP8('F_0',X459)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7574,plain,(
  ( ! [X463,X465,X462,X464] : (sP9(X462,X463) | sP8('F_0',X463) | p(s(bool,X464)) | p(s(bool,X465)) | ~sP9(X464,X465)) )),
  inference(resolution,[],[f3642,f1000])).
fof(f7573,plain,(
  ( ! [X461,X459,X460,X458] : (sP9(X458,X459) | sP8('F_0',X459) | p(s(bool,X460)) | p(s(bool,X461)) | ~sP9(X460,X461)) )),
  inference(resolution,[],[f3642,f996])).
fof(f7567,plain,(
  ( ! [X445,X447,X446] : (sP9(X445,X446) | sP8('T_0',X446) | ~p(s(bool,X447))) )),
  inference(resolution,[],[f3642,f1076])).
fof(f7566,plain,(
  ( ! [X443,X444,X442] : (sP9(X442,X443) | sP8('T_0',X443) | ~p(s(bool,X444))) )),
  inference(resolution,[],[f3642,f1074])).
fof(f7565,plain,(
  ( ! [X441,X440,X439] : (sP9(X439,X440) | sP8('T_0',X440) | ~p(s(bool,X441))) )),
  inference(resolution,[],[f3642,f1052])).
fof(f7564,plain,(
  ( ! [X436,X438,X437] : (sP9(X436,X437) | sP8('T_0',X437) | ~p(s(bool,X438))) )),
  inference(resolution,[],[f3642,f1048])).
fof(f7563,plain,(
  ( ! [X434,X433,X435] : (sP9(X433,X434) | sP8('T_0',X434) | ~p(s(bool,X435))) )),
  inference(resolution,[],[f3642,f1046])).
fof(f7562,plain,(
  ( ! [X430,X432,X431] : (sP9(X430,X431) | sP8('T_0',X431) | p(s(bool,X432))) )),
  inference(resolution,[],[f3642,f1045])).
fof(f7561,plain,(
  ( ! [X428,X429,X427] : (sP9(X427,X428) | sP8('T_0',X428) | p(s(bool,X429))) )),
  inference(resolution,[],[f3642,f1042])).
fof(f7560,plain,(
  ( ! [X424,X426,X425] : (sP9(X424,X425) | sP8('T_0',X425) | ~p(s(bool,X426))) )),
  inference(resolution,[],[f3642,f1040])).
fof(f7558,plain,(
  ( ! [X420,X421,X419] : (sP9(X419,X420) | sP8('T_0',X420) | p(s(bool,X421))) )),
  inference(resolution,[],[f3642,f1036])).
fof(f7653,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X2,X1)) ) | $spl78),
  inference(cnf_transformation,[],[f7653_D])).
fof(f7653_D,plain,(
  ( ! [X4,X2,X0,X1] : (sP9(X0,X1) | s(bool,X2) != s(bool,X4) | ~p(s(bool,X4)) | sP8(X2,X1)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7549,plain,(
  ( ! [X383,X385,X382,X384,X386] : (sP9(X382,X383) | sP8(X384,X383) | p(s(bool,X385)) | ~p(s(bool,X386)) | s(bool,X384) != s(bool,X386)) )),
  inference(resolution,[],[f3642,f1102])).
fof(f7472,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP9(X42,X43) | sP8(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f3642,f892])).
fof(f7470,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP9(X30,X31) | sP8(X32,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f3642,f891])).
fof(f7468,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP9(X18,X19) | sP8(X20,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3642,f890])).
fof(f7467,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP9(X12,X13) | sP8(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f3642,f888])).
fof(f7465,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1) | sP8(X2,X1) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X2) != s(bool,X4)) )),
  inference(resolution,[],[f3642,f886])).
fof(f6429,plain,(
  ( ! [X6,X7,X5] : (sP7(X6,X7,X5) | ~sP5(X7,X5,X6) | ~sP6(X6,X5,X7) | ~p(s(bool,X5))) ) | $spl72),
  inference(cnf_transformation,[],[f6429_D])).
fof(f6429_D,plain,(
  ( ! [X6,X7,X5] : (sP7(X6,X7,X5) | ~sP5(X7,X5,X6) | ~sP6(X6,X5,X7) | ~p(s(bool,X5))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7122,plain,(
  ( ! [X4] : (sP8(X4,sK14)) ) | $spl76),
  inference(cnf_transformation,[],[f7122_D])).
fof(f7122_D,plain,(
  ( ! [X4] : (sP8(X4,sK14)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7462,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,sK14) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f7449,f1726])).
fof(f1726,plain,(
  ( ! [X165,X164] : (~sP8(X165,X164) | ~p(s(bool,X165)) | sP8(X164,sK14)) )),
  inference(resolution,[],[f1168,f981])).
fof(f1168,plain,(
  ( ! [X115] : (p(s(bool,X115)) | sP8(X115,sK14)) )),
  inference(resolution,[],[f811,f985])).
fof(f7449,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,sK14) | sP8(X5,X4) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(resolution,[],[f5866,f931])).
fof(f931,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2)) | ~sP6(X0,X2,X1) | ~sP5(X1,X2,X0) | ~sP4(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f729])).
fof(f729,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP4(X2,X1,X0) & sP5(X1,X2,X0) & sP6(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X2,X1,X0) | ~sP5(X1,X2,X0) | ~sP6(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f728])).
fof(f728,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP7(X2,X1,X0)))),
  inference(flattening,[],[f727])).
fof(f727,plain,(
  ! [X2,X1,X0] : ((~sP7(X2,X1,X0) | (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | ~sP5(X1,X0,X2) | ~sP6(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP7(X2,X1,X0)))),
  inference(nnf_transformation,[],[f696])).
fof(f696,plain,(
  ! [X2,X1,X0] : (sP7(X2,X1,X0) <=> (sP4(X0,X1,X2) & sP5(X1,X0,X2) & sP6(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f5866,plain,(
  ( ! [X206,X204,X207,X205] : (sP8(X204,sK14) | sP8(X205,X204) | sP4(X205,X206,X207)) )),
  inference(resolution,[],[f1729,f945])).
fof(f945,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f736])).
fof(f736,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f693])).
fof(f693,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f1729,plain,(
  ( ! [X171,X170] : (sP8(X170,sK14) | sP8(X171,X170) | p(s(bool,X171))) )),
  inference(resolution,[],[f1168,f985])).
fof(f7121,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,sK14) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f7108,f1727])).
fof(f1727,plain,(
  ( ! [X167,X166] : (~sP8(X166,X167) | ~p(s(bool,X167)) | sP8(X166,sK14)) )),
  inference(resolution,[],[f1168,f981])).
fof(f7108,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,sK14) | sP8(X4,X5) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(resolution,[],[f5645,f931])).
fof(f5645,plain,(
  ( ! [X206,X204,X207,X205] : (sP8(X204,sK14) | sP8(X204,X205) | sP4(X205,X206,X207)) )),
  inference(resolution,[],[f1728,f945])).
fof(f1728,plain,(
  ( ! [X169,X168] : (sP8(X168,sK14) | sP8(X168,X169) | p(s(bool,X169))) )),
  inference(resolution,[],[f1168,f982])).
fof(f6430,plain,(
  ( ! [X4] : (sP8(sK14,X4)) ) | $spl74),
  inference(cnf_transformation,[],[f6430_D])).
fof(f6430_D,plain,(
  ( ! [X4] : (sP8(sK14,X4)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f6781,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(sK14,X4) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f6766,f1528])).
fof(f1528,plain,(
  ( ! [X165,X164] : (~sP8(X165,X164) | ~p(s(bool,X165)) | sP8(sK14,X164)) )),
  inference(resolution,[],[f1167,f981])).
fof(f6766,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(sK14,X4) | sP8(X5,X4) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(resolution,[],[f4246,f931])).
fof(f4246,plain,(
  ( ! [X206,X204,X207,X205] : (sP8(sK14,X204) | sP8(X205,X204) | sP4(X205,X206,X207)) )),
  inference(resolution,[],[f1531,f945])).
fof(f6428,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(sK14,X4) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f6413,f1529])).
fof(f6413,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(sK14,X4) | sP8(X4,X5) | sP7(X6,X7,X5) | ~p(s(bool,X5)) | ~sP6(X6,X5,X7) | ~sP5(X7,X5,X6)) )),
  inference(resolution,[],[f3025,f931])).
fof(f3025,plain,(
  ( ! [X206,X204,X207,X205] : (sP8(sK14,X204) | sP8(X204,X205) | sP4(X205,X206,X207)) )),
  inference(resolution,[],[f1530,f945])).
fof(f6004,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,sK14) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) | $spl70),
  inference(cnf_transformation,[],[f6004_D])).
fof(f6004_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,sK14) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5899,plain,(
  ( ! [X294,X296,X295,X297] : (sP8(X294,sK14) | sP8(X295,X294) | p(s(bool,X296)) | p(s(bool,X297)) | s(bool,X295) != s(bool,X297)) )),
  inference(resolution,[],[f1729,f1102])).
fof(f5824,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,sK14) | sP8(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1729,f892])).
fof(f5821,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,sK14) | sP8(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1729,f891])).
fof(f5820,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,sK14) | sP8(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1729,f890])).
fof(f5819,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,sK14) | sP8(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1729,f888])).
fof(f5817,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,sK14) | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1729,f886])).
fof(f5784,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,sK14) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) | $spl68),
  inference(cnf_transformation,[],[f5784_D])).
fof(f5784_D,plain,(
  ( ! [X0,X3,X1] : (sP8(X0,sK14) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5678,plain,(
  ( ! [X294,X296,X295,X297] : (sP8(X294,sK14) | sP8(X294,X295) | p(s(bool,X296)) | p(s(bool,X297)) | s(bool,X295) != s(bool,X297)) )),
  inference(resolution,[],[f1728,f1102])).
fof(f5603,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(X35,sK14) | sP8(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1728,f892])).
fof(f5600,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(X20,sK14) | sP8(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1728,f891])).
fof(f5599,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(X15,sK14) | sP8(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1728,f890])).
fof(f5598,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(X10,sK14) | sP8(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1728,f888])).
fof(f5596,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,sK14) | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1728,f886])).
fof(f5246,plain,(
  ( ! [X0,X1] : (sP8(X0,sK14) | p(s(bool,X1)) | sP8(X1,X0)) ) | $spl66),
  inference(cnf_transformation,[],[f5246_D])).
fof(f5246_D,plain,(
  ( ! [X0,X1] : (sP8(X0,sK14) | p(s(bool,X1)) | sP8(X1,X0)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5282,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X4,X3) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f5269,f1055])).
fof(f5269,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f5136,f996])).
fof(f5136,plain,(
  ( ! [X70,X68,X69] : (sP8(X68,sK14) | sP8(X69,X68) | sP9(X69,X70)) )),
  inference(subsumption_resolution,[],[f5135,f1105])).
fof(f1105,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f1026])).
fof(f1026,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f771])).
fof(f771,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f617])).
fof(f617,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X31] : (s(X6,X31) = s(X6,X31) <=> p(s(bool,'T_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f5135,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | sP8(X68,sK14) | sP8(X69,X68) | sP9(X69,X70)) )),
  inference(forward_demodulation,[],[f5099,f1193])).
fof(f5099,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,sK14)) | sP8(X68,sK14) | sP8(X69,X68) | sP9(X69,X70)) )),
  inference(resolution,[],[f1726,f4258])).
fof(f5281,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X1,X0) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f5268,f1055])).
fof(f5268,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(resolution,[],[f5136,f1000])).
fof(f5247,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f5233,f1055])).
fof(f5233,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(resolution,[],[f5134,f996])).
fof(f5134,plain,(
  ( ! [X66,X67,X65] : (sP8(X65,sK14) | sP8(X66,X65) | sP9(X67,X66)) )),
  inference(subsumption_resolution,[],[f5133,f1105])).
fof(f5133,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | sP8(X65,sK14) | sP8(X66,X65) | sP9(X67,X66)) )),
  inference(forward_demodulation,[],[f5098,f1193])).
fof(f5098,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,sK14)) | sP8(X65,sK14) | sP8(X66,X65) | sP9(X67,X66)) )),
  inference(resolution,[],[f1726,f4257])).
fof(f5245,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f5232,f1055])).
fof(f5232,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(resolution,[],[f5134,f1000])).
fof(f5176,plain,(
  ( ! [X0,X1] : (sP8(X0,sK14) | p(s(bool,X1)) | sP8(X0,X1)) ) | $spl64),
  inference(cnf_transformation,[],[f5176_D])).
fof(f5176_D,plain,(
  ( ! [X0,X1] : (sP8(X0,sK14) | p(s(bool,X1)) | sP8(X0,X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5211,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f5199,f1055])).
fof(f5199,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f5130,f996])).
fof(f5130,plain,(
  ( ! [X61,X59,X60] : (sP8(X59,sK14) | sP8(X59,X60) | sP9(X60,X61)) )),
  inference(subsumption_resolution,[],[f5129,f1105])).
fof(f5129,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'T_0')) | sP8(X59,sK14) | sP8(X59,X60) | sP9(X60,X61)) )),
  inference(forward_demodulation,[],[f5095,f1193])).
fof(f5095,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,sK14)) | sP8(X59,sK14) | sP8(X59,X60) | sP9(X60,X61)) )),
  inference(resolution,[],[f1726,f3037])).
fof(f5210,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f5198,f1055])).
fof(f5198,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(resolution,[],[f5130,f1000])).
fof(f5177,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f5164,f1055])).
fof(f5164,plain,(
  ( ! [X4,X5,X3] : (sP8(X3,sK14) | sP8(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(resolution,[],[f5128,f996])).
fof(f5128,plain,(
  ( ! [X57,X58,X56] : (sP8(X56,sK14) | sP8(X56,X57) | sP9(X58,X57)) )),
  inference(subsumption_resolution,[],[f5127,f1105])).
fof(f5127,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'T_0')) | sP8(X56,sK14) | sP8(X56,X57) | sP9(X58,X57)) )),
  inference(forward_demodulation,[],[f5094,f1193])).
fof(f5094,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,sK14)) | sP8(X56,sK14) | sP8(X56,X57) | sP9(X58,X57)) )),
  inference(resolution,[],[f1726,f3036])).
fof(f5175,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f5163,f1055])).
fof(f5163,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,sK14) | sP8(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(resolution,[],[f5128,f1000])).
fof(f4607,plain,(
  sP8('F_0','T_0') | $spl62),
  inference(cnf_transformation,[],[f4607_D])).
fof(f4607_D,plain,(
  sP8('F_0','T_0') <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4594,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | sP8('F_0','T_0')) )),
  inference(duplicate_literal_removal,[],[f4562])).
fof(f4562,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | sP8('F_0','T_0') | ~p(s(bool,X35))) )),
  inference(superposition,[],[f4399,f1083])).
fof(f4399,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP8('F_0',X0)) )),
  inference(resolution,[],[f4295,f1166])).
fof(f4295,plain,(
  ( ! [X330] : (sP8(sK14,X330) | sP8('F_0',X330)) )),
  inference(resolution,[],[f1531,f1095])).
fof(f1095,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f653])).
fof(f653,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb',aHL_FALSITY)).
fof(f4595,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP8('F_0','T_0')) )),
  inference(duplicate_literal_removal,[],[f4561])).
fof(f4561,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP8('F_0','T_0') | ~p(s(bool,X34))) )),
  inference(superposition,[],[f4399,f1085])).
fof(f3434,plain,(
  sP8('F_0','F_0') | $spl50),
  inference(cnf_transformation,[],[f3434_D])).
fof(f3434_D,plain,(
  sP8('F_0','F_0') <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4522,plain,(
  ( ! [X241,X240] : (sP8('F_0','F_0') | p(s(bool,X240)) | p(s(bool,X241)) | ~sP9(X240,X241)) )),
  inference(resolution,[],[f4399,f1000])).
fof(f4521,plain,(
  ( ! [X239,X238] : (sP8('F_0','F_0') | p(s(bool,X238)) | p(s(bool,X239)) | ~sP9(X238,X239)) )),
  inference(resolution,[],[f4399,f996])).
fof(f4515,plain,(
  ( ! [X237] : (sP8('F_0','T_0') | ~p(s(bool,X237))) )),
  inference(resolution,[],[f4399,f1076])).
fof(f4514,plain,(
  ( ! [X236] : (sP8('F_0','T_0') | ~p(s(bool,X236))) )),
  inference(resolution,[],[f4399,f1074])).
fof(f4513,plain,(
  ( ! [X235] : (sP8('F_0','T_0') | ~p(s(bool,X235))) )),
  inference(resolution,[],[f4399,f1052])).
fof(f4512,plain,(
  ( ! [X234] : (sP8('F_0','T_0') | ~p(s(bool,X234))) )),
  inference(resolution,[],[f4399,f1048])).
fof(f4511,plain,(
  ( ! [X233] : (sP8('F_0','T_0') | ~p(s(bool,X233))) )),
  inference(resolution,[],[f4399,f1046])).
fof(f4510,plain,(
  ( ! [X232] : (sP8('F_0','T_0') | p(s(bool,X232))) )),
  inference(resolution,[],[f4399,f1045])).
fof(f4509,plain,(
  ( ! [X231] : (sP8('F_0','T_0') | p(s(bool,X231))) )),
  inference(resolution,[],[f4399,f1042])).
fof(f4508,plain,(
  ( ! [X230] : (sP8('F_0','T_0') | ~p(s(bool,X230))) )),
  inference(resolution,[],[f4399,f1040])).
fof(f4506,plain,(
  ( ! [X229] : (sP8('F_0','T_0') | p(s(bool,X229))) )),
  inference(resolution,[],[f4399,f1036])).
fof(f4601,plain,(
  ( ! [X2,X0] : (sP8('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl60),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X2,X0] : (sP8('F_0',X0) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4499,plain,(
  ( ! [X216,X215,X214] : (sP8('F_0',X214) | p(s(bool,X215)) | ~p(s(bool,X216)) | s(bool,X214) != s(bool,X216)) )),
  inference(resolution,[],[f4399,f1102])).
fof(f4422,plain,(
  ( ! [X30,X28,X31,X29] : (sP8('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4399,f892])).
fof(f4420,plain,(
  ( ! [X23,X21,X22,X20] : (sP8('F_0',X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4399,f891])).
fof(f4418,plain,(
  ( ! [X14,X12,X15,X13] : (sP8('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4399,f890])).
fof(f4417,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4399,f888])).
fof(f4415,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('F_0',X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f4399,f886])).
fof(f4387,plain,(
  ( ! [X308] : (sP8(sK14,X308) | sP8('F_0',X308)) ) | $spl58),
  inference(cnf_transformation,[],[f4387_D])).
fof(f4387_D,plain,(
  ( ! [X308] : (sP8(sK14,X308) | sP8('F_0',X308)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4367,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | sP8(sK14,X61) | sP8('F_0',X61)) )),
  inference(duplicate_literal_removal,[],[f4337])).
fof(f4337,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | sP8(sK14,X61) | sP8('F_0',X61) | p(s(bool,X60))) )),
  inference(superposition,[],[f1531,f1033])).
fof(f1033,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f776])).
fof(f776,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f775])).
fof(f775,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f360])).
fof(f360,axiom,(
  ! [X30] : ((p(s(bool,X30)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X30))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_F_EQ_F')).
fof(f4368,plain,(
  ( ! [X59,X58] : (p(s(bool,X58)) | sP8(sK14,X59) | sP8('F_0',X59)) )),
  inference(duplicate_literal_removal,[],[f4336])).
fof(f4336,plain,(
  ( ! [X59,X58] : (p(s(bool,X58)) | sP8(sK14,X59) | sP8('F_0',X59) | p(s(bool,X58))) )),
  inference(superposition,[],[f1531,f1079])).
fof(f1079,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f805])).
fof(f805,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f638])).
fof(f638,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X30] : (s(bool,'F_0') = s(bool,X30) <=> ~p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f4369,plain,(
  ( ! [X57,X56] : (p(s(bool,X56)) | sP8(sK14,X57) | sP8('F_0',X57)) )),
  inference(duplicate_literal_removal,[],[f4335])).
fof(f4335,plain,(
  ( ! [X57,X56] : (p(s(bool,X56)) | sP8(sK14,X57) | sP8('F_0',X57) | p(s(bool,X56))) )),
  inference(superposition,[],[f1531,f1081])).
fof(f1081,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f640])).
fof(f640,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X30] : (s(bool,'F_0') = s(bool,X30) <=> ~p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f4294,plain,(
  ( ! [X329,X328] : (sP8(sK14,X328) | sP8('F_0',X328) | p(s(bool,X329))) )),
  inference(resolution,[],[f1531,f1091])).
fof(f1091,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f648])).
fof(f648,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f331])).
fof(f331,axiom,(
  ! [X30] : (p(s(bool,'F_0')) => p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.FALSITY')).
fof(f4293,plain,(
  ( ! [X327,X326] : (sP8(sK14,X326) | sP8('F_0',X326) | p(s(bool,X327))) )),
  inference(resolution,[],[f1531,f1070])).
fof(f1070,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f800])).
fof(f800,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f799])).
fof(f799,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f635])).
fof(f635,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f338])).
fof(f338,axiom,(
  ! [X30] : ((p(s(bool,X30)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.AND_CLAUSES|split|3')).
fof(f4292,plain,(
  ( ! [X325,X324] : (sP8(sK14,X324) | sP8('F_0',X324) | p(s(bool,X325))) )),
  inference(resolution,[],[f1531,f1068])).
fof(f1068,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f798])).
fof(f798,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f797])).
fof(f797,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f337])).
fof(f337,axiom,(
  ! [X30] : ((p(s(bool,'F_0')) & p(s(bool,X30))) <=> p(s(bool,'F_0')))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f4291,plain,(
  ( ! [X323,X322] : (sP8(sK14,X322) | sP8('F_0',X322) | p(s(bool,X323))) )),
  inference(resolution,[],[f1531,f1060])).
fof(f1060,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f794])).
fof(f794,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f793])).
fof(f793,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f632])).
fof(f632,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f365])).
fof(f365,axiom,(
  ! [X30] : ((p(s(bool,'F_0')) | p(s(bool,X30))) <=> p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.OR_CLAUSES|split|2')).
fof(f4290,plain,(
  ( ! [X321,X320] : (sP8(sK14,X320) | sP8('F_0',X320) | p(s(bool,X321))) )),
  inference(resolution,[],[f1531,f1055])).
fof(f4388,plain,(
  ( ! [X318,X319] : (sP8(sK14,X318) | sP8('F_0',X318) | p(s(bool,X319))) )),
  inference(subsumption_resolution,[],[f4289,f1105])).
fof(f4289,plain,(
  ( ! [X318,X319] : (sP8(sK14,X318) | sP8('F_0',X318) | p(s(bool,X319)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1531,f1041])).
fof(f1041,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f780])).
fof(f1656,plain,(
  ( ! [X224] : (s(bool,'F_0') = s(bool,X224)) ) | $spl30),
  inference(cnf_transformation,[],[f1656_D])).
fof(f1656_D,plain,(
  ( ! [X224] : (s(bool,'F_0') = s(bool,X224)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4288,plain,(
  ( ! [X316,X317] : (sP8(sK14,X316) | sP8('F_0',X316) | s(bool,'F_0') = s(bool,X317)) )),
  inference(resolution,[],[f1531,f1034])).
fof(f1034,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f776])).
fof(f4287,plain,(
  ( ! [X315,X314] : (sP8(sK14,X314) | sP8('F_0',X314) | ~p(s(bool,X315))) )),
  inference(resolution,[],[f1531,f1031])).
fof(f1031,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f774])).
fof(f774,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f773])).
fof(f773,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f676])).
fof(f676,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f621])).
fof(f621,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f330])).
fof(f330,axiom,(
  ! [X30] : ((p(s(bool,X30)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X30)))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.IMP_CLAUSES|split|4')).
fof(f1653,plain,(
  ( ! [X220,X219] : (sP9(X219,X220)) ) | $spl26),
  inference(cnf_transformation,[],[f1653_D])).
fof(f1653_D,plain,(
  ( ! [X220,X219] : (sP9(X219,X220)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4286,plain,(
  ( ! [X313,X312,X311] : (sP8(sK14,X311) | sP8('F_0',X311) | sP9(X312,X313)) )),
  inference(resolution,[],[f1531,f1004])).
fof(f1004,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f761])).
fof(f4285,plain,(
  ( ! [X308,X310,X309] : (sP8(sK14,X308) | sP8('F_0',X308) | sP9(X309,X310)) )),
  inference(resolution,[],[f1531,f999])).
fof(f999,plain,(
  ( ! [X0,X1] : (sP9(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f759])).
fof(f4383,plain,(
  ( ! [X0,X3,X1] : (sP8(sK14,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) | $spl56),
  inference(cnf_transformation,[],[f4383_D])).
fof(f4383_D,plain,(
  ( ! [X0,X3,X1] : (sP8(sK14,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X1,X0)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4279,plain,(
  ( ! [X294,X296,X295,X297] : (sP8(sK14,X294) | sP8(X295,X294) | p(s(bool,X296)) | p(s(bool,X297)) | s(bool,X295) != s(bool,X297)) )),
  inference(resolution,[],[f1531,f1102])).
fof(f4204,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(sK14,X35) | sP8(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1531,f892])).
fof(f4201,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(sK14,X20) | sP8(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1531,f891])).
fof(f4200,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(sK14,X15) | sP8(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1531,f890])).
fof(f4199,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(sK14,X10) | sP8(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1531,f888])).
fof(f4197,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK14,X0) | sP8(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1531,f886])).
fof(f4104,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3993,f1222])).
fof(f3993,plain,(
  ( ! [X253,X254] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X253),s('type.num.num',X254))))) ) | $spl10),
  inference(resolution,[],[f1222,f1010])).
fof(f1010,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f764])).
fof(f764,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f611])).
fof(f611,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f610])).
fof(f610,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f401])).
fof(f401,axiom,(
  ! [X16,X53] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X53)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X16)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.numeral.numeral_lte|split|5')).
fof(f4103,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3991,f1222])).
fof(f3991,plain,(
  ( ! [X251,X252] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X251),s('type.num.num',X252))))) ) | $spl10),
  inference(resolution,[],[f1222,f978])).
fof(f978,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f751])).
fof(f751,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(flattening,[],[f600])).
fof(f600,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(rectify,[],[f445])).
fof(f445,axiom,(
  ! [X53,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X16)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num',X53)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.arithmetic.NOT_LEQ')).
fof(f4102,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3988,f1222])).
fof(f3988,plain,(
  ( ! [X249,X250] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X249),s('type.num.num',X250))))) ) | $spl10),
  inference(resolution,[],[f1222,f816])).
fof(f4101,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3986,f1222])).
fof(f3986,plain,(
  ( ! [X244,X245] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X244))),s('type.num.num',X245))))) ) | $spl10),
  inference(resolution,[],[f1222,f978])).
fof(f4100,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3985,f1222])).
fof(f3985,plain,(
  ( ! [X243,X242] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X242))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X243))))))) ) | $spl10),
  inference(resolution,[],[f1222,f1010])).
fof(f4099,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3982,f1222])).
fof(f3982,plain,(
  ( ! [X237,X236] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X236),s('type.num.num',X237))))) ) | $spl10),
  inference(resolution,[],[f1222,f819])).
fof(f819,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f705])).
fof(f705,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f525])).
fof(f525,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f496])).
fof(f496,axiom,(
  ! [X16,X53] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X53)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X53),s('type.num.num',X16)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.numeral.numeral_lt|split|5')).
fof(f4098,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3981,f1222])).
fof(f3981,plain,(
  ( ! [X235,X234] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X234),s('type.num.num',X235))))) ) | $spl10),
  inference(resolution,[],[f1222,f816])).
fof(f4097,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3980,f1222])).
fof(f3980,plain,(
  ( ! [X233,X232] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X232))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X233))))))) ) | $spl10),
  inference(resolution,[],[f1222,f819])).
fof(f4092,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3969,f1222])).
fof(f3969,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f1222,f1098])).
fof(f1098,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f810])).
fof(f810,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f655])).
fof(f655,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f376])).
fof(f376,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f4091,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3968,f1222])).
fof(f3968,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f1222,f1096])).
fof(f1096,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f809])).
fof(f809,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f654])).
fof(f654,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f370])).
fof(f370,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f4090,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3962,f1222])).
fof(f3962,plain,(
  ( ! [X219] : (p(s(bool,X219))) ) | $spl10),
  inference(resolution,[],[f1222,f1045])).
fof(f4089,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3961,f1222])).
fof(f3961,plain,(
  ( ! [X218] : (p(s(bool,X218))) ) | $spl10),
  inference(resolution,[],[f1222,f1042])).
fof(f4088,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3959,f1222])).
fof(f3959,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f1222,f1039])).
fof(f1039,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f780])).
fof(f4087,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3958,f1222])).
fof(f3958,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | $spl10),
  inference(resolution,[],[f1222,f1036])).
fof(f3970,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f1100])).
fof(f3971,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f1101])).
fof(f3972,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f1105])).
fof(f3989,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f1093])).
fof(f1093,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f651])).
fof(f651,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f334])).
fof(f334,axiom,(
  ! [X53] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X53))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.arithmetic.LESS_EQ_REFL')).
fof(f3990,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f877])).
fof(f3992,plain,(
  $false | $spl10),
  inference(resolution,[],[f1222,f1094])).
fof(f1094,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f431])).
fof(f431,axiom,(
  ! [X16] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.arithmetic.ZERO_LESS_EQ')).
fof(f3846,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) | $spl54),
  inference(cnf_transformation,[],[f3846_D])).
fof(f3846_D,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3840,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3753])).
fof(f3753,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50)) | ~p(s(bool,X50))) ) | $spl8),
  inference(superposition,[],[f1221,f1085])).
fof(f3841,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3752])).
fof(f3752,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48)) | ~p(s(bool,X48))) ) | $spl8),
  inference(superposition,[],[f1221,f1083])).
fof(f3618,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) | $spl52),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3611,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3523])).
fof(f3523,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46)) | p(s(bool,X46))) ) | $spl6),
  inference(superposition,[],[f1220,f1081])).
fof(f3612,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3522])).
fof(f3522,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | p(s(bool,X44))) ) | $spl6),
  inference(superposition,[],[f1220,f1079])).
fof(f3613,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3520])).
fof(f3520,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl6),
  inference(superposition,[],[f1220,f1033])).
fof(f3428,plain,(
  sP8('T_0','F_0') | $spl46),
  inference(cnf_transformation,[],[f3428_D])).
fof(f3428_D,plain,(
  sP8('T_0','F_0') <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3415,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | sP8('T_0','F_0')) )),
  inference(duplicate_literal_removal,[],[f3383])).
fof(f3383,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | sP8('T_0','F_0') | ~p(s(bool,X35))) )),
  inference(superposition,[],[f3176,f1083])).
fof(f3176,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP8(X0,'F_0')) )),
  inference(resolution,[],[f3073,f1166])).
fof(f3073,plain,(
  ( ! [X328] : (sP8(sK14,X328) | sP8(X328,'F_0')) )),
  inference(resolution,[],[f1530,f1095])).
fof(f3416,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP8('T_0','F_0')) )),
  inference(duplicate_literal_removal,[],[f3382])).
fof(f3382,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | sP8('T_0','F_0') | ~p(s(bool,X34))) )),
  inference(superposition,[],[f3176,f1085])).
fof(f3343,plain,(
  ( ! [X239,X238] : (sP8('F_0','F_0') | p(s(bool,X238)) | p(s(bool,X239)) | ~sP9(X238,X239)) )),
  inference(resolution,[],[f3176,f1000])).
fof(f3342,plain,(
  ( ! [X237,X236] : (sP8('F_0','F_0') | p(s(bool,X236)) | p(s(bool,X237)) | ~sP9(X236,X237)) )),
  inference(resolution,[],[f3176,f996])).
fof(f3336,plain,(
  ( ! [X235] : (sP8('T_0','F_0') | ~p(s(bool,X235))) )),
  inference(resolution,[],[f3176,f1076])).
fof(f3335,plain,(
  ( ! [X234] : (sP8('T_0','F_0') | ~p(s(bool,X234))) )),
  inference(resolution,[],[f3176,f1074])).
fof(f3334,plain,(
  ( ! [X233] : (sP8('T_0','F_0') | ~p(s(bool,X233))) )),
  inference(resolution,[],[f3176,f1052])).
fof(f3333,plain,(
  ( ! [X232] : (sP8('T_0','F_0') | ~p(s(bool,X232))) )),
  inference(resolution,[],[f3176,f1048])).
fof(f3332,plain,(
  ( ! [X231] : (sP8('T_0','F_0') | ~p(s(bool,X231))) )),
  inference(resolution,[],[f3176,f1046])).
fof(f3331,plain,(
  ( ! [X230] : (sP8('T_0','F_0') | p(s(bool,X230))) )),
  inference(resolution,[],[f3176,f1045])).
fof(f3330,plain,(
  ( ! [X229] : (sP8('T_0','F_0') | p(s(bool,X229))) )),
  inference(resolution,[],[f3176,f1042])).
fof(f3329,plain,(
  ( ! [X228] : (sP8('T_0','F_0') | ~p(s(bool,X228))) )),
  inference(resolution,[],[f3176,f1040])).
fof(f3327,plain,(
  ( ! [X227] : (sP8('T_0','F_0') | p(s(bool,X227))) )),
  inference(resolution,[],[f3176,f1036])).
fof(f3422,plain,(
  ( ! [X2,X0] : (sP8(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl44),
  inference(cnf_transformation,[],[f3422_D])).
fof(f3422_D,plain,(
  ( ! [X2,X0] : (sP8(X0,'F_0') | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3321,plain,(
  ( ! [X216,X215,X214] : (sP8(X214,'F_0') | p(s(bool,X215)) | ~p(s(bool,X216)) | s(bool,X214) != s(bool,X216)) )),
  inference(resolution,[],[f3176,f1102])).
fof(f3244,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f3176,f892])).
fof(f3242,plain,(
  ( ! [X23,X21,X22,X20] : (sP8(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f3176,f891])).
fof(f3240,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f3176,f890])).
fof(f3239,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f3176,f888])).
fof(f3237,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'F_0') | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f3176,f886])).
fof(f3228,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3227,f1101])).
fof(f3227,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3209,f1193])).
fof(f3209,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(duplicate_literal_removal,[],[f3208])).
fof(f3208,plain,(
  ~p(s(bool,sK14)) | ~p(s(bool,sK14)) | $spl40),
  inference(resolution,[],[f3193,f981])).
fof(f3224,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3223,f1101])).
fof(f3223,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3212,f1193])).
fof(f3212,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(duplicate_literal_removal,[],[f3205])).
fof(f3205,plain,(
  ~p(s(bool,sK14)) | ~p(s(bool,sK14)) | $spl40),
  inference(resolution,[],[f3193,f988])).
fof(f988,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X1,X0)) )),
  inference(cnf_transformation,[],[f755])).
fof(f755,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP8(X1,X0)) & (~sP8(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f699])).
fof(f699,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP8(X1,X0))),
  inference(definition_folding,[],[f603,f698])).
fof(f603,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f602])).
fof(f602,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f472])).
fof(f472,axiom,(
  ! [X132,X126] : ((p(s(bool,X126)) <=> ~p(s(bool,X132))) <=> ((p(s(bool,X126)) | p(s(bool,X132))) & (~p(s(bool,X132)) | ~p(s(bool,X126)))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.sat.sat_29_prove_ltb')).
fof(f3222,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3221,f1101])).
fof(f3221,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3220,f1193])).
fof(f3220,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(subsumption_resolution,[],[f3204,f1165])).
fof(f3204,plain,(
  ~p(s(bool,sK14)) | sP8(sK14,sK14) | $spl40),
  inference(resolution,[],[f3193,f1528])).
fof(f3219,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3218,f1101])).
fof(f3218,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3217,f1193])).
fof(f3217,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(subsumption_resolution,[],[f3203,f1165])).
fof(f3203,plain,(
  ~p(s(bool,sK14)) | sP8(sK14,sK14) | $spl40),
  inference(resolution,[],[f3193,f1529])).
fof(f3216,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3215,f1101])).
fof(f3215,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3202,f1193])).
fof(f3202,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(resolution,[],[f3193,f1165])).
fof(f3214,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f3213,f1101])).
fof(f3213,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(forward_demodulation,[],[f3201,f1193])).
fof(f3201,plain,(
  ~p(s(bool,sK14)) | $spl40),
  inference(resolution,[],[f3193,f1166])).
fof(f3193,plain,(
  sP8(sK14,sK14) | $spl40),
  inference(cnf_transformation,[],[f3193_D])).
fof(f3193_D,plain,(
  sP8(sK14,sK14) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3194,plain,(
  ( ! [X1] : (sP8(X1,'F_0') | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f3194_D])).
fof(f3194_D,plain,(
  ( ! [X1] : (sP8(X1,'F_0') | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3177,plain,(
  ( ! [X1] : (sP8(X1,'F_0') | ~p(s(bool,X1)) | sP8(sK14,sK14)) )),
  inference(resolution,[],[f3073,f1529])).
fof(f3164,plain,(
  ( ! [X306] : (sP8(sK14,X306) | sP8(X306,'F_0')) ) | $spl38),
  inference(cnf_transformation,[],[f3164_D])).
fof(f3164_D,plain,(
  ( ! [X306] : (sP8(sK14,X306) | sP8(X306,'F_0')) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3144,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | sP8(sK14,X61) | sP8(X61,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3115])).
fof(f3115,plain,(
  ( ! [X61,X60] : (p(s(bool,X60)) | sP8(sK14,X61) | sP8(X61,'F_0') | p(s(bool,X60))) )),
  inference(superposition,[],[f1530,f1033])).
fof(f3145,plain,(
  ( ! [X59,X58] : (p(s(bool,X58)) | sP8(sK14,X59) | sP8(X59,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3114])).
fof(f3114,plain,(
  ( ! [X59,X58] : (p(s(bool,X58)) | sP8(sK14,X59) | sP8(X59,'F_0') | p(s(bool,X58))) )),
  inference(superposition,[],[f1530,f1079])).
fof(f3146,plain,(
  ( ! [X57,X56] : (p(s(bool,X56)) | sP8(sK14,X57) | sP8(X57,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3113])).
fof(f3113,plain,(
  ( ! [X57,X56] : (p(s(bool,X56)) | sP8(sK14,X57) | sP8(X57,'F_0') | p(s(bool,X56))) )),
  inference(superposition,[],[f1530,f1081])).
fof(f3072,plain,(
  ( ! [X327,X326] : (sP8(sK14,X326) | sP8(X326,'F_0') | p(s(bool,X327))) )),
  inference(resolution,[],[f1530,f1091])).
fof(f3071,plain,(
  ( ! [X325,X324] : (sP8(sK14,X324) | sP8(X324,'F_0') | p(s(bool,X325))) )),
  inference(resolution,[],[f1530,f1070])).
fof(f3070,plain,(
  ( ! [X323,X322] : (sP8(sK14,X322) | sP8(X322,'F_0') | p(s(bool,X323))) )),
  inference(resolution,[],[f1530,f1068])).
fof(f3069,plain,(
  ( ! [X321,X320] : (sP8(sK14,X320) | sP8(X320,'F_0') | p(s(bool,X321))) )),
  inference(resolution,[],[f1530,f1060])).
fof(f3068,plain,(
  ( ! [X318,X319] : (sP8(sK14,X318) | sP8(X318,'F_0') | p(s(bool,X319))) )),
  inference(resolution,[],[f1530,f1055])).
fof(f3165,plain,(
  ( ! [X316,X317] : (sP8(sK14,X316) | sP8(X316,'F_0') | p(s(bool,X317))) )),
  inference(subsumption_resolution,[],[f3067,f1105])).
fof(f3067,plain,(
  ( ! [X316,X317] : (sP8(sK14,X316) | sP8(X316,'F_0') | p(s(bool,X317)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1530,f1041])).
fof(f3066,plain,(
  ( ! [X315,X314] : (sP8(sK14,X314) | sP8(X314,'F_0') | s(bool,'F_0') = s(bool,X315)) )),
  inference(resolution,[],[f1530,f1034])).
fof(f3065,plain,(
  ( ! [X313,X312] : (sP8(sK14,X312) | sP8(X312,'F_0') | ~p(s(bool,X313))) )),
  inference(resolution,[],[f1530,f1031])).
fof(f3064,plain,(
  ( ! [X310,X309,X311] : (sP8(sK14,X309) | sP8(X309,'F_0') | sP9(X310,X311)) )),
  inference(resolution,[],[f1530,f1004])).
fof(f3063,plain,(
  ( ! [X308,X306,X307] : (sP8(sK14,X306) | sP8(X306,'F_0') | sP9(X307,X308)) )),
  inference(resolution,[],[f1530,f999])).
fof(f3160,plain,(
  ( ! [X0,X3,X1] : (sP8(sK14,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) | $spl36),
  inference(cnf_transformation,[],[f3160_D])).
fof(f3160_D,plain,(
  ( ! [X0,X3,X1] : (sP8(sK14,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP8(X0,X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3058,plain,(
  ( ! [X294,X296,X295,X297] : (sP8(sK14,X294) | sP8(X294,X295) | p(s(bool,X296)) | p(s(bool,X297)) | s(bool,X295) != s(bool,X297)) )),
  inference(resolution,[],[f1530,f1102])).
fof(f2983,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP8(sK14,X35) | sP8(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1530,f892])).
fof(f2980,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP8(sK14,X20) | sP8(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1530,f891])).
fof(f2979,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP8(sK14,X15) | sP8(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1530,f890])).
fof(f2978,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP8(sK14,X10) | sP8(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1530,f888])).
fof(f2976,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(sK14,X0) | sP8(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1530,f886])).
fof(f2076,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1984,f1223])).
fof(f1984,plain,(
  ( ! [X255,X256] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X255),s('type.num.num',X256))))) ) | $spl12),
  inference(resolution,[],[f1223,f1009])).
fof(f1009,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f764])).
fof(f2075,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1983,f1223])).
fof(f1983,plain,(
  ( ! [X253,X254] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X253),s('type.num.num',X254))))) ) | $spl12),
  inference(resolution,[],[f1223,f979])).
fof(f979,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f751])).
fof(f2072,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1978,f1223])).
fof(f1978,plain,(
  ( ! [X239,X240] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X239))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X240))))))) ) | $spl12),
  inference(resolution,[],[f1223,f1009])).
fof(f2071,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1977,f1223])).
fof(f1977,plain,(
  ( ! [X237,X238] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X237))),s('type.num.num',X238))))) ) | $spl12),
  inference(resolution,[],[f1223,f979])).
fof(f2070,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1974,f1223])).
fof(f1974,plain,(
  ( ! [X229,X230] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X229),s('type.num.num',X230))))) ) | $spl12),
  inference(resolution,[],[f1223,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f704])).
fof(f2069,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1972,f1223])).
fof(f1972,plain,(
  ( ! [X227,X228] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X227),s('type.num.num',X228))))) ) | $spl12),
  inference(resolution,[],[f1223,f818])).
fof(f818,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f705])).
fof(f2068,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1971,f1223])).
fof(f1971,plain,(
  ( ! [X225,X226] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X225),s('type.num.num',X226))))) ) | $spl12),
  inference(resolution,[],[f1223,f817])).
fof(f2067,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1970,f1223])).
fof(f1970,plain,(
  ( ! [X223,X224] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X223))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X224))))))) ) | $spl12),
  inference(resolution,[],[f1223,f818])).
fof(f2066,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1969,f1223])).
fof(f1969,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1223,f1099])).
fof(f1099,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f810])).
fof(f2065,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1968,f1223])).
fof(f1968,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1223,f1097])).
fof(f1097,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f809])).
fof(f2064,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1959,f1223])).
fof(f1959,plain,(
  ( ! [X215] : (~p(s(bool,X215))) ) | $spl12),
  inference(resolution,[],[f1223,f1031])).
fof(f1967,plain,(
  $false | $spl12),
  inference(resolution,[],[f1223,f1095])).
fof(f1973,plain,(
  $false | $spl12),
  inference(resolution,[],[f1223,f812])).
fof(f1852,plain,(
  sP8('F_0',sK14) | $spl34),
  inference(cnf_transformation,[],[f1852_D])).
fof(f1852_D,plain,(
  sP8('F_0',sK14) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1840,plain,(
  ( ! [X38] : (p(s(bool,X38)) | sP8('F_0',sK14)) )),
  inference(duplicate_literal_removal,[],[f1812])).
fof(f1812,plain,(
  ( ! [X38] : (p(s(bool,X38)) | sP8('F_0',sK14) | p(s(bool,X38))) )),
  inference(superposition,[],[f1168,f1033])).
fof(f1841,plain,(
  ( ! [X37] : (p(s(bool,X37)) | sP8('F_0',sK14)) )),
  inference(duplicate_literal_removal,[],[f1811])).
fof(f1811,plain,(
  ( ! [X37] : (p(s(bool,X37)) | sP8('F_0',sK14) | p(s(bool,X37))) )),
  inference(superposition,[],[f1168,f1079])).
fof(f1842,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP8('F_0',sK14)) )),
  inference(duplicate_literal_removal,[],[f1810])).
fof(f1810,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP8('F_0',sK14) | p(s(bool,X36))) )),
  inference(superposition,[],[f1168,f1081])).
fof(f1770,plain,(
  ( ! [X230] : (sP8('F_0',sK14) | p(s(bool,X230))) )),
  inference(resolution,[],[f1168,f1091])).
fof(f1769,plain,(
  ( ! [X229] : (sP8('F_0',sK14) | p(s(bool,X229))) )),
  inference(resolution,[],[f1168,f1070])).
fof(f1768,plain,(
  ( ! [X228] : (sP8('F_0',sK14) | p(s(bool,X228))) )),
  inference(resolution,[],[f1168,f1068])).
fof(f1767,plain,(
  ( ! [X227] : (sP8('F_0',sK14) | p(s(bool,X227))) )),
  inference(resolution,[],[f1168,f1060])).
fof(f1766,plain,(
  ( ! [X226] : (sP8('F_0',sK14) | p(s(bool,X226))) )),
  inference(resolution,[],[f1168,f1055])).
fof(f1853,plain,(
  ( ! [X225] : (sP8('F_0',sK14) | p(s(bool,X225))) )),
  inference(subsumption_resolution,[],[f1765,f1105])).
fof(f1765,plain,(
  ( ! [X225] : (sP8('F_0',sK14) | p(s(bool,X225)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1168,f1041])).
fof(f1764,plain,(
  ( ! [X224] : (sP8('F_0',sK14) | s(bool,'F_0') = s(bool,X224)) )),
  inference(resolution,[],[f1168,f1034])).
fof(f1763,plain,(
  ( ! [X223] : (sP8('F_0',sK14) | ~p(s(bool,X223))) )),
  inference(resolution,[],[f1168,f1031])).
fof(f1762,plain,(
  ( ! [X222,X221] : (sP8('F_0',sK14) | sP9(X221,X222)) )),
  inference(resolution,[],[f1168,f1004])).
fof(f1761,plain,(
  ( ! [X220,X219] : (sP8('F_0',sK14) | sP9(X219,X220)) )),
  inference(resolution,[],[f1168,f999])).
fof(f1847,plain,(
  ( ! [X2,X0] : (sP8(X0,sK14) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f1847_D])).
fof(f1847_D,plain,(
  ( ! [X2,X0] : (sP8(X0,sK14) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1756,plain,(
  ( ! [X212,X213,X214] : (sP8(X212,sK14) | p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X212) != s(bool,X214)) )),
  inference(resolution,[],[f1168,f1102])).
fof(f1681,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(X28,sK14) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1168,f892])).
fof(f1678,plain,(
  ( ! [X19,X17,X18,X16] : (sP8(X16,sK14) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1168,f891])).
fof(f1677,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(X12,sK14) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1168,f890])).
fof(f1676,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,sK14) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1168,f888])).
fof(f1674,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,sK14) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1168,f886])).
fof(f1655,plain,(
  sP8(sK14,'F_0') | $spl28),
  inference(cnf_transformation,[],[f1655_D])).
fof(f1655_D,plain,(
  sP8(sK14,'F_0') <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1642,plain,(
  ( ! [X38] : (p(s(bool,X38)) | sP8(sK14,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1614])).
fof(f1614,plain,(
  ( ! [X38] : (p(s(bool,X38)) | sP8(sK14,'F_0') | p(s(bool,X38))) )),
  inference(superposition,[],[f1167,f1033])).
fof(f1643,plain,(
  ( ! [X37] : (p(s(bool,X37)) | sP8(sK14,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1613])).
fof(f1613,plain,(
  ( ! [X37] : (p(s(bool,X37)) | sP8(sK14,'F_0') | p(s(bool,X37))) )),
  inference(superposition,[],[f1167,f1079])).
fof(f1644,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP8(sK14,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1612])).
fof(f1612,plain,(
  ( ! [X36] : (p(s(bool,X36)) | sP8(sK14,'F_0') | p(s(bool,X36))) )),
  inference(superposition,[],[f1167,f1081])).
fof(f1572,plain,(
  ( ! [X230] : (sP8(sK14,'F_0') | p(s(bool,X230))) )),
  inference(resolution,[],[f1167,f1091])).
fof(f1571,plain,(
  ( ! [X229] : (sP8(sK14,'F_0') | p(s(bool,X229))) )),
  inference(resolution,[],[f1167,f1070])).
fof(f1570,plain,(
  ( ! [X228] : (sP8(sK14,'F_0') | p(s(bool,X228))) )),
  inference(resolution,[],[f1167,f1068])).
fof(f1569,plain,(
  ( ! [X227] : (sP8(sK14,'F_0') | p(s(bool,X227))) )),
  inference(resolution,[],[f1167,f1060])).
fof(f1568,plain,(
  ( ! [X226] : (sP8(sK14,'F_0') | p(s(bool,X226))) )),
  inference(resolution,[],[f1167,f1055])).
fof(f1657,plain,(
  ( ! [X225] : (sP8(sK14,'F_0') | p(s(bool,X225))) )),
  inference(subsumption_resolution,[],[f1567,f1105])).
fof(f1567,plain,(
  ( ! [X225] : (sP8(sK14,'F_0') | p(s(bool,X225)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1167,f1041])).
fof(f1566,plain,(
  ( ! [X224] : (sP8(sK14,'F_0') | s(bool,'F_0') = s(bool,X224)) )),
  inference(resolution,[],[f1167,f1034])).
fof(f1565,plain,(
  ( ! [X223] : (sP8(sK14,'F_0') | ~p(s(bool,X223))) )),
  inference(resolution,[],[f1167,f1031])).
fof(f1564,plain,(
  ( ! [X222,X221] : (sP8(sK14,'F_0') | sP9(X221,X222)) )),
  inference(resolution,[],[f1167,f1004])).
fof(f1563,plain,(
  ( ! [X220,X219] : (sP8(sK14,'F_0') | sP9(X219,X220)) )),
  inference(resolution,[],[f1167,f999])).
fof(f1649,plain,(
  ( ! [X2,X0] : (sP8(sK14,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ( ! [X2,X0] : (sP8(sK14,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1558,plain,(
  ( ! [X212,X213,X214] : (sP8(sK14,X212) | p(s(bool,X213)) | p(s(bool,X214)) | s(bool,X212) != s(bool,X214)) )),
  inference(resolution,[],[f1167,f1102])).
fof(f1483,plain,(
  ( ! [X30,X28,X31,X29] : (sP8(sK14,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1167,f892])).
fof(f1480,plain,(
  ( ! [X19,X17,X18,X16] : (sP8(sK14,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1167,f891])).
fof(f1479,plain,(
  ( ! [X14,X12,X15,X13] : (sP8(sK14,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1167,f890])).
fof(f1478,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(sK14,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1167,f888])).
fof(f1476,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(sK14,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1167,f886])).
fof(f1414,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | $spl23),
  inference(cnf_transformation,[],[f1414_D])).
fof(f1414_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1412,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | $spl21),
  inference(cnf_transformation,[],[f1412_D])).
fof(f1412_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1362,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))),
  inference(superposition,[],[f812,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2))))))) )),
  inference(cnf_transformation,[],[f658])).
fof(f658,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))) | s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(ennf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))) => s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(flattening,[],[f545])).
fof(f545,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))) => s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : ! [X11,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X8))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X11)))))) => s('type.stackSem.state'(X6,X7),X8) = s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X11))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.stackSem.stackSem_5|48|5_prove_ltb|split|1')).
fof(f1250,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) | $spl15),
  inference(cnf_transformation,[],[f1250_D])).
fof(f1250_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1361,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))))))),
  inference(superposition,[],[f812,f883])).
fof(f883,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(X0,X1),X2))))))))) )),
  inference(cnf_transformation,[],[f659])).
fof(f659,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))))) | s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(ennf_transformation,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))))) => s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(flattening,[],[f579])).
fof(f579,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),X3))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X0,X1),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(X0,X1),X2)))))))) => s('type.stackSem.state'(X0,X1),X3) = s('type.stackSem.state'(X0,X1),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X0,X1),X3),s('type.stackSem.state'(X0,X1),X2))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X11,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),X8))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(X6,X7),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(X6,X7),X11)))))))) => s('type.stackSem.state'(X6,X7),X8) = s('type.stackSem.state'(X6,X7),'const.stackSem.check_clock_2'(s('type.stackSem.state'(X6,X7),X8),s('type.stackSem.state'(X6,X7),X11))))),
  file('thm.stackSem.stackSem_5|48|7_prove_ltb','thm.stackSem.stackSem_5|48|6_prove_ltb')).
fof(f1406,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) | $spl18),
  inference(cnf_transformation,[],[f1406_D])).
fof(f1406_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1354,plain,(
  ( ! [X130,X131] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X131) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) )),
  inference(resolution,[],[f812,f1102])).
fof(f1277,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f812,f892])).
fof(f1275,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) )),
  inference(resolution,[],[f812,f891])).
fof(f1273,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) )),
  inference(resolution,[],[f812,f890])).
fof(f1272,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f812,f888])).
fof(f1270,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))) )),
  inference(resolution,[],[f812,f886])).
fof(f1252,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) | $spl17),
  inference(cnf_transformation,[],[f1252_D])).
fof(f1252_D,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1237,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK12))),s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.dec_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13)))))))) | $spl1),
  inference(superposition,[],[f1110,f883])).
fof(f1219,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK14) != s(bool,X1)) ) | $spl4),
  inference(cnf_transformation,[],[f1219_D])).
fof(f1219_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK14) != s(bool,X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1195,plain,(
  ( ! [X130,X131] : (p(s(bool,X130)) | p(s(bool,X131)) | s(bool,sK14) != s(bool,X131)) )),
  inference(resolution,[],[f811,f1102])).
fof(f1120,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK14) != s(bool,X21) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f811,f892])).
fof(f1117,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK14) != s(bool,X14)) )),
  inference(resolution,[],[f811,f891])).
fof(f1116,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK14) != s(bool,X10)) )),
  inference(resolution,[],[f811,f890])).
fof(f1115,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK14) != s(bool,X6) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f811,f888])).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK14) != s(bool,X1)) )),
  inference(resolution,[],[f811,f886])).
fof(f1110,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))) | $spl1),
  inference(cnf_transformation,[],[f1110_D])).
fof(f1110_D,plain,(
  s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1112,plain,(
  ~p(s(bool,sK14)) | $spl3),
  inference(cnf_transformation,[],[f1112_D])).
fof(f1112_D,plain,(
  ~p(s(bool,sK14)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f813,plain,(
  ~p(s(bool,sK14)) | s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),sK13))) != s('type.num.num','const.stackSem.state_clock_1'(s('type.stackSem.state'(sK10,sK11),'const.stackSem.check_clock_2'(s('type.stackSem.state'(sK10,sK11),sK12),s('type.stackSem.state'(sK10,sK11),sK13)))))),
  inference(cnf_transformation,[],[f702])).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 426
% Passive clauses: 3321
% Generated clauses: 9692
% Final active clauses: 411
% Final passive clauses: 2379
% Input formulas: 517
% Initial clauses: 291
% 
% Selected by SInE selection: 91
% SInE iterations: 1
% 
% Duplicate literals: 345
% Trivial inequalities: 9
% Fw subsumption resolutions: 980
% Fw demodulations: 483
% Bw demodulations: 2
% 
% Simple tautologies: 412
% Equational tautologies: 1
% Forward subsumptions: 3725
% 
% Binary resolution: 5497
% Factoring: 72
% Forward superposition: 1869
% Backward superposition: 226
% Equality resolution: 2
% 
% Split clauses: 299
% Split components: 622
% SAT solver clauses: 343
% SAT solver unit clauses: 43
% SAT solver binary clauses: 198
% 
% Sat splits: 299
% Sat splitting refutations: 43
% 
% Memory used [KB]: 4477
% Time elapsed: 0.171 s
% ------------------------------
% ------------------------------
% Success in time 11.956 s
