Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (68248)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 361
% Passive clauses: 2906
% Generated clauses: 25383
% Final active clauses: 311
% Final passive clauses: 2330
% Input formulas: 260
% Initial clauses: 404
% 
% Duplicate literals: 409
% Trivial inequalities: 176
% Fw subsumption resolutions: 1452
% Fw demodulations: 5778
% Bw demodulations: 47
% 
% Simple tautologies: 254
% Equational tautologies: 35
% Forward subsumptions: 14344
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 542
% Factoring: 6
% Forward superposition: 8528
% Backward superposition: 7780
% Self superposition: 128
% Equality factoring: 46
% Equality resolution: 26
% 
% Split clauses: 135
% Split components: 276
% SAT solver clauses: 221
% SAT solver unit clauses: 11
% SAT solver binary clauses: 60
% 
% Sat splits: 135
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 220
% TWLsolver calls for satisfiability: 77
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (68267)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 175
% Passive clauses: 747
% Generated clauses: 4709
% Final active clauses: 117
% Final passive clauses: 325
% Input formulas: 260
% Initial clauses: 404
% 
% Duplicate literals: 23
% Trivial inequalities: 24
% Fw subsumption resolutions: 150
% Fw demodulations: 495
% Bw demodulations: 283
% Condensations: 24
% Global subsumptions: 21
% 
% Simple tautologies: 120
% Equational tautologies: 8
% Forward subsumptions: 2260
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 221
% Factoring: 2
% Forward superposition: 971
% Backward superposition: 1207
% Self superposition: 30
% Equality factoring: 2
% Equality resolution: 2
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 1180
% SAT solver unit clauses: 601
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 42
% SAT solver learnt literals: 20
% 
% Sat splits: 15
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1096
% TWLsolver calls for satisfiability: 1871
% 
% Memory used [KB]: 34668
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% (68288)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1256
% Passive clauses: 11346
% Generated clauses: 160532
% Final active clauses: 940
% Final passive clauses: 2589
% Input formulas: 260
% Initial clauses: 404
% Discarded non-redundant clauses: 81876
% 
% Duplicate literals: 172
% Trivial inequalities: 62
% Fw subsumption resolutions: 640
% Fw demodulations: 35361
% Bw demodulations: 2072
% Condensations: 87
% Global subsumptions: 654
% 
% Simple tautologies: 104
% Equational tautologies: 79
% Forward subsumptions: 70256
% Fw demodulations to eq. taut.: 16026
% Bw demodulations to eq. taut.: 127
% 
% Unit resulting resolution: 55791
% Factoring: 1
% Forward superposition: 38389
% Backward superposition: 26572
% Self superposition: 256
% Equality factoring: 26
% Equality resolution: 19
% 
% Split clauses: 25
% Split components: 50
% SAT solver clauses: 13012
% SAT solver unit clauses: 5310
% SAT solver binary clauses: 5092
% 
% Sat splits: 25
% Sat splitting refutations: 959
% 
% Memory used [KB]: 231808
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f1679,plain,(
  $false),
  inference(equality_resolution,[],[f1648])).
fof(f1648,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X0))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X1)))) )),
  inference(superposition,[],[f1585,f481])).
fof(f481,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t',X0) = s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'),s('type.infer_t.infer_t',X0)))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : s('type.infer_t.infer_t',X0) = s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'),s('type.infer_t.infer_t',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X17] : s('type.infer_t.infer_t',X17) = s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'),s('type.infer_t.infer_t',X17)))),
  file('thm.inferProps.infer_subst_FEMPTY|split|1','thm.inferProps.infer_subst_FEMPTY|split|0')).
fof(f1585,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'),s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X0)))))) )),
  inference(superposition,[],[f1546,f479])).
fof(f479,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : s('type.infer_t.infer_t','const.infer.infer_subst_2'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X9),s('type.infer_t.infer_t',X10))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X9))),s('type.infer_t.infer_t',X10)))),
  file('thm.inferProps.infer_subst_FEMPTY|split|1','ARITY_1_2_const.infer.infer_subst')).
fof(f1546,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X0)))))) )),
  inference(superposition,[],[f1001,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X2))),s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X2))),s('type.list.list'('type.infer_t.infer_t'),X0))),s('type.ast.tctor',X1)))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X2))),s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X2))),s('type.list.list'('type.infer_t.infer_t'),X0))),s('type.ast.tctor',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X8] : s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X8))),s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X6),s('type.ast.tctor',X7))))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),X8))),s('type.list.list'('type.infer_t.infer_t'),X6))),s('type.ast.tctor',X7)))),
  file('thm.inferProps.infer_subst_FEMPTY|split|1','thm.infer.infer_subst_def|split|1')).
fof(f1001,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),sK7),s('type.ast.tctor',X0))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),sK7))),s('type.ast.tctor',X1)))) )),
  inference(equality_resolution,[],[f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.infer_t.infer_t'),sK7) != s('type.list.list'('type.infer_t.infer_t'),X0) | s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),sK7))),s('type.ast.tctor',X2)))) )),
  inference(superposition,[],[f429,f550])).
fof(f550,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.infer_t.infer_t'),X0) = s('type.list.list'('type.infer_t.infer_t'),X2) | s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3)))) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2,X3] : ((s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.infer_t.infer_t'),X0) = s('type.list.list'('type.infer_t.infer_t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & (s('type.list.list'('type.infer_t.infer_t'),X0) != s('type.list.list'('type.infer_t.infer_t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3) | s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3)))))),
  inference(flattening,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2,X3] : ((s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.infer_t.infer_t'),X0) = s('type.list.list'('type.infer_t.infer_t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & ((s('type.list.list'('type.infer_t.infer_t'),X0) != s('type.list.list'('type.infer_t.infer_t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3)) | s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3)))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2,X3] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.ast.tctor',X1))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.ast.tctor',X3))) <=> (s('type.list.list'('type.infer_t.infer_t'),X0) = s('type.list.list'('type.infer_t.infer_t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3)))),
  inference(rectify,[],[f219])).
fof(f219,axiom,(
  ! [X49,X50,X76,X77] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X49),s('type.ast.tctor',X50))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tapp_2'(s('type.list.list'('type.infer_t.infer_t'),X76),s('type.ast.tctor',X77))) <=> (s('type.list.list'('type.infer_t.infer_t'),X49) = s('type.list.list'('type.infer_t.infer_t'),X76) & s('type.ast.tctor',X50) = s('type.ast.tctor',X77)))),
  file('thm.inferProps.infer_subst_FEMPTY|split|1','thm.infer_t.infer_t_|49||49||split|1')).
fof(f429,plain,(
  s('type.list.list'('type.infer_t.infer_t'),sK7) != s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),sK7)))),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  s('type.list.list'('type.infer_t.infer_t'),sK7) != s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),sK7)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f332])).
fof(f332,plain,(
  ? [X0] : s('type.list.list'('type.infer_t.infer_t'),X0) != s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),X0)))),
  inference(ennf_transformation,[],[f262])).
fof(f262,plain,(
  ~! [X0] : s('type.list.list'('type.infer_t.infer_t'),X0) = s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.list.list'('type.infer_t.infer_t'),X6) = s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.list.list'('type.infer_t.infer_t'),X6) = s('type.list.list'('type.infer_t.infer_t'),'const.list.MAP_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_subst_1'(s('type.finite_map.fmap'('type.num.num','type.num.num'),'const.finite_map.FEMPTY_0'))),s('type.list.list'('type.infer_t.infer_t'),X6)))),
  file('thm.inferProps.infer_subst_FEMPTY|split|1','thm.inferProps.infer_subst_FEMPTY|split|1')).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 123
% Passive clauses: 581
% Generated clauses: 1083
% Final active clauses: 122
% Final passive clauses: 535
% Input formulas: 260
% Initial clauses: 136
% 
% Selected by SInE selection: 48
% SInE iterations: 1
% 
% Duplicate literals: 7
% Trivial inequalities: 57
% Fw subsumption resolutions: 11
% Fw demodulations: 43
% 
% Simple tautologies: 36
% Equational tautologies: 2
% Forward subsumptions: 383
% 
% Forward superposition: 977
% Equality resolution: 8
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 14
% SAT solver unit clauses: 1
% SAT solver binary clauses: 12
% 
% Sat splits: 12
% Sat splitting refutations: 1
% 
% Memory used [KB]: 2686
% Time elapsed: 0.048 s
% ------------------------------
% ------------------------------
% Success in time 11.8 s
