Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ImplicitQueue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ImplicitQueue
% SZS output start Proof for thm.ImplicitQueue
fof(f2175,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2159,f174,f176,f121,f179,f538,f541,f605,f1719,f539,f1720])).
fof(f1720,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X5),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0')))) ) | $spl42),
  inference(cnf_transformation,[],[f1720_D])).
fof(f1720_D,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X5),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0')))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f539,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1719,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X5),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0')))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1710,f184])).
fof(f184,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f179,f126])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb',aHL_BOOL_CASES)).
fof(f1710,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X5),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0')))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f1701,f182])).
fof(f182,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f126,f126])).
fof(f1701,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X0),s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.empty_0'))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1700,f179])).
fof(f1700,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X0),s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.empty_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f1699,f122])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.two_def|split|1')).
fof(f1699,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X0),s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.empty_0')))) | p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f1698,f123])).
fof(f123,plain,(
  ( ! [X0] : (s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0') = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0') = s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0') = s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.empty_def')).
fof(f1698,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X0),s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0')))))) | p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0'))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1687,f176])).
fof(f1687,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X0),s('type.ImplicitQueue.queue'(X1),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0')))))) | p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0'))))) )),
  inference(superposition,[],[f161,f1102])).
fof(f1102,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(forward_demodulation,[],[f1083,f145])).
fof(f145,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X21] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X21),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.list.EVERY_DEF|split|0')).
fof(f1083,plain,(
  ( ! [X0,X1] : (s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))) = s(bool,'const.list.EVERY_2'(s(fun('type.ImplicitQueue.times'(X0),bool),'const.ImplicitQueue.depth_1'(s('type.num.num',X1))),s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.list.NIL_0')))) )),
  inference(superposition,[],[f153,f125])).
fof(f125,plain,(
  ( ! [X0] : (s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.list.NIL_0') = s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.list.NIL_0') = s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : s('type.list.list'('type.ImplicitQueue.times'(X6)),'const.list.NIL_0') = s('type.list.list'('type.ImplicitQueue.times'(X6)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.only_digits_def|split|0')).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))) = s(bool,'const.list.EVERY_2'(s(fun('type.ImplicitQueue.times'(X0),bool),'const.ImplicitQueue.depth_1'(s('type.num.num',X1))),s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))) = s(bool,'const.list.EVERY_2'(s(fun('type.ImplicitQueue.times'(X0),bool),'const.ImplicitQueue.depth_1'(s('type.num.num',X1))),s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X0),X2)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))) = s(bool,'const.list.EVERY_2'(s(fun('type.ImplicitQueue.times'(X0),bool),'const.ImplicitQueue.depth_1'(s('type.num.num',X1))),s('type.list.list'('type.ImplicitQueue.times'(X0)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X0),X2)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X12,X13] : s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X12),s('type.ImplicitQueue.digit'(X6),X13))) = s(bool,'const.list.EVERY_2'(s(fun('type.ImplicitQueue.times'(X6),bool),'const.ImplicitQueue.depth_1'(s('type.num.num',X12))),s('type.list.list'('type.ImplicitQueue.times'(X6)),'const.ImplicitQueue.only_digits_1'(s('type.ImplicitQueue.digit'(X6),X13)))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.ddepth_def')).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2)))) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) | p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) | (~p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) & p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2)))))) & (p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2)))) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) | (~p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) & p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2)))))) & ((p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))))) | p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2))))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) <=> (~p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) & p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X2)))))) <=> (~p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X0),X2)))) & p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X1),s('type.ImplicitQueue.digit'(X0),X2))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X12,X7] : (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num',X12),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X7)))))) <=> (~p(s(bool,'const.ImplicitQueue.two_1'(s('type.ImplicitQueue.digit'(X6),X7)))) & p(s(bool,'const.ImplicitQueue.ddepth_2'(s('type.num.num',X12),s('type.ImplicitQueue.digit'(X6),X7))))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.queue_ok_def|split|0')).
fof(f605,plain,(
  $false | ($spl0 | $spl3 | $spl12)),
  inference(subsumption_resolution,[],[f551,f176])).
fof(f551,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f539,f179])).
fof(f541,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) | $spl14),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f538,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f480])).
fof(f480,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f382,f126])).
fof(f382,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f296,f176])).
fof(f296,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f114,f182])).
fof(f114,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(sK0),'const.ImplicitQueue.empty_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f88])).
fof(f88,plain,(
  ? [X0] : ~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0'))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ~! [X0] : p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.empty_0'))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb')).
fof(f179,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f121,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb',aHL_FALSITY)).
fof(f176,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f174,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f2159,plain,(
  $false | ($spl0 | $spl3 | $spl14 | $spl42)),
  inference(subsumption_resolution,[],[f2143,f179])).
fof(f2143,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl14 | $spl42)),
  inference(superposition,[],[f2129,f541])).
fof(f2129,plain,(
  ( ! [X0] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),'const.list.NIL_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0'))))) ) | ($spl0 | $spl42)),
  inference(equality_resolution,[],[f2120])).
fof(f2120,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0'))))) ) | ($spl0 | $spl42)),
  inference(subsumption_resolution,[],[f2119,f176])).
fof(f2119,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1)) ) | $spl42),
  inference(forward_demodulation,[],[f2118,f1720])).
fof(f2118,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1)) )),
  inference(forward_demodulation,[],[f2117,f123])).
fof(f2117,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))))) )),
  inference(forward_demodulation,[],[f2116,f123])).
fof(f2116,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))))) )),
  inference(forward_demodulation,[],[f2108,f217])).
fof(f217,plain,(
  ( ! [X0] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) )),
  inference(forward_demodulation,[],[f216,f124])).
fof(f124,plain,(
  ( ! [X0] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0')))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.digits_def|split|0')).
fof(f216,plain,(
  ( ! [X0] : (s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0')))) )),
  inference(superposition,[],[f147,f123])).
fof(f147,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X0),X1))) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s('type.list.list'(X6),'const.ImplicitQueue.digits_1'(s('type.ImplicitQueue.digit'(X6),X7))) = s('type.list.list'(X6),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X6),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X6),X7)))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.flatten_def|split|0')).
fof(f2108,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.empty_0'))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),'const.ImplicitQueue.Shallow_1'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))))) )),
  inference(superposition,[],[f158,f123])).
fof(f158,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | ~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))))) & (~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) | (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2))))) & ((~p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))) | p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2))))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X0),X1),s('type.ImplicitQueue.queue'(X0),X2)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X0),X2)))) & s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X0),X2)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X19,X10] : (p(s(bool,'const.ImplicitQueue.queue_inv_2'(s('type.list.list'(X6),X19),s('type.ImplicitQueue.queue'(X6),X10)))) <=> (p(s(bool,'const.ImplicitQueue.queue_ok_2'(s('type.num.num','const.num.|48|_0'),s('type.ImplicitQueue.queue'(X6),X10)))) & s('type.list.list'(X6),X19) = s('type.list.list'(X6),'const.ImplicitQueue.flatten_1'(s('type.ImplicitQueue.queue'(X6),X10)))))),
  file('thm.ImplicitQueue.ImplicitQueue_|49||48|9_prove_ltb','thm.ImplicitQueue.queue_inv_def')).
% SZS output end Proof for thm.ImplicitQueue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 75
% Passive clauses: 279
% Generated clauses: 2017
% Final active clauses: 57
% Final passive clauses: 147
% Input formulas: 35
% Initial clauses: 60
% 
% Duplicate literals: 189
% Trivial inequalities: 19
% Fw subsumption resolutions: 127
% Fw demodulations: 198
% Bw demodulations: 13
% 
% Simple tautologies: 48
% Equational tautologies: 23
% Forward subsumptions: 1041
% 
% Binary resolution: 35
% Forward superposition: 758
% Backward superposition: 494
% Self superposition: 50
% Equality factoring: 38
% Equality resolution: 3
% 
% Split clauses: 87
% Split components: 174
% SAT solver clauses: 143
% SAT solver unit clauses: 4
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 87
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 134
% TWLsolver calls for satisfiability: 23
% 
% Memory used [KB]: 2046
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
% Success in time 0.063 s
