Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvpSem
% (66322)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 183
% Passive clauses: 3600
% Generated clauses: 25891
% Final active clauses: 153
% Final passive clauses: 3160
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 909
% Trivial inequalities: 239
% Fw subsumption resolutions: 2443
% Fw demodulations: 3889
% Bw demodulations: 24
% 
% Simple tautologies: 207
% Equational tautologies: 51
% Forward subsumptions: 13670
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 626
% Factoring: 2
% Forward superposition: 8825
% Backward superposition: 8355
% Self superposition: 193
% Equality factoring: 112
% Equality resolution: 17
% 
% Split clauses: 285
% Split components: 570
% SAT solver clauses: 775
% SAT solver unit clauses: 48
% SAT solver binary clauses: 143
% 
% Sat splits: 285
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 773
% TWLsolver calls for satisfiability: 166
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvpSem
% (66337)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1151
% Generated clauses: 16315
% Final active clauses: 118
% Final passive clauses: 920
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 622
% Trivial inequalities: 62
% Fw subsumption resolutions: 811
% Fw demodulations: 1742
% Bw demodulations: 112
% Condensations: 1238
% 
% Simple tautologies: 113
% Equational tautologies: 41
% Forward subsumptions: 8581
% 
% Binary resolution: 348
% Factoring: 1
% Forward superposition: 5001
% Backward superposition: 4821
% Self superposition: 140
% Equality factoring: 82
% Equality resolution: 6
% 
% Split clauses: 64
% Split components: 128
% SAT solver clauses: 2054
% SAT solver unit clauses: 259
% SAT solver binary clauses: 29
% 
% Sat splits: 64
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2054
% TWLsolver calls for satisfiability: 327
% 
% Memory used [KB]: 19061
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvpSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvpSem
% SZS output start Proof for thm.bvpSem
fof(f45756,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f560,f563,f590,f595,f592,f613,f815,f3615,f3620,f4772,f4773,f4777,f4778,f4782,f4783,f4901,f5011,f5012,f5033,f5037,f5041,f5043,f5047,f5049,f5083,f5084,f6776,f6759,f6778,f8581,f8583,f8566,f8590,f8585,f8911,f8843,f8913,f8951,f9153,f9156,f9162,f9154,f11620,f11621,f11622,f11623,f11666,f11624,f11663,f11667,f11626,f11627,f11628,f11671,f11629,f11670,f11672,f11662,f11660,f11673,f11665,f11669,f11686,f12477,f12401,f11659,f12521,f18193,f18153,f18143,f18266,f18664,f18665,f17974,f18819,f18818,f19426,f19475,f19476,f19477,f19478,f19479,f19480,f23780,f23799,f23775,f23802,f23773,f23807,f23772,f23808,f23766,f23809,f23765,f23812,f23764,f19961,f6665,f23885,f6086,f5279,f824,f2973,f3621,f23903,f23882,f1476,f823,f23898,f23909,f23931,f24663,f24664,f24666,f24667,f24668,f25225,f26768,f27637,f26769,f33227,f23813,f33330,f37119,f37210,f37479,f37480,f37481,f37482,f37483,f37530,f37531,f37526,f37532,f37533,f37529,f37534,f37489,f37490,f37537,f37491,f37492,f37538,f37539,f37525,f37523,f37540,f37528,f37536,f37553,f38809,f38812,f38813,f38814,f38719,f38822,f591,f38824,f37522,f40219,f40307,f561,f45748,f45750])).
fof(f45750,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f45749,f401])).
fof(f401,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f47])).
fof(f47,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bool.TRUTH')).
fof(f45749,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f45746,f8200])).
fof(f8200,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f8043,f1479])).
fof(f1479,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) )),
  inference(forward_demodulation,[],[f1478,f408])).
fof(f408,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X37] : ! [X39] : s(X37,X39) = s(X37,happ(s(fun(X37,X37),'_i_0_0'),s(X37,X39)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_i0_ax')).
fof(f1478,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | s(bool,X7) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X6))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1447,f401])).
fof(f1447,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | s(bool,X7) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X6))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f476,f571])).
fof(f571,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f408,f407])).
fof(f407,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f476,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X21,X22] : ((p(s(bool,X21)) => p(s(bool,X22))) => ((p(s(bool,X22)) => p(s(bool,X21))) => s(bool,X21) = s(bool,X22)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f8043,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7891,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X14,X6] : ! [X16,X20] : s(X6,X16) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X16),s(X14,X20)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.pair.FST')).
fof(f7891,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f486,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X59,X60,X61] : (p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X59,fun(X59,bool)),'=_0'),s(X59,X60))),s(X59,X61)))) <=> s(X59,X60) = s(X59,X61))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f45746,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),fun(fun(sK0,'type.bvpSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f45683,f14599])).
fof(f14599,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),fun(fun(sK0,'type.bvpSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f400,f499])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f399])).
fof(f400,plain,(
  s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f378])).
fof(f378,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) != s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) = s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) = s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))))),s(fun(X6,'type.bvpSem.state'(X7)),X8))) = s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10))),s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))),s(fun(X6,'type.bvpSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))))),s(fun(X6,'type.bvpSem.state'(X7)),X8))) = s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10))),s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))),s(fun(X6,'type.bvpSem.state'(X7)),X8)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_fupdfupds_comp|split|3')).
fof(f45683,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))))),s(fun(X3,'type.bvpSem.state'(X0)),X4))) = s(fun(X3,'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s(fun(X3,'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s(fun(X3,'type.bvpSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f506,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))))) = s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10))),s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_fupdfupds_comp|split|2')).
fof(f506,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14,X6,X19,X15] : ! [X10,X9,X8] : s(fun(X19,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X19,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X19,X15),X8))))) = s(fun(X19,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X19,X15),X8)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
fof(f45748,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f45747,f401])).
fof(f45747,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f45745,f8200])).
fof(f45745,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),fun(fun(sK0,'type.bvpSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f45683,f14601])).
fof(f14601,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvpSem.state'(sK1)),fun(fun(sK0,'type.bvpSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f400,f499])).
fof(f561,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f40307,plain,(
  ( ! [X82] : (s(bool,X82) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X82)) ) | $spl102),
  inference(cnf_transformation,[],[f40307_D])).
fof(f40307_D,plain,(
  ( ! [X82] : (s(bool,X82) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X82)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f40219,plain,(
  ( ! [X81,X82] : (s(bool,X82) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81)) ) | ($spl3 | $spl22 | $spl95)),
  inference(duplicate_literal_removal,[],[f40039])).
fof(f40039,plain,(
  ( ! [X81,X82] : (s(bool,X82) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X81)) ) | ($spl3 | $spl22 | $spl95)),
  inference(superposition,[],[f26037,f39131])).
fof(f39131,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X12)) ) | $spl95),
  inference(trivial_inequality_removal,[],[f39124])).
fof(f39124,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X12) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X12)) ) | $spl95),
  inference(superposition,[],[f37535,f23893])).
fof(f23893,plain,(
  ( ! [X53,X51] : (s(bool,X51) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f23892,f401])).
fof(f23892,plain,(
  ( ! [X53,X51] : (~p(s(bool,'T_0')) | s(bool,X51) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X51)) )),
  inference(forward_demodulation,[],[f6292,f898])).
fof(f898,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f877,f557])).
fof(f557,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f403,f407])).
fof(f403,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f877,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f401,f816])).
fof(f816,plain,(
  ( ! [X4,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | ~p(s(bool,X5))) )),
  inference(superposition,[],[f434,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X37,X35] : ! [X38,X40] : s(X35,X38) = s(X35,happ(s(fun(X37,X35),'_dst_x0_1'(s(X35,X38))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0_ax')).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f379])).
fof(f379,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X25,X18] : (! [X16] : (s(X6,X16) = s(X6,X18) => p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X16))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X25)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bool.JRH_INDUCT_UTIL')).
fof(f6292,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(X52,bool),'_dst_x0_1'(s(bool,'T_0')))))) | s(bool,X51) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f1479,f1535])).
fof(f1535,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1518,f401])).
fof(f1518,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f1479,f898])).
fof(f37535,plain,(
  s(bool,'F_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | $spl95),
  inference(cnf_transformation,[],[f37535_D])).
fof(f37535_D,plain,(
  s(bool,'F_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f26037,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f18002,f23839])).
fof(f23839,plain,(
  ( ! [X110] : (s(bool,X110) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X110)) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f20906,f401])).
fof(f20906,plain,(
  ( ! [X110] : (~p(s(bool,'T_0')) | s(bool,X110) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X110)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f1479,f8998])).
fof(f8998,plain,(
  s(bool,'T_0') = s(bool,sK8('_i_0_0',bool)) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f563,f8984,f1479])).
fof(f8984,plain,(
  p(s(bool,sK8('_i_0_0',bool))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f8983,f408])).
fof(f8983,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK8('_i_0_0',bool))))) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f8960,f401])).
fof(f8960,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK8('_i_0_0',bool))))) | ($spl3 | $spl22)),
  inference(superposition,[],[f479,f6793])).
fof(f6793,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f563,f6778,f1479])).
fof(f479,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0)))))) )),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f393])).
fof(f393,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f392])).
fof(f392,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X59,X69] : (p(s(bool,happ(s(fun(fun(X59,bool),bool),'?_1'),s(fun(X59,bool),X69)))) <=> ? [X60] : p(s(bool,happ(s(fun(X59,bool),X69),s(X59,X60)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','?_HL_QUANT')).
fof(f18002,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl22)),
  inference(condensation,[],[f17980])).
fof(f17980,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK8('_i_0_0',bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl3 | $spl22)),
  inference(equality_factoring,[],[f9095])).
fof(f9095,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,sK8('_i_0_0',bool)) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f9046,f401])).
fof(f9046,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,sK8('_i_0_0',bool)) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f1409,f8998])).
fof(f1409,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1164,f401])).
fof(f1164,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f476,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f407,f407])).
fof(f37522,plain,(
  s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | $spl85),
  inference(cnf_transformation,[],[f37522_D])).
fof(f37522_D,plain,(
  s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f38824,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl94)),
  inference(condensation,[],[f38823])).
fof(f38823,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,X85) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) ) | ($spl3 | $spl94)),
  inference(subsumption_resolution,[],[f38539,f31745])).
fof(f31745,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f563,f23893])).
fof(f38539,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,X85) = s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X85)) ) | ($spl3 | $spl94)),
  inference(superposition,[],[f31745,f37649])).
fof(f37649,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl94)),
  inference(subsumption_resolution,[],[f37557,f563])).
fof(f37557,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl94),
  inference(superposition,[],[f37536,f555])).
fof(f591,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f38822,plain,(
  ( ! [X84] : (s(bool,X84) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X84)) ) | $spl100),
  inference(cnf_transformation,[],[f38822_D])).
fof(f38822_D,plain,(
  ( ! [X84] : (s(bool,X84) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X84)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f38719,plain,(
  ( ! [X83,X84] : (s(bool,X84) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X83) | s(bool,'F_0') = s(bool,X84)) ) | ($spl3 | $spl22 | $spl94)),
  inference(duplicate_literal_removal,[],[f38538])).
fof(f38538,plain,(
  ( ! [X83,X84] : (s(bool,X84) != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X83) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl3 | $spl22 | $spl94)),
  inference(superposition,[],[f26037,f37649])).
fof(f38814,plain,(
  ( ! [X58] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X58)) ) | ($spl3 | $spl68 | $spl94)),
  inference(forward_demodulation,[],[f38721,f13101])).
fof(f13101,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X0,bool),bool),'!_1')))) )),
  inference(unit_resulting_resolution,[],[f12677,f2467])).
fof(f2467,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f2377,f557])).
fof(f2377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f434,f807])).
fof(f807,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f806,f403])).
fof(f806,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK5(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f433,f407])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f386])).
fof(f12677,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f401,f8567])).
fof(f8567,plain,(
  ( ! [X37,X36] : (~p(s(bool,X37)) | p(s(bool,happ(s(fun(fun(X36,bool),bool),'!_1'),s(fun(X36,bool),'_dst_x0_1'(s(bool,X37))))))) )),
  inference(superposition,[],[f478,f487])).
fof(f478,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f390])).
fof(f390,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f389])).
fof(f389,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X59,X69] : (p(s(bool,happ(s(fun(fun(X59,bool),bool),'!_1'),s(fun(X59,bool),X69)))) <=> ! [X60] : p(s(bool,happ(s(fun(X59,bool),X69),s(X59,X60)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','!_HL_QUANT')).
fof(f38721,plain,(
  ( ! [X59,X58] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun(fun(X59,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl3 | $spl68 | $spl94)),
  inference(duplicate_literal_removal,[],[f38523])).
fof(f38523,plain,(
  ( ! [X59,X58] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun(fun(X59,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl3 | $spl68 | $spl94)),
  inference(superposition,[],[f23808,f37649])).
fof(f38813,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X56)) ) | ($spl3 | $spl66 | $spl94)),
  inference(forward_demodulation,[],[f38722,f10419])).
fof(f10419,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X0,bool),bool),'?_1')))) )),
  inference(unit_resulting_resolution,[],[f10278,f2467])).
fof(f10278,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f401,f6760])).
fof(f6760,plain,(
  ( ! [X50,X49] : (~p(s(bool,X50)) | p(s(bool,happ(s(fun(fun(X49,bool),bool),'?_1'),s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))))))) )),
  inference(superposition,[],[f480,f487])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f394])).
fof(f38722,plain,(
  ( ! [X57,X56] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun(fun(X57,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl3 | $spl66 | $spl94)),
  inference(duplicate_literal_removal,[],[f38522])).
fof(f38522,plain,(
  ( ! [X57,X56] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun(fun(X57,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | ($spl3 | $spl66 | $spl94)),
  inference(superposition,[],[f23807,f37649])).
fof(f38812,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X54)) ) | ($spl3 | $spl12 | $spl64 | $spl94)),
  inference(forward_demodulation,[],[f38723,f3620])).
fof(f38723,plain,(
  ( ! [X54,X55] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X55),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl3 | $spl64 | $spl94)),
  inference(duplicate_literal_removal,[],[f38521])).
fof(f38521,plain,(
  ( ! [X54,X55] : (s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X55),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl3 | $spl64 | $spl94)),
  inference(superposition,[],[f23802,f37649])).
fof(f38809,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl94)),
  inference(forward_demodulation,[],[f38808,f8998])).
fof(f38808,plain,(
  ( ! [X52] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,X52) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl22 | $spl94)),
  inference(subsumption_resolution,[],[f38519,f18675])).
fof(f18675,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,X142) | s(bool,X142) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f17697,f563])).
fof(f17697,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,X142) | s(bool,X142) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f563,f9095])).
fof(f38519,plain,(
  ( ! [X52] : (s(bool,'T_0') != s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,X52) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X52)) ) | ($spl3 | $spl22 | $spl94)),
  inference(superposition,[],[f18675,f37649])).
fof(f37553,plain,(
  $false | ($spl3 | $spl84 | $spl94)),
  inference(subsumption_resolution,[],[f37552,f563])).
fof(f37552,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl84 | $spl94)),
  inference(backward_demodulation,[],[f37523,f37536])).
fof(f37536,plain,(
  s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | $spl94),
  inference(cnf_transformation,[],[f37536_D])).
fof(f37536_D,plain,(
  s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f37528,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | $spl91),
  inference(cnf_transformation,[],[f37528_D])).
fof(f37528_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f37540,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl27 | $spl76)),
  inference(forward_demodulation,[],[f37494,f23882])).
fof(f37494,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl27 | $spl76)),
  inference(superposition,[],[f29117,f25125])).
fof(f25125,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X14)) ) | $spl76),
  inference(subsumption_resolution,[],[f25064,f401])).
fof(f25064,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | s(bool,X14) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X14)) ) | $spl76),
  inference(superposition,[],[f1479,f23882])).
fof(f29117,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))))) | $spl27),
  inference(unit_resulting_resolution,[],[f29083,f478])).
fof(f29083,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),sK7('!_1',fun(bool,bool)))))) | $spl27),
  inference(unit_resulting_resolution,[],[f25311,f478])).
fof(f25311,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'!_1'),s(fun(fun(bool,bool),bool),'!_1')))) | $spl27),
  inference(unit_resulting_resolution,[],[f8584,f477])).
fof(f477,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f8584,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl27),
  inference(cnf_transformation,[],[f8584_D])).
fof(f8584_D,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f37523,plain,(
  s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | $spl84),
  inference(cnf_transformation,[],[f37523_D])).
fof(f37523_D,plain,(
  s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f37525,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) | $spl87),
  inference(cnf_transformation,[],[f37525_D])).
fof(f37525_D,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f37539,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl6 | $spl27 | $spl29)),
  inference(forward_demodulation,[],[f37493,f24648])).
fof(f24648,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | ($spl6 | $spl29)),
  inference(unit_resulting_resolution,[],[f8590,f592])).
fof(f37493,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,sK7('_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl27 | $spl29)),
  inference(superposition,[],[f29117,f24669])).
fof(f24669,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl29),
  inference(subsumption_resolution,[],[f24656,f401])).
fof(f24656,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl29),
  inference(superposition,[],[f8590,f555])).
fof(f37538,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) ) | $spl98),
  inference(cnf_transformation,[],[f37538_D])).
fof(f37538_D,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f37492,plain,(
  ( ! [X15] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X15)))) | s(bool,'F_0') = s(bool,X15) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | $spl27),
  inference(superposition,[],[f29117,f23893])).
fof(f37491,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X14)))) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X14)) ) | $spl27),
  inference(superposition,[],[f29117,f23893])).
fof(f37537,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X12)) ) | $spl96),
  inference(cnf_transformation,[],[f37537_D])).
fof(f37537_D,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f37490,plain,(
  ( ! [X13] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X13) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X13)) ) | $spl27),
  inference(superposition,[],[f29117,f23893])).
fof(f37489,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X12) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X12)) ) | $spl27),
  inference(superposition,[],[f29117,f23893])).
fof(f37534,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl3 | $spl22 | $spl27)),
  inference(forward_demodulation,[],[f37488,f8998])).
fof(f37488,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,sK8('_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ($spl3 | $spl22 | $spl27)),
  inference(superposition,[],[f29117,f23839])).
fof(f37529,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl92),
  inference(cnf_transformation,[],[f37529_D])).
fof(f37529_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f37533,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,X8) | ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X8)))) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl27)),
  inference(forward_demodulation,[],[f37487,f8998])).
fof(f37487,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X8)))) | s(bool,X8) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl27)),
  inference(superposition,[],[f29117,f9095])).
fof(f37532,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X7)))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl3 | $spl22 | $spl27)),
  inference(forward_demodulation,[],[f37486,f8998])).
fof(f37486,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X7)))) | s(bool,sK8('_i_0_0',bool)) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X7)) ) | ($spl3 | $spl22 | $spl27)),
  inference(superposition,[],[f29117,f9095])).
fof(f37526,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl88),
  inference(cnf_transformation,[],[f37526_D])).
fof(f37526_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f37531,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X6) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl27)),
  inference(forward_demodulation,[],[f37485,f8998])).
fof(f37485,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X6) = s(bool,sK8('_i_0_0',bool)) | s(bool,X6) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl27)),
  inference(superposition,[],[f29117,f9095])).
fof(f37530,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X5) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X5)) ) | ($spl3 | $spl22 | $spl27)),
  inference(forward_demodulation,[],[f37484,f8998])).
fof(f37484,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,sK8('_i_0_0',bool))))) | s(bool,X5) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X5)) ) | ($spl3 | $spl22 | $spl27)),
  inference(superposition,[],[f29117,f9095])).
fof(f37483,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X3) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl27),
  inference(superposition,[],[f29117,f555])).
fof(f37482,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl27),
  inference(superposition,[],[f29117,f555])).
fof(f37481,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool))) ) | $spl27),
  inference(superposition,[],[f29117,f555])).
fof(f37480,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl27),
  inference(superposition,[],[f29117,f555])).
fof(f37479,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),sK7('!_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(sK7('!_1',fun(bool,bool)),bool)) | $spl27),
  inference(superposition,[],[f29117,f407])).
fof(f37210,plain,(
  ( ! [X127,X126] : (s(bool,X127) != s(bool,'|dollar|exists_1'(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X127)) ) | $spl82),
  inference(cnf_transformation,[],[f37210_D])).
fof(f37210_D,plain,(
  ( ! [X127,X126] : (s(bool,X127) != s(bool,'|dollar|exists_1'(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X127)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f37119,plain,(
  ( ! [X127,X125,X126] : (s(bool,X127) != s(bool,'|dollar|exists_1'(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X125)) ) | ($spl3 | $spl22)),
  inference(duplicate_literal_removal,[],[f36936])).
fof(f36936,plain,(
  ( ! [X127,X125,X126] : (s(bool,X127) != s(bool,'|dollar|exists_1'(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') = s(bool,X125)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f26037,f1535])).
fof(f33330,plain,(
  $false | ($spl37 | $spl47)),
  inference(subsumption_resolution,[],[f11668,f33294])).
fof(f33294,plain,(
  s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl37),
  inference(condensation,[],[f33293])).
fof(f33293,plain,(
  ( ! [X237] : (s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X237)) ) | $spl37),
  inference(subsumption_resolution,[],[f31700,f407])).
fof(f31700,plain,(
  ( ! [X237] : (s(bool,'T_0') != s(bool,X237) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X237)) ) | $spl37),
  inference(superposition,[],[f11659,f23893])).
fof(f11668,plain,(
  s(bool,'F_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl47),
  inference(cnf_transformation,[],[f11668_D])).
fof(f11668_D,plain,(
  s(bool,'F_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f23813,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X51)) ) | $spl74),
  inference(cnf_transformation,[],[f23813_D])).
fof(f23813_D,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X51)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f33227,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27)) ) | ($spl3 | $spl22 | $spl46)),
  inference(duplicate_literal_removal,[],[f33160])).
fof(f33160,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) ) | ($spl3 | $spl22 | $spl46)),
  inference(superposition,[],[f26037,f19769])).
fof(f19769,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X15)) ) | ($spl3 | $spl22 | $spl46)),
  inference(forward_demodulation,[],[f19768,f8998])).
fof(f19768,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,sK8('_i_0_0',bool)) | s(bool,X15) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl46)),
  inference(subsumption_resolution,[],[f19733,f563])).
fof(f19733,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,sK8('_i_0_0',bool)) | s(bool,X15) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl22 | $spl46)),
  inference(superposition,[],[f9095,f11669])).
fof(f26769,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool))) ) | $spl80),
  inference(cnf_transformation,[],[f26769_D])).
fof(f26769_D,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f27637,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool))) ) | ($spl3 | $spl10 | $spl22)),
  inference(forward_demodulation,[],[f27636,f8998])).
fof(f27636,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool)) | s(bool,X71) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl10 | $spl22)),
  inference(subsumption_resolution,[],[f27427,f18675])).
fof(f27427,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool)) | s(bool,X71) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X71)) ) | ($spl3 | $spl10 | $spl22)),
  inference(superposition,[],[f18675,f27038])).
fof(f27038,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,sK5(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f27018,f401])).
fof(f27018,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,sK5(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f824,f555])).
fof(f26768,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool))) ) | ($spl3 | $spl22 | $spl78)),
  inference(forward_demodulation,[],[f26767,f8998])).
fof(f26767,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool)) | s(bool,X71) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl22 | $spl78)),
  inference(subsumption_resolution,[],[f26558,f18675])).
fof(f26558,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,sK5(X72,'_i_0_0',bool)) | s(bool,X71) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X71)) ) | ($spl3 | $spl22 | $spl78)),
  inference(superposition,[],[f18675,f24959])).
fof(f24959,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK5(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl3 | $spl78)),
  inference(subsumption_resolution,[],[f24897,f563])).
fof(f24897,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK5(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl78),
  inference(superposition,[],[f23898,f555])).
fof(f25225,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f25171,f25166])).
fof(f25166,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl26),
  inference(forward_demodulation,[],[f25137,f408])).
fof(f25137,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8585,f477])).
fof(f25171,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f401,f403,f481])).
fof(f481,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X60,X61] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X60))),s(bool,X61)))) <=> (p(s(bool,X60)) => p(s(bool,X61))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','==>_HL_0')).
fof(f24668,plain,(
  $false | ($spl6 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f24648,f8582])).
fof(f8582,plain,(
  s(bool,'F_0') != s(bool,sK7('_i_0_0',bool)) | $spl25),
  inference(cnf_transformation,[],[f8582_D])).
fof(f8582_D,plain,(
  s(bool,'F_0') != s(bool,sK7('_i_0_0',bool)) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f24667,plain,(
  $false | ($spl3 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f24649,f8582])).
fof(f24649,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | ($spl3 | $spl29)),
  inference(unit_resulting_resolution,[],[f563,f8590,f593])).
fof(f593,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f585,f403])).
fof(f585,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f473,f407])).
fof(f473,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f381])).
fof(f24666,plain,(
  $false | ($spl25 | $spl29 | $spl37 | $spl46)),
  inference(subsumption_resolution,[],[f24665,f8582])).
fof(f24665,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | ($spl29 | $spl37 | $spl46)),
  inference(forward_demodulation,[],[f24650,f11669])).
fof(f24650,plain,(
  s(bool,sK7('_i_0_0',bool)) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | ($spl29 | $spl37)),
  inference(unit_resulting_resolution,[],[f11659,f8590,f593])).
fof(f24664,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f24651,f8582])).
fof(f24651,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | $spl29),
  inference(unit_resulting_resolution,[],[f403,f8590,f473])).
fof(f24663,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f24653,f8582])).
fof(f24653,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | $spl29),
  inference(unit_resulting_resolution,[],[f403,f8590,f473])).
fof(f23931,plain,(
  $false | ($spl25 | $spl31)),
  inference(subsumption_resolution,[],[f8582,f20853])).
fof(f20853,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | $spl31),
  inference(unit_resulting_resolution,[],[f401,f8913,f1479])).
fof(f23909,plain,(
  $false | ($spl9 | $spl76)),
  inference(subsumption_resolution,[],[f23908,f401])).
fof(f23908,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl76)),
  inference(backward_demodulation,[],[f23882,f822])).
fof(f822,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl9),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f23898,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) ) | $spl78),
  inference(cnf_transformation,[],[f23898_D])).
fof(f23898_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f823,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(cnf_transformation,[],[f823_D])).
fof(f823_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1476,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1443,f401])).
fof(f1443,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f434,f571])).
fof(f23882,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl76),
  inference(cnf_transformation,[],[f23882_D])).
fof(f23882_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f23903,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f2961,f401])).
fof(f2961,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f2467,f571])).
fof(f3621,plain,(
  ( ! [X32] : (~p(s(bool,X32))) ) | $spl14),
  inference(cnf_transformation,[],[f3621_D])).
fof(f3621_D,plain,(
  ( ! [X32] : (~p(s(bool,X32))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2973,plain,(
  ( ! [X44] : (~p(s(bool,X44)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) )),
  inference(superposition,[],[f2467,f408])).
fof(f824,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool)))) ) | $spl10),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5279,plain,(
  ( ! [X35] : (~p(s(bool,sK5(X35,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f434,f408])).
fof(f6086,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f6026,f401])).
fof(f6026,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f434,f571])).
fof(f23885,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f6652,f401])).
fof(f6652,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f2467,f571])).
fof(f6665,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) )),
  inference(superposition,[],[f2467,f408])).
fof(f19961,plain,(
  ( ! [X59] : (~p(s(bool,sK5(X59,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f434,f408])).
fof(f23764,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51)) ) | ($spl3 | $spl8 | $spl22 | $spl40)),
  inference(duplicate_literal_removal,[],[f23709])).
fof(f23709,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | ($spl3 | $spl8 | $spl22 | $spl40)),
  inference(superposition,[],[f20339,f11663])).
fof(f20339,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f18002,f9019])).
fof(f9019,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X5)) ) | ($spl3 | $spl8 | $spl22)),
  inference(subsumption_resolution,[],[f9004,f403])).
fof(f9004,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X5)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f8984,f1556])).
fof(f1556,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1534,f1534])).
fof(f1534,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X24)) ) | $spl8),
  inference(subsumption_resolution,[],[f1517,f401])).
fof(f1517,plain,(
  ( ! [X24] : (~p(s(bool,'T_0')) | s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X24)) ) | $spl8),
  inference(superposition,[],[f1479,f829])).
fof(f829,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f823,f557])).
fof(f23812,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X46)) ) | $spl72),
  inference(cnf_transformation,[],[f23812_D])).
fof(f23812_D,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X46)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f23765,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f23707])).
fof(f23707,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) != s(bool,X46) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X44)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f20339,f9095])).
fof(f23809,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,sK7('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl70),
  inference(cnf_transformation,[],[f23809_D])).
fof(f23809_D,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,sK7('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X32)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f23766,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | ($spl3 | $spl8 | $spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f23702])).
fof(f23702,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) ) | ($spl3 | $spl8 | $spl22 | $spl24)),
  inference(superposition,[],[f20339,f8653])).
fof(f8653,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl24)),
  inference(subsumption_resolution,[],[f8594,f563])).
fof(f8594,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f8583,f555])).
fof(f23808,plain,(
  ( ! [X54,X53] : (s(bool,X54) != s(bool,'|dollar|exists_1'(s(fun(fun(X53,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X54)) ) | $spl68),
  inference(cnf_transformation,[],[f23808_D])).
fof(f23808_D,plain,(
  ( ! [X54,X53] : (s(bool,X54) != s(bool,'|dollar|exists_1'(s(fun(fun(X53,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X54)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f23772,plain,(
  ( ! [X54,X52,X53] : (s(bool,X54) != s(bool,'|dollar|exists_1'(s(fun(fun(X53,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X52)) ) | ($spl3 | $spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f23638])).
fof(f23638,plain,(
  ( ! [X54,X52,X53] : (s(bool,X54) != s(bool,'|dollar|exists_1'(s(fun(fun(X53,bool),bool),'!_1'))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X52)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f20339,f13258])).
fof(f13258,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun(fun(X12,bool),bool),'!_1'))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f13178,f563])).
fof(f13178,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun(fun(X12,bool),bool),'!_1'))) | s(bool,'F_0') = s(bool,X13)) ) | $spl8),
  inference(superposition,[],[f13101,f1556])).
fof(f23807,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,'|dollar|exists_1'(s(fun(fun(X48,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X49)) ) | $spl66),
  inference(cnf_transformation,[],[f23807_D])).
fof(f23807_D,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,'|dollar|exists_1'(s(fun(fun(X48,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X49)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f23773,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) != s(bool,'|dollar|exists_1'(s(fun(fun(X48,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47)) ) | ($spl3 | $spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f23636])).
fof(f23636,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) != s(bool,'|dollar|exists_1'(s(fun(fun(X48,bool),bool),'?_1'))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f20339,f10564])).
fof(f10564,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,'|dollar|exists_1'(s(fun(fun(X13,bool),bool),'?_1'))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f10490,f563])).
fof(f10490,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,'|dollar|exists_1'(s(fun(fun(X13,bool),bool),'?_1'))) | s(bool,'F_0') = s(bool,X14)) ) | $spl8),
  inference(superposition,[],[f10419,f1556])).
fof(f23802,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X30)) ) | $spl64),
  inference(cnf_transformation,[],[f23802_D])).
fof(f23802_D,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f23775,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X28)) ) | ($spl3 | $spl8 | $spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f23629])).
fof(f23629,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X28) | s(bool,'F_0') = s(bool,X28)) ) | ($spl3 | $spl8 | $spl12 | $spl22)),
  inference(superposition,[],[f20339,f5177])).
fof(f5177,plain,(
  ( ! [X31,X32] : (s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X32)) ) | ($spl3 | $spl8 | $spl12)),
  inference(trivial_inequality_removal,[],[f5167])).
fof(f5167,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X32)) ) | ($spl3 | $spl8 | $spl12)),
  inference(superposition,[],[f1940,f3620])).
fof(f1940,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl3 | $spl8)),
  inference(superposition,[],[f563,f1556])).
fof(f23799,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl62),
  inference(cnf_transformation,[],[f23799_D])).
fof(f23799_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f23780,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,X14)) ) | ($spl3 | $spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f23624])).
fof(f23624,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X14)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f20339,f1534])).
fof(f19480,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19048,f19473])).
fof(f19473,plain,(
  ( ! [X110] : (p(s(bool,X110))) ) | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19472,f401])).
fof(f19472,plain,(
  ( ! [X110] : (~p(s(bool,'T_0')) | p(s(bool,X110))) ) | ($spl3 | $spl22 | $spl58)),
  inference(forward_demodulation,[],[f19471,f8998])).
fof(f19471,plain,(
  ( ! [X110] : (~p(s(bool,sK8('_i_0_0',bool))) | p(s(bool,X110))) ) | ($spl3 | $spl22 | $spl58)),
  inference(forward_demodulation,[],[f19470,f18818])).
fof(f19470,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X107,bool),X109),s(X107,X108))),s(bool,X110)))) | p(s(bool,X110))) ) | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19469,f401])).
fof(f19469,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X107,bool),X109),s(X107,X108))),s(bool,X110)))) | p(s(bool,X110))) ) | ($spl3 | $spl22 | $spl58)),
  inference(forward_demodulation,[],[f18881,f8998])).
fof(f18881,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,sK8('_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X107,bool),X109),s(X107,X108))),s(bool,X110)))) | p(s(bool,X110))) ) | $spl58),
  inference(backward_demodulation,[],[f18818,f11527])).
fof(f11527,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X107,bool),X109),s(X107,X108))),s(bool,X110)))) | ~p(s(bool,happ(s(fun(X107,bool),X109),s(X107,X108)))) | p(s(bool,X110))) )),
  inference(backward_demodulation,[],[f11522,f9755])).
fof(f9755,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X107,bool),bool),'_dstsel_0_x0_1'(s(X107,X108))),s(fun(X107,bool),X109))),s(bool,X110)))) | ~p(s(bool,happ(s(fun(X107,bool),X109),s(X107,X108)))) | p(s(bool,X110))) )),
  inference(forward_demodulation,[],[f9654,f502])).
fof(f502,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X14,X6,X15] : ! [X10,X9,X16] : s(X14,'const.combin.o_3'(s(fun(X6,X14),X10),s(fun(X15,X6),X9),s(X15,X16))) = s(X14,happ(s(fun(X6,X14),X10),s(X6,happ(s(fun(X15,X6),X9),s(X15,X16)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.combin.o_THM')).
fof(f9654,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(X107,bool),X109),s(X107,X108)))) | p(s(bool,X110)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X107,bool),bool),'_dstsel_0_x0_1'(s(X107,X108))),s(fun(X107,bool),X109))))),s(bool,X110))))) )),
  inference(superposition,[],[f481,f491])).
fof(f491,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X35,X37] : ! [X39,X40] : s(X35,happ(s(fun(X37,X35),X40),s(X37,X39))) = s(X35,happ(s(fun(fun(X37,X35),X35),'_dstsel_0_x0_1'(s(X37,X39))),s(fun(X37,X35),X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dstsel_0_x0_ax')).
fof(f11522,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(X22,'const.combin.o_3'(s(fun(X18,X22),X23),s(fun(X19,X18),X21),s(X19,X20))) = s(X22,'const.combin.o_3'(s(fun(X18,X22),X23),s(fun(fun(X19,X18),X18),'_dstsel_0_x0_1'(s(X19,X20))),s(fun(X19,X18),X21)))) )),
  inference(forward_demodulation,[],[f11319,f502])).
fof(f11319,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(X22,happ(s(fun(X18,X22),X23),s(X18,happ(s(fun(X19,X18),X21),s(X19,X20))))) = s(X22,'const.combin.o_3'(s(fun(X18,X22),X23),s(fun(fun(X19,X18),X18),'_dstsel_0_x0_1'(s(X19,X20))),s(fun(X19,X18),X21)))) )),
  inference(superposition,[],[f502,f491])).
fof(f19048,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | ($spl3 | $spl58)),
  inference(backward_demodulation,[],[f18818,f9992])).
fof(f9992,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))))),s(bool,'T_0')))) | $spl3),
  inference(forward_demodulation,[],[f9941,f508])).
fof(f508,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X35,X36,X55,X37] : ! [X38,X39,X56,X40] : s(X35,happ(s(fun(X37,X35),'_dst_x0x0x1_3'(s(fun(X36,fun(X55,X35)),X38),s(X36,X39),s(fun(X37,X55),X56))),s(X37,X40))) = s(X35,happ(s(fun(X55,X35),happ(s(fun(X36,fun(X55,X35)),X38),s(X36,X39))),s(X55,happ(s(fun(X37,X55),X56),s(X37,X40)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0x0x1_ax')).
fof(f9941,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f401,f7534,f481])).
fof(f7534,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0')))) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f499])).
fof(f19479,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19047,f19473])).
fof(f19047,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl58),
  inference(backward_demodulation,[],[f18818,f9835])).
fof(f9835,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))))),s(bool,'F_0'))))),
  inference(forward_demodulation,[],[f9784,f508])).
fof(f9784,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f401,f9101,f481])).
fof(f9101,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f401,f403,f481])).
fof(f19478,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19032,f19473])).
fof(f19032,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl58),
  inference(backward_demodulation,[],[f18818,f14582])).
fof(f14582,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f14545,f502])).
fof(f14545,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f13774,f481])).
fof(f13774,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))) )),
  inference(unit_resulting_resolution,[],[f13654,f480])).
fof(f13654,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f13595,f477])).
fof(f13595,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))) )),
  inference(unit_resulting_resolution,[],[f13100,f479])).
fof(f13100,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))))) )),
  inference(unit_resulting_resolution,[],[f12677,f480])).
fof(f19477,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19014,f19473])).
fof(f19014,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl58),
  inference(backward_demodulation,[],[f18818,f11655])).
fof(f11655,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))),s(X0,sK8(sK8('?_1',fun(X0,bool)),X0)))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f11611,f502])).
fof(f11611,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X0,bool),sK8('?_1',fun(X0,bool))),s(X0,sK8(sK8('?_1',fun(X0,bool)),X0)))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f10683,f481])).
fof(f10683,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(X0,bool),sK8('?_1',fun(X0,bool))),s(X0,sK8(sK8('?_1',fun(X0,bool)),X0)))))) )),
  inference(unit_resulting_resolution,[],[f10630,f479])).
fof(f10630,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))))) )),
  inference(unit_resulting_resolution,[],[f10418,f479])).
fof(f10418,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))) )),
  inference(unit_resulting_resolution,[],[f10278,f480])).
fof(f19476,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f18993,f19473])).
fof(f18993,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl58),
  inference(backward_demodulation,[],[f18818,f14599])).
fof(f19475,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f18992,f19473])).
fof(f18992,plain,(
  ~p(s(bool,sK8('_i_0_0',bool))) | $spl58),
  inference(backward_demodulation,[],[f18818,f14601])).
fof(f19426,plain,(
  $false | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f19258,f401])).
fof(f19258,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl58)),
  inference(backward_demodulation,[],[f19257,f403])).
fof(f19257,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | ($spl3 | $spl22 | $spl58)),
  inference(global_subsumption,[],[f400,f401,f404,f405,f406,f407,f408,f411,f412,f413,f414,f415,f416,f417,f418,f419,f420,f421,f422,f423,f424,f425,f426,f427,f428,f429,f430,f431,f432,f437,f438,f439,f440,f441,f442,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f471,f551,f472,f476,f473,f484,f485,f486,f487,f488,f490,f491,f492,f493,f494,f495,f496,f497,f498,f501,f502,f503,f552,f505,f506,f507,f508,f553,f554,f511,f512,f513,f514,f515,f516,f517,f518,f519,f520,f521,f522,f523,f524,f525,f526,f527,f528,f529,f530,f531,f532,f533,f534,f535,f536,f537,f538,f539,f540,f541,f542,f555,f557,f560,f563,f590,f593,f595,f807,f815,f877,f898,f1021,f1022,f1024,f1025,f1026,f1043,f1053,f1054,f1056,f1057,f1058,f1143,f1409,f1413,f1526,f2352,f2353,f2999,f3100,f3221,f3294,f3359,f3350,f3474,f3617,f3615,f4143,f4338,f4526,f4775,f4776,f4779,f4781,f4787,f4786,f4785,f4780,f4771,f4770,f4600,f4599,f4598,f4597,f4595,f4318,f4317,f4316,f4315,f4313,f4123,f4122,f4121,f4120,f4118,f3972,f3971,f3970,f3969,f3967,f3108,f2613,f2612,f2611,f2610,f2609,f2607,f2606,f2525,f2524,f2523,f2522,f2520,f2395,f2351,f1535,f1533,f1482,f1479,f1475,f1474,f1473,f1472,f764,f724,f688,f656,f571,f403,f4596,f4314,f4119,f3968,f2521,f1435,f1055,f1023,f911,f470,f4756,f2400,f2397,f2394,f2389,f2387,f931,f2363,f2399,f2585,f2582,f2398,f1405,f892,f816,f2996,f2471,f2468,f2414,f2367,f2366,f2365,f2364,f2362,f2361,f805,f804,f803,f802,f801,f433,f3346,f3345,f4476,f3425,f3331,f3344,f944,f500,f499,f3326,f3343,f3342,f3327,f3095,f3082,f3030,f483,f482,f481,f479,f478,f2413,f821,f820,f819,f818,f817,f434,f1414,f477,f480,f2467,f2918,f1192,f1108,f813,f920,f2975,f6457,f6476,f6770,f6771,f6772,f6694,f6776,f6779,f6759,f6760,f6761,f6778,f6793,f6810,f6783,f6782,f6885,f6980,f6993,f7010,f6983,f6982,f7096,f7097,f7098,f7201,f7130,f7137,f7217,f7218,f7165,f7166,f7290,f7534,f7533,f7667,f7542,f7543,f7544,f7545,f7556,f7560,f7677,f7566,f7567,f7708,f7599,f7600,f7601,f7623,f7624,f7744,f7753,f7793,f7794,f7799,f7809,f7811,f7839,f8031,f8032,f8033,f8034,f8035,f8041,f8042,f8043,f8046,f8054,f8055,f8056,f8057,f8058,f8059,f8060,f8061,f8084,f8085,f8086,f8087,f8088,f8089,f8090,f8091,f8097,f7911,f7912,f7913,f7914,f7925,f7929,f8107,f7935,f7936,f8138,f7968,f7969,f7970,f7971,f7982,f7986,f8148,f7992,f7993,f8174,f8200,f8290,f8190,f8293,f8373,f8569,f8570,f8571,f8572,f8574,f8575,f8576,f8537,f8538,f8539,f8540,f8581,f8586,f8566,f8567,f8568,f8975,f8976,f8977,f8978,f8980,f8981,f8982,f8959,f8984,f8985,f8998,f9095,f9101,f9147,f9153,f9156,f9158,f9160,f9164,f9165,f9166,f9185,f9186,f9187,f9146,f9735,f9737,f9616,f9617,f9618,f9745,f9747,f9749,f9751,f9646,f9647,f9648,f9649,f9755,f9757,f9758,f9663,f9670,f9698,f9773,f9775,f9776,f9778,f9779,f9704,f9791,f9834,f9835,f9836,f9793,f9803,f9841,f9854,f9856,f9889,f9948,f9991,f9992,f9993,f9998,f10011,f10015,f10050,f10127,f10272,f10278,f10328,f10340,f10341,f10327,f10432,f10431,f10464,f10465,f10419,f10418,f10576,f10645,f10644,f10668,f10669,f10630,f10698,f10697,f10720,f10721,f10683,f10735,f10736,f10737,f10750,f10751,f10752,f11195,f11196,f11105,f11204,f11211,f11505,f11506,f11316,f11508,f11317,f11516,f11318,f11522,f11523,f11524,f11525,f11526,f11527,f11320,f11531,f11322,f11323,f11324,f11325,f11326,f11327,f11341,f11342,f11343,f11344,f11345,f11346,f11347,f11348,f11350,f11351,f11352,f11353,f11367,f11368,f11369,f11370,f11392,f11393,f11396,f11400,f11414,f11421,f11439,f11440,f11460,f11599,f11601,f11602,f11603,f11466,f11619,f11618,f11654,f11655,f11606,f11620,f11621,f11622,f11623,f11624,f11828,f11829,f11791,f11808,f11932,f12053,f12054,f12032,f12033,f12584,f12585,f12546,f12564,f12677,f12737,f12750,f12757,f12736,f13047,f13048,f12900,f12901,f12902,f12903,f12918,f12938,f12939,f12940,f12941,f12955,f12962,f13001,f13091,f13094,f13095,f13007,f13116,f13115,f13114,f13152,f13153,f13101,f13100,f13270,f13543,f13544,f13393,f13394,f13395,f13396,f13411,f13432,f13433,f13434,f13435,f13449,f13456,f13495,f13587,f13590,f13591,f13501,f13612,f13611,f13610,f13638,f13639,f13595,f13671,f13670,f13669,f13697,f13698,f13654,f13762,f13768,f13772,f13791,f13790,f13789,f13912,f13913,f13775,f13774,f13895,f14120,f14281,f14282,f14283,f14236,f14286,f14290,f14293,f14296,f14242,f14299,f14251,f14500,f14501,f14337,f14412,f14510,f14517,f14454,f14555,f14554,f14553,f14581,f14582,f14601,f14600,f14599,f14598,f15080,f15081,f14979,f15089,f15096,f15044,f15169,f15170,f15144,f15145,f15382,f15383,f15384,f15388,f15442,f15484,f15584,f15586,f15587,f15490,f15684,f15685,f15693,f15700,f15650,f15720,f15719,f15721,f15722,f15799,f15800,f15801,f15802,f15812,f15843,f15844,f15845,f15846,f15953,f15954,f15962,f15969,f15919,f16084,f16085,f16093,f16100,f16050,f16122,f16121,f16123,f16124,f16204,f16205,f16206,f16207,f16217,f16249,f16250,f16251,f16252,f16746,f16747,f16755,f16762,f16712,f16956,f16957,f16965,f16972,f16922,f18193,f18153,f17244,f17245,f18143,f18675,f17974,f18002,f18818,f19256])).
fof(f19256,plain,(
  ( ! [X6,X4,X5,X3] : (s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK10(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl3 | $spl22 | $spl58)),
  inference(subsumption_resolution,[],[f18826,f8998])).
fof(f18826,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,sK8('_i_0_0',bool)) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK10(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl58),
  inference(backward_demodulation,[],[f18818,f10751])).
fof(f17245,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,sK8('_i_0_0',bool)) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f9095,f555])).
fof(f17244,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,sK8('_i_0_0',bool)) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,sK8('_i_0_0',bool)) | s(bool,X0) = s(bool,X3)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f9095,f9095])).
fof(f16922,plain,(
  ( ! [X74,X75,X73] : (s(X74,happ(s(fun(bool,X74),X75),s(bool,'T_0'))) = s(X74,'const.combin.o_3'(s(fun(bool,X74),X75),s(fun(fun(X73,bool),bool),'!_1'),s(fun(X73,bool),sK8('!_1',fun(X73,bool)))))) ) | $spl3),
  inference(superposition,[],[f502,f13669])).
fof(f16972,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,happ(s(fun(fun(X34,bool),bool),'!_1'),s(fun(X34,bool),sK8('!_1',fun(X34,bool))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl3),
  inference(subsumption_resolution,[],[f16899,f401])).
fof(f16899,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X35) = s(bool,happ(s(fun(fun(X34,bool),bool),'!_1'),s(fun(X34,bool),sK8('!_1',fun(X34,bool))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl3),
  inference(superposition,[],[f1409,f13669])).
fof(f16965,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),sK8('!_1',fun(X11,bool))))),s(bool,X12)))) | p(s(bool,X12))) ) | $spl3),
  inference(forward_demodulation,[],[f16964,f502])).
fof(f16964,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),sK8('!_1',fun(X11,bool))))))),s(bool,X12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f16887,f401])).
fof(f16887,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),sK8('!_1',fun(X11,bool))))))),s(bool,X12))))) ) | $spl3),
  inference(superposition,[],[f481,f13669])).
fof(f16957,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))) ) | $spl3),
  inference(forward_demodulation,[],[f16861,f502])).
fof(f16861,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13669,f500])).
fof(f16956,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f16862,f502])).
fof(f16862,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13669,f500])).
fof(f16712,plain,(
  ( ! [X74,X72,X73] : (s(X73,happ(s(fun(bool,X73),X74),s(bool,'T_0'))) = s(X73,'const.combin.o_3'(s(fun(bool,X73),X74),s(fun(fun(X72,bool),bool),'?_1'),s(fun(X72,bool),sK8('?_1',fun(X72,bool)))))) ) | $spl3),
  inference(superposition,[],[f502,f10697])).
fof(f16762,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,happ(s(fun(fun(X33,bool),bool),'?_1'),s(fun(X33,bool),sK8('?_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f16689,f401])).
fof(f16689,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,happ(s(fun(fun(X33,bool),bool),'?_1'),s(fun(X33,bool),sK8('?_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f1409,f10697])).
fof(f16755,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),sK8('?_1',fun(X10,bool))))),s(bool,X11)))) | p(s(bool,X11))) ) | $spl3),
  inference(forward_demodulation,[],[f16754,f502])).
fof(f16754,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),sK8('?_1',fun(X10,bool))))))),s(bool,X11))))) ) | $spl3),
  inference(subsumption_resolution,[],[f16677,f401])).
fof(f16677,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),sK8('?_1',fun(X10,bool))))))),s(bool,X11))))) ) | $spl3),
  inference(superposition,[],[f481,f10697])).
fof(f16747,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))))) ) | $spl3),
  inference(forward_demodulation,[],[f16651,f502])).
fof(f16651,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10697,f500])).
fof(f16746,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f16652,f502])).
fof(f16652,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10697,f500])).
fof(f16252,plain,(
  ( ! [X187,X185,X186,X184,X183] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X183),bool),happ(s(fun('type.bvpSem.state'(X183),fun('type.bvpSem.state'(X183),bool)),'=_0'),s('type.bvpSem.state'(X183),X187))),s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X184),s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X185),s('type.bvpSem.state'(X183),X186)))))))) | s('type.bvpSem.state'(X183),X187) != s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X184),s(fun('type.num.num','type.num.num'),X185))),s('type.bvpSem.state'(X183),X186)))) )),
  inference(superposition,[],[f500,f439])).
fof(f16251,plain,(
  ( ! [X182,X180,X178,X181,X179] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X178),bool),happ(s(fun('type.bvpSem.state'(X178),fun('type.bvpSem.state'(X178),bool)),'=_0'),s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X179),s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X180),s('type.bvpSem.state'(X178),X181))))))),s('type.bvpSem.state'(X178),X182)))) | s('type.bvpSem.state'(X178),X182) != s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X179),s(fun('type.num.num','type.num.num'),X180))),s('type.bvpSem.state'(X178),X181)))) )),
  inference(superposition,[],[f500,f439])).
fof(f16250,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X173),bool),happ(s(fun('type.bvpSem.state'(X173),fun('type.bvpSem.state'(X173),bool)),'=_0'),s('type.bvpSem.state'(X173),X177))),s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X174),s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X175),s('type.bvpSem.state'(X173),X176)))))))) | s('type.bvpSem.state'(X173),X177) = s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X174),s(fun('type.num.num','type.num.num'),X175))),s('type.bvpSem.state'(X173),X176)))) )),
  inference(superposition,[],[f499,f439])).
fof(f16249,plain,(
  ( ! [X171,X169,X172,X170,X168] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X168),bool),happ(s(fun('type.bvpSem.state'(X168),fun('type.bvpSem.state'(X168),bool)),'=_0'),s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X169),s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X170),s('type.bvpSem.state'(X168),X171))))))),s('type.bvpSem.state'(X168),X172)))) | s('type.bvpSem.state'(X168),X172) = s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X169),s(fun('type.num.num','type.num.num'),X170))),s('type.bvpSem.state'(X168),X171)))) )),
  inference(superposition,[],[f499,f439])).
fof(f16217,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X3))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3)))))))) )),
  inference(superposition,[],[f412,f439])).
fof(f16207,plain,(
  ( ! [X187,X185,X186,X184,X183] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X183),bool),happ(s(fun('type.bvpSem.state'(X183),fun('type.bvpSem.state'(X183),bool)),'=_0'),s('type.bvpSem.state'(X183),X187))),s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X184),s(fun('type.num.num','type.num.num'),X185))),s('type.bvpSem.state'(X183),X186)))))) | s('type.bvpSem.state'(X183),X187) != s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X184),s('type.bvpSem.state'(X183),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X185),s('type.bvpSem.state'(X183),X186)))))) )),
  inference(superposition,[],[f500,f439])).
fof(f16206,plain,(
  ( ! [X182,X180,X178,X181,X179] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X178),bool),happ(s(fun('type.bvpSem.state'(X178),fun('type.bvpSem.state'(X178),bool)),'=_0'),s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X179),s(fun('type.num.num','type.num.num'),X180))),s('type.bvpSem.state'(X178),X181))))),s('type.bvpSem.state'(X178),X182)))) | s('type.bvpSem.state'(X178),X182) != s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X179),s('type.bvpSem.state'(X178),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X180),s('type.bvpSem.state'(X178),X181)))))) )),
  inference(superposition,[],[f500,f439])).
fof(f16205,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X173),bool),happ(s(fun('type.bvpSem.state'(X173),fun('type.bvpSem.state'(X173),bool)),'=_0'),s('type.bvpSem.state'(X173),X177))),s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X174),s(fun('type.num.num','type.num.num'),X175))),s('type.bvpSem.state'(X173),X176)))))) | s('type.bvpSem.state'(X173),X177) = s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X174),s('type.bvpSem.state'(X173),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X175),s('type.bvpSem.state'(X173),X176)))))) )),
  inference(superposition,[],[f499,f439])).
fof(f16204,plain,(
  ( ! [X171,X169,X172,X170,X168] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X168),bool),happ(s(fun('type.bvpSem.state'(X168),fun('type.bvpSem.state'(X168),bool)),'=_0'),s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X169),s(fun('type.num.num','type.num.num'),X170))),s('type.bvpSem.state'(X168),X171))))),s('type.bvpSem.state'(X168),X172)))) | s('type.bvpSem.state'(X168),X172) = s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X169),s('type.bvpSem.state'(X168),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X170),s('type.bvpSem.state'(X168),X171)))))) )),
  inference(superposition,[],[f499,f439])).
fof(f16124,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.bvpSem.state'(X5),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.bvpSem.state'(X5),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X6),s('type.bvpSem.state'(X5),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X7),s('type.bvpSem.state'(X5),X8))))))) = s('type.bvpSem.state'(X5),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X9),s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X6),s(fun('type.num.num','type.num.num'),X7))))),s('type.bvpSem.state'(X5),X8)))) )),
  inference(superposition,[],[f439,f439])).
fof(f16123,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X4),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X4),s(fun('type.num.num','type.num.num'),X1))),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3)))))) )),
  inference(superposition,[],[f439,f439])).
fof(f16121,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3))))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))))) )),
  inference(unit_resulting_resolution,[],[f439,f500])).
fof(f16122,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3))))))))) )),
  inference(unit_resulting_resolution,[],[f439,f500])).
fof(f16050,plain,(
  ( ! [X72,X71,X73] : (s(X72,happ(s(fun(bool,X72),X73),s(bool,'T_0'))) = s(X72,'const.combin.o_3'(s(fun(bool,X72),X73),s(fun(fun(fun(X71,bool),bool),bool),'?_1'),s(fun(fun(X71,bool),bool),'!_1')))) ) | $spl3),
  inference(superposition,[],[f502,f13610])).
fof(f16100,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,happ(s(fun(fun(fun(X33,bool),bool),bool),'?_1'),s(fun(fun(X33,bool),bool),'!_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f16028,f401])).
fof(f16028,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,happ(s(fun(fun(fun(X33,bool),bool),bool),'?_1'),s(fun(fun(X33,bool),bool),'!_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f1409,f13610])).
fof(f16093,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'!_1'))),s(bool,X11)))) | p(s(bool,X11))) ) | $spl3),
  inference(forward_demodulation,[],[f16092,f502])).
fof(f16092,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'!_1'))))),s(bool,X11))))) ) | $spl3),
  inference(subsumption_resolution,[],[f16016,f401])).
fof(f16016,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'!_1'))))),s(bool,X11))))) ) | $spl3),
  inference(superposition,[],[f481,f13610])).
fof(f16085,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))))) ) | $spl3),
  inference(forward_demodulation,[],[f15991,f502])).
fof(f15991,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13610,f500])).
fof(f16084,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f15992,f502])).
fof(f15992,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13610,f500])).
fof(f15919,plain,(
  ( ! [X72,X71,X73] : (s(X72,happ(s(fun(bool,X72),X73),s(bool,'T_0'))) = s(X72,'const.combin.o_3'(s(fun(bool,X72),X73),s(fun(fun(fun(X71,bool),bool),bool),'?_1'),s(fun(fun(X71,bool),bool),'?_1')))) ) | $spl3),
  inference(superposition,[],[f502,f10644])).
fof(f15969,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,happ(s(fun(fun(fun(X33,bool),bool),bool),'?_1'),s(fun(fun(X33,bool),bool),'?_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f15897,f401])).
fof(f15897,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,happ(s(fun(fun(fun(X33,bool),bool),bool),'?_1'),s(fun(fun(X33,bool),bool),'?_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f1409,f10644])).
fof(f15962,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'?_1'))),s(bool,X11)))) | p(s(bool,X11))) ) | $spl3),
  inference(forward_demodulation,[],[f15961,f502])).
fof(f15961,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'?_1'))))),s(bool,X11))))) ) | $spl3),
  inference(subsumption_resolution,[],[f15885,f401])).
fof(f15885,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X10,bool),bool),bool),'?_1'),s(fun(fun(X10,bool),bool),'?_1'))))),s(bool,X11))))) ) | $spl3),
  inference(superposition,[],[f481,f10644])).
fof(f15954,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))) ) | $spl3),
  inference(forward_demodulation,[],[f15860,f502])).
fof(f15860,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10644,f500])).
fof(f15953,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f15861,f502])).
fof(f15861,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10644,f500])).
fof(f15846,plain,(
  ( ! [X177,X180,X178,X181,X179] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X177),bool),happ(s(fun('type.bvpSem.state'(X177),fun('type.bvpSem.state'(X177),bool)),'=_0'),s('type.bvpSem.state'(X177),X181))),s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X178),s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X179),s('type.bvpSem.state'(X177),X180)))))))) | s('type.bvpSem.state'(X177),X181) != s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X178),s(fun('type.num.num','type.num.num'),X179))),s('type.bvpSem.state'(X177),X180)))) )),
  inference(superposition,[],[f500,f438])).
fof(f15845,plain,(
  ( ! [X175,X173,X176,X174,X172] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X172),bool),happ(s(fun('type.bvpSem.state'(X172),fun('type.bvpSem.state'(X172),bool)),'=_0'),s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X173),s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X174),s('type.bvpSem.state'(X172),X175))))))),s('type.bvpSem.state'(X172),X176)))) | s('type.bvpSem.state'(X172),X176) != s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X173),s(fun('type.num.num','type.num.num'),X174))),s('type.bvpSem.state'(X172),X175)))) )),
  inference(superposition,[],[f500,f438])).
fof(f15844,plain,(
  ( ! [X167,X171,X169,X170,X168] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X167),bool),happ(s(fun('type.bvpSem.state'(X167),fun('type.bvpSem.state'(X167),bool)),'=_0'),s('type.bvpSem.state'(X167),X171))),s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X168),s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X169),s('type.bvpSem.state'(X167),X170)))))))) | s('type.bvpSem.state'(X167),X171) = s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X168),s(fun('type.num.num','type.num.num'),X169))),s('type.bvpSem.state'(X167),X170)))) )),
  inference(superposition,[],[f499,f438])).
fof(f15843,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X162),bool),happ(s(fun('type.bvpSem.state'(X162),fun('type.bvpSem.state'(X162),bool)),'=_0'),s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X163),s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X164),s('type.bvpSem.state'(X162),X165))))))),s('type.bvpSem.state'(X162),X166)))) | s('type.bvpSem.state'(X162),X166) = s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X163),s(fun('type.num.num','type.num.num'),X164))),s('type.bvpSem.state'(X162),X165)))) )),
  inference(superposition,[],[f499,f438])).
fof(f15812,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X3))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3)))))))) )),
  inference(superposition,[],[f414,f438])).
fof(f15802,plain,(
  ( ! [X177,X180,X178,X181,X179] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X177),bool),happ(s(fun('type.bvpSem.state'(X177),fun('type.bvpSem.state'(X177),bool)),'=_0'),s('type.bvpSem.state'(X177),X181))),s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X178),s(fun('type.num.num','type.num.num'),X179))),s('type.bvpSem.state'(X177),X180)))))) | s('type.bvpSem.state'(X177),X181) != s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X178),s('type.bvpSem.state'(X177),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X179),s('type.bvpSem.state'(X177),X180)))))) )),
  inference(superposition,[],[f500,f438])).
fof(f15801,plain,(
  ( ! [X175,X173,X176,X174,X172] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X172),bool),happ(s(fun('type.bvpSem.state'(X172),fun('type.bvpSem.state'(X172),bool)),'=_0'),s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X173),s(fun('type.num.num','type.num.num'),X174))),s('type.bvpSem.state'(X172),X175))))),s('type.bvpSem.state'(X172),X176)))) | s('type.bvpSem.state'(X172),X176) != s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X173),s('type.bvpSem.state'(X172),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X174),s('type.bvpSem.state'(X172),X175)))))) )),
  inference(superposition,[],[f500,f438])).
fof(f15800,plain,(
  ( ! [X167,X171,X169,X170,X168] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X167),bool),happ(s(fun('type.bvpSem.state'(X167),fun('type.bvpSem.state'(X167),bool)),'=_0'),s('type.bvpSem.state'(X167),X171))),s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X168),s(fun('type.num.num','type.num.num'),X169))),s('type.bvpSem.state'(X167),X170)))))) | s('type.bvpSem.state'(X167),X171) = s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X168),s('type.bvpSem.state'(X167),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X169),s('type.bvpSem.state'(X167),X170)))))) )),
  inference(superposition,[],[f499,f438])).
fof(f15799,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,happ(s(fun('type.bvpSem.state'(X162),bool),happ(s(fun('type.bvpSem.state'(X162),fun('type.bvpSem.state'(X162),bool)),'=_0'),s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X163),s(fun('type.num.num','type.num.num'),X164))),s('type.bvpSem.state'(X162),X165))))),s('type.bvpSem.state'(X162),X166)))) | s('type.bvpSem.state'(X162),X166) = s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X163),s('type.bvpSem.state'(X162),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X164),s('type.bvpSem.state'(X162),X165)))))) )),
  inference(superposition,[],[f499,f438])).
fof(f15722,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.bvpSem.state'(X5),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.bvpSem.state'(X5),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X6),s('type.bvpSem.state'(X5),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X7),s('type.bvpSem.state'(X5),X8))))))) = s('type.bvpSem.state'(X5),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X9),s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X6),s(fun('type.num.num','type.num.num'),X7))))),s('type.bvpSem.state'(X5),X8)))) )),
  inference(superposition,[],[f438,f438])).
fof(f15721,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X4),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X4),s(fun('type.num.num','type.num.num'),X1))),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3)))))) )),
  inference(superposition,[],[f438,f438])).
fof(f15719,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3))))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))))) )),
  inference(unit_resulting_resolution,[],[f438,f500])).
fof(f15720,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X1),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X3))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X3))))))))) )),
  inference(unit_resulting_resolution,[],[f438,f500])).
fof(f15650,plain,(
  ( ! [X28,X29] : (s(X28,happ(s(fun(bool,X28),X29),s(bool,'T_0'))) = s(X28,'const.combin.o_3'(s(fun(bool,X28),X29),s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1')))) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f502,f6993])).
fof(f15700,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f15628,f401])).
fof(f15628,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X12) = s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f1409,f6993])).
fof(f15693,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))),s(bool,X3)))) | p(s(bool,X3))) ) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f15692,f502])).
fof(f15692,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))))),s(bool,X3))))) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f15616,f401])).
fof(f15616,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))))),s(bool,X3))))) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f481,f6993])).
fof(f15685,plain,(
  p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1')))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f15591,f502])).
fof(f15591,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1')))))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f6993,f500])).
fof(f15684,plain,(
  p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))),s(bool,'T_0')))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f15592,f502])).
fof(f15592,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))))),s(bool,'T_0')))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f6993,f500])).
fof(f15490,plain,(
  ( ! [X600,X596,X598,X594,X597,X599,X595] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X594,bool),X596))) | s(X594,X600) = s(X594,sK5(X600,'const.combin.o_3'(s(fun(X595,fun(X594,bool)),'_dst_x0_1'(s(fun(X594,bool),X596))),s(fun(X597,X595),X598),s(X597,X599)),X594))) )),
  inference(superposition,[],[f807,f11346])).
fof(f15587,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (p(s(bool,happ(s(fun(X581,bool),X583),s(X581,sK8('const.combin.o_3'(s(fun(X582,fun(X581,bool)),'_dst_x0_1'(s(fun(X581,bool),X583))),s(fun(X584,X582),X585),s(X584,X586)),X581))))) | ~p(s(bool,happ(s(fun(fun(X581,bool),bool),'?_1'),s(fun(X581,bool),X583))))) )),
  inference(forward_demodulation,[],[f15488,f11346])).
fof(f15488,plain,(
  ( ! [X584,X586,X582,X585,X581,X583] : (~p(s(bool,happ(s(fun(fun(X581,bool),bool),'?_1'),s(fun(X581,bool),X583)))) | p(s(bool,happ(s(fun(X581,bool),'const.combin.o_3'(s(fun(X582,fun(X581,bool)),'_dst_x0_1'(s(fun(X581,bool),X583))),s(fun(X584,X582),X585),s(X584,X586))),s(X581,sK8('const.combin.o_3'(s(fun(X582,fun(X581,bool)),'_dst_x0_1'(s(fun(X581,bool),X583))),s(fun(X584,X582),X585),s(X584,X586)),X581)))))) )),
  inference(superposition,[],[f479,f11346])).
fof(f15586,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (p(s(bool,happ(s(fun(fun(X575,bool),bool),'!_1'),s(fun(X575,bool),X577)))) | ~p(s(bool,happ(s(fun(X575,bool),X577),s(X575,sK7('const.combin.o_3'(s(fun(X576,fun(X575,bool)),'_dst_x0_1'(s(fun(X575,bool),X577))),s(fun(X578,X576),X579),s(X578,X580)),X575)))))) )),
  inference(forward_demodulation,[],[f15487,f11346])).
fof(f15487,plain,(
  ( ! [X575,X580,X576,X578,X577,X579] : (~p(s(bool,happ(s(fun(X575,bool),X577),s(X575,sK7('const.combin.o_3'(s(fun(X576,fun(X575,bool)),'_dst_x0_1'(s(fun(X575,bool),X577))),s(fun(X578,X576),X579),s(X578,X580)),X575))))) | p(s(bool,happ(s(fun(fun(X575,bool),bool),'!_1'),s(fun(X575,bool),'const.combin.o_3'(s(fun(X576,fun(X575,bool)),'_dst_x0_1'(s(fun(X575,bool),X577))),s(fun(X578,X576),X579),s(X578,X580))))))) )),
  inference(superposition,[],[f478,f11346])).
fof(f15584,plain,(
  ( ! [X565,X567,X561,X563,X564,X566,X562] : (p(s(bool,'|dollar|exists_1'(s(fun(X561,bool),X563)))) | ~p(s(bool,happ(s(fun(X561,bool),X563),s(X561,sK5(X567,'const.combin.o_3'(s(fun(X562,fun(X561,bool)),'_dst_x0_1'(s(fun(X561,bool),X563))),s(fun(X564,X562),X565),s(X564,X566)),X561)))))) )),
  inference(forward_demodulation,[],[f15485,f11346])).
fof(f15485,plain,(
  ( ! [X565,X567,X561,X563,X564,X566,X562] : (~p(s(bool,happ(s(fun(X561,bool),X563),s(X561,sK5(X567,'const.combin.o_3'(s(fun(X562,fun(X561,bool)),'_dst_x0_1'(s(fun(X561,bool),X563))),s(fun(X564,X562),X565),s(X564,X566)),X561))))) | p(s(bool,'|dollar|exists_1'(s(fun(X561,bool),'const.combin.o_3'(s(fun(X562,fun(X561,bool)),'_dst_x0_1'(s(fun(X561,bool),X563))),s(fun(X564,X562),X565),s(X564,X566))))))) )),
  inference(superposition,[],[f434,f11346])).
fof(f15484,plain,(
  ( ! [X557,X559,X555,X560,X556,X558,X554] : (p(s(bool,'|dollar|exists_1'(s(fun(X554,bool),X556)))) | s(X554,X560) = s(X554,sK5(X560,'const.combin.o_3'(s(fun(X555,fun(X554,bool)),'_dst_x0_1'(s(fun(X554,bool),X556))),s(fun(X557,X555),X558),s(X557,X559)),X554))) )),
  inference(superposition,[],[f433,f11346])).
fof(f15442,plain,(
  ( ! [X292,X288,X290,X287,X289,X291] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X292,bool),'_dst_x0_1'(s(bool,X288))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X287,bool),'_dst_x0_1'(s(bool,X288))),s(fun(X289,X287),X290),s(X289,X291)))) )),
  inference(superposition,[],[f1482,f11346])).
fof(f15388,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X45,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X40,bool),'_dst_x0_1'(s(bool,X41))),s(fun(X42,X40),X43),s(X42,X44)))) )),
  inference(superposition,[],[f11346,f1535])).
fof(f15384,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X18) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))),s(fun(X19,X17),X20),s(X19,X21)))) )),
  inference(superposition,[],[f11346,f555])).
fof(f15383,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,X12) | s(bool,X16) = s(bool,'const.combin.o_3'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12))),s(fun(X13,X11),X14),s(X13,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f11346,f555])).
fof(f15382,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X6,bool),'_dst_x0_1'(s(bool,X7))),s(fun(X8,X6),X9),s(X8,X10)))) )),
  inference(superposition,[],[f11346,f407])).
fof(f15145,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun('type.modSem.environment',X2),X3),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))))) = s(X2,'const.combin.o_3'(s(fun('type.modSem.environment',X2),X3),s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1)))) )),
  inference(superposition,[],[f502,f472])).
fof(f15144,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2),s('type.modSem.environment',happ(s(fun(X3,'type.modSem.environment'),X4),s(X3,X5))))) = s('type.modSem.environment','const.combin.o_3'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))),s(fun(X3,'type.modSem.environment'),X4),s(X3,X5)))) )),
  inference(superposition,[],[f472,f502])).
fof(f15170,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.modSem.environment',bool),'_dst_x0x0x1_3'(s(fun('type.modSem.environment',fun('type.modSem.environment',bool)),'=_0'),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))),s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))))),s('type.modSem.environment',X1))))) )),
  inference(forward_demodulation,[],[f15117,f508])).
fof(f15117,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.modSem.environment',bool),happ(s(fun('type.modSem.environment',fun('type.modSem.environment',bool)),'=_0'),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))))),s('type.modSem.environment',happ(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1))))))) )),
  inference(unit_resulting_resolution,[],[f472,f500])).
fof(f15169,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.modSem.environment',bool),'const.combin.o_3'(s(fun('type.modSem.environment',fun('type.modSem.environment',bool)),'=_0'),s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1))),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))))))) )),
  inference(forward_demodulation,[],[f15118,f502])).
fof(f15118,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.modSem.environment',bool),happ(s(fun('type.modSem.environment',fun('type.modSem.environment',bool)),'=_0'),s('type.modSem.environment',happ(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1))))),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))))))) )),
  inference(unit_resulting_resolution,[],[f472,f500])).
fof(f15044,plain,(
  ( ! [X74,X72,X73] : (s(X73,happ(s(fun(bool,X73),X74),s(bool,'T_0'))) = s(X73,'const.combin.o_3'(s(fun(bool,X73),X74),s(fun(fun(X72,bool),bool),'!_1'),s(fun(X72,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | $spl3),
  inference(superposition,[],[f502,f13114])).
fof(f15096,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,happ(s(fun(fun(X34,bool),bool),'!_1'),s(fun(X34,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl3),
  inference(subsumption_resolution,[],[f15022,f401])).
fof(f15022,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X35) = s(bool,happ(s(fun(fun(X34,bool),bool),'!_1'),s(fun(X34,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl3),
  inference(superposition,[],[f1409,f13114])).
fof(f15089,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,X12)))) | p(s(bool,X12))) ) | $spl3),
  inference(forward_demodulation,[],[f15088,f502])).
fof(f15088,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f15010,f401])).
fof(f15010,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X11,bool),bool),'!_1'),s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X12))))) ) | $spl3),
  inference(superposition,[],[f481,f13114])).
fof(f14979,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f13114,f555])).
fof(f15081,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f14977,f502])).
fof(f14977,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13114,f500])).
fof(f15080,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f14978,f502])).
fof(f14978,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13114,f500])).
fof(f14598,plain,(
  s('type.bvpSem.state'(sK1),f(s(sK0,sK10('const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)),'type.bvpSem.state'(sK1),sK0)))) != s('type.bvpSem.state'(sK1),g(s(sK0,sK10('const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)),'type.bvpSem.state'(sK1),sK0))))),
  inference(unit_resulting_resolution,[],[f400,f492])).
fof(f14600,plain,(
  s('type.bvpSem.state'(sK1),f(s(sK0,sK10('const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))),'type.bvpSem.state'(sK1),sK0)))) != s('type.bvpSem.state'(sK1),g(s(sK0,sK10('const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK4))),s(fun(sK0,'type.bvpSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(sK1),'type.bvpSem.state'(sK1)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),sK3))),s(fun(sK0,'type.bvpSem.state'(sK1)),sK2)))),'type.bvpSem.state'(sK1),sK0))))),
  inference(unit_resulting_resolution,[],[f400,f492])).
fof(f14581,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))) )),
  inference(forward_demodulation,[],[f14546,f502])).
fof(f14546,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))) )),
  inference(unit_resulting_resolution,[],[f13774,f483])).
fof(f14553,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f13774,f1479])).
fof(f14554,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))))) )),
  inference(unit_resulting_resolution,[],[f13774,f6760])).
fof(f14555,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))))) )),
  inference(unit_resulting_resolution,[],[f13774,f8567])).
fof(f14454,plain,(
  ( ! [X114,X112,X115,X113] : (s(X114,happ(s(fun(bool,X114),X115),s(bool,'T_0'))) = s(X114,'const.combin.o_3'(s(fun(bool,X114),X115),s(fun(X112,bool),sK8('!_1',fun(X112,bool))),s(X112,X113)))) ) | $spl3),
  inference(superposition,[],[f502,f13789])).
fof(f14517,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,X57) = s(bool,happ(s(fun(X55,bool),sK8('!_1',fun(X55,bool))),s(X55,X56))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | $spl3),
  inference(subsumption_resolution,[],[f14433,f401])).
fof(f14433,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | s(bool,X57) = s(bool,happ(s(fun(X55,bool),sK8('!_1',fun(X55,bool))),s(X55,X56))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | $spl3),
  inference(superposition,[],[f1409,f13789])).
fof(f14510,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X18,bool),sK8('!_1',fun(X18,bool))),s(X18,X19))),s(bool,X20)))) | p(s(bool,X20))) ) | $spl3),
  inference(forward_demodulation,[],[f14509,f502])).
fof(f14509,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X20)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X18,bool),sK8('!_1',fun(X18,bool))),s(X18,X19))))),s(bool,X20))))) ) | $spl3),
  inference(subsumption_resolution,[],[f14421,f401])).
fof(f14421,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X18,bool),sK8('!_1',fun(X18,bool))),s(X18,X19))))),s(bool,X20))))) ) | $spl3),
  inference(superposition,[],[f481,f13789])).
fof(f14412,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X45,bool),sK8('!_1',fun(X45,bool))),s(fun(X46,X45),X47),s(X46,X48)))) ) | $spl3),
  inference(superposition,[],[f13789,f502])).
fof(f14337,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,happ(s(fun(bool,bool),sK8('!_1',fun(bool,bool))),s(bool,X77))) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | $spl3),
  inference(superposition,[],[f13789,f555])).
fof(f14501,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))))) ) | $spl3),
  inference(forward_demodulation,[],[f14306,f502])).
fof(f14306,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13789,f500])).
fof(f14500,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f14307,f502])).
fof(f14307,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f13789,f500])).
fof(f14251,plain,(
  ( ! [X94,X92,X97,X95,X93,X98,X96] : (s(X97,'const.combin.o_3'(s(fun(fun(fun(X92,X93),X93),X97),X98),s(fun(X94,fun(fun(X92,X93),X93)),'_dstsel_1_x10_1'(s(fun(X94,X92),X95))),s(X94,X96))) = s(X97,happ(s(fun(fun(fun(X92,X93),X93),X97),X98),s(fun(fun(X92,X93),X93),'_dstsel_0_x0_1'(s(X92,happ(s(fun(X94,X92),X95),s(X94,X96)))))))) )),
  inference(superposition,[],[f502,f501])).
fof(f14299,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),'_dstsel_0_x0_1'(s(X33,happ(s(fun(X34,X33),X35),s(X34,X36))))))) | ~p(s(bool,'const.combin.o_3'(s(fun(X33,bool),X37),s(fun(X34,X33),X35),s(X34,X36))))) )),
  inference(forward_demodulation,[],[f14298,f501])).
fof(f14298,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,'const.combin.o_3'(s(fun(X33,bool),X37),s(fun(X34,X33),X35),s(X34,X36)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),happ(s(fun(X34,fun(fun(X33,bool),bool)),'_dstsel_1_x10_1'(s(fun(X34,X33),X35))),s(X34,X36)))))) )),
  inference(forward_demodulation,[],[f14297,f502])).
fof(f14297,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(X33,bool),X37),s(X33,happ(s(fun(X34,X33),X35),s(X34,X36)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),happ(s(fun(X34,fun(fun(X33,bool),bool)),'_dstsel_1_x10_1'(s(fun(X34,X33),X35))),s(X34,X36)))))) )),
  inference(forward_demodulation,[],[f14243,f491])).
fof(f14243,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(fun(X33,bool),bool),'_dstsel_0_x0_1'(s(X33,happ(s(fun(X34,X33),X35),s(X34,X36))))),s(fun(X33,bool),X37)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),happ(s(fun(X34,fun(fun(X33,bool),bool)),'_dstsel_1_x10_1'(s(fun(X34,X33),X35))),s(X34,X36)))))) )),
  inference(superposition,[],[f2467,f501])).
fof(f14242,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X28,bool),bool),'_dstsel_0_x0_1'(s(X28,happ(s(fun(X29,X28),X30),s(X29,X31))))))) | s(fun(X28,bool),X32) = s(fun(X28,bool),sK5(X32,happ(s(fun(X29,fun(fun(X28,bool),bool)),'_dstsel_1_x10_1'(s(fun(X29,X28),X30))),s(X29,X31)),fun(X28,bool)))) )),
  inference(superposition,[],[f807,f501])).
fof(f14296,plain,(
  ( ! [X26,X24,X23,X27,X25] : (p(s(bool,'const.combin.o_3'(s(fun(fun(fun(X23,bool),bool),bool),'?_1'),s(fun(X24,fun(fun(X23,bool),bool)),'_dstsel_1_x10_1'(s(fun(X24,X23),X25))),s(X24,X26)))) | ~p(s(bool,'const.combin.o_3'(s(fun(X23,bool),X27),s(fun(X24,X23),X25),s(X24,X26))))) )),
  inference(forward_demodulation,[],[f14295,f502])).
fof(f14295,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,'const.combin.o_3'(s(fun(X23,bool),X27),s(fun(X24,X23),X25),s(X24,X26)))) | p(s(bool,happ(s(fun(fun(fun(X23,bool),bool),bool),'?_1'),s(fun(fun(X23,bool),bool),happ(s(fun(X24,fun(fun(X23,bool),bool)),'_dstsel_1_x10_1'(s(fun(X24,X23),X25))),s(X24,X26))))))) )),
  inference(forward_demodulation,[],[f14294,f502])).
fof(f14294,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(X23,bool),X27),s(X23,happ(s(fun(X24,X23),X25),s(X24,X26)))))) | p(s(bool,happ(s(fun(fun(fun(X23,bool),bool),bool),'?_1'),s(fun(fun(X23,bool),bool),happ(s(fun(X24,fun(fun(X23,bool),bool)),'_dstsel_1_x10_1'(s(fun(X24,X23),X25))),s(X24,X26))))))) )),
  inference(forward_demodulation,[],[f14241,f491])).
fof(f14241,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(fun(X23,bool),bool),'_dstsel_0_x0_1'(s(X23,happ(s(fun(X24,X23),X25),s(X24,X26))))),s(fun(X23,bool),X27)))) | p(s(bool,happ(s(fun(fun(fun(X23,bool),bool),bool),'?_1'),s(fun(fun(X23,bool),bool),happ(s(fun(X24,fun(fun(X23,bool),bool)),'_dstsel_1_x10_1'(s(fun(X24,X23),X25))),s(X24,X26))))))) )),
  inference(superposition,[],[f480,f501])).
fof(f14293,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.combin.o_3'(s(fun(X19,bool),sK8(happ(s(fun(X20,fun(fun(X19,bool),bool)),'_dstsel_1_x10_1'(s(fun(X20,X19),X21))),s(X20,X22)),fun(X19,bool))),s(fun(X20,X19),X21),s(X20,X22)))) | ~p(s(bool,happ(s(fun(fun(fun(X19,bool),bool),bool),'?_1'),s(fun(fun(X19,bool),bool),'_dstsel_0_x0_1'(s(X19,happ(s(fun(X20,X19),X21),s(X20,X22))))))))) )),
  inference(forward_demodulation,[],[f14292,f502])).
fof(f14292,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun(X19,bool),sK8(happ(s(fun(X20,fun(fun(X19,bool),bool)),'_dstsel_1_x10_1'(s(fun(X20,X19),X21))),s(X20,X22)),fun(X19,bool))),s(X19,happ(s(fun(X20,X19),X21),s(X20,X22)))))) | ~p(s(bool,happ(s(fun(fun(fun(X19,bool),bool),bool),'?_1'),s(fun(fun(X19,bool),bool),'_dstsel_0_x0_1'(s(X19,happ(s(fun(X20,X19),X21),s(X20,X22))))))))) )),
  inference(forward_demodulation,[],[f14291,f491])).
fof(f14291,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun(fun(X19,bool),bool),'_dstsel_0_x0_1'(s(X19,happ(s(fun(X20,X19),X21),s(X20,X22))))),s(fun(X19,bool),sK8(happ(s(fun(X20,fun(fun(X19,bool),bool)),'_dstsel_1_x10_1'(s(fun(X20,X19),X21))),s(X20,X22)),fun(X19,bool)))))) | ~p(s(bool,happ(s(fun(fun(fun(X19,bool),bool),bool),'?_1'),s(fun(fun(X19,bool),bool),'_dstsel_0_x0_1'(s(X19,happ(s(fun(X20,X19),X21),s(X20,X22))))))))) )),
  inference(forward_demodulation,[],[f14240,f501])).
fof(f14240,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(fun(fun(X19,bool),bool),bool),'?_1'),s(fun(fun(X19,bool),bool),'_dstsel_0_x0_1'(s(X19,happ(s(fun(X20,X19),X21),s(X20,X22)))))))) | p(s(bool,happ(s(fun(fun(X19,bool),bool),happ(s(fun(X20,fun(fun(X19,bool),bool)),'_dstsel_1_x10_1'(s(fun(X20,X19),X21))),s(X20,X22))),s(fun(X19,bool),sK8(happ(s(fun(X20,fun(fun(X19,bool),bool)),'_dstsel_1_x10_1'(s(fun(X20,X19),X21))),s(X20,X22)),fun(X19,bool))))))) )),
  inference(superposition,[],[f479,f501])).
fof(f14290,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.combin.o_3'(s(fun(fun(fun(X15,bool),bool),bool),'!_1'),s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18)))) | ~p(s(bool,'const.combin.o_3'(s(fun(X15,bool),sK7(happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18)),fun(X15,bool))),s(fun(X16,X15),X17),s(X16,X18))))) )),
  inference(forward_demodulation,[],[f14289,f502])).
fof(f14289,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.combin.o_3'(s(fun(X15,bool),sK7(happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18)),fun(X15,bool))),s(fun(X16,X15),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(fun(fun(X15,bool),bool),bool),'!_1'),s(fun(fun(X15,bool),bool),happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18))))))) )),
  inference(forward_demodulation,[],[f14288,f502])).
fof(f14288,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X15,bool),sK7(happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18)),fun(X15,bool))),s(X15,happ(s(fun(X16,X15),X17),s(X16,X18)))))) | p(s(bool,happ(s(fun(fun(fun(X15,bool),bool),bool),'!_1'),s(fun(fun(X15,bool),bool),happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18))))))) )),
  inference(forward_demodulation,[],[f14239,f491])).
fof(f14239,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(fun(X15,bool),bool),'_dstsel_0_x0_1'(s(X15,happ(s(fun(X16,X15),X17),s(X16,X18))))),s(fun(X15,bool),sK7(happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18)),fun(X15,bool)))))) | p(s(bool,happ(s(fun(fun(fun(X15,bool),bool),bool),'!_1'),s(fun(fun(X15,bool),bool),happ(s(fun(X16,fun(fun(X15,bool),bool)),'_dstsel_1_x10_1'(s(fun(X16,X15),X17))),s(X16,X18))))))) )),
  inference(superposition,[],[f478,f501])).
fof(f14286,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(fun(X5,bool),bool),'_dstsel_0_x0_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8)))))))) | ~p(s(bool,'const.combin.o_3'(s(fun(X5,bool),sK5(X9,happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8)),fun(X5,bool))),s(fun(X6,X5),X7),s(X6,X8))))) )),
  inference(forward_demodulation,[],[f14285,f501])).
fof(f14285,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.combin.o_3'(s(fun(X5,bool),sK5(X9,happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8)),fun(X5,bool))),s(fun(X6,X5),X7),s(X6,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(fun(X5,bool),bool),happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8))))))) )),
  inference(forward_demodulation,[],[f14284,f502])).
fof(f14284,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),sK5(X9,happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8)),fun(X5,bool))),s(X5,happ(s(fun(X6,X5),X7),s(X6,X8)))))) | p(s(bool,'|dollar|exists_1'(s(fun(fun(X5,bool),bool),happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8))))))) )),
  inference(forward_demodulation,[],[f14237,f491])).
fof(f14237,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(fun(X5,bool),bool),'_dstsel_0_x0_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))))),s(fun(X5,bool),sK5(X9,happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8)),fun(X5,bool)))))) | p(s(bool,'|dollar|exists_1'(s(fun(fun(X5,bool),bool),happ(s(fun(X6,fun(fun(X5,bool),bool)),'_dstsel_1_x10_1'(s(fun(X6,X5),X7))),s(X6,X8))))))) )),
  inference(superposition,[],[f434,f501])).
fof(f14236,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(fun(X0,bool),bool),'_dstsel_0_x0_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))))))) | s(fun(X0,bool),X4) = s(fun(X0,bool),sK5(X4,happ(s(fun(X1,fun(fun(X0,bool),bool)),'_dstsel_1_x10_1'(s(fun(X1,X0),X2))),s(X1,X3)),fun(X0,bool)))) )),
  inference(superposition,[],[f433,f501])).
fof(f14283,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s(fun(fun(X5,X6),X6),'_dstsel_0_x0_1'(s(X5,'const.combin.o_3'(s(fun(X7,X5),X8),s(fun(X9,X7),X10),s(X9,X11))))) = s(fun(fun(X5,X6),X6),'const.combin.o_3'(s(fun(X7,fun(fun(X5,X6),X6)),'_dstsel_1_x10_1'(s(fun(X7,X5),X8))),s(fun(X9,X7),X10),s(X9,X11)))) )),
  inference(forward_demodulation,[],[f14235,f502])).
fof(f14235,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s(fun(fun(X5,X6),X6),'_dstsel_0_x0_1'(s(X5,happ(s(fun(X7,X5),X8),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))))) = s(fun(fun(X5,X6),X6),'const.combin.o_3'(s(fun(X7,fun(fun(X5,X6),X6)),'_dstsel_1_x10_1'(s(fun(X7,X5),X8))),s(fun(X9,X7),X10),s(X9,X11)))) )),
  inference(superposition,[],[f501,f502])).
fof(f14282,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(fun(fun(X0,X1),X1),bool),happ(s(fun(fun(fun(X0,X1),X1),fun(fun(fun(X0,X1),X1),bool)),'=_0'),s(fun(fun(X0,X1),X1),'_dstsel_0_x0_1'(s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))),s(fun(X2,fun(fun(X0,X1),X1)),'_dstsel_1_x10_1'(s(fun(X2,X0),X3))),s(X2,X4))))) )),
  inference(forward_demodulation,[],[f14132,f502])).
fof(f14132,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(fun(X0,X1),X1),bool),happ(s(fun(fun(fun(X0,X1),X1),fun(fun(fun(X0,X1),X1),bool)),'=_0'),s(fun(fun(X0,X1),X1),'_dstsel_0_x0_1'(s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))),s(fun(fun(X0,X1),X1),happ(s(fun(X2,fun(fun(X0,X1),X1)),'_dstsel_1_x10_1'(s(fun(X2,X0),X3))),s(X2,X4))))))) )),
  inference(unit_resulting_resolution,[],[f501,f500])).
fof(f14281,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(fun(X0,X1),X1),bool),'const.combin.o_3'(s(fun(fun(fun(X0,X1),X1),fun(fun(fun(X0,X1),X1),bool)),'=_0'),s(fun(X2,fun(fun(X0,X1),X1)),'_dstsel_1_x10_1'(s(fun(X2,X0),X3))),s(X2,X4))),s(fun(fun(X0,X1),X1),'_dstsel_0_x0_1'(s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))))) )),
  inference(forward_demodulation,[],[f14133,f502])).
fof(f14133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(fun(fun(X0,X1),X1),bool),happ(s(fun(fun(fun(X0,X1),X1),fun(fun(fun(X0,X1),X1),bool)),'=_0'),s(fun(fun(X0,X1),X1),happ(s(fun(X2,fun(fun(X0,X1),X1)),'_dstsel_1_x10_1'(s(fun(X2,X0),X3))),s(X2,X4))))),s(fun(fun(X0,X1),X1),'_dstsel_0_x0_1'(s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))))) )),
  inference(unit_resulting_resolution,[],[f501,f500])).
fof(f14120,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(X33,bool),sK8('!_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f14054,f401])).
fof(f14054,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(X33,bool),sK8('!_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1409,f13775])).
fof(f13895,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,'const.combin.o_3'(s(fun(X45,bool),sK8('!_1',fun(X45,bool))),s(fun(X46,X45),X47),s(X46,X48))))) )),
  inference(superposition,[],[f13654,f502])).
fof(f13775,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),sK8('!_1',fun(X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f13654,f2467])).
fof(f13913,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f13781,f502])).
fof(f13781,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f13654,f481])).
fof(f13912,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(X1,bool),sK8('!_1',fun(X1,bool))),s(X1,X2))))) )),
  inference(forward_demodulation,[],[f13782,f502])).
fof(f13782,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(X1,bool),sK8('!_1',fun(X1,bool))),s(X1,X2))))))) )),
  inference(unit_resulting_resolution,[],[f13654,f483])).
fof(f13789,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),sK8('!_1',fun(X0,bool))),s(X0,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f13654,f1479])).
fof(f13790,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(X1,bool),sK8('!_1',fun(X1,bool))),s(X1,X2))))))))) )),
  inference(unit_resulting_resolution,[],[f13654,f6760])).
fof(f13791,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(X1,bool),sK8('!_1',fun(X1,bool))),s(X1,X2))))))))) )),
  inference(unit_resulting_resolution,[],[f13654,f8567])).
fof(f13772,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.list.list'('type.bvpSem.stack'),X66) != s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),X65))))) | p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X66),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),X65))))))) )),
  inference(forward_demodulation,[],[f13771,f430])).
fof(f13771,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X66),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),X65)))))) | s('type.list.list'('type.bvpSem.stack'),X66) != s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64),s('type.bvpSem.state'(X63),X65)))))) )),
  inference(forward_demodulation,[],[f13755,f508])).
fof(f13755,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X66))),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),X65)))))))) | s('type.list.list'('type.bvpSem.stack'),X66) != s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X63),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X64),s('type.bvpSem.state'(X63),X65)))))) )),
  inference(superposition,[],[f500,f430])).
fof(f13768,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.list.list'('type.bvpSem.stack'),X58) = s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),X57))))) | ~p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X58),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),X57))))))) )),
  inference(forward_demodulation,[],[f13767,f430])).
fof(f13767,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X58),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),X57)))))) | s('type.list.list'('type.bvpSem.stack'),X58) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56),s('type.bvpSem.state'(X55),X57)))))) )),
  inference(forward_demodulation,[],[f13753,f508])).
fof(f13753,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),X58))),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),X57)))))))) | s('type.list.list'('type.bvpSem.stack'),X58) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X55),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X56),s('type.bvpSem.state'(X55),X57)))))) )),
  inference(superposition,[],[f499,f430])).
fof(f13762,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))))),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f13761,f430])).
fof(f13761,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),'_dst_x0x0x1_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f13714,f508])).
fof(f13714,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X2))))))))) )),
  inference(unit_resulting_resolution,[],[f430,f500])).
fof(f13698,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f13661,f502])).
fof(f13661,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f13595,f481])).
fof(f13697,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))) )),
  inference(forward_demodulation,[],[f13662,f502])).
fof(f13662,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))) )),
  inference(unit_resulting_resolution,[],[f13595,f483])).
fof(f13669,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),sK8('!_1',fun(X0,bool)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f13595,f1479])).
fof(f13670,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))))) )),
  inference(unit_resulting_resolution,[],[f13595,f6760])).
fof(f13671,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),sK8('!_1',fun(X1,bool))))))))))) )),
  inference(unit_resulting_resolution,[],[f13595,f8567])).
fof(f13639,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f13602,f502])).
fof(f13602,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1'))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f13100,f481])).
fof(f13638,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'!_1'))))) )),
  inference(forward_demodulation,[],[f13603,f502])).
fof(f13603,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'!_1'))))))) )),
  inference(unit_resulting_resolution,[],[f13100,f483])).
fof(f13610,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'!_1')))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f13100,f1479])).
fof(f13611,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'!_1'))))))))) )),
  inference(unit_resulting_resolution,[],[f13100,f6760])).
fof(f13612,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'!_1'))))))))) )),
  inference(unit_resulting_resolution,[],[f13100,f8567])).
fof(f13501,plain,(
  ( ! [X588,X590,X586,X589,X591,X585,X587] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X585,bool),'const.combin.o_3'(s(fun(X587,fun(X585,bool)),X588),s(fun(X586,X587),X589),s(X586,X590))))) | s(X585,X591) = s(X585,sK5(X591,happ(s(fun(X586,fun(X585,bool)),'_dst_x0x1_2'(s(fun(X587,fun(X585,bool)),X588),s(fun(X586,X587),X589))),s(X586,X590)),X585))) )),
  inference(superposition,[],[f807,f552])).
fof(f13591,plain,(
  ( ! [X573,X575,X576,X572,X574,X577] : (p(s(bool,happ(s(fun(X572,bool),'const.combin.o_3'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576),s(X573,X577))),s(X572,sK8(happ(s(fun(X573,fun(X572,bool)),'_dst_x0x1_2'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576))),s(X573,X577)),X572))))) | ~p(s(bool,happ(s(fun(fun(X572,bool),bool),'?_1'),s(fun(X572,bool),'const.combin.o_3'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576),s(X573,X577))))))) )),
  inference(forward_demodulation,[],[f13499,f552])).
fof(f13499,plain,(
  ( ! [X573,X575,X576,X572,X574,X577] : (~p(s(bool,happ(s(fun(fun(X572,bool),bool),'?_1'),s(fun(X572,bool),'const.combin.o_3'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576),s(X573,X577)))))) | p(s(bool,happ(s(fun(X572,bool),happ(s(fun(X573,fun(X572,bool)),'_dst_x0x1_2'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576))),s(X573,X577))),s(X572,sK8(happ(s(fun(X573,fun(X572,bool)),'_dst_x0x1_2'(s(fun(X574,fun(X572,bool)),X575),s(fun(X573,X574),X576))),s(X573,X577)),X572)))))) )),
  inference(superposition,[],[f479,f552])).
fof(f13590,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (p(s(bool,happ(s(fun(fun(X566,bool),bool),'!_1'),s(fun(X566,bool),'const.combin.o_3'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570),s(X567,X571)))))) | ~p(s(bool,happ(s(fun(X566,bool),'const.combin.o_3'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570),s(X567,X571))),s(X566,sK7(happ(s(fun(X567,fun(X566,bool)),'_dst_x0x1_2'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570))),s(X567,X571)),X566)))))) )),
  inference(forward_demodulation,[],[f13589,f13411])).
fof(f13589,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X566,bool),bool),'!_1'),s(fun(X567,fun(X566,bool)),'_dst_x0x1_2'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570))),s(X567,X571)))) | ~p(s(bool,happ(s(fun(X566,bool),'const.combin.o_3'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570),s(X567,X571))),s(X566,sK7(happ(s(fun(X567,fun(X566,bool)),'_dst_x0x1_2'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570))),s(X567,X571)),X566)))))) )),
  inference(forward_demodulation,[],[f13498,f502])).
fof(f13498,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (~p(s(bool,happ(s(fun(X566,bool),'const.combin.o_3'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570),s(X567,X571))),s(X566,sK7(happ(s(fun(X567,fun(X566,bool)),'_dst_x0x1_2'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570))),s(X567,X571)),X566))))) | p(s(bool,happ(s(fun(fun(X566,bool),bool),'!_1'),s(fun(X566,bool),happ(s(fun(X567,fun(X566,bool)),'_dst_x0x1_2'(s(fun(X568,fun(X566,bool)),X569),s(fun(X567,X568),X570))),s(X567,X571))))))) )),
  inference(superposition,[],[f478,f552])).
fof(f13587,plain,(
  ( ! [X557,X553,X555,X556,X558,X552,X554] : (p(s(bool,'|dollar|exists_1'(s(fun(X552,bool),'const.combin.o_3'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556),s(X553,X557)))))) | ~p(s(bool,happ(s(fun(X552,bool),'const.combin.o_3'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556),s(X553,X557))),s(X552,sK5(X558,happ(s(fun(X553,fun(X552,bool)),'_dst_x0x1_2'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556))),s(X553,X557)),X552)))))) )),
  inference(forward_demodulation,[],[f13496,f552])).
fof(f13496,plain,(
  ( ! [X557,X553,X555,X556,X558,X552,X554] : (~p(s(bool,happ(s(fun(X552,bool),'const.combin.o_3'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556),s(X553,X557))),s(X552,sK5(X558,happ(s(fun(X553,fun(X552,bool)),'_dst_x0x1_2'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556))),s(X553,X557)),X552))))) | p(s(bool,'|dollar|exists_1'(s(fun(X552,bool),happ(s(fun(X553,fun(X552,bool)),'_dst_x0x1_2'(s(fun(X554,fun(X552,bool)),X555),s(fun(X553,X554),X556))),s(X553,X557))))))) )),
  inference(superposition,[],[f434,f552])).
fof(f13495,plain,(
  ( ! [X546,X549,X551,X545,X547,X548,X550] : (p(s(bool,'|dollar|exists_1'(s(fun(X545,bool),'const.combin.o_3'(s(fun(X547,fun(X545,bool)),X548),s(fun(X546,X547),X549),s(X546,X550)))))) | s(X545,X551) = s(X545,sK5(X551,happ(s(fun(X546,fun(X545,bool)),'_dst_x0x1_2'(s(fun(X547,fun(X545,bool)),X548),s(fun(X546,X547),X549))),s(X546,X550)),X545))) )),
  inference(superposition,[],[f433,f552])).
fof(f13456,plain,(
  ( ! [X300,X302,X298,X301,X297,X299] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X302,bool),'_dst_x0_1'(s(bool,'const.combin.o_3'(s(fun(X298,bool),X299),s(fun(X297,X298),X300),s(X297,X301))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X297,bool),'_dst_x0x1_2'(s(fun(X298,bool),X299),s(fun(X297,X298),X300))),s(X297,X301)))) )),
  inference(superposition,[],[f1482,f552])).
fof(f13449,plain,(
  ( ! [X257,X259,X260,X256,X258] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.combin.o_3'(s(fun(X257,bool),X258),s(fun(X256,X257),X259),s(X256,X260))))) | s(bool,'F_0') = s(bool,happ(s(fun(X256,bool),'_dst_x0x1_2'(s(fun(X257,bool),X258),s(fun(X256,X257),X259))),s(X256,X260)))) )),
  inference(superposition,[],[f571,f552])).
fof(f13435,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'const.combin.o_3'(s(fun(X177,bool),X178),s(fun(X176,X177),X179),s(X176,sK7('_dst_x0x1_2'(s(fun(X177,bool),X178),s(fun(X176,X177),X179)),X176))))) | p(s(bool,happ(s(fun(fun(X176,bool),bool),'!_1'),s(fun(X176,bool),'_dst_x0x1_2'(s(fun(X177,bool),X178),s(fun(X176,X177),X179))))))) )),
  inference(superposition,[],[f478,f552])).
fof(f13434,plain,(
  ( ! [X175,X173,X171,X174,X172] : (~p(s(bool,'const.combin.o_3'(s(fun(X172,bool),X173),s(fun(X171,X172),X174),s(X171,X175)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X171,bool),'_dst_x0x1_2'(s(fun(X172,bool),X173),s(fun(X171,X172),X174)))))) )),
  inference(superposition,[],[f2467,f552])).
fof(f13433,plain,(
  ( ! [X167,X169,X166,X170,X168] : (~p(s(bool,'const.combin.o_3'(s(fun(X167,bool),X168),s(fun(X166,X167),X169),s(X166,X170)))) | p(s(bool,happ(s(fun(fun(X166,bool),bool),'?_1'),s(fun(X166,bool),'_dst_x0x1_2'(s(fun(X167,bool),X168),s(fun(X166,X167),X169))))))) )),
  inference(superposition,[],[f480,f552])).
fof(f13432,plain,(
  ( ! [X165,X163,X161,X164,X162] : (~p(s(bool,'const.combin.o_3'(s(fun(X162,bool),X163),s(fun(X161,X162),X164),s(X161,sK5(X165,'_dst_x0x1_2'(s(fun(X162,bool),X163),s(fun(X161,X162),X164)),X161))))) | p(s(bool,'|dollar|exists_1'(s(fun(X161,bool),'_dst_x0x1_2'(s(fun(X162,bool),X163),s(fun(X161,X162),X164))))))) )),
  inference(superposition,[],[f434,f552])).
fof(f13411,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))) = s(X6,'const.combin.o_3'(s(fun(X0,X6),X7),s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5)))) )),
  inference(superposition,[],[f502,f552])).
fof(f13396,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X30) = s(bool,'const.combin.o_3'(s(fun(X26,bool),X27),s(fun(X25,X26),X28),s(X25,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X25,X26),X28))),s(X25,X29)))) )),
  inference(superposition,[],[f552,f555])).
fof(f13395,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X20,bool),X21),s(fun(X19,X20),X22),s(X19,X23))) | s(bool,X24) = s(bool,happ(s(fun(X19,bool),'_dst_x0x1_2'(s(fun(X20,bool),X21),s(fun(X19,X20),X22))),s(X19,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f552,f555])).
fof(f13394,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X15,bool),X16),s(fun(X14,X15),X17),s(X14,X18))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),'_dst_x0x1_2'(s(fun(X15,bool),X16),s(fun(X14,X15),X17))),s(X14,X18)))) )),
  inference(superposition,[],[f552,f407])).
fof(f13393,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X6,'const.combin.o_3'(s(fun(X8,X6),X9),s(fun(X7,X8),X10),s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))) = s(X6,'const.combin.o_3'(s(fun(X7,X6),'_dst_x0x1_2'(s(fun(X8,X6),X9),s(fun(X7,X8),X10))),s(fun(X11,X7),X12),s(X11,X13)))) )),
  inference(superposition,[],[f552,f502])).
fof(f13544,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))) )),
  inference(forward_demodulation,[],[f13282,f502])).
fof(f13282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(X0,happ(s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f552,f500])).
fof(f13543,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))))) )),
  inference(forward_demodulation,[],[f13283,f502])).
fof(f13283,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))))) )),
  inference(unit_resulting_resolution,[],[f552,f500])).
fof(f13270,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),'!_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f13205,f401])).
fof(f13205,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(fun(X33,bool),bool),'!_1'))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1409,f13101])).
fof(f13153,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f13106,f502])).
fof(f13106,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f12677,f481])).
fof(f13152,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(forward_demodulation,[],[f13107,f502])).
fof(f13107,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f12677,f483])).
fof(f13114,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f12677,f1479])).
fof(f13115,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f12677,f6760])).
fof(f13116,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'!_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f12677,f8567])).
fof(f13007,plain,(
  ( ! [X575,X580,X576,X578,X581,X577,X579] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X575,bool),'const.combin.o_3'(s(fun(X577,fun(X575,bool)),X578),s(fun(X576,X577),X579),s(X576,X580))))) | s(X575,X581) = s(X575,sK5(X581,happ(s(fun(X576,fun(X575,bool)),'const.combin.o_2'(s(fun(X577,fun(X575,bool)),X578),s(fun(X576,X577),X579))),s(X576,X580)),X575))) )),
  inference(superposition,[],[f807,f503])).
fof(f13095,plain,(
  ( ! [X565,X567,X563,X564,X566,X562] : (p(s(bool,happ(s(fun(X562,bool),'const.combin.o_3'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566),s(X563,X567))),s(X562,sK8(happ(s(fun(X563,fun(X562,bool)),'const.combin.o_2'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566))),s(X563,X567)),X562))))) | ~p(s(bool,happ(s(fun(fun(X562,bool),bool),'?_1'),s(fun(X562,bool),'const.combin.o_3'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566),s(X563,X567))))))) )),
  inference(forward_demodulation,[],[f13005,f503])).
fof(f13005,plain,(
  ( ! [X565,X567,X563,X564,X566,X562] : (~p(s(bool,happ(s(fun(fun(X562,bool),bool),'?_1'),s(fun(X562,bool),'const.combin.o_3'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566),s(X563,X567)))))) | p(s(bool,happ(s(fun(X562,bool),happ(s(fun(X563,fun(X562,bool)),'const.combin.o_2'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566))),s(X563,X567))),s(X562,sK8(happ(s(fun(X563,fun(X562,bool)),'const.combin.o_2'(s(fun(X564,fun(X562,bool)),X565),s(fun(X563,X564),X566))),s(X563,X567)),X562)))))) )),
  inference(superposition,[],[f479,f503])).
fof(f13094,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,happ(s(fun(fun(X556,bool),bool),'!_1'),s(fun(X556,bool),'const.combin.o_3'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560),s(X557,X561)))))) | ~p(s(bool,happ(s(fun(X556,bool),'const.combin.o_3'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560),s(X557,X561))),s(X556,sK7(happ(s(fun(X557,fun(X556,bool)),'const.combin.o_2'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560))),s(X557,X561)),X556)))))) )),
  inference(forward_demodulation,[],[f13093,f12918])).
fof(f13093,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X556,bool),bool),'!_1'),s(fun(X557,fun(X556,bool)),'const.combin.o_2'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560))),s(X557,X561)))) | ~p(s(bool,happ(s(fun(X556,bool),'const.combin.o_3'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560),s(X557,X561))),s(X556,sK7(happ(s(fun(X557,fun(X556,bool)),'const.combin.o_2'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560))),s(X557,X561)),X556)))))) )),
  inference(forward_demodulation,[],[f13004,f502])).
fof(f13004,plain,(
  ( ! [X561,X557,X559,X560,X556,X558] : (~p(s(bool,happ(s(fun(X556,bool),'const.combin.o_3'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560),s(X557,X561))),s(X556,sK7(happ(s(fun(X557,fun(X556,bool)),'const.combin.o_2'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560))),s(X557,X561)),X556))))) | p(s(bool,happ(s(fun(fun(X556,bool),bool),'!_1'),s(fun(X556,bool),happ(s(fun(X557,fun(X556,bool)),'const.combin.o_2'(s(fun(X558,fun(X556,bool)),X559),s(fun(X557,X558),X560))),s(X557,X561))))))) )),
  inference(superposition,[],[f478,f503])).
fof(f13091,plain,(
  ( ! [X546,X542,X545,X547,X543,X548,X544] : (p(s(bool,'|dollar|exists_1'(s(fun(X542,bool),'const.combin.o_3'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546),s(X543,X547)))))) | ~p(s(bool,happ(s(fun(X542,bool),'const.combin.o_3'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546),s(X543,X547))),s(X542,sK5(X548,happ(s(fun(X543,fun(X542,bool)),'const.combin.o_2'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546))),s(X543,X547)),X542)))))) )),
  inference(forward_demodulation,[],[f13002,f503])).
fof(f13002,plain,(
  ( ! [X546,X542,X545,X547,X543,X548,X544] : (~p(s(bool,happ(s(fun(X542,bool),'const.combin.o_3'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546),s(X543,X547))),s(X542,sK5(X548,happ(s(fun(X543,fun(X542,bool)),'const.combin.o_2'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546))),s(X543,X547)),X542))))) | p(s(bool,'|dollar|exists_1'(s(fun(X542,bool),happ(s(fun(X543,fun(X542,bool)),'const.combin.o_2'(s(fun(X544,fun(X542,bool)),X545),s(fun(X543,X544),X546))),s(X543,X547))))))) )),
  inference(superposition,[],[f434,f503])).
fof(f13001,plain,(
  ( ! [X540,X536,X538,X541,X537,X539,X535] : (p(s(bool,'|dollar|exists_1'(s(fun(X535,bool),'const.combin.o_3'(s(fun(X537,fun(X535,bool)),X538),s(fun(X536,X537),X539),s(X536,X540)))))) | s(X535,X541) = s(X535,sK5(X541,happ(s(fun(X536,fun(X535,bool)),'const.combin.o_2'(s(fun(X537,fun(X535,bool)),X538),s(fun(X536,X537),X539))),s(X536,X540)),X535))) )),
  inference(superposition,[],[f433,f503])).
fof(f12962,plain,(
  ( ! [X292,X288,X290,X287,X289,X291] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X292,bool),'_dst_x0_1'(s(bool,'const.combin.o_3'(s(fun(X288,bool),X289),s(fun(X287,X288),X290),s(X287,X291))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X287,bool),'const.combin.o_2'(s(fun(X288,bool),X289),s(fun(X287,X288),X290))),s(X287,X291)))) )),
  inference(superposition,[],[f1482,f503])).
fof(f12955,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.combin.o_3'(s(fun(X247,bool),X248),s(fun(X246,X247),X249),s(X246,X250))))) | s(bool,'F_0') = s(bool,happ(s(fun(X246,bool),'const.combin.o_2'(s(fun(X247,bool),X248),s(fun(X246,X247),X249))),s(X246,X250)))) )),
  inference(superposition,[],[f571,f503])).
fof(f12941,plain,(
  ( ! [X167,X169,X166,X168] : (~p(s(bool,'const.combin.o_3'(s(fun(X167,bool),X168),s(fun(X166,X167),X169),s(X166,sK7('const.combin.o_2'(s(fun(X167,bool),X168),s(fun(X166,X167),X169)),X166))))) | p(s(bool,happ(s(fun(fun(X166,bool),bool),'!_1'),s(fun(X166,bool),'const.combin.o_2'(s(fun(X167,bool),X168),s(fun(X166,X167),X169))))))) )),
  inference(superposition,[],[f478,f503])).
fof(f12940,plain,(
  ( ! [X165,X163,X161,X164,X162] : (~p(s(bool,'const.combin.o_3'(s(fun(X162,bool),X163),s(fun(X161,X162),X164),s(X161,X165)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X161,bool),'const.combin.o_2'(s(fun(X162,bool),X163),s(fun(X161,X162),X164)))))) )),
  inference(superposition,[],[f2467,f503])).
fof(f12939,plain,(
  ( ! [X158,X156,X159,X157,X160] : (~p(s(bool,'const.combin.o_3'(s(fun(X157,bool),X158),s(fun(X156,X157),X159),s(X156,X160)))) | p(s(bool,happ(s(fun(fun(X156,bool),bool),'?_1'),s(fun(X156,bool),'const.combin.o_2'(s(fun(X157,bool),X158),s(fun(X156,X157),X159))))))) )),
  inference(superposition,[],[f480,f503])).
fof(f12938,plain,(
  ( ! [X154,X152,X151,X155,X153] : (~p(s(bool,'const.combin.o_3'(s(fun(X152,bool),X153),s(fun(X151,X152),X154),s(X151,sK5(X155,'const.combin.o_2'(s(fun(X152,bool),X153),s(fun(X151,X152),X154)),X151))))) | p(s(bool,'|dollar|exists_1'(s(fun(X151,bool),'const.combin.o_2'(s(fun(X152,bool),X153),s(fun(X151,X152),X154))))))) )),
  inference(superposition,[],[f434,f503])).
fof(f12918,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))) = s(X6,'const.combin.o_3'(s(fun(X0,X6),X7),s(fun(X1,X0),'const.combin.o_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5)))) )),
  inference(superposition,[],[f502,f503])).
fof(f12903,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X30) = s(bool,'const.combin.o_3'(s(fun(X26,bool),X27),s(fun(X25,X26),X28),s(X25,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'const.combin.o_2'(s(fun(X26,bool),X27),s(fun(X25,X26),X28))),s(X25,X29)))) )),
  inference(superposition,[],[f503,f555])).
fof(f12902,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X20,bool),X21),s(fun(X19,X20),X22),s(X19,X23))) | s(bool,X24) = s(bool,happ(s(fun(X19,bool),'const.combin.o_2'(s(fun(X20,bool),X21),s(fun(X19,X20),X22))),s(X19,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f503,f555])).
fof(f12901,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X15,bool),X16),s(fun(X14,X15),X17),s(X14,X18))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),'const.combin.o_2'(s(fun(X15,bool),X16),s(fun(X14,X15),X17))),s(X14,X18)))) )),
  inference(superposition,[],[f503,f407])).
fof(f12900,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X6,'const.combin.o_3'(s(fun(X8,X6),X9),s(fun(X7,X8),X10),s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))) = s(X6,'const.combin.o_3'(s(fun(X7,X6),'const.combin.o_2'(s(fun(X8,X6),X9),s(fun(X7,X8),X10))),s(fun(X11,X7),X12),s(X11,X13)))) )),
  inference(superposition,[],[f503,f502])).
fof(f13048,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))) )),
  inference(forward_demodulation,[],[f12793,f502])).
fof(f12793,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(X0,happ(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f503,f500])).
fof(f13047,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),'const.combin.o_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))))) )),
  inference(forward_demodulation,[],[f12794,f502])).
fof(f12794,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'const.combin.o_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s(X0,'const.combin.o_3'(s(fun(X2,X0),X3),s(fun(X1,X2),X4),s(X1,X5))))))) )),
  inference(unit_resulting_resolution,[],[f503,f500])).
fof(f12736,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,X146)) | p(s(bool,happ(s(fun(fun(X148,bool),bool),'!_1'),s(fun(X148,bool),'_dst_x0_1'(s(bool,sK5(X146,X147,bool))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f8567,f807])).
fof(f12757,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,happ(s(fun(fun(X81,bool),bool),'!_1'),s(fun(X81,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,X80)))))))))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f12709,f401])).
fof(f12709,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X81,bool),bool),'!_1'),s(fun(X81,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,X80)))))))))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f8567,f1482])).
fof(f12750,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,happ(s(fun(fun(X69,bool),bool),'!_1'),s(fun(X69,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))))))) | s(X67,X70) = s(X67,sK5(X70,X68,X67))) )),
  inference(subsumption_resolution,[],[f12705,f401])).
fof(f12705,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X69,bool),bool),'!_1'),s(fun(X69,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))))))) | s(X67,X70) = s(X67,sK5(X70,X68,X67))) )),
  inference(superposition,[],[f8567,f807])).
fof(f12737,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),'_dst_x0_1'(s(bool,X2)))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f12679,f401])).
fof(f12679,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),'_dst_x0_1'(s(bool,X2)))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f8567,f555])).
fof(f12564,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(X72,happ(s(fun('type.bvpSem.state'(X69),X72),X73),s('type.bvpSem.state'(X69),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X70),s('type.bvpSem.state'(X69),X71))))) = s(X72,'const.combin.o_3'(s(fun('type.bvpSem.state'(X69),X72),X73),s(fun('type.bvpSem.state'(X69),'type.bvpSem.state'(X69)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X70))),s('type.bvpSem.state'(X69),X71)))) )),
  inference(superposition,[],[f502,f429])).
fof(f12546,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.bvpSem.state'(X3),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4),s('type.bvpSem.state'(X3),happ(s(fun(X5,'type.bvpSem.state'(X3)),X6),s(X5,X7))))) = s('type.bvpSem.state'(X3),'const.combin.o_3'(s(fun('type.bvpSem.state'(X3),'type.bvpSem.state'(X3)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X4))),s(fun(X5,'type.bvpSem.state'(X3)),X6),s(X5,X7)))) )),
  inference(superposition,[],[f429,f502])).
fof(f12585,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),'_dst_x0x0x1_3'(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))))),s('type.bvpSem.state'(X0),X2))))) )),
  inference(forward_demodulation,[],[f12522,f508])).
fof(f12522,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))),s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f429,f500])).
fof(f12584,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),'const.combin.o_3'(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2))),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f12523,f502])).
fof(f12523,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f429,f500])).
fof(f12033,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,happ(s(fun('type.ffi.ffi_state'(X0),X3),X4),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))))) = s(X3,'const.combin.o_3'(s(fun('type.ffi.ffi_state'(X0),X3),X4),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2)))) )),
  inference(superposition,[],[f502,f428])).
fof(f12032,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.ffi.ffi_state'(X3),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4),s('type.ffi.ffi_state'(X3),happ(s(fun(X5,'type.ffi.ffi_state'(X3)),X6),s(X5,X7))))) = s('type.ffi.ffi_state'(X3),'const.combin.o_3'(s(fun('type.ffi.ffi_state'(X3),'type.ffi.ffi_state'(X3)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4))),s(fun(X5,'type.ffi.ffi_state'(X3)),X6),s(X5,X7)))) )),
  inference(superposition,[],[f428,f502])).
fof(f12054,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),'_dst_x0x0x1_3'(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))))),s('type.ffi.ffi_state'(X0),X2))))) )),
  inference(forward_demodulation,[],[f12008,f508])).
fof(f12008,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),happ(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))))),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f428,f500])).
fof(f12053,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),'const.combin.o_3'(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2))),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f12009,f502])).
fof(f12009,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),happ(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2))))),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f428,f500])).
fof(f11932,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(X33,bool),sK8('?_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f11870,f401])).
fof(f11870,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,'|dollar|exists_1'(s(fun(X33,bool),sK8('?_1',fun(X33,bool))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1409,f11606])).
fof(f11808,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(X68,happ(s(fun('type.bvpSem.state'(X65),X68),X69),s('type.bvpSem.state'(X65),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X66),s('type.bvpSem.state'(X65),X67))))) = s(X68,'const.combin.o_3'(s(fun('type.bvpSem.state'(X65),X68),X69),s(fun('type.bvpSem.state'(X65),'type.bvpSem.state'(X65)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X66))),s('type.bvpSem.state'(X65),X67)))) )),
  inference(superposition,[],[f502,f427])).
fof(f11791,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.bvpSem.state'(X3),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4),s('type.bvpSem.state'(X3),happ(s(fun(X5,'type.bvpSem.state'(X3)),X6),s(X5,X7))))) = s('type.bvpSem.state'(X3),'const.combin.o_3'(s(fun('type.bvpSem.state'(X3),'type.bvpSem.state'(X3)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4))),s(fun(X5,'type.bvpSem.state'(X3)),X6),s(X5,X7)))) )),
  inference(superposition,[],[f427,f502])).
fof(f11829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),'_dst_x0x0x1_3'(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))))),s('type.bvpSem.state'(X0),X2))))) )),
  inference(forward_demodulation,[],[f11768,f508])).
fof(f11768,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))))),s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f427,f500])).
fof(f11828,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),'const.combin.o_3'(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2))),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f11769,f502])).
fof(f11769,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.bvpSem.state'(X0),bool),happ(s(fun('type.bvpSem.state'(X0),fun('type.bvpSem.state'(X0),bool)),'=_0'),s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2))))),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f427,f500])).
fof(f11606,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),sK8('?_1',fun(X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f10683,f2467])).
fof(f11654,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(X1,bool),sK8('?_1',fun(X1,bool))),s(X1,sK8(sK8('?_1',fun(X1,bool)),X1)))))) )),
  inference(forward_demodulation,[],[f11612,f502])).
fof(f11612,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(X1,bool),sK8('?_1',fun(X1,bool))),s(X1,sK8(sK8('?_1',fun(X1,bool)),X1)))))))) )),
  inference(unit_resulting_resolution,[],[f10683,f483])).
fof(f11618,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),sK8('?_1',fun(X0,bool))),s(X0,sK8(sK8('?_1',fun(X0,bool)),X0))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f10683,f1479])).
fof(f11619,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(X1,bool),sK8('?_1',fun(X1,bool))),s(X1,sK8(sK8('?_1',fun(X1,bool)),X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f10683,f6760])).
fof(f11466,plain,(
  ( ! [X584,X580,X582,X585,X581,X583,X579] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X579,bool),'const.combin.o_3'(s(fun(X580,fun(X579,bool)),X581),s(fun(X582,X580),X583),s(X582,X584))))) | s(X579,X585) = s(X579,sK5(X585,happ(s(fun(X580,fun(X579,bool)),X581),s(X580,happ(s(fun(X582,X580),X583),s(X582,X584)))),X579))) )),
  inference(superposition,[],[f807,f502])).
fof(f11603,plain,(
  ( ! [X573,X575,X576,X578,X572,X574,X577] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X572,bool),bool),'?_1'),s(fun(X573,fun(X572,bool)),X574),s(X573,happ(s(fun(X575,X573),X576),s(X575,X577)))))) | ~p(s(bool,happ(s(fun(X572,bool),'const.combin.o_3'(s(fun(X573,fun(X572,bool)),X574),s(fun(X575,X573),X576),s(X575,X577))),s(X572,X578))))) )),
  inference(forward_demodulation,[],[f11465,f502])).
fof(f11465,plain,(
  ( ! [X573,X575,X576,X578,X572,X574,X577] : (~p(s(bool,happ(s(fun(X572,bool),'const.combin.o_3'(s(fun(X573,fun(X572,bool)),X574),s(fun(X575,X573),X576),s(X575,X577))),s(X572,X578)))) | p(s(bool,happ(s(fun(fun(X572,bool),bool),'?_1'),s(fun(X572,bool),happ(s(fun(X573,fun(X572,bool)),X574),s(X573,happ(s(fun(X575,X573),X576),s(X575,X577))))))))) )),
  inference(superposition,[],[f480,f502])).
fof(f11602,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (p(s(bool,happ(s(fun(X566,bool),'const.combin.o_3'(s(fun(X567,fun(X566,bool)),X568),s(fun(X569,X567),X570),s(X569,X571))),s(X566,sK8(happ(s(fun(X567,fun(X566,bool)),X568),s(X567,happ(s(fun(X569,X567),X570),s(X569,X571)))),X566))))) | ~p(s(bool,happ(s(fun(fun(X566,bool),bool),'?_1'),s(fun(X566,bool),'const.combin.o_3'(s(fun(X567,fun(X566,bool)),X568),s(fun(X569,X567),X570),s(X569,X571))))))) )),
  inference(forward_demodulation,[],[f11464,f502])).
fof(f11464,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (~p(s(bool,happ(s(fun(fun(X566,bool),bool),'?_1'),s(fun(X566,bool),'const.combin.o_3'(s(fun(X567,fun(X566,bool)),X568),s(fun(X569,X567),X570),s(X569,X571)))))) | p(s(bool,happ(s(fun(X566,bool),happ(s(fun(X567,fun(X566,bool)),X568),s(X567,happ(s(fun(X569,X567),X570),s(X569,X571))))),s(X566,sK8(happ(s(fun(X567,fun(X566,bool)),X568),s(X567,happ(s(fun(X569,X567),X570),s(X569,X571)))),X566)))))) )),
  inference(superposition,[],[f479,f502])).
fof(f11601,plain,(
  ( ! [X565,X561,X563,X564,X560,X562] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X560,bool),bool),'!_1'),s(fun(X561,fun(X560,bool)),X562),s(X561,happ(s(fun(X563,X561),X564),s(X563,X565)))))) | ~p(s(bool,happ(s(fun(X560,bool),'const.combin.o_3'(s(fun(X561,fun(X560,bool)),X562),s(fun(X563,X561),X564),s(X563,X565))),s(X560,sK7(happ(s(fun(X561,fun(X560,bool)),X562),s(X561,happ(s(fun(X563,X561),X564),s(X563,X565)))),X560)))))) )),
  inference(forward_demodulation,[],[f11463,f502])).
fof(f11463,plain,(
  ( ! [X565,X561,X563,X564,X560,X562] : (~p(s(bool,happ(s(fun(X560,bool),'const.combin.o_3'(s(fun(X561,fun(X560,bool)),X562),s(fun(X563,X561),X564),s(X563,X565))),s(X560,sK7(happ(s(fun(X561,fun(X560,bool)),X562),s(X561,happ(s(fun(X563,X561),X564),s(X563,X565)))),X560))))) | p(s(bool,happ(s(fun(fun(X560,bool),bool),'!_1'),s(fun(X560,bool),happ(s(fun(X561,fun(X560,bool)),X562),s(X561,happ(s(fun(X563,X561),X564),s(X563,X565))))))))) )),
  inference(superposition,[],[f478,f502])).
fof(f11599,plain,(
  ( ! [X546,X549,X551,X547,X552,X548,X550] : (p(s(bool,'|dollar|exists_1'(s(fun(X546,bool),'const.combin.o_3'(s(fun(X547,fun(X546,bool)),X548),s(fun(X549,X547),X550),s(X549,X551)))))) | ~p(s(bool,happ(s(fun(X546,bool),'const.combin.o_3'(s(fun(X547,fun(X546,bool)),X548),s(fun(X549,X547),X550),s(X549,X551))),s(X546,sK5(X552,happ(s(fun(X547,fun(X546,bool)),X548),s(X547,happ(s(fun(X549,X547),X550),s(X549,X551)))),X546)))))) )),
  inference(forward_demodulation,[],[f11461,f502])).
fof(f11461,plain,(
  ( ! [X546,X549,X551,X547,X552,X548,X550] : (~p(s(bool,happ(s(fun(X546,bool),'const.combin.o_3'(s(fun(X547,fun(X546,bool)),X548),s(fun(X549,X547),X550),s(X549,X551))),s(X546,sK5(X552,happ(s(fun(X547,fun(X546,bool)),X548),s(X547,happ(s(fun(X549,X547),X550),s(X549,X551)))),X546))))) | p(s(bool,'|dollar|exists_1'(s(fun(X546,bool),happ(s(fun(X547,fun(X546,bool)),X548),s(X547,happ(s(fun(X549,X547),X550),s(X549,X551))))))))) )),
  inference(superposition,[],[f434,f502])).
fof(f11460,plain,(
  ( ! [X540,X542,X545,X541,X543,X539,X544] : (p(s(bool,'|dollar|exists_1'(s(fun(X539,bool),'const.combin.o_3'(s(fun(X540,fun(X539,bool)),X541),s(fun(X542,X540),X543),s(X542,X544)))))) | s(X539,X545) = s(X539,sK5(X545,happ(s(fun(X540,fun(X539,bool)),X541),s(X540,happ(s(fun(X542,X540),X543),s(X542,X544)))),X539))) )),
  inference(superposition,[],[f433,f502])).
fof(f11440,plain,(
  ( ! [X408,X405,X407,X406] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X405,bool),'const.combin.o_3'(s(fun(X405,fun(X405,bool)),'=_0'),s(fun(X406,X405),X407),s(X406,X408)))))) )),
  inference(superposition,[],[f8190,f502])).
fof(f11439,plain,(
  ( ! [X403,X404,X402] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X402,bool),X403),s(X402,X404)))))) )),
  inference(superposition,[],[f3350,f502])).
fof(f11421,plain,(
  ( ! [X302,X304,X306,X303,X305,X307] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X307,bool),'_dst_x0_1'(s(bool,'const.combin.o_3'(s(fun(X302,bool),X303),s(fun(X304,X302),X305),s(X304,X306))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X302,bool),X303),s(X302,happ(s(fun(X304,X302),X305),s(X304,X306)))))) )),
  inference(superposition,[],[f1482,f502])).
fof(f11414,plain,(
  ( ! [X261,X263,X265,X262,X264] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.combin.o_3'(s(fun(X261,bool),X262),s(fun(X263,X261),X264),s(X263,X265))))) | s(bool,'F_0') = s(bool,happ(s(fun(X261,bool),X262),s(X261,happ(s(fun(X263,X261),X264),s(X263,X265)))))) )),
  inference(superposition,[],[f571,f502])).
fof(f11400,plain,(
  ( ! [X182,X184,X183,X181] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(X181,bool),bool),'?_1'),s(fun(X182,fun(X181,bool)),X183),s(X182,X184)))) | p(s(bool,happ(s(fun(X181,bool),happ(s(fun(X182,fun(X181,bool)),X183),s(X182,X184))),s(X181,sK8(happ(s(fun(X182,fun(X181,bool)),X183),s(X182,X184)),X181)))))) )),
  inference(superposition,[],[f479,f502])).
fof(f11396,plain,(
  ( ! [X167,X169,X166,X168] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X166))),s(fun(X167,bool),X168),s(X167,X169)))) | p(s(bool,X166))) )),
  inference(superposition,[],[f482,f502])).
fof(f11393,plain,(
  ( ! [X154,X152,X151,X155,X153] : (~p(s(bool,'const.combin.o_3'(s(fun(X151,bool),X152),s(fun(X153,X151),X154),s(X153,X155)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X151,bool),X152)))) )),
  inference(superposition,[],[f2467,f502])).
fof(f11392,plain,(
  ( ! [X146,X149,X147,X150,X148] : (~p(s(bool,'const.combin.o_3'(s(fun(X146,bool),X147),s(fun(X148,X146),X149),s(X148,X150)))) | p(s(bool,happ(s(fun(fun(X146,bool),bool),'?_1'),s(fun(X146,bool),X147))))) )),
  inference(superposition,[],[f480,f502])).
fof(f11370,plain,(
  ( ! [X134,X136,X135,X133] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X133),s('type.num.num',happ(s(fun(X134,'type.num.num'),X135),s(X134,X136))))) = s('type.num.num','const.combin.o_3'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X133))),s(fun(X134,'type.num.num'),X135),s(X134,X136)))) )),
  inference(superposition,[],[f502,f471])).
fof(f11369,plain,(
  ( ! [X132,X130,X131] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',happ(s(fun(X130,'type.bvlSem.v'),X131),s(X130,X132))))) = s('type.num.num','const.combin.o_3'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s(fun(X130,'type.bvlSem.v'),X131),s(X130,X132)))) )),
  inference(superposition,[],[f502,f406])).
fof(f11368,plain,(
  ( ! [X127,X128,X129] : (s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',happ(s(fun(X127,'type.bvp.prog'),X128),s(X127,X129))))) = s('type.num.num','const.combin.o_3'(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s(fun(X127,'type.bvp.prog'),X128),s(X127,X129)))) )),
  inference(superposition,[],[f502,f405])).
fof(f11367,plain,(
  ( ! [X125,X126,X124] : (s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',happ(s(fun(X124,'type.bvpSem.stack'),X125),s(X124,X126))))) = s('type.num.num','const.combin.o_3'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s(fun(X124,'type.bvpSem.stack'),X125),s(X124,X126)))) )),
  inference(superposition,[],[f502,f404])).
fof(f11353,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,X49) = s(bool,'const.combin.o_3'(s(fun(X44,bool),X45),s(fun(X46,X44),X47),s(X46,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48)))))) )),
  inference(superposition,[],[f502,f555])).
fof(f11352,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(X38,bool),X39),s(fun(X40,X38),X41),s(X40,X42))) | s(bool,X43) = s(bool,happ(s(fun(X38,bool),X39),s(X38,happ(s(fun(X40,X38),X41),s(X40,X42))))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f502,f555])).
fof(f11351,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(X33,bool),X34),s(fun(X35,X33),X36),s(X35,X37))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),X34),s(X33,happ(s(fun(X35,X33),X36),s(X35,X37)))))) )),
  inference(superposition,[],[f502,f407])).
fof(f11350,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(bool,bool),'_i_0_0'),s(fun(X30,bool),X31),s(X30,X32))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) )),
  inference(superposition,[],[f502,f571])).
fof(f11348,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(X22,'const.pair.FST_1'(s('type.pair.prod'(X22,X23),happ(s(fun(X24,'type.pair.prod'(X22,X23)),X25),s(X24,X26))))) = s(X22,'const.combin.o_3'(s(fun('type.pair.prod'(X22,X23),X22),'const.pair.FST_0'),s(fun(X24,'type.pair.prod'(X22,X23)),X25),s(X24,X26)))) )),
  inference(superposition,[],[f502,f484])).
fof(f11347,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(X16,happ(s(fun(X17,X16),happ(s(fun(X19,fun(X17,X16)),X20),s(X19,X21))),s(X17,X18))) = s(X16,'const.combin.o_3'(s(fun(fun(X17,X16),X16),'_dstsel_0_x0_1'(s(X17,X18))),s(fun(X19,fun(X17,X16)),X20),s(X19,X21)))) )),
  inference(superposition,[],[f502,f491])).
fof(f11346,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X4,X6) = s(X4,'const.combin.o_3'(s(fun(X5,X4),'_dst_x0_1'(s(X4,X6))),s(fun(X7,X5),X8),s(X7,X9)))) )),
  inference(superposition,[],[f502,f487])).
fof(f11345,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,'const.combin.o_3'(s(fun(X0,X0),'_i_0_0'),s(fun(X1,X0),X2),s(X1,X3)))) )),
  inference(superposition,[],[f502,f408])).
fof(f11344,plain,(
  ( ! [X142,X140,X143,X141] : (s(X142,happ(s(fun('type.num.num',X142),X143),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X140),s('type.num.num',X141))))) = s(X142,'const.combin.o_3'(s(fun('type.num.num',X142),X143),s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X140))),s('type.num.num',X141)))) )),
  inference(superposition,[],[f502,f471])).
fof(f11343,plain,(
  ( ! [X138,X139,X137] : (s(X138,happ(s(fun('type.num.num',X138),X139),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X137))))) = s(X138,'const.combin.o_3'(s(fun('type.num.num',X138),X139),s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X137)))) )),
  inference(superposition,[],[f502,f406])).
fof(f11342,plain,(
  ( ! [X134,X136,X135] : (s(X135,happ(s(fun('type.num.num',X135),X136),s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X134))))) = s(X135,'const.combin.o_3'(s(fun('type.num.num',X135),X136),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X134)))) )),
  inference(superposition,[],[f502,f405])).
fof(f11341,plain,(
  ( ! [X132,X133,X131] : (s(X132,happ(s(fun('type.num.num',X132),X133),s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X131))))) = s(X132,'const.combin.o_3'(s(fun('type.num.num',X132),X133),s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X131)))) )),
  inference(superposition,[],[f502,f404])).
fof(f11327,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(X55,happ(s(fun(bool,X55),X56),s(bool,X54))) = s(X55,'const.combin.o_3'(s(fun(bool,X55),X56),s(fun(X51,bool),X52),s(X51,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun(X51,bool),X52),s(X51,X53)))) )),
  inference(superposition,[],[f502,f555])).
fof(f11326,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(X48,happ(s(fun(bool,X48),X49),s(bool,'T_0'))) = s(X48,'const.combin.o_3'(s(fun(bool,X48),X49),s(fun(X45,bool),X46),s(X45,X47))) | s(bool,X50) = s(bool,happ(s(fun(X45,bool),X46),s(X45,X47))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f502,f555])).
fof(f11325,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(X43,happ(s(fun(bool,X43),X44),s(bool,'F_0'))) = s(X43,'const.combin.o_3'(s(fun(bool,X43),X44),s(fun(X40,bool),X41),s(X40,X42))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) )),
  inference(superposition,[],[f502,f407])).
fof(f11324,plain,(
  ( ! [X39,X37,X38] : (s(X38,happ(s(fun(bool,X38),X39),s(bool,'T_0'))) = s(X38,'const.combin.o_3'(s(fun(bool,X38),X39),s(fun(fun(X37,bool),bool),'?_1'),s(fun(X37,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | $spl3),
  inference(superposition,[],[f502,f10431])).
fof(f11323,plain,(
  ( ! [X35,X36] : (s(X35,happ(s(fun(bool,X35),X36),s(bool,'T_0'))) = s(X35,'const.combin.o_3'(s(fun(bool,X35),X36),s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f502,f6793])).
fof(f11322,plain,(
  ( ! [X33,X34,X32] : (s(X33,happ(s(fun(bool,X33),X34),s(bool,'T_0'))) = s(X33,'const.combin.o_3'(s(fun(bool,X33),X34),s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f502,f571])).
fof(f11531,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | ~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(backward_demodulation,[],[f11320,f8571])).
fof(f11320,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(X27,happ(s(fun(X24,X27),X28),s(X24,'const.pair.FST_1'(s('type.pair.prod'(X24,X25),X26))))) = s(X27,'const.combin.o_3'(s(fun(X24,X27),X28),s(fun('type.pair.prod'(X24,X25),X24),'const.pair.FST_0'),s('type.pair.prod'(X24,X25),X26)))) )),
  inference(superposition,[],[f502,f484])).
fof(f11526,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),X2),s(X1,X3))),s(fun(X1,X0),X2),s(X1,X3))))) )),
  inference(backward_demodulation,[],[f11522,f9737])).
fof(f11525,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X304,bool),bool),'!_1'),s(fun(X305,fun(X304,bool)),X307),s(X305,X306)))) | ~p(s(bool,happ(s(fun(X304,bool),happ(s(fun(X305,fun(X304,bool)),X307),s(X305,X306))),s(X304,sK7(happ(s(fun(fun(X305,fun(X304,bool)),fun(X304,bool)),'_dstsel_0_x0_1'(s(X305,X306))),s(fun(X305,fun(X304,bool)),X307)),X304)))))) )),
  inference(backward_demodulation,[],[f11522,f9776])).
fof(f11524,plain,(
  ( ! [X313,X315,X316,X312,X314] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X312,bool),bool),'?_1'),s(fun(X313,fun(X312,bool)),X315),s(X313,X314)))) | ~p(s(bool,happ(s(fun(X312,bool),happ(s(fun(X313,fun(X312,bool)),X315),s(X313,X314))),s(X312,X316))))) )),
  inference(backward_demodulation,[],[f11522,f9779])).
fof(f11523,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X122))),s(fun(X119,bool),X121),s(X119,X120)))) | ~p(s(bool,happ(s(fun(X119,bool),X121),s(X119,X120))))) )),
  inference(backward_demodulation,[],[f11522,f9758])).
fof(f11318,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X16,happ(s(fun(X10,X16),X17),s(X10,'const.combin.o_3'(s(fun(X11,X10),X12),s(fun(X13,X11),X14),s(X13,X15))))) = s(X16,'const.combin.o_3'(s(fun(X10,X16),X17),s(fun(X11,X10),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f502,f502])).
fof(f11516,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(backward_demodulation,[],[f11317,f8570])).
fof(f11317,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X8,happ(s(fun(X4,X8),X9),s(X4,X6))) = s(X8,'const.combin.o_3'(s(fun(X4,X8),X9),s(fun(X5,X4),'_dst_x0_1'(s(X4,X6))),s(X5,X7)))) )),
  inference(superposition,[],[f502,f487])).
fof(f11508,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(backward_demodulation,[],[f11316,f8569])).
fof(f11316,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f502,f408])).
fof(f11506,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) )),
  inference(forward_demodulation,[],[f11220,f502])).
fof(f11220,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))),s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f502,f500])).
fof(f11505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))))) )),
  inference(forward_demodulation,[],[f11221,f502])).
fof(f11221,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))),s(X0,'const.combin.o_3'(s(fun(X1,X0),X2),s(fun(X3,X1),X4),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f502,f500])).
fof(f11211,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,happ(s(fun(fun(X33,bool),bool),'?_1'),s(fun(X33,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f11146,f401])).
fof(f11146,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X34) = s(bool,happ(s(fun(fun(X33,bool),bool),'?_1'),s(fun(X33,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f1409,f10431])).
fof(f11204,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,X11)))) | p(s(bool,X11))) ) | $spl3),
  inference(forward_demodulation,[],[f11203,f502])).
fof(f11203,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X11))))) ) | $spl3),
  inference(subsumption_resolution,[],[f11134,f401])).
fof(f11134,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X10,bool),bool),'?_1'),s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X11))))) ) | $spl3),
  inference(superposition,[],[f481,f10431])).
fof(f11105,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f10431,f555])).
fof(f11196,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f11103,f502])).
fof(f11103,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10431,f500])).
fof(f11195,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f11104,f502])).
fof(f11104,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,'T_0'))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f10431,f500])).
fof(f10752,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK10(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK10(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f492,f555])).
fof(f10751,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK10(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK10(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f492,f555])).
fof(f10750,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK10(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK10(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f492,f407])).
fof(f10737,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK10(X7,X8,X9,bool))) )),
  inference(superposition,[],[f492,f555])).
fof(f10736,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK10(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f492,f555])).
fof(f10735,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK10(X0,X1,X2,bool))) )),
  inference(superposition,[],[f492,f407])).
fof(f10721,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f10690,f502])).
fof(f10690,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f10630,f481])).
fof(f10720,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('?_1',fun(X1,bool))))))) )),
  inference(forward_demodulation,[],[f10691,f502])).
fof(f10691,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('?_1',fun(X1,bool))))))))) )),
  inference(unit_resulting_resolution,[],[f10630,f483])).
fof(f10697,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK8('?_1',fun(X0,bool)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f10630,f1479])).
fof(f10698,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),sK8('?_1',fun(X1,bool))))))))))) )),
  inference(unit_resulting_resolution,[],[f10630,f6760])).
fof(f10669,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f10637,f502])).
fof(f10637,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f10418,f481])).
fof(f10668,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'?_1'))))) )),
  inference(forward_demodulation,[],[f10638,f502])).
fof(f10638,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'?_1'))))))) )),
  inference(unit_resulting_resolution,[],[f10418,f483])).
fof(f10644,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1')))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f10418,f1479])).
fof(f10645,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(fun(X1,bool),bool),bool),'?_1'),s(fun(fun(X1,bool),bool),'?_1'))))))))) )),
  inference(unit_resulting_resolution,[],[f10418,f6760])).
fof(f10576,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(fun(X34,bool),bool),'?_1'))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f10516,f401])).
fof(f10516,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(fun(X34,bool),bool),'?_1'))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1409,f10419])).
fof(f10465,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f10424,f502])).
fof(f10424,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,'F_0'))))) )),
  inference(unit_resulting_resolution,[],[f403,f10278,f481])).
fof(f10464,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(forward_demodulation,[],[f10425,f502])).
fof(f10425,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f10278,f483])).
fof(f10431,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f10278,f1479])).
fof(f10432,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(X1,bool),bool),'?_1'),s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f10278,f6760])).
fof(f10327,plain,(
  ( ! [X123,X124,X122] : (~p(s(bool,X122)) | p(s(bool,happ(s(fun(fun(X124,bool),bool),'?_1'),s(fun(X124,bool),'_dst_x0_1'(s(bool,sK5(X122,X123,bool))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X123)))) )),
  inference(superposition,[],[f6760,f807])).
fof(f10341,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,happ(s(fun(fun(X69,bool),bool),'?_1'),s(fun(X69,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))))))) | s(X67,X70) = s(X67,sK5(X70,X68,X67))) )),
  inference(subsumption_resolution,[],[f10305,f401])).
fof(f10305,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X69,bool),bool),'?_1'),s(fun(X69,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))))))) | s(X67,X70) = s(X67,sK5(X70,X68,X67))) )),
  inference(superposition,[],[f6760,f807])).
fof(f10340,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,happ(s(fun(fun(X66,bool),bool),'?_1'),s(fun(X66,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X64,bool),'_dst_x0_1'(s(bool,X65)))))))))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f10304,f401])).
fof(f10304,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X66,bool),bool),'?_1'),s(fun(X66,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X64,bool),'_dst_x0_1'(s(bool,X65)))))))))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f6760,f1482])).
fof(f10328,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X3,bool),'_dst_x0_1'(s(bool,X2)))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f10280,f401])).
fof(f10280,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X3,bool),'_dst_x0_1'(s(bool,X2)))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6760,f555])).
fof(f10272,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X55) = s(bool,'|dollar|exists_1'(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),'=_0'),s(X53,X54))))) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f10200,f401])).
fof(f10200,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X55) = s(bool,'|dollar|exists_1'(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),'=_0'),s(X53,X54))))) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f1409,f8190])).
fof(f10127,plain,(
  ( ! [X57,X56] : (s(bool,X56) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X56))))) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f8190,f555])).
fof(f10050,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f10026,f401])).
fof(f10026,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f9993,f555])).
fof(f10015,plain,(
  ( ! [X4] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f9993,f407])).
fof(f10011,plain,(
  s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f403,f9993,f473])).
fof(f9998,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f9971,f401])).
fof(f9971,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f7534,f555])).
fof(f9993,plain,(
  ~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0')))) | $spl3),
  inference(forward_demodulation,[],[f9940,f502])).
fof(f9940,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f7534,f477])).
fof(f9991,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))),s(bool,X0))))) ) | $spl3),
  inference(forward_demodulation,[],[f9942,f502])).
fof(f9942,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))),s(bool,X0))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f7534,f482])).
fof(f9948,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f403,f7534,f473])).
fof(f9889,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f9865,f401])).
fof(f9865,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f9836,f555])).
fof(f9856,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f9836,f555])).
fof(f9854,plain,(
  s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f403,f9836,f473])).
fof(f9841,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f9814,f401])).
fof(f9814,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f9101,f555])).
fof(f9803,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f9101,f407])).
fof(f9793,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0')))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f9101,f555])).
fof(f9836,plain,(
  ~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))))),
  inference(forward_demodulation,[],[f9783,f502])).
fof(f9783,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f9101,f477])).
fof(f9834,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(forward_demodulation,[],[f9785,f502])).
fof(f9785,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f9101,f482])).
fof(f9791,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f403,f9101,f473])).
fof(f9704,plain,(
  ( ! [X321,X318,X320,X317,X319] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),X320),s(X318,X319))))) | s(X317,X321) = s(X317,sK5(X321,happ(s(fun(fun(X318,fun(X317,bool)),fun(X317,bool)),'_dstsel_0_x0_1'(s(X318,X319))),s(fun(X318,fun(X317,bool)),X320)),X317))) )),
  inference(superposition,[],[f807,f491])).
fof(f9779,plain,(
  ( ! [X313,X315,X316,X312,X314] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X312,bool),bool),'?_1'),s(fun(fun(X313,fun(X312,bool)),fun(X312,bool)),'_dstsel_0_x0_1'(s(X313,X314))),s(fun(X313,fun(X312,bool)),X315)))) | ~p(s(bool,happ(s(fun(X312,bool),happ(s(fun(X313,fun(X312,bool)),X315),s(X313,X314))),s(X312,X316))))) )),
  inference(forward_demodulation,[],[f9703,f502])).
fof(f9703,plain,(
  ( ! [X313,X315,X316,X312,X314] : (~p(s(bool,happ(s(fun(X312,bool),happ(s(fun(X313,fun(X312,bool)),X315),s(X313,X314))),s(X312,X316)))) | p(s(bool,happ(s(fun(fun(X312,bool),bool),'?_1'),s(fun(X312,bool),happ(s(fun(fun(X313,fun(X312,bool)),fun(X312,bool)),'_dstsel_0_x0_1'(s(X313,X314))),s(fun(X313,fun(X312,bool)),X315))))))) )),
  inference(superposition,[],[f480,f491])).
fof(f9778,plain,(
  ( ! [X308,X310,X309,X311] : (p(s(bool,happ(s(fun(X308,bool),happ(s(fun(X309,fun(X308,bool)),X311),s(X309,X310))),s(X308,sK8(happ(s(fun(fun(X309,fun(X308,bool)),fun(X308,bool)),'_dstsel_0_x0_1'(s(X309,X310))),s(fun(X309,fun(X308,bool)),X311)),X308))))) | ~p(s(bool,'const.combin.o_3'(s(fun(fun(X308,bool),bool),'?_1'),s(fun(X309,fun(X308,bool)),X311),s(X309,X310))))) )),
  inference(forward_demodulation,[],[f9777,f491])).
fof(f9777,plain,(
  ( ! [X308,X310,X309,X311] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(X308,bool),bool),'?_1'),s(fun(X309,fun(X308,bool)),X311),s(X309,X310)))) | p(s(bool,happ(s(fun(X308,bool),happ(s(fun(fun(X309,fun(X308,bool)),fun(X308,bool)),'_dstsel_0_x0_1'(s(X309,X310))),s(fun(X309,fun(X308,bool)),X311))),s(X308,sK8(happ(s(fun(fun(X309,fun(X308,bool)),fun(X308,bool)),'_dstsel_0_x0_1'(s(X309,X310))),s(fun(X309,fun(X308,bool)),X311)),X308)))))) )),
  inference(forward_demodulation,[],[f9702,f502])).
fof(f9702,plain,(
  ( ! [X308,X310,X309,X311] : (~p(s(bool,happ(s(fun(fun(X308,bool),bool),'?_1'),s(fun(X308,bool),happ(s(fun(X309,fun(X308,bool)),X311),s(X309,X310)))))) | p(s(bool,happ(s(fun(X308,bool),happ(s(fun(fun(X309,fun(X308,bool)),fun(X308,bool)),'_dstsel_0_x0_1'(s(X309,X310))),s(fun(X309,fun(X308,bool)),X311))),s(X308,sK8(happ(s(fun(fun(X309,fun(X308,bool)),fun(X308,bool)),'_dstsel_0_x0_1'(s(X309,X310))),s(fun(X309,fun(X308,bool)),X311)),X308)))))) )),
  inference(superposition,[],[f479,f491])).
fof(f9776,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X304,bool),bool),'!_1'),s(fun(fun(X305,fun(X304,bool)),fun(X304,bool)),'_dstsel_0_x0_1'(s(X305,X306))),s(fun(X305,fun(X304,bool)),X307)))) | ~p(s(bool,happ(s(fun(X304,bool),happ(s(fun(X305,fun(X304,bool)),X307),s(X305,X306))),s(X304,sK7(happ(s(fun(fun(X305,fun(X304,bool)),fun(X304,bool)),'_dstsel_0_x0_1'(s(X305,X306))),s(fun(X305,fun(X304,bool)),X307)),X304)))))) )),
  inference(forward_demodulation,[],[f9701,f502])).
fof(f9701,plain,(
  ( ! [X304,X306,X305,X307] : (~p(s(bool,happ(s(fun(X304,bool),happ(s(fun(X305,fun(X304,bool)),X307),s(X305,X306))),s(X304,sK7(happ(s(fun(fun(X305,fun(X304,bool)),fun(X304,bool)),'_dstsel_0_x0_1'(s(X305,X306))),s(fun(X305,fun(X304,bool)),X307)),X304))))) | p(s(bool,happ(s(fun(fun(X304,bool),bool),'!_1'),s(fun(X304,bool),happ(s(fun(fun(X305,fun(X304,bool)),fun(X304,bool)),'_dstsel_0_x0_1'(s(X305,X306))),s(fun(X305,fun(X304,bool)),X307))))))) )),
  inference(superposition,[],[f478,f491])).
fof(f9775,plain,(
  ( ! [X300,X302,X301,X303,X299] : (p(s(bool,happ(s(fun(X299,bool),happ(s(fun(X300,fun(X299,bool)),X302),s(X300,X301))),s(X299,X303)))) | ~p(s(bool,'const.combin.o_3'(s(fun(fun(X299,bool),bool),'!_1'),s(fun(X300,fun(X299,bool)),X302),s(X300,X301))))) )),
  inference(forward_demodulation,[],[f9774,f491])).
fof(f9774,plain,(
  ( ! [X300,X302,X301,X303,X299] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(X299,bool),bool),'!_1'),s(fun(X300,fun(X299,bool)),X302),s(X300,X301)))) | p(s(bool,happ(s(fun(X299,bool),happ(s(fun(fun(X300,fun(X299,bool)),fun(X299,bool)),'_dstsel_0_x0_1'(s(X300,X301))),s(fun(X300,fun(X299,bool)),X302))),s(X299,X303))))) )),
  inference(forward_demodulation,[],[f9700,f502])).
fof(f9700,plain,(
  ( ! [X300,X302,X301,X303,X299] : (~p(s(bool,happ(s(fun(fun(X299,bool),bool),'!_1'),s(fun(X299,bool),happ(s(fun(X300,fun(X299,bool)),X302),s(X300,X301)))))) | p(s(bool,happ(s(fun(X299,bool),happ(s(fun(fun(X300,fun(X299,bool)),fun(X299,bool)),'_dstsel_0_x0_1'(s(X300,X301))),s(fun(X300,fun(X299,bool)),X302))),s(X299,X303))))) )),
  inference(superposition,[],[f477,f491])).
fof(f9773,plain,(
  ( ! [X294,X296,X298,X295,X297] : (p(s(bool,'|dollar|exists_1'(s(fun(X294,bool),happ(s(fun(X295,fun(X294,bool)),X297),s(X295,X296)))))) | ~p(s(bool,happ(s(fun(X294,bool),happ(s(fun(X295,fun(X294,bool)),X297),s(X295,X296))),s(X294,sK5(X298,happ(s(fun(fun(X295,fun(X294,bool)),fun(X294,bool)),'_dstsel_0_x0_1'(s(X295,X296))),s(fun(X295,fun(X294,bool)),X297)),X294)))))) )),
  inference(forward_demodulation,[],[f9699,f491])).
fof(f9699,plain,(
  ( ! [X294,X296,X298,X295,X297] : (~p(s(bool,happ(s(fun(X294,bool),happ(s(fun(X295,fun(X294,bool)),X297),s(X295,X296))),s(X294,sK5(X298,happ(s(fun(fun(X295,fun(X294,bool)),fun(X294,bool)),'_dstsel_0_x0_1'(s(X295,X296))),s(fun(X295,fun(X294,bool)),X297)),X294))))) | p(s(bool,'|dollar|exists_1'(s(fun(X294,bool),happ(s(fun(fun(X295,fun(X294,bool)),fun(X294,bool)),'_dstsel_0_x0_1'(s(X295,X296))),s(fun(X295,fun(X294,bool)),X297))))))) )),
  inference(superposition,[],[f434,f491])).
fof(f9698,plain,(
  ( ! [X292,X290,X293,X289,X291] : (p(s(bool,'|dollar|exists_1'(s(fun(X289,bool),happ(s(fun(X290,fun(X289,bool)),X292),s(X290,X291)))))) | s(X289,X293) = s(X289,sK5(X293,happ(s(fun(fun(X290,fun(X289,bool)),fun(X289,bool)),'_dstsel_0_x0_1'(s(X290,X291))),s(fun(X290,fun(X289,bool)),X292)),X289))) )),
  inference(superposition,[],[f433,f491])).
fof(f9670,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X172,bool),'_dst_x0_1'(s(bool,happ(s(fun(X169,bool),X171),s(X169,X170))))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X169,bool),bool),'_dstsel_0_x0_1'(s(X169,X170))),s(fun(X169,bool),X171)))) )),
  inference(superposition,[],[f1482,f491])).
fof(f9663,plain,(
  ( ! [X144,X142,X143] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X142,bool),X144),s(X142,X143))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X142,bool),bool),'_dstsel_0_x0_1'(s(X142,X143))),s(fun(X142,bool),X144)))) )),
  inference(superposition,[],[f571,f491])).
fof(f9758,plain,(
  ( ! [X121,X122,X120,X119] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X122))),s(fun(fun(X119,bool),bool),'_dstsel_0_x0_1'(s(X119,X120))),s(fun(X119,bool),X121)))) | ~p(s(bool,happ(s(fun(X119,bool),X121),s(X119,X120))))) )),
  inference(forward_demodulation,[],[f9657,f502])).
fof(f9657,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,happ(s(fun(X119,bool),X121),s(X119,X120)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X122))),s(bool,happ(s(fun(fun(X119,bool),bool),'_dstsel_0_x0_1'(s(X119,X120))),s(fun(X119,bool),X121))))))) )),
  inference(superposition,[],[f483,f491])).
fof(f9757,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,happ(s(fun(X111,bool),X113),s(X111,X112)))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X111,bool),X113),s(X111,X112))),s(bool,X114))))) )),
  inference(forward_demodulation,[],[f9756,f491])).
fof(f9756,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X111,bool),X113),s(X111,X112))),s(bool,X114)))) | p(s(bool,happ(s(fun(fun(X111,bool),bool),'_dstsel_0_x0_1'(s(X111,X112))),s(fun(X111,bool),X113))))) )),
  inference(forward_demodulation,[],[f9655,f502])).
fof(f9655,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X111,bool),X113),s(X111,X112))))),s(bool,X114)))) | p(s(bool,happ(s(fun(fun(X111,bool),bool),'_dstsel_0_x0_1'(s(X111,X112))),s(fun(X111,bool),X113))))) )),
  inference(superposition,[],[f482,f491])).
fof(f9649,plain,(
  ( ! [X90,X91] : (~p(s(bool,happ(s(fun(X90,bool),sK7('_dstsel_0_x0_1'(s(X90,X91)),fun(X90,bool))),s(X90,X91)))) | p(s(bool,happ(s(fun(fun(fun(X90,bool),bool),bool),'!_1'),s(fun(fun(X90,bool),bool),'_dstsel_0_x0_1'(s(X90,X91))))))) )),
  inference(superposition,[],[f478,f491])).
fof(f9648,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,happ(s(fun(X87,bool),X89),s(X87,X88)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X87,bool),bool),'_dstsel_0_x0_1'(s(X87,X88)))))) )),
  inference(superposition,[],[f2467,f491])).
fof(f9647,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,happ(s(fun(X84,bool),X86),s(X84,X85)))) | p(s(bool,happ(s(fun(fun(fun(X84,bool),bool),bool),'?_1'),s(fun(fun(X84,bool),bool),'_dstsel_0_x0_1'(s(X84,X85))))))) )),
  inference(superposition,[],[f480,f491])).
fof(f9646,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,happ(s(fun(X81,bool),sK5(X83,'_dstsel_0_x0_1'(s(X81,X82)),fun(X81,bool))),s(X81,X82)))) | p(s(bool,'|dollar|exists_1'(s(fun(fun(X81,bool),bool),'_dstsel_0_x0_1'(s(X81,X82))))))) )),
  inference(superposition,[],[f434,f491])).
fof(f9751,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(X71,X75) != s(X71,happ(s(fun(X72,X71),X74),s(X72,X73))) | p(s(bool,'const.combin.o_3'(s(fun(X71,bool),happ(s(fun(X71,fun(X71,bool)),'=_0'),s(X71,X75))),s(fun(X72,X71),X74),s(X72,X73))))) )),
  inference(forward_demodulation,[],[f9750,f491])).
fof(f9750,plain,(
  ( ! [X74,X72,X71,X75,X73] : (p(s(bool,'const.combin.o_3'(s(fun(X71,bool),happ(s(fun(X71,fun(X71,bool)),'=_0'),s(X71,X75))),s(fun(X72,X71),X74),s(X72,X73)))) | s(X71,X75) != s(X71,happ(s(fun(fun(X72,X71),X71),'_dstsel_0_x0_1'(s(X72,X73))),s(fun(X72,X71),X74)))) )),
  inference(forward_demodulation,[],[f9644,f502])).
fof(f9644,plain,(
  ( ! [X74,X72,X71,X75,X73] : (p(s(bool,happ(s(fun(X71,bool),happ(s(fun(X71,fun(X71,bool)),'=_0'),s(X71,X75))),s(X71,happ(s(fun(X72,X71),X74),s(X72,X73)))))) | s(X71,X75) != s(X71,happ(s(fun(fun(X72,X71),X71),'_dstsel_0_x0_1'(s(X72,X73))),s(fun(X72,X71),X74)))) )),
  inference(superposition,[],[f500,f491])).
fof(f9749,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(X66,X70) != s(X66,happ(s(fun(X67,X66),X69),s(X67,X68))) | p(s(bool,happ(s(fun(X66,bool),'const.combin.o_3'(s(fun(X66,fun(X66,bool)),'=_0'),s(fun(X67,X66),X69),s(X67,X68))),s(X66,X70))))) )),
  inference(forward_demodulation,[],[f9748,f491])).
fof(f9748,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,happ(s(fun(X66,bool),'const.combin.o_3'(s(fun(X66,fun(X66,bool)),'=_0'),s(fun(X67,X66),X69),s(X67,X68))),s(X66,X70)))) | s(X66,X70) != s(X66,happ(s(fun(fun(X67,X66),X66),'_dstsel_0_x0_1'(s(X67,X68))),s(fun(X67,X66),X69)))) )),
  inference(forward_demodulation,[],[f9643,f502])).
fof(f9643,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,happ(s(fun(X66,bool),happ(s(fun(X66,fun(X66,bool)),'=_0'),s(X66,happ(s(fun(X67,X66),X69),s(X67,X68))))),s(X66,X70)))) | s(X66,X70) != s(X66,happ(s(fun(fun(X67,X66),X66),'_dstsel_0_x0_1'(s(X67,X68))),s(fun(X67,X66),X69)))) )),
  inference(superposition,[],[f500,f491])).
fof(f9747,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(X61,X65) = s(X61,happ(s(fun(X62,X61),X64),s(X62,X63))) | ~p(s(bool,'const.combin.o_3'(s(fun(X61,bool),happ(s(fun(X61,fun(X61,bool)),'=_0'),s(X61,X65))),s(fun(X62,X61),X64),s(X62,X63))))) )),
  inference(forward_demodulation,[],[f9746,f491])).
fof(f9746,plain,(
  ( ! [X61,X64,X62,X65,X63] : (~p(s(bool,'const.combin.o_3'(s(fun(X61,bool),happ(s(fun(X61,fun(X61,bool)),'=_0'),s(X61,X65))),s(fun(X62,X61),X64),s(X62,X63)))) | s(X61,X65) = s(X61,happ(s(fun(fun(X62,X61),X61),'_dstsel_0_x0_1'(s(X62,X63))),s(fun(X62,X61),X64)))) )),
  inference(forward_demodulation,[],[f9642,f502])).
fof(f9642,plain,(
  ( ! [X61,X64,X62,X65,X63] : (~p(s(bool,happ(s(fun(X61,bool),happ(s(fun(X61,fun(X61,bool)),'=_0'),s(X61,X65))),s(X61,happ(s(fun(X62,X61),X64),s(X62,X63)))))) | s(X61,X65) = s(X61,happ(s(fun(fun(X62,X61),X61),'_dstsel_0_x0_1'(s(X62,X63))),s(fun(X62,X61),X64)))) )),
  inference(superposition,[],[f499,f491])).
fof(f9745,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(X56,X60) = s(X56,happ(s(fun(X57,X56),X59),s(X57,X58))) | ~p(s(bool,happ(s(fun(X56,bool),'const.combin.o_3'(s(fun(X56,fun(X56,bool)),'=_0'),s(fun(X57,X56),X59),s(X57,X58))),s(X56,X60))))) )),
  inference(forward_demodulation,[],[f9744,f491])).
fof(f9744,plain,(
  ( ! [X59,X57,X60,X58,X56] : (~p(s(bool,happ(s(fun(X56,bool),'const.combin.o_3'(s(fun(X56,fun(X56,bool)),'=_0'),s(fun(X57,X56),X59),s(X57,X58))),s(X56,X60)))) | s(X56,X60) = s(X56,happ(s(fun(fun(X57,X56),X56),'_dstsel_0_x0_1'(s(X57,X58))),s(fun(X57,X56),X59)))) )),
  inference(forward_demodulation,[],[f9641,f502])).
fof(f9641,plain,(
  ( ! [X59,X57,X60,X58,X56] : (~p(s(bool,happ(s(fun(X56,bool),happ(s(fun(X56,fun(X56,bool)),'=_0'),s(X56,happ(s(fun(X57,X56),X59),s(X57,X58))))),s(X56,X60)))) | s(X56,X60) = s(X56,happ(s(fun(fun(X57,X56),X56),'_dstsel_0_x0_1'(s(X57,X58))),s(fun(X57,X56),X59)))) )),
  inference(superposition,[],[f499,f491])).
fof(f9618,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun(X11,bool),X13),s(X11,X12))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X11,bool),bool),'_dstsel_0_x0_1'(s(X11,X12))),s(fun(X11,bool),X13)))) )),
  inference(superposition,[],[f491,f555])).
fof(f9617,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X9),s(X7,X8))) | s(bool,X10) = s(bool,happ(s(fun(fun(X7,bool),bool),'_dstsel_0_x0_1'(s(X7,X8))),s(fun(X7,bool),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f491,f555])).
fof(f9616,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X6),s(X4,X5))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X4,bool),bool),'_dstsel_0_x0_1'(s(X4,X5))),s(fun(X4,bool),X6)))) )),
  inference(superposition,[],[f491,f407])).
fof(f9737,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),X2),s(X1,X3))),s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X3))),s(fun(X1,X0),X2))))) )),
  inference(forward_demodulation,[],[f9736,f502])).
fof(f9736,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))),s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X3))),s(fun(X1,X0),X2))))) )),
  inference(forward_demodulation,[],[f9541,f502])).
fof(f9541,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))),s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X3))),s(fun(X1,X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f491,f500])).
fof(f9735,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3))),s(fun(X1,X0),X3),s(X1,X2))))) )),
  inference(forward_demodulation,[],[f9734,f502])).
fof(f9734,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3))))),s(fun(X1,X0),X3),s(X1,X2))))) )),
  inference(forward_demodulation,[],[f9542,f502])).
fof(f9542,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3))))),s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))))))) )),
  inference(unit_resulting_resolution,[],[f491,f500])).
fof(f9146,plain,(
  ( ! [X111,X112,X113] : (~p(s(bool,X111)) | p(s(bool,X113)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,sK5(X111,X112,bool)))),s(bool,X113)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) )),
  inference(superposition,[],[f481,f807])).
fof(f9187,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94))),s(bool,X95)))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94)))) | p(s(bool,X95))) )),
  inference(forward_demodulation,[],[f9141,f502])).
fof(f9141,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94)))) | p(s(bool,X95)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94))))),s(bool,X95))))) )),
  inference(superposition,[],[f481,f484])).
fof(f9186,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(X89,bool),'_dst_x0_1'(s(bool,X90))),s(X89,X91))),s(bool,X92)))) | ~p(s(bool,X90)) | p(s(bool,X92))) )),
  inference(forward_demodulation,[],[f9140,f502])).
fof(f9140,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X90)) | p(s(bool,X92)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(X89,bool),'_dst_x0_1'(s(bool,X90))),s(X89,X91))))),s(bool,X92))))) )),
  inference(superposition,[],[f481,f487])).
fof(f9185,plain,(
  ( ! [X88,X87] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X87))),s(bool,X88)))) | ~p(s(bool,X87)) | p(s(bool,X88))) )),
  inference(forward_demodulation,[],[f9139,f502])).
fof(f9139,plain,(
  ( ! [X88,X87] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))))),s(bool,X88))))) )),
  inference(superposition,[],[f481,f408])).
fof(f9166,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66))))),s(bool,X67)))) | s(X65,X68) = s(X65,sK5(X68,X66,X65))) )),
  inference(subsumption_resolution,[],[f9127,f401])).
fof(f9127,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66))))),s(bool,X67)))) | s(X65,X68) = s(X65,sK5(X68,X66,X65))) )),
  inference(superposition,[],[f481,f807])).
fof(f9165,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,X64)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X62,bool),'_dst_x0_1'(s(bool,X63))))))),s(bool,X64)))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f9126,f401])).
fof(f9126,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X62,bool),'_dst_x0_1'(s(bool,X63))))))),s(bool,X64)))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f481,f1482])).
fof(f9164,plain,(
  ( ! [X61] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X61)))) | p(s(bool,X61))) )),
  inference(forward_demodulation,[],[f9163,f898])).
fof(f9163,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X60,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X61))))) )),
  inference(subsumption_resolution,[],[f9125,f401])).
fof(f9125,plain,(
  ( ! [X61,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X60,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X61))))) )),
  inference(superposition,[],[f481,f898])).
fof(f9160,plain,(
  ( ! [X56] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))),s(bool,X56)))) | p(s(bool,X56))) ) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f9159,f502])).
fof(f9159,plain,(
  ( ! [X56] : (p(s(bool,X56)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X56))))) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f9123,f401])).
fof(f9123,plain,(
  ( ! [X56] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X56))))) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f481,f6793])).
fof(f9158,plain,(
  ( ! [X54,X55] : (~p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X54))),s(bool,X55)))) | p(s(bool,X55)) | s(bool,'F_0') = s(bool,X54)) )),
  inference(forward_demodulation,[],[f9157,f502])).
fof(f9157,plain,(
  ( ! [X54,X55] : (p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X54))))),s(bool,X55)))) | s(bool,'F_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f9122,f401])).
fof(f9122,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X54))))),s(bool,X55)))) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f481,f571])).
fof(f9147,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f9103,f401])).
fof(f9103,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f481,f555])).
fof(f8985,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK8(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8962,f401])).
fof(f8962,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK8(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f479,f555])).
fof(f8959,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X25)))) | p(s(bool,happ(s(fun(X24,bool),sK5(X25,X26,fun(X24,bool))),s(X24,sK8(sK5(X25,X26,fun(X24,bool)),X24))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X24,bool),bool),X26)))) )),
  inference(superposition,[],[f479,f807])).
fof(f8982,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X22),s(X20,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23)))),X20))))) | ~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X22))))) )),
  inference(forward_demodulation,[],[f8958,f486])).
fof(f8958,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X22)))) | p(s(bool,happ(s(fun(X20,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23))))),s(X20,sK8('const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23)))),X20)))))) )),
  inference(superposition,[],[f479,f486])).
fof(f8981,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),X19),s(X16,sK8('const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))))) | ~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X19))))) )),
  inference(forward_demodulation,[],[f8957,f488])).
fof(f8957,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X19)))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19))))),s(X16,sK8('const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16)))))) )),
  inference(superposition,[],[f479,f488])).
fof(f8980,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK8('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,'const.combin.o_3'(s(fun(fun(X12,bool),bool),'?_1'),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))))) )),
  inference(forward_demodulation,[],[f8979,f437])).
fof(f8979,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(X12,bool),bool),'?_1'),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))) | p(s(bool,happ(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s(X12,sK8('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f8956,f502])).
fof(f8956,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | p(s(bool,happ(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s(X12,sK8('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(superposition,[],[f479,f437])).
fof(f8978,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK8('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10))))) )),
  inference(forward_demodulation,[],[f8955,f411])).
fof(f8955,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10)))) | p(s(bool,happ(s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))),s(X9,sK8('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(superposition,[],[f479,f411])).
fof(f8977,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK8(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(forward_demodulation,[],[f8954,f484])).
fof(f8954,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK8(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(superposition,[],[f479,f484])).
fof(f8976,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK8(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | ~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4))))) )),
  inference(forward_demodulation,[],[f8953,f487])).
fof(f8953,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4)))) | p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))),s(X2,sK8(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(superposition,[],[f479,f487])).
fof(f8975,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f8952,f408])).
fof(f8952,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))),s(X0,sK8(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(superposition,[],[f479,f408])).
fof(f8568,plain,(
  ( ! [X38] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X38),sK7('const.pair.FST_0','type.pair.prod'(bool,X38)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X38),bool),bool),'!_1'),s(fun('type.pair.prod'(bool,X38),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f478,f484])).
fof(f8586,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8554,f401])).
fof(f8554,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f478,f555])).
fof(f8540,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK7(X3,bool))) )),
  inference(superposition,[],[f478,f555])).
fof(f8539,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK7(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f478,f555])).
fof(f8538,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK7(X0,bool))) )),
  inference(superposition,[],[f478,f407])).
fof(f8537,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun(X24,bool),X25),s(X24,sK7(sK5(X25,X26,fun(X24,bool)),X24))))) | p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),sK5(X25,X26,fun(X24,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X24,bool),bool),X26)))) )),
  inference(superposition,[],[f478,f807])).
fof(f8576,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),X22)))) | ~p(s(bool,happ(s(fun(X20,bool),X22),s(X20,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23)))),X20)))))) )),
  inference(forward_demodulation,[],[f8536,f486])).
fof(f8536,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X22),s(X20,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23)))),X20))))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X20,bool),X21),'const.pair.,_2'(s(fun(X20,bool),X22),s(X21,X23))))))))) )),
  inference(superposition,[],[f478,f486])).
fof(f8575,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),X19)))) | ~p(s(bool,happ(s(fun(X16,bool),X19),s(X16,sK7('const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f8535,f488])).
fof(f8535,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X19),s(X16,sK7('const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))))) | p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19))))))))) )),
  inference(superposition,[],[f478,f488])).
fof(f8574,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X12,bool),bool),'!_1'),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f8573,f502])).
fof(f8573,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f8534,f437])).
fof(f8534,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))))) )),
  inference(superposition,[],[f478,f437])).
fof(f8572,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f8533,f411])).
fof(f8533,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f478,f411])).
fof(f8571,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X6,bool),bool),'!_1'),s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)))) | ~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f8532,f502])).
fof(f8532,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(superposition,[],[f478,f484])).
fof(f8570,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)))) | ~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(forward_demodulation,[],[f8531,f502])).
fof(f8531,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))))))) )),
  inference(superposition,[],[f478,f487])).
fof(f8569,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X0,bool),bool),'!_1'),s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f8530,f502])).
fof(f8530,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f478,f408])).
fof(f8373,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.o_3'(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f8293,f1479])).
fof(f8293,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f8189,f502])).
fof(f8189,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f8043,f480])).
fof(f8190,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1)))))) )),
  inference(unit_resulting_resolution,[],[f8043,f2467])).
fof(f8290,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),'=_0'),s(X1,X2))),s(X1,X2))))) )),
  inference(forward_demodulation,[],[f8195,f502])).
fof(f8195,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),'=_0'),s(X1,X2))),s(X1,X2))))))) )),
  inference(unit_resulting_resolution,[],[f8043,f483])).
fof(f8174,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(subsumption_resolution,[],[f8018,f403])).
fof(f8018,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(superposition,[],[f500,f407])).
fof(f7993,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X100))),s(bool,'T_0')))) | s(bool,X100) != s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(X98,X101) = s(X98,sK5(X101,X99,X98))) )),
  inference(superposition,[],[f500,f807])).
fof(f7992,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X97))),s(bool,'T_0')))) | s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun(X95,bool),'_dst_x0_1'(s(bool,X96))))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f500,f1482])).
fof(f8148,plain,(
  ( ! [X92,X91] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X92))),s(bool,X91)))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(forward_demodulation,[],[f7990,f898])).
fof(f7990,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X92))),s(bool,X91)))) | s(bool,X92) != s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f500,f1535])).
fof(f7986,plain,(
  ( ! [X83,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X84))),s(bool,X83)))) | s(bool,'F_0') != s(bool,X84) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f500,f407])).
fof(f7982,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X72))),s(bool,X71)))) | s(bool,'T_0') != s(bool,X72) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f500,f555])).
fof(f7971,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,X41)))) | s(bool,X40) != s(bool,X42) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f500,f555])).
fof(f7970,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X38))),s(bool,'T_0')))) | s(bool,X37) != s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f500,f555])).
fof(f7969,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X36))),s(bool,'F_0')))) | s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f500,f407])).
fof(f7968,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X34))),s(X31,X32)))) | s(X31,X34) != s(X31,sK5(X32,X33,X31)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X33)))) )),
  inference(superposition,[],[f500,f807])).
fof(f8138,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f8137,f437])).
fof(f8137,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7965,f508])).
fof(f7965,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f500,f437])).
fof(f7936,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X100)))) | s(bool,X100) != s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(X98,X101) = s(X98,sK5(X101,X99,X98))) )),
  inference(superposition,[],[f500,f807])).
fof(f7935,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X97)))) | s(bool,X97) != s(bool,'|dollar|exists_1'(s(fun(X95,bool),'_dst_x0_1'(s(bool,X96))))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f500,f1482])).
fof(f8107,plain,(
  ( ! [X92,X91] : (s(bool,'T_0') != s(bool,X92) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X92)))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(forward_demodulation,[],[f7933,f898])).
fof(f7933,plain,(
  ( ! [X92,X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X92)))) | s(bool,X92) != s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f500,f1535])).
fof(f7929,plain,(
  ( ! [X83,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X83))),s(bool,X84)))) | s(bool,'F_0') != s(bool,X84) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f500,f407])).
fof(f7925,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X71))),s(bool,X72)))) | s(bool,'T_0') != s(bool,X72) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f500,f555])).
fof(f7914,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X41))),s(bool,X42)))) | s(bool,X40) != s(bool,X42) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f500,f555])).
fof(f7913,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,X37) != s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f500,f555])).
fof(f7912,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X36)))) | s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f500,f407])).
fof(f7911,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X32))),s(X31,X34)))) | s(X31,X34) != s(X31,sK5(X32,X33,X31)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X33)))) )),
  inference(superposition,[],[f500,f807])).
fof(f8097,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun(X16,bool),'const.combin.o_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s('type.num.num',X19))),s(X16,X20))))) )),
  inference(forward_demodulation,[],[f8096,f437])).
fof(f8096,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),'const.combin.o_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s('type.num.num',X19))),s(X16,X20)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7908,f502])).
fof(f7908,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f500,f437])).
fof(f8091,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f7844,f508])).
fof(f7844,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f408,f500])).
fof(f8090,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X2,X0),'_dst_x0_1'(s(X0,X1))),s(X2,X3))))) )),
  inference(forward_demodulation,[],[f7845,f502])).
fof(f7845,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X2,X0),'_dst_x0_1'(s(X0,X1))),s(X2,X3))))))) )),
  inference(unit_resulting_resolution,[],[f487,f500])).
fof(f8089,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))) )),
  inference(forward_demodulation,[],[f7846,f502])).
fof(f7846,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(unit_resulting_resolution,[],[f484,f500])).
fof(f8088,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.bvpSem.stack',bool),'_dst_x0x0x1_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))),s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'))),s('type.bvpSem.stack',X0))))) )),
  inference(forward_demodulation,[],[f7847,f508])).
fof(f7847,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))))),s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0))))))) )),
  inference(unit_resulting_resolution,[],[f404,f500])).
fof(f8087,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.bvp.prog',bool),'_dst_x0x0x1_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.bvp.prog',X0))))) )),
  inference(forward_demodulation,[],[f7848,f508])).
fof(f7848,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))))),s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0))))))) )),
  inference(unit_resulting_resolution,[],[f405,f500])).
fof(f8086,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.bvlSem.v',bool),'_dst_x0x0x1_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))),s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))),s('type.bvlSem.v',X0))))) )),
  inference(forward_demodulation,[],[f7849,f508])).
fof(f7849,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))))),s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0))))))) )),
  inference(unit_resulting_resolution,[],[f406,f500])).
fof(f8085,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))))),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f7850,f508])).
fof(f7850,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f471,f500])).
fof(f8084,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'),s(fun(fun(bool,bool),bool),'?_1'))),s(fun(bool,bool),'_i_0_0')))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f7851,f508])).
fof(f7851,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f6793,f500])).
fof(f8061,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f7874,f502])).
fof(f7874,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f408,f500])).
fof(f8060,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),'_dst_x0_1'(s(X0,X2))),s(X1,X3))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7875,f502])).
fof(f7875,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0_1'(s(X0,X2))),s(X1,X3))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f487,f500])).
fof(f8059,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7876,f502])).
fof(f7876,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(unit_resulting_resolution,[],[f484,f500])).
fof(f8058,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'const.combin.o_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0))),s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))))))) )),
  inference(forward_demodulation,[],[f7877,f502])).
fof(f7877,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0))))),s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))))))) )),
  inference(unit_resulting_resolution,[],[f404,f500])).
fof(f8057,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'const.combin.o_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0))),s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))))))) )),
  inference(forward_demodulation,[],[f7878,f502])).
fof(f7878,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0))))),s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))))))) )),
  inference(unit_resulting_resolution,[],[f405,f500])).
fof(f8056,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'const.combin.o_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0))),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))))))) )),
  inference(forward_demodulation,[],[f7879,f502])).
fof(f7879,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0))))),s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))))))) )),
  inference(unit_resulting_resolution,[],[f406,f500])).
fof(f8055,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'const.combin.o_3'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f7880,f502])).
fof(f7880,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f471,f500])).
fof(f8054,plain,(
  p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))),s(bool,'T_0')))) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f7881,f502])).
fof(f7881,plain,(
  p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,'T_0')))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f6793,f500])).
fof(f8046,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f8045,f437])).
fof(f8045,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f7889,f508])).
fof(f7889,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(unit_resulting_resolution,[],[f437,f500])).
fof(f8042,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),happ(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))))),s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7892,f418])).
fof(f7892,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.ffi.ffi_state'(X0),bool),happ(s(fun('type.ffi.ffi_state'(X0),fun('type.ffi.ffi_state'(X0),bool)),'=_0'),s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f418,f500])).
fof(f8041,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))))),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7893,f416])).
fof(f7893,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f416,f500])).
fof(f8035,plain,(
  ( ! [X3,X1] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X3))))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X3))))))) )),
  inference(forward_demodulation,[],[f7899,f490])).
fof(f7899,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'('type.bvpSem.stack'),bool),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.list.list'('type.bvpSem.stack'),bool)),'=_0'),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),X3))))))),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X3))))))) )),
  inference(unit_resulting_resolution,[],[f490,f500])).
fof(f8034,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.option.option'('type.num.num'),bool),happ(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))))),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7900,f420])).
fof(f7900,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.option.option'('type.num.num'),bool),happ(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f420,f500])).
fof(f8033,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))))),s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7901,f413])).
fof(f7901,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f413,f500])).
fof(f8032,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),bool),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'=_0'),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))))),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7902,f422])).
fof(f7902,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.sptree.spt'('type.bvlSem.v'),bool),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'=_0'),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f422,f500])).
fof(f8031,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))))),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f7903,f415])).
fof(f7903,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))))))),s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f415,f500])).
fof(f7839,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f7820,f401])).
fof(f7820,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f7794,f555])).
fof(f7811,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0)))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f7794,f555])).
fof(f7809,plain,(
  s(bool,'F_0') = s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f403,f7794,f473])).
fof(f7799,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f7776,f401])).
fof(f7776,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f7533,f555])).
fof(f7794,plain,(
  ~p(s(bool,'const.combin.o_3'(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0')))) | $spl3),
  inference(forward_demodulation,[],[f7747,f502])).
fof(f7747,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f7533,f477])).
fof(f7793,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))),s(bool,X0))))) ) | $spl3),
  inference(forward_demodulation,[],[f7748,f502])).
fof(f7748,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),s(bool,X0))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f7533,f482])).
fof(f7753,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | $spl3),
  inference(unit_resulting_resolution,[],[f403,f7533,f473])).
fof(f7744,plain,(
  ( ! [X6,X4,X5,X3] : (s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7650,f401])).
fof(f7650,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f499,f555])).
fof(f7624,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X100))),s(bool,'T_0')))) | s(bool,X100) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(X98,X101) = s(X98,sK5(X101,X99,X98))) )),
  inference(superposition,[],[f499,f807])).
fof(f7623,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X97))),s(bool,'T_0')))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun(X95,bool),'_dst_x0_1'(s(bool,X96))))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f499,f1482])).
fof(f7601,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X38))),s(bool,'T_0')))) | s(bool,X37) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f499,f555])).
fof(f7600,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X36))),s(bool,'F_0')))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f499,f407])).
fof(f7599,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X34))),s(X31,X32)))) | s(X31,X34) = s(X31,sK5(X32,X33,X31)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X33)))) )),
  inference(superposition,[],[f499,f807])).
fof(f7708,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f7707,f437])).
fof(f7707,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7596,f508])).
fof(f7596,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f499,f437])).
fof(f7567,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X100)))) | s(bool,X100) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(X98,X101) = s(X98,sK5(X101,X99,X98))) )),
  inference(superposition,[],[f499,f807])).
fof(f7566,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X97)))) | s(bool,X97) = s(bool,'|dollar|exists_1'(s(fun(X95,bool),'_dst_x0_1'(s(bool,X96))))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f499,f1482])).
fof(f7677,plain,(
  ( ! [X92,X91] : (s(bool,'T_0') = s(bool,X92) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X92)))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(forward_demodulation,[],[f7564,f898])).
fof(f7564,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X92)))) | s(bool,X92) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f499,f1535])).
fof(f7560,plain,(
  ( ! [X83,X84] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X83))),s(bool,X84)))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f499,f407])).
fof(f7556,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X71))),s(bool,X72)))) | s(bool,'T_0') = s(bool,X72) | s(bool,X71) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f499,f555])).
fof(f7545,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X41))),s(bool,X42)))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f499,f555])).
fof(f7544,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,X37) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f499,f555])).
fof(f7543,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X36)))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f499,f407])).
fof(f7542,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),happ(s(fun(X31,fun(X31,bool)),'=_0'),s(X31,X32))),s(X31,X34)))) | s(X31,X34) = s(X31,sK5(X32,X33,X31)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X33)))) )),
  inference(superposition,[],[f499,f807])).
fof(f7667,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun(X16,bool),'const.combin.o_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s('type.num.num',X19))),s(X16,X20))))) )),
  inference(forward_demodulation,[],[f7666,f437])).
fof(f7666,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),'const.combin.o_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s('type.num.num',X19))),s(X16,X20)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7539,f502])).
fof(f7539,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f499,f437])).
fof(f7533,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f499])).
fof(f7290,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,'|dollar|exists_1'(s(fun(fun(fun(bool,bool),bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(subsumption_resolution,[],[f7246,f401])).
fof(f7246,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X12) = s(bool,'|dollar|exists_1'(s(fun(fun(fun(bool,bool),bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(superposition,[],[f1409,f6983])).
fof(f7166,plain,(
  ( ! [X261,X259,X260,X262,X258] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X258,bool),happ(s(fun('type.num.num',fun(X258,bool)),X260),s('type.num.num',X261))))) | s(X258,X262) = s(X258,sK5(X262,'const.ind_type.FCONS_3'(s(fun(X258,bool),X259),s(fun('type.num.num',fun(X258,bool)),X260),s('type.num.num','const.num.SUC_1'(s('type.num.num',X261)))),X258))) )),
  inference(superposition,[],[f807,f437])).
fof(f7165,plain,(
  ( ! [X255,X253,X257,X254,X256] : (p(s(bool,'|dollar|exists_1'(s(fun(X253,bool),happ(s(fun('type.num.num',fun(X253,bool)),X255),s('type.num.num',X256)))))) | s(X253,X257) = s(X253,sK5(X257,'const.ind_type.FCONS_3'(s(fun(X253,bool),X254),s(fun('type.num.num',fun(X253,bool)),X255),s('type.num.num','const.num.SUC_1'(s('type.num.num',X256)))),X253))) )),
  inference(superposition,[],[f433,f437])).
fof(f7218,plain,(
  ( ! [X251,X249,X252,X250,X248] : (p(s(bool,'|dollar|exists_1'(s(fun(X248,bool),happ(s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num',X251)))))) | ~p(s(bool,happ(s(fun(X248,bool),happ(s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num',X251))),s(X248,sK5(X252,'const.ind_type.FCONS_3'(s(fun(X248,bool),X249),s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num','const.num.SUC_1'(s('type.num.num',X251)))),X248)))))) )),
  inference(forward_demodulation,[],[f7164,f437])).
fof(f7164,plain,(
  ( ! [X251,X249,X252,X250,X248] : (~p(s(bool,happ(s(fun(X248,bool),happ(s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num',X251))),s(X248,sK5(X252,'const.ind_type.FCONS_3'(s(fun(X248,bool),X249),s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num','const.num.SUC_1'(s('type.num.num',X251)))),X248))))) | p(s(bool,'|dollar|exists_1'(s(fun(X248,bool),'const.ind_type.FCONS_3'(s(fun(X248,bool),X249),s(fun('type.num.num',fun(X248,bool)),X250),s('type.num.num','const.num.SUC_1'(s('type.num.num',X251))))))))) )),
  inference(superposition,[],[f434,f437])).
fof(f7217,plain,(
  ( ! [X243,X246,X247,X245] : (p(s(bool,happ(s(fun(X243,bool),happ(s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num',X246))),s(X243,X247)))) | ~p(s(bool,'const.combin.o_3'(s(fun(fun(X243,bool),bool),'!_1'),s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num',X246))))) )),
  inference(forward_demodulation,[],[f7216,f437])).
fof(f7216,plain,(
  ( ! [X243,X246,X244,X247,X245] : (~p(s(bool,'const.combin.o_3'(s(fun(fun(X243,bool),bool),'!_1'),s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num',X246)))) | p(s(bool,happ(s(fun(X243,bool),'const.ind_type.FCONS_3'(s(fun(X243,bool),X244),s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num','const.num.SUC_1'(s('type.num.num',X246))))),s(X243,X247))))) )),
  inference(forward_demodulation,[],[f7163,f502])).
fof(f7163,plain,(
  ( ! [X243,X246,X244,X247,X245] : (~p(s(bool,happ(s(fun(fun(X243,bool),bool),'!_1'),s(fun(X243,bool),happ(s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num',X246)))))) | p(s(bool,happ(s(fun(X243,bool),'const.ind_type.FCONS_3'(s(fun(X243,bool),X244),s(fun('type.num.num',fun(X243,bool)),X245),s('type.num.num','const.num.SUC_1'(s('type.num.num',X246))))),s(X243,X247))))) )),
  inference(superposition,[],[f477,f437])).
fof(f7137,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X131,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),X129),s('type.num.num',X130))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X128),s(fun('type.num.num',bool),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))))) )),
  inference(superposition,[],[f1482,f437])).
fof(f7130,plain,(
  ( ! [X103,X101,X102] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X102),s('type.num.num',X103))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X101),s(fun('type.num.num',bool),X102),s('type.num.num','const.num.SUC_1'(s('type.num.num',X103)))))) )),
  inference(superposition,[],[f571,f437])).
fof(f7201,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,happ(s(fun('type.num.num',bool),X71),s('type.num.num',X72)))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X71),s('type.num.num',X72))),s(bool,X73))))) )),
  inference(forward_demodulation,[],[f7200,f437])).
fof(f7200,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.o_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X71),s('type.num.num',X72))),s(bool,X73)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X70),s(fun('type.num.num',bool),X71),s('type.num.num','const.num.SUC_1'(s('type.num.num',X72))))))) )),
  inference(forward_demodulation,[],[f7122,f502])).
fof(f7122,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),X71),s('type.num.num',X72))))),s(bool,X73)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X70),s(fun('type.num.num',bool),X71),s('type.num.num','const.num.SUC_1'(s('type.num.num',X72))))))) )),
  inference(superposition,[],[f482,f437])).
fof(f7098,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f437,f555])).
fof(f7097,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,X10) = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f437,f555])).
fof(f7096,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f437,f407])).
fof(f6982,plain,(
  p(s(bool,happ(s(fun(fun(fun(fun(bool,bool),bool),bool),bool),'?_1'),s(fun(fun(fun(bool,bool),bool),bool),'?_1')))) | $spl22),
  inference(unit_resulting_resolution,[],[f6782,f480])).
fof(f6983,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(fun(bool,bool),bool),bool),'?_1'))) | $spl22),
  inference(unit_resulting_resolution,[],[f6782,f2467])).
fof(f7010,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))))) ) | $spl22),
  inference(forward_demodulation,[],[f6988,f502])).
fof(f6988,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))))))) ) | $spl22),
  inference(unit_resulting_resolution,[],[f6782,f483])).
fof(f6993,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1'))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f563,f6782,f1479])).
fof(f6980,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f6939,f401])).
fof(f6939,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f1409,f6793])).
fof(f6885,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,'|dollar|exists_1'(s(fun(fun(bool,bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(subsumption_resolution,[],[f6844,f401])).
fof(f6844,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X12) = s(bool,'|dollar|exists_1'(s(fun(fun(bool,bool),bool),'?_1'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl22),
  inference(superposition,[],[f1409,f6783])).
fof(f6782,plain,(
  p(s(bool,happ(s(fun(fun(fun(bool,bool),bool),bool),'?_1'),s(fun(fun(bool,bool),bool),'?_1')))) | $spl22),
  inference(unit_resulting_resolution,[],[f6778,f480])).
fof(f6783,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(bool,bool),bool),'?_1'))) | $spl22),
  inference(unit_resulting_resolution,[],[f6778,f2467])).
fof(f6810,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(fun(bool,bool),bool),'?_1'))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl22),
  inference(forward_demodulation,[],[f6788,f508])).
fof(f6788,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl22),
  inference(unit_resulting_resolution,[],[f6778,f483])).
fof(f6761,plain,(
  ( ! [X52,X53] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X52),bool),bool),'?_1'),s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f480,f484])).
fof(f6779,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(fun(X5,bool),bool),'?_1'),s(fun(X5,bool),X6)))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f6748,f401])).
fof(f6748,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X5,bool),bool),'?_1'),s(fun(X5,bool),X6)))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f480,f555])).
fof(f6694,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | p(s(bool,happ(s(fun(fun(X26,bool),bool),'?_1'),s(fun(X26,bool),sK5(X27,X28,fun(X26,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f480,f807])).
fof(f6772,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X8,bool),bool),'?_1'),s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)))) | ~p(s(bool,happ(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))),s(X8,X11))))) )),
  inference(forward_demodulation,[],[f6690,f502])).
fof(f6690,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))),s(X8,X11)))) | p(s(bool,happ(s(fun(fun(X8,bool),bool),'?_1'),s(fun(X8,bool),happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10))))))) )),
  inference(superposition,[],[f480,f484])).
fof(f6771,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,X7))))) )),
  inference(forward_demodulation,[],[f6689,f502])).
fof(f6689,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,X7)))) | p(s(bool,happ(s(fun(fun(X3,bool),bool),'?_1'),s(fun(X3,bool),happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6))))))) )),
  inference(superposition,[],[f480,f487])).
fof(f6770,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(fun(X0,bool),bool),'?_1'),s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f6688,f502])).
fof(f6688,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f480,f408])).
fof(f6476,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6459,f401])).
fof(f6459,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f477,f555])).
fof(f6457,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(fun(X26,bool),bool),'!_1'),s(fun(X26,bool),X27)))) | p(s(bool,happ(s(fun(X26,bool),sK5(X27,X28,fun(X26,bool))),s(X26,X29)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f477,f807])).
fof(f2975,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),X49)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X48),bool),'const.pair.FST_0')))) )),
  inference(superposition,[],[f2467,f484])).
fof(f920,plain,(
  ( ! [X37,X38] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),sK5(X38,'const.pair.FST_0','type.pair.prod'(bool,X37)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f434,f484])).
fof(f813,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f434,f407])).
fof(f1108,plain,(
  ( ! [X97,X95,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X97)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,sK5(X95,X96,bool))) )),
  inference(superposition,[],[f434,f555])).
fof(f1192,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,happ(s(fun(bool,bool),X93),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) | s(bool,X94) = s(bool,sK5(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f434,f555])).
fof(f2918,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),sK5(X27,X28,fun(X26,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f2467,f807])).
fof(f1414,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32)))) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31)))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1175,f401])).
fof(f1175,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32)))) | s(bool,X34) = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31)))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f434,f555])).
fof(f817,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f808,f408])).
fof(f808,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f434,f408])).
fof(f818,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | ~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3)))))) )),
  inference(forward_demodulation,[],[f809,f487])).
fof(f809,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6))))))) )),
  inference(superposition,[],[f434,f487])).
fof(f819,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f810,f411])).
fof(f810,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f434,f411])).
fof(f820,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X15)))) | ~p(s(bool,happ(s(fun(X12,bool),X15),s(X12,sK5(X16,'const.pair.SND_1'(s('type.pair.prod'(X13,fun(X12,bool)),'const.pair.,_2'(s(X13,X14),s(fun(X12,bool),X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f811,f488])).
fof(f811,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X15),s(X12,sK5(X16,'const.pair.SND_1'(s('type.pair.prod'(X13,fun(X12,bool)),'const.pair.,_2'(s(X13,X14),s(fun(X12,bool),X15)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.pair.SND_1'(s('type.pair.prod'(X13,fun(X12,bool)),'const.pair.,_2'(s(X13,X14),s(fun(X12,bool),X15))))))))) )),
  inference(superposition,[],[f434,f488])).
fof(f821,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | ~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK5(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f812,f486])).
fof(f812,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK5(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20))))))))) )),
  inference(superposition,[],[f434,f486])).
fof(f2413,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,happ(s(fun(X121,bool),X122),s(X121,sK5(X124,sK5(X122,X123,fun(X121,bool)),X121))))) | p(s(bool,'|dollar|exists_1'(s(fun(X121,bool),sK5(X122,X123,fun(X121,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X121,bool),bool),X123)))) )),
  inference(superposition,[],[f434,f807])).
fof(f482,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f396])).
fof(f483,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f3030,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X84))),s(bool,X86)))) | p(s(bool,sK5(X84,X85,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X85)))) )),
  inference(superposition,[],[f482,f807])).
fof(f3082,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(condensation,[],[f3081])).
fof(f3081,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f3002,f557])).
fof(f3002,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f482,f555])).
fof(f3095,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(subsumption_resolution,[],[f3062,f403])).
fof(f3062,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(superposition,[],[f482,f407])).
fof(f3327,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3296,f401])).
fof(f3296,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f483,f555])).
fof(f3342,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X60))),s(bool,'|dollar|exists_1'(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59)))))))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f3318,f401])).
fof(f3318,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X60))),s(bool,'|dollar|exists_1'(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59)))))))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f483,f1482])).
fof(f3343,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X63))),s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62)))))) | s(X61,X64) = s(X61,sK5(X64,X62,X61))) )),
  inference(subsumption_resolution,[],[f3319,f401])).
fof(f3319,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X63))),s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62)))))) | s(X61,X64) = s(X61,sK5(X64,X62,X61))) )),
  inference(superposition,[],[f483,f807])).
fof(f3326,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,X85)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X87))),s(bool,sK5(X85,X86,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X86)))) )),
  inference(superposition,[],[f483,f807])).
fof(f944,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'|dollar|exists_1'(s(fun(X64,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X64,bool),X65),X66)))))) | ~p(s(bool,happ(s(fun(X64,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X64,bool),X65),X66))),s(X64,sK5(X67,happ(s(fun('type.pair.prod'(fun(X64,bool),X65),fun(X64,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X64,bool),X65),X66)),X64)))))) )),
  inference(forward_demodulation,[],[f930,f484])).
fof(f930,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X64,bool),X65),X66))),s(X64,sK5(X67,happ(s(fun('type.pair.prod'(fun(X64,bool),X65),fun(X64,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X64,bool),X65),X66)),X64))))) | p(s(bool,'|dollar|exists_1'(s(fun(X64,bool),happ(s(fun('type.pair.prod'(fun(X64,bool),X65),fun(X64,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X64,bool),X65),X66))))))) )),
  inference(superposition,[],[f434,f484])).
fof(f3344,plain,(
  ( ! [X66,X65] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X66),s(fun(bool,bool),'_i_0_0'))),s(bool,X65)))) | ~p(s(bool,X65))) )),
  inference(forward_demodulation,[],[f3320,f508])).
fof(f3320,plain,(
  ( ! [X66,X65] : (~p(s(bool,X65)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X66))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X65))))))) )),
  inference(superposition,[],[f483,f408])).
fof(f3331,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X48),s(fun(bool,bool),'_i_0_0'))),s(bool,X47)))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(forward_demodulation,[],[f3330,f508])).
fof(f3330,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X48))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X47)))))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f3312,f401])).
fof(f3312,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X48))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X47)))))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f483,f571])).
fof(f3425,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f3355,f508])).
fof(f3355,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f483])).
fof(f4476,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X32))),s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f4475,f502])).
fof(f4475,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X32))),s(bool,happ(s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4390,f401])).
fof(f4390,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X32))),s(bool,happ(s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f483,f724])).
fof(f3345,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X70))),s(fun(X67,bool),'_dst_x0_1'(s(bool,X68))),s(X67,X69)))) | ~p(s(bool,X68))) )),
  inference(forward_demodulation,[],[f3321,f502])).
fof(f3321,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X70))),s(bool,happ(s(fun(X67,bool),'_dst_x0_1'(s(bool,X68))),s(X67,X69))))))) )),
  inference(superposition,[],[f483,f487])).
fof(f3346,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.combin.o_3'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X73))),s(fun('type.pair.prod'(bool,X71),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X71),X72)))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X71),X72))))) )),
  inference(forward_demodulation,[],[f3322,f502])).
fof(f3322,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X71),X72)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X73))),s(bool,happ(s(fun('type.pair.prod'(bool,X71),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X71),X72))))))) )),
  inference(superposition,[],[f483,f484])).
fof(f801,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f433,f408])).
fof(f802,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X7) = s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f433,f487])).
fof(f803,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | s(X8,X11) = s(X8,sK5(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) )),
  inference(superposition,[],[f433,f411])).
fof(f804,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X15)))) | s(X12,X16) = s(X12,sK5(X16,'const.pair.SND_1'(s('type.pair.prod'(X13,fun(X12,bool)),'const.pair.,_2'(s(X13,X14),s(fun(X12,bool),X15)))),X12))) )),
  inference(superposition,[],[f433,f488])).
fof(f805,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | s(X17,X21) = s(X17,sK5(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))) )),
  inference(superposition,[],[f433,f486])).
fof(f2361,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK5(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f807,f408])).
fof(f2362,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5))) | s(X3,X7) = s(X3,sK5(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f807,f487])).
fof(f2364,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK5(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f807,f411])).
fof(f2365,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),X19))) | s(X16,X20) = s(X16,sK5(X20,'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))) )),
  inference(superposition,[],[f807,f488])).
fof(f2366,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X25) = s(X21,sK5(X25,'const.pair.FST_1'(s('type.pair.prod'(fun(X21,bool),X22),'const.pair.,_2'(s(fun(X21,bool),X23),s(X22,X24)))),X21))) )),
  inference(superposition,[],[f807,f486])).
fof(f2367,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))) | s(X26,X29) = s(X26,sK5(X29,sK5(X27,X28,fun(X26,bool)),X26)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f807,f807])).
fof(f2414,plain,(
  ( ! [X127,X125,X128,X126] : (p(s(bool,'|dollar|exists_1'(s(fun(X125,bool),X126)))) | s(X125,X128) = s(X125,sK5(X128,sK5(X126,X127,fun(X125,bool)),X125)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X125,bool),bool),X127)))) )),
  inference(superposition,[],[f433,f807])).
fof(f2468,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))) | s(X24,X27) = s(X24,sK5(X27,X25,X24))) )),
  inference(subsumption_resolution,[],[f2432,f401])).
fof(f2432,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))) | s(X24,X27) = s(X24,sK5(X27,X25,X24))) )),
  inference(superposition,[],[f476,f807])).
fof(f2471,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58))))))) | s(X57,X60) = s(X57,sK5(X60,X58,X57))) )),
  inference(subsumption_resolution,[],[f2440,f401])).
fof(f2440,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58))))))) | s(X57,X60) = s(X57,sK5(X60,X58,X57))) )),
  inference(superposition,[],[f892,f807])).
fof(f2996,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2963,f401])).
fof(f2963,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2467,f555])).
fof(f892,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) )),
  inference(subsumption_resolution,[],[f886,f403])).
fof(f886,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))))) )),
  inference(superposition,[],[f816,f407])).
fof(f1405,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,X22)))))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1151,f401])).
fof(f1151,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,X22)))))) | ~p(s(bool,'T_0')) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f816,f555])).
fof(f2398,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'|dollar|exists_1'(s(fun(X74,bool),'_dst_x0_1'(s(bool,X72)))))) | ~p(s(bool,sK5(X72,X73,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f816,f807])).
fof(f2582,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X20)) | s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f2540,f401])).
fof(f2540,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(superposition,[],[f476,f1482])).
fof(f2585,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X45,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X43,bool),'_dst_x0_1'(s(bool,X44))))))))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f2548,f401])).
fof(f2548,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X45,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X43,bool),'_dst_x0_1'(s(bool,X44))))))))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f892,f1482])).
fof(f2399,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,X75)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X77,bool),'_dst_x0_1'(s(bool,sK5(X75,X76,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76)))) )),
  inference(superposition,[],[f892,f807])).
fof(f2363,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK5(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f807,f484])).
fof(f931,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,'|dollar|exists_1'(s(fun(X68,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X68,bool),X69),X70)))))) | s(X68,X71) = s(X68,sK5(X71,happ(s(fun('type.pair.prod'(fun(X68,bool),X69),fun(X68,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X68,bool),X69),X70)),X68))) )),
  inference(superposition,[],[f433,f484])).
fof(f2387,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | p(s(bool,X44)) | s(bool,X44) = s(bool,sK5(X42,X43,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f473,f807])).
fof(f2389,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X50)) | s(bool,X50) = s(bool,sK5(X48,X49,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f476,f807])).
fof(f2394,plain,(
  ( ! [X64,X63] : (~p(s(bool,X63)) | s(bool,'T_0') = s(bool,sK5(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f557,f807])).
fof(f2397,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | s(bool,X71) = s(bool,sK5(X69,X70,bool)) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f593,f807])).
fof(f2400,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,X78)) | s(bool,X80) = s(bool,sK5(X78,X79,bool)) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) )),
  inference(superposition,[],[f1479,f807])).
fof(f4756,plain,(
  ( ! [X152,X151,X149,X150] : (~p(s(bool,X149)) | s(bool,X151) = s(bool,sK5(X149,X150,bool)) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150)))) )),
  inference(superposition,[],[f1409,f807])).
fof(f470,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun('type.sptree.spt'('type.bvlSem.v'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'_dst_x00x11_2'(s(fun(fun('type.list.list'('type.bvpSem.stack'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),bool))),'_dst_x000x111_2'(s(fun(fun('type.option.option'('type.num.num'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),bool)))),'_dst_x0000x1111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',bool))))),'_dst_x00000x11111_2'(s(fun(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool)))))),'_dst_x000000x111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',bool))))))),'_dst_x0000000x1111111_2'(s(fun(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool)))))))),'_dst_x00000000x11111111_2'(s(fun(fun('type.ffi.ffi_state'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),bool))))))))),'_dst_x000000000x111111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num',bool)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'_i_0_0'),s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))),'_dst_x00000000x10000000x01111111_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.pair.,_0'),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'_i_0_0'),s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))),'_dst_x0000000x1000000x0111111_3'(s(fun('type.option.option'('type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),'const.pair.,_0'),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'_i_0_0'),s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'_dst_x000000x100000x011111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'_dst_x00000x10000x01111_3'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'const.pair.,_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'_i_0_0'),s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'_dst_x0000x1000x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'),'type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'_i_0_0'),s(fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'const.pair.,_0'))))))))))))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.ind_type.BOTTOM_0'))))))))))))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'))))))),s(fun('type.bvpSem.state'(X0),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),sK6(X0)))))) )),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun('type.sptree.spt'('type.bvlSem.v'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'_dst_x00x11_2'(s(fun(fun('type.list.list'('type.bvpSem.stack'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),bool))),'_dst_x000x111_2'(s(fun(fun('type.option.option'('type.num.num'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),bool)))),'_dst_x0000x1111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',bool))))),'_dst_x00000x11111_2'(s(fun(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool)))))),'_dst_x000000x111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',bool))))))),'_dst_x0000000x1111111_2'(s(fun(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool)))))))),'_dst_x00000000x11111111_2'(s(fun(fun('type.ffi.ffi_state'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),bool))))))))),'_dst_x000000000x111111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num',bool)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'_i_0_0'),s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))),'_dst_x00000000x10000000x01111111_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.pair.,_0'),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'_i_0_0'),s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))),'_dst_x0000000x1000000x0111111_3'(s(fun('type.option.option'('type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),'const.pair.,_0'),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'_i_0_0'),s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'_dst_x000000x100000x011111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'_dst_x00000x10000x01111_3'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'const.pair.,_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'_i_0_0'),s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'_dst_x0000x1000x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'),'type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'_i_0_0'),s(fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'const.pair.,_0'))))))))))))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.ind_type.BOTTOM_0'))))))))))))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'))))))),s(fun('type.bvpSem.state'(X0),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),sK6(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f258])).
fof(f258,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun('type.sptree.spt'('type.bvlSem.v'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'_dst_x00x11_2'(s(fun(fun('type.list.list'('type.bvpSem.stack'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),bool))),'_dst_x000x111_2'(s(fun(fun('type.option.option'('type.num.num'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),bool)))),'_dst_x0000x1111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',bool))))),'_dst_x00000x11111_2'(s(fun(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool)))))),'_dst_x000000x111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',bool))))))),'_dst_x0000000x1111111_2'(s(fun(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool)))))))),'_dst_x00000000x11111111_2'(s(fun(fun('type.ffi.ffi_state'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),bool))))))))),'_dst_x000000000x111111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num',bool)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'_i_0_0'),s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))))),'_dst_x00000000x10000000x01111111_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.pair.,_0'),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'_i_0_0'),s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))))),'_dst_x0000000x1000000x0111111_3'(s(fun('type.option.option'('type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))),'const.pair.,_0'),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'_i_0_0'),s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))))),'_dst_x000000x100000x011111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'_dst_x00000x10000x01111_3'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))),'const.pair.,_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'_i_0_0'),s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))),'_dst_x0000x1000x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'),'type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'_i_0_0'),s(fun('type.ffi.ffi_state'(X0),fun('type.num.num','type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))),'const.pair.,_0'))))))))))))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'const.ind_type.BOTTOM_0'))))))))))))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),'_i_0_0'))))))),s(fun('type.bvpSem.state'(X0),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X0),'type.num.num')))))))))),X1))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7] : ? [X34] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),'_dst_x0x1_2'(s(fun(fun('type.sptree.spt'('type.bvlSem.v'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),bool)),'_dst_x00x11_2'(s(fun(fun('type.list.list'('type.bvpSem.stack'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),bool))),'_dst_x000x111_2'(s(fun(fun('type.option.option'('type.num.num'),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),bool)))),'_dst_x0000x1111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',bool))))),'_dst_x00000x11111_2'(s(fun(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),bool)))))),'_dst_x000000x111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',bool))))))),'_dst_x0000000x1111111_2'(s(fun(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),bool)))))))),'_dst_x00000000x11111111_2'(s(fun(fun('type.ffi.ffi_state'(X7),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),bool))))))))),'_dst_x000000000x111111111_2'(s(fun(fun('type.num.num',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',bool)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),'_i_0_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))),'type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'_i_0_0'),s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))))))))),'_dst_x00000000x10000000x01111111_3'(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),'const.pair.,_0'),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'_i_0_0'),s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))))))),'_dst_x0000000x1000000x0111111_3'(s(fun('type.option.option'('type.num.num'),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))),'const.pair.,_0'),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'_i_0_0'),s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))))),'_dst_x000000x100000x011111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),'_dst_x00000x10000x01111_3'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))),'type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))),'const.pair.,_0'),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'_i_0_0'),s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))),'_dst_x0000x1000x0111_3'(s(fun('type.num.num',fun('type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))),'const.pair.,_0'),s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))),'_dst_x000x100x011_3'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'),'type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))),'const.pair.,_0'),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'_i_0_0'),s(fun('type.ffi.ffi_state'(X7),fun('type.num.num','type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))),'const.pair.,_0'))))))))))))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),'const.ind_type.BOTTOM_0'))))))))))))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num'))))))))),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),'_i_0_0'))))))),s(fun('type.bvpSem.state'(X7),'type.ind_type.recspace'('type.pair.prod'('type.sptree.spt'('type.bvlSem.v'),'type.pair.prod'('type.list.list'('type.bvpSem.stack'),'type.pair.prod'('type.option.option'('type.num.num'),'type.pair.prod'('type.num.num','type.pair.prod'('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.pair.prod'('type.num.num','type.pair.prod'('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.pair.prod'('type.ffi.ffi_state'(X7),'type.num.num')))))))))),X34))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_TY_DEF')).
fof(f911,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1)))) )),
  inference(superposition,[],[f484,f407])).
fof(f1023,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X29),X30))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X29),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X29),X30)))) )),
  inference(superposition,[],[f555,f484])).
fof(f1055,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,X23) = s(bool,happ(s(fun('type.pair.prod'(bool,X21),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X21),X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f555,f484])).
fof(f1435,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X20),X21))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X20),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X20),X21)))) )),
  inference(superposition,[],[f571,f484])).
fof(f2521,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X62,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),X61))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X60),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X60),X61)))) )),
  inference(superposition,[],[f1482,f484])).
fof(f3968,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83))),s(fun('type.num.num',bool),X84),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),X83)))) )),
  inference(superposition,[],[f656,f484])).
fof(f4119,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X121),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),X120))),s(X121,X122))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X119),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X119),X120)))) )),
  inference(superposition,[],[f688,f484])).
fof(f4314,plain,(
  ( ! [X127,X125,X126,X124] : (s(bool,'T_0') = s(bool,happ(s(fun(X126,bool),'_dst_x0_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X124),X125))))),s(X126,X127))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X124),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X124),X125)))) )),
  inference(superposition,[],[f724,f484])).
fof(f4596,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X131,bool),'const.pair.,_2'(s(X131,X132),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f764,f484])).
fof(f656,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f411,f407])).
fof(f688,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f486,f407])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f487,f407])).
fof(f764,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f488,f407])).
fof(f1472,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))),s(X17,X19)))) )),
  inference(forward_demodulation,[],[f1434,f408])).
fof(f1434,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18))) | s(bool,'F_0') = s(bool,happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))),s(X17,X19)))) )),
  inference(superposition,[],[f571,f487])).
fof(f1473,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X22),s(fun('type.num.num',bool),X23),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1436,f408])).
fof(f1436,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X22),s(fun('type.num.num',bool),X23),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f571,f411])).
fof(f1474,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,X25),s(bool,X26)))))) )),
  inference(forward_demodulation,[],[f1437,f408])).
fof(f1437,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,X25),s(bool,X26)))))) )),
  inference(superposition,[],[f571,f488])).
fof(f1475,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29)))))) )),
  inference(forward_demodulation,[],[f1438,f408])).
fof(f1438,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29)))))) )),
  inference(superposition,[],[f571,f486])).
fof(f1482,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,X20))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(forward_demodulation,[],[f1481,f408])).
fof(f1481,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1454,f401])).
fof(f1454,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))))))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f892,f571])).
fof(f1533,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1516,f401])).
fof(f1516,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1479,f571])).
fof(f2351,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,sK5(X3,X4,bool))) )),
  inference(superposition,[],[f807,f407])).
fof(f2395,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X65))) | s(bool,'F_0') = s(bool,sK5(X65,X66,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66)))) )),
  inference(superposition,[],[f571,f807])).
fof(f2520,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),'_dst_x0_1'(s(bool,X57))))) | s(bool,'F_0') = s(bool,happ(s(fun(X56,bool),'_dst_x0_1'(s(bool,X57))),s(X56,X58)))) )),
  inference(superposition,[],[f1482,f487])).
fof(f2522,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X65,bool),'_dst_x0_1'(s(bool,X63))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X63),s(fun('type.num.num',bool),X64),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1482,f411])).
fof(f2523,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X69,bool),'_dst_x0_1'(s(bool,X68))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X66,bool),'const.pair.,_2'(s(X66,X67),s(bool,X68)))))) )),
  inference(superposition,[],[f1482,f488])).
fof(f2524,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X73,bool),'_dst_x0_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72)))))) )),
  inference(superposition,[],[f1482,f486])).
fof(f2525,plain,(
  ( ! [X76,X74,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X76,bool),'_dst_x0_1'(s(bool,X74))))) | s(bool,'F_0') = s(bool,sK5(X74,X75,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f1482,f807])).
fof(f2606,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,'|dollar|exists_1'(s(fun(X61,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) )),
  inference(superposition,[],[f1535,f408])).
fof(f2607,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X63) = s(bool,'|dollar|exists_1'(s(fun(X65,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(X62,bool),'_dst_x0_1'(s(bool,X63))),s(X62,X64)))) )),
  inference(superposition,[],[f1535,f487])).
fof(f2609,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,'|dollar|exists_1'(s(fun(X71,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1535,f411])).
fof(f2610,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X74) = s(bool,'|dollar|exists_1'(s(fun(X75,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X72,bool),'const.pair.,_2'(s(X72,X73),s(bool,X74)))))) )),
  inference(superposition,[],[f1535,f488])).
fof(f2611,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X77),s(X76,X78)))))) )),
  inference(superposition,[],[f1535,f486])).
fof(f2612,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,'|dollar|exists_1'(s(fun(X82,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,sK5(X80,X81,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X81)))) )),
  inference(superposition,[],[f1535,f807])).
fof(f2613,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1535,f555])).
fof(f3108,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f2999,f557])).
fof(f3967,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X79),s(fun('type.num.num',bool),X81),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(X78,bool),'_dst_x0_1'(s(bool,X79))),s(X78,X80)))) )),
  inference(superposition,[],[f656,f487])).
fof(f3969,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X87),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f656,f411])).
fof(f3970,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90)))))) )),
  inference(superposition,[],[f656,f488])).
fof(f3971,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X95),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X92),'const.pair.,_2'(s(bool,X93),s(X92,X94)))))) )),
  inference(superposition,[],[f656,f486])).
fof(f3972,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK5(X96,X97,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X97)))) )),
  inference(superposition,[],[f656,f807])).
fof(f4118,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X117),'const.pair.,_2'(s(bool,X115),s(X117,X118))))) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),'_dst_x0_1'(s(bool,X115))),s(X114,X116)))) )),
  inference(superposition,[],[f688,f487])).
fof(f4120,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X125),'const.pair.,_2'(s(bool,X123),s(X125,X126))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f688,f411])).
fof(f4121,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,X129),s(X130,X131))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X129)))))) )),
  inference(superposition,[],[f688,f488])).
fof(f4122,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X135),'const.pair.,_2'(s(bool,X133),s(X135,X136))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))))) )),
  inference(superposition,[],[f688,f486])).
fof(f4123,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X139),'const.pair.,_2'(s(bool,X137),s(X139,X140))))) | s(bool,'F_0') = s(bool,sK5(X137,X138,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X138)))) )),
  inference(superposition,[],[f688,f807])).
fof(f4313,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,happ(s(fun(X122,bool),'_dst_x0_1'(s(bool,X120))),s(X122,X123))) | s(bool,'F_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0_1'(s(bool,X120))),s(X119,X121)))) )),
  inference(superposition,[],[f724,f487])).
fof(f4315,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,X128))),s(X130,X131))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X128),s(fun('type.num.num',bool),X129),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f724,f411])).
fof(f4316,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'T_0') = s(bool,happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,X134))),s(X135,X136))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X132,bool),'const.pair.,_2'(s(X132,X133),s(bool,X134)))))) )),
  inference(superposition,[],[f724,f488])).
fof(f4317,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'T_0') = s(bool,happ(s(fun(X140,bool),'_dst_x0_1'(s(bool,X138))),s(X140,X141))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X138),s(X137,X139)))))) )),
  inference(superposition,[],[f724,f486])).
fof(f4318,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,'T_0') = s(bool,happ(s(fun(X144,bool),'_dst_x0_1'(s(bool,X142))),s(X144,X145))) | s(bool,'F_0') = s(bool,sK5(X142,X143,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X143)))) )),
  inference(superposition,[],[f724,f807])).
fof(f4595,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X125))))) | s(bool,'F_0') = s(bool,happ(s(fun(X124,bool),'_dst_x0_1'(s(bool,X125))),s(X124,X126)))) )),
  inference(superposition,[],[f764,f487])).
fof(f4597,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X135,bool),'const.pair.,_2'(s(X135,X136),s(bool,X133))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f764,f411])).
fof(f4598,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,X139))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X137,bool),'const.pair.,_2'(s(X137,X138),s(bool,X139)))))) )),
  inference(superposition,[],[f764,f488])).
fof(f4599,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X145,bool),'const.pair.,_2'(s(X145,X146),s(bool,X143))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,X143),s(X142,X144)))))) )),
  inference(superposition,[],[f764,f486])).
fof(f4600,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X149,bool),'const.pair.,_2'(s(X149,X150),s(bool,X147))))) | s(bool,'F_0') = s(bool,sK5(X147,X148,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X148)))) )),
  inference(superposition,[],[f764,f807])).
fof(f4770,plain,(
  ( ! [X70,X68,X69] : (s(bool,X69) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f4735,f401])).
fof(f4735,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | s(bool,X69) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f1409,f571])).
fof(f4771,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X74) = s(bool,happ(s(fun(X71,bool),'_dst_x0_1'(s(bool,X72))),s(X71,X73))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f4736,f401])).
fof(f4736,plain,(
  ( ! [X74,X72,X71,X75,X73] : (~p(s(bool,'T_0')) | s(bool,X74) = s(bool,happ(s(fun(X71,bool),'_dst_x0_1'(s(bool,X72))),s(X71,X73))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X72)) )),
  inference(superposition,[],[f1409,f724])).
fof(f4780,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X95) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))))) | s(bool,X95) = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,'F_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f4743,f401])).
fof(f4743,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | s(bool,X95) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))))) | s(bool,X95) = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,'F_0') = s(bool,X94)) )),
  inference(superposition,[],[f1409,f1482])).
fof(f4785,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X111) = s(bool,'const.ind_type.FCONS_3'(s(bool,X109),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.|48|_0'))) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f4747,f401])).
fof(f4747,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | s(bool,X111) = s(bool,'const.ind_type.FCONS_3'(s(bool,X109),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.|48|_0'))) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f1409,f656])).
fof(f4786,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,X116) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X115))))) | s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X117) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f4748,f401])).
fof(f4748,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | s(bool,X116) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X115))))) | s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X117) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f1409,f764])).
fof(f4787,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X121) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120))))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122) | s(bool,'F_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f4749,f401])).
fof(f4749,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | s(bool,X121) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120))))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f1409,f688])).
fof(f4781,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(bool,X99) = s(bool,X100) | s(bool,'T_0') = s(bool,X100) | s(X97,X101) = s(X97,sK5(X101,X98,X97))) )),
  inference(subsumption_resolution,[],[f4744,f401])).
fof(f4744,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(bool,X99) = s(bool,X100) | s(bool,'T_0') = s(bool,X100) | s(X97,X101) = s(X97,sK5(X101,X98,X97))) )),
  inference(superposition,[],[f1409,f807])).
fof(f4779,plain,(
  ( ! [X92,X90,X91] : (s(bool,X91) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X91) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f4742,f401])).
fof(f4742,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | s(bool,X91) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X91) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f1409,f898])).
fof(f4776,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X83))))) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f4740,f401])).
fof(f4740,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'T_0')) | s(bool,X84) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X83))))) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f1409,f3350])).
fof(f4775,plain,(
  ( ! [X81,X82] : (s(bool,X81) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4739,f401])).
fof(f4739,plain,(
  ( ! [X81,X82] : (~p(s(bool,'T_0')) | s(bool,X81) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f1409,f3100])).
fof(f4526,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X64),s(bool,X65))))) | s(bool,'F_0') = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f764,f555])).
fof(f4338,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X65))),s(bool,X64))) | s(bool,'F_0') = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f724,f555])).
fof(f4143,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X65),s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f688,f555])).
fof(f3617,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0'))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3605,f401])).
fof(f3605,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0'))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2975,f555])).
fof(f3474,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X20))))) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f3350,f555])).
fof(f3350,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f3294,f2467])).
fof(f3359,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0')))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f563,f3294,f1479])).
fof(f3294,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f401,f483])).
fof(f3221,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X4))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f3100,f407])).
fof(f3100,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f2999,f2467])).
fof(f2999,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f403,f482])).
fof(f2353,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK5(X8,X9,bool))) )),
  inference(superposition,[],[f807,f555])).
fof(f2352,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6))) | s(bool,X7) = s(bool,sK5(X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f807,f555])).
fof(f1526,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1508,f401])).
fof(f1508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1479,f555])).
fof(f1413,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),'_dst_x0_1'(s(bool,X21))))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f1169,f401])).
fof(f1169,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),'_dst_x0_1'(s(bool,X21))))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f892,f555])).
fof(f1143,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f898,f555])).
fof(f1058,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X34) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f555,f486])).
fof(f1057,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f555,f488])).
fof(f1056,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f555,f411])).
fof(f1054,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X20) = s(bool,happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))),s(X17,X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f555,f487])).
fof(f1053,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f555,f408])).
fof(f1043,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f555,f555])).
fof(f1026,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X39),'const.pair.,_2'(s(bool,X40),s(X39,X41)))))) )),
  inference(superposition,[],[f555,f486])).
fof(f1025,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X37)))))) )),
  inference(superposition,[],[f555,f488])).
fof(f1024,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X32),s(fun('type.num.num',bool),X33),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f555,f411])).
fof(f1022,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26))),s(X25,X27)))) )),
  inference(superposition,[],[f555,f487])).
fof(f1021,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X23)))) )),
  inference(superposition,[],[f555,f408])).
fof(f542,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X15,X13,X11,X9,X16] : (s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(X1,happ(s(fun(X2,X1),X14),s(X2,X16))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(fun(X2,X1),X14),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))),s(X2,X16)))) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(X1,happ(s(fun(X2,X1),X14),s(X2,X16))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(fun(X2,X1),X14),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))),s(X2,X16)))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ! [X13,X14,X15,X16] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(X1,happ(s(fun(X2,X1),X14),s(X2,X16))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0))))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X0)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun(X1,fun(X12,X0)),X13),s(fun(X2,X1),X14),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,fun(X11,X12))))))))),X15))),s(X2,X16)))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X35,X67,X37,X36,X55,X57,X58,X63,X64,X65,X62,X66,X68] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,fun(X66,X35))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X67,fun(X68,X35)),X38),s(X67,happ(s(fun(X37,X67),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,fun(X66,X68))))))))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,fun(X66,X35))))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,fun(X66,X35)))))))))),'_dst_x0000000000x1000000000x0111111111_3'(s(fun(X67,fun(X68,X35)),X38),s(fun(X37,X67),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,fun(X66,X68))))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000000000x1000000000x0111111111_ax')).
fof(f541,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X15,X13,X11,X9] : (s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14),s(X11,X15))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14))),s(X11,X15)))) )),
  inference(cnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14),s(X11,X15))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14))),s(X11,X15)))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ! [X12,X13,X14,X15] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14),s(X11,X15))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),happ(s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X1,fun(X10,X0)),X12),s(X1,X13),s(fun(X11,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))))),X14))),s(X11,X15)))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X35,X66,X36,X55,X57,X58,X63,X64,X65,X62,X67,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X66,fun(X67,X35)),X38),s(X66,X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X67)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X67))))))))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35))))))))),'_dst_x000000000x000000000x111111111_3'(s(fun(X66,fun(X67,X35)),X38),s(X66,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X67))))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000000x000000000x111111111_ax')).
fof(f540,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X15,X13,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X10,X15))),s(X11,X14))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X11,X14))),s(X10,X15)))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X10,X15))),s(X11,X14))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X11,X14))),s(X10,X15)))),
  inference(flattening,[],[f372])).
fof(f372,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ! [X12,X13,X14,X15] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X10,X15))),s(X11,X14))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun(X9,fun(X11,X0)),X12),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X13),s(X11,X14))),s(X10,X15)))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X62,X66,X37,X67] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X66,fun(X67,X35)),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66))))))))),X39),s(X37,X40))),s(X67,X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35))))))))),'_dst_x000000000x111111111x000000000_3'(s(fun(X66,fun(X67,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66))))))))),X39),s(X67,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000000x111111111x000000000_ax')).
fof(f539,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X15,X13,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X11,X0)),X12),s(X1,happ(s(fun(X2,X1),X13),s(X2,X15))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun(X1,fun(X11,X0)),X12),s(fun(X2,X1),X13),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))),s(X2,X15)))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X11,X0)),X12),s(X1,happ(s(fun(X2,X1),X13),s(X2,X15))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun(X1,fun(X11,X0)),X12),s(fun(X2,X1),X13),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))),s(X2,X15)))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ! [X12,X13,X14,X15] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X11,X0)),X12),s(X1,happ(s(fun(X2,X1),X13),s(X2,X15))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0)))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X0))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun(X1,fun(X11,X0)),X12),s(fun(X2,X1),X13),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,fun(X10,X11)))))))),X14))),s(X2,X15)))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X35,X66,X37,X36,X55,X57,X58,X63,X64,X65,X62,X67] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X66,fun(X67,X35)),X38),s(X66,happ(s(fun(X37,X66),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X67)))))))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35))))))))),'_dst_x000000000x100000000x011111111_3'(s(fun(X66,fun(X67,X35)),X38),s(fun(X37,X66),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X67)))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000000x100000000x011111111_ax')).
fof(f538,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13),s(X10,X14))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13))),s(X10,X14)))) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13),s(X10,X14))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13))),s(X10,X14)))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ! [X11,X12,X13,X14] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13),s(X10,X14))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),happ(s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X1,fun(X9,X0)),X11),s(X1,X12),s(fun(X10,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),X13))),s(X10,X14)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X35,X62,X36,X55,X57,X58,X63,X64,X65,X66,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X62,fun(X66,X35)),X38),s(X62,X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X66))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X66)))))))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35)))))))),'_dst_x00000000x00000000x11111111_3'(s(fun(X62,fun(X66,X35)),X38),s(X62,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X66)))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000000x00000000x11111111_ax')).
fof(f537,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X9,X14))),s(X10,X13))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X10,X13))),s(X9,X14)))) )),
  inference(cnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X9,X14))),s(X10,X13))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X10,X13))),s(X9,X14)))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ! [X11,X12,X13,X14] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X9,X14))),s(X10,X13))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X8,fun(X10,X0)),X11),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X12),s(X10,X13))),s(X9,X14)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X62,X37,X66] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X62,fun(X66,X35)),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62)))))))),X39),s(X37,X40))),s(X66,X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35)))))))),'_dst_x00000000x11111111x00000000_3'(s(fun(X62,fun(X66,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62)))))))),X39),s(X66,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000000x11111111x00000000_ax')).
fof(f536,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X10,X0)),X11),s(X1,happ(s(fun(X2,X1),X12),s(X2,X14))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x10000000x01111111_3'(s(fun(X1,fun(X10,X0)),X11),s(fun(X2,X1),X12),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))),s(X2,X14)))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X10,X0)),X11),s(X1,happ(s(fun(X2,X1),X12),s(X2,X14))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x10000000x01111111_3'(s(fun(X1,fun(X10,X0)),X11),s(fun(X2,X1),X12),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))),s(X2,X14)))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ! [X11,X12,X13,X14] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X10,X0)),X11),s(X1,happ(s(fun(X2,X1),X12),s(X2,X14))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0))))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X0)))))))),'_dst_x00000000x10000000x01111111_3'(s(fun(X1,fun(X10,X0)),X11),s(fun(X2,X1),X12),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,fun(X9,X10))))))),X13))),s(X2,X14)))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X35,X62,X37,X36,X55,X57,X58,X63,X64,X65,X66] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X62,fun(X66,X35)),X38),s(X62,happ(s(fun(X37,X62),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X66))))))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35)))))))),'_dst_x00000000x10000000x01111111_3'(s(fun(X62,fun(X66,X35)),X38),s(fun(X37,X62),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X66))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000000x10000000x01111111_ax')).
fof(f535,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X9,X0),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12),s(X10,X13))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111_2'(s(fun(X9,X0),X11),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12))),s(X10,X13)))) )),
  inference(cnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X9,X0),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12),s(X10,X13))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111_2'(s(fun(X9,X0),X11),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12))),s(X10,X13)))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ! [X11,X12,X13] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X9,X0),X11),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12),s(X10,X13))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))))),happ(s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))))),'_dst_x000000000x111111111_2'(s(fun(X9,X0),X11),s(fun(X10,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9))))))))),X12))),s(X10,X13)))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X62,X66,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),'_dst_x00000000x11111111_2'(s(fun(X66,X35),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66))))))))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35)))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X35))))))))),'_dst_x000000000x111111111_2'(s(fun(X66,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,fun(X62,X66))))))))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000000x111111111_ax')).
fof(f534,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X8,X13))),s(X9,X12))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X9,X12))),s(X8,X13)))) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X8,X13))),s(X9,X12))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X9,X12))),s(X8,X13)))),
  inference(flattening,[],[f360])).
fof(f360,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ! [X10,X11,X12,X13] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X8,X13))),s(X9,X12))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X7,fun(X9,X0)),X10),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X11),s(X9,X12))),s(X8,X13)))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X37,X62] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),'_dst_x000000x111111x000000_3'(s(fun(X65,fun(X62,X35)),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65))))))),X39),s(X37,X40))),s(X62,X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35))))))),'_dst_x0000000x1111111x0000000_3'(s(fun(X65,fun(X62,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65))))))),X39),s(X62,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000000x1111111x0000000_ax')).
fof(f533,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12),s(X9,X13))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12))),s(X9,X13)))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12),s(X9,X13))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12))),s(X9,X13)))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ! [X10,X11,X12,X13] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12),s(X9,X13))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),happ(s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X1,fun(X8,X0)),X10),s(X1,X11),s(fun(X9,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),X12))),s(X9,X13)))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X35,X65,X36,X55,X57,X58,X63,X64,X62,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),'_dst_x000000x000000x111111_3'(s(fun(X65,fun(X62,X35)),X38),s(X65,X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X62)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X62))))))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35))))))),'_dst_x0000000x0000000x1111111_3'(s(fun(X65,fun(X62,X35)),X38),s(X65,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X62))))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000000x0000000x1111111_ax')).
fof(f532,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X9,X0)),X10),s(X1,happ(s(fun(X2,X1),X11),s(X2,X13))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x1000000x0111111_3'(s(fun(X1,fun(X9,X0)),X10),s(fun(X2,X1),X11),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))),s(X2,X13)))) )),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X9,X0)),X10),s(X1,happ(s(fun(X2,X1),X11),s(X2,X13))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x1000000x0111111_3'(s(fun(X1,fun(X9,X0)),X10),s(fun(X2,X1),X11),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))),s(X2,X13)))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ! [X10,X11,X12,X13] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X9,X0)),X10),s(X1,happ(s(fun(X2,X1),X11),s(X2,X13))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0)))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X0))))))),'_dst_x0000000x1000000x0111111_3'(s(fun(X1,fun(X9,X0)),X10),s(fun(X2,X1),X11),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,fun(X8,X9)))))),X12))),s(X2,X13)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X35,X65,X37,X36,X55,X57,X58,X63,X64,X62] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),'_dst_x000000x000000x111111_3'(s(fun(X65,fun(X62,X35)),X38),s(X65,happ(s(fun(X37,X65),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X62)))))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35))))))),'_dst_x0000000x1000000x0111111_3'(s(fun(X65,fun(X62,X35)),X38),s(fun(X37,X65),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X62)))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000000x1000000x0111111_ax')).
fof(f531,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X8,X0),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11),s(X9,X12))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X8,X0),X10),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11))),s(X9,X12)))) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X8,X0),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11),s(X9,X12))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X8,X0),X10),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11))),s(X9,X12)))),
  inference(flattening,[],[f354])).
fof(f354,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ! [X10,X11,X12] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X8,X0),X10),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11),s(X9,X12))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))))),happ(s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))))),'_dst_x00000000x11111111_2'(s(fun(X8,X0),X10),s(fun(X9,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8)))))))),X11))),s(X9,X12)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X62,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),'_dst_x0000000x1111111_2'(s(fun(X62,X35),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62)))))))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35))))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X35)))))))),'_dst_x00000000x11111111_2'(s(fun(X62,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,fun(X65,X62)))))))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000000x11111111_ax')).
fof(f530,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X7,X12))),s(X8,X11))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X8,X11))),s(X7,X12)))) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X7,X12))),s(X8,X11))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X8,X11))),s(X7,X12)))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ! [X9,X10,X11,X12] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X7,X12))),s(X8,X11))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111x000000_3'(s(fun(X6,fun(X8,X0)),X9),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X10),s(X8,X11))),s(X7,X12)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X37,X65] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),'_dst_x00000x11111x00000_3'(s(fun(X64,fun(X65,X35)),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64)))))),X39),s(X37,X40))),s(X65,X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35)))))),'_dst_x000000x111111x000000_3'(s(fun(X64,fun(X65,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64)))))),X39),s(X65,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000x111111x000000_ax')).
fof(f529,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11),s(X8,X12))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11))),s(X8,X12)))) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11),s(X8,X12))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11))),s(X8,X12)))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ! [X9,X10,X11,X12] : s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11),s(X8,X12))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),happ(s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x000000x111111_3'(s(fun(X1,fun(X7,X0)),X9),s(X1,X10),s(fun(X8,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),X11))),s(X8,X12)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X35,X64,X36,X55,X57,X58,X63,X65,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),'_dst_x00000x00000x11111_3'(s(fun(X64,fun(X65,X35)),X38),s(X64,X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X65))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X65)))))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35)))))),'_dst_x000000x000000x111111_3'(s(fun(X64,fun(X65,X35)),X38),s(X64,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X65)))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000x000000x111111_ax')).
fof(f528,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X8,X0)),X9),s(X1,happ(s(fun(X2,X1),X10),s(X2,X12))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x100000x011111_3'(s(fun(X1,fun(X8,X0)),X9),s(fun(X2,X1),X10),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))),s(X2,X12)))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X8,X0)),X9),s(X1,happ(s(fun(X2,X1),X10),s(X2,X12))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x100000x011111_3'(s(fun(X1,fun(X8,X0)),X9),s(fun(X2,X1),X10),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))),s(X2,X12)))),
  inference(flattening,[],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ! [X9,X10,X11,X12] : s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X8,X0)),X9),s(X1,happ(s(fun(X2,X1),X10),s(X2,X12))),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))) = s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0))))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X0)))))),'_dst_x000000x100000x011111_3'(s(fun(X1,fun(X8,X0)),X9),s(fun(X2,X1),X10),s(fun(X3,fun(X4,fun(X5,fun(X6,fun(X7,X8))))),X11))),s(X2,X12)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X35,X64,X37,X36,X55,X57,X58,X63,X65] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),'_dst_x00000x00000x11111_3'(s(fun(X64,fun(X65,X35)),X38),s(X64,happ(s(fun(X37,X64),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X65))))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35)))))),'_dst_x000000x100000x011111_3'(s(fun(X64,fun(X65,X35)),X38),s(fun(X37,X64),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X65))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000x100000x011111_ax')).
fof(f527,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111_2'(s(fun(X7,X0),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10),s(X8,X11))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X7,X0),X9),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10))),s(X8,X11)))) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111_2'(s(fun(X7,X0),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10),s(X8,X11))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X7,X0),X9),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10))),s(X8,X11)))),
  inference(flattening,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ! [X9,X10,X11] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),'_dst_x000000x111111_2'(s(fun(X7,X0),X9),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10),s(X8,X11))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0)))))),happ(s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))))),'_dst_x0000000x1111111_2'(s(fun(X7,X0),X9),s(fun(X8,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X7))))))),X10))),s(X8,X11)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X65,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),'_dst_x000000x111111_2'(s(fun(X65,X35),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65))))))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35)))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X35))))))),'_dst_x0000000x1111111_2'(s(fun(X65,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,fun(X64,X65))))))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000000x1111111_ax')).
fof(f526,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X2,fun(X3,fun(X4,fun(X5,X6)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10),s(X7,X11))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10))),s(X7,X11)))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X2,fun(X3,fun(X4,fun(X5,X6)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10),s(X7,X11))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10))),s(X7,X11)))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11] : s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X2,fun(X3,fun(X4,fun(X5,X6)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10),s(X7,X11))))) = s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),happ(s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x00000x11111_3'(s(fun(X1,fun(X6,X0)),X8),s(X1,X9),s(fun(X7,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),X10))),s(X7,X11)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X35,X63,X36,X55,X57,X58,X64,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),'_dst_x0000x0000x1111_3'(s(fun(X63,fun(X64,X35)),X38),s(X63,X39),s(fun(X36,fun(X55,fun(X57,fun(X58,X64)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X64))))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X35))))),'_dst_x00000x00000x11111_3'(s(fun(X63,fun(X64,X35)),X38),s(X63,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X64))))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000x00000x11111_ax')).
fof(f525,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X6,X11))),s(X7,X10))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X7,X10))),s(X6,X11)))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X6,X11))),s(X7,X10))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X7,X10))),s(X6,X11)))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X6,X11))),s(X7,X10))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111x00000_3'(s(fun(X5,fun(X7,X0)),X8),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X9),s(X7,X10))),s(X6,X11)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X37,X64] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),'_dst_x0000x1111x0000_3'(s(fun(X63,fun(X64,X35)),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,X63)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X63))))),X39),s(X37,X40))),s(X64,X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X35))))),'_dst_x00000x11111x00000_3'(s(fun(X63,fun(X64,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X63))))),X39),s(X64,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000x11111x00000_ax')).
fof(f524,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(X1,happ(s(fun(X2,X1),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10000x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,X1),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(X1,happ(s(fun(X2,X1),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10000x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,X1),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(flattening,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10,X11] : s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X7,X0)),X8),s(X1,happ(s(fun(X2,X1),X9),s(X2,X11))),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))) = s(fun(X3,fun(X4,fun(X5,fun(X6,X0)))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,fun(X6,X0))))),'_dst_x00000x10000x01111_3'(s(fun(X1,fun(X7,X0)),X8),s(fun(X2,X1),X9),s(fun(X3,fun(X4,fun(X5,fun(X6,X7)))),X10))),s(X2,X11)))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X35,X63,X37,X36,X55,X57,X58,X64] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),'_dst_x0000x0000x1111_3'(s(fun(X63,fun(X64,X35)),X38),s(X63,happ(s(fun(X37,X63),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,fun(X58,X64)))),X56))) = s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X35))))),'_dst_x00000x10000x01111_3'(s(fun(X63,fun(X64,X35)),X38),s(fun(X37,X63),X39),s(fun(X36,fun(X55,fun(X57,fun(X58,X64)))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000x10000x01111_ax')).
fof(f523,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111_2'(s(fun(X6,X0),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9),s(X7,X10))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111_2'(s(fun(X6,X0),X8),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9))),s(X7,X10)))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111_2'(s(fun(X6,X0),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9),s(X7,X10))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111_2'(s(fun(X6,X0),X8),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9))),s(X7,X10)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ! [X8,X9,X10] : s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),'_dst_x00000x11111_2'(s(fun(X6,X0),X8),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9),s(X7,X10))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0))))),happ(s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X0)))))),'_dst_x000000x111111_2'(s(fun(X6,X0),X8),s(fun(X7,fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,X6)))))),X9))),s(X7,X10)))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X64,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),'_dst_x00000x11111_2'(s(fun(X64,X35),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64)))))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35))))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X35)))))),'_dst_x000000x111111_2'(s(fun(X64,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,fun(X63,X64)))))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000000x111111_ax')).
fof(f522,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X2,fun(X3,fun(X4,X5))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9),s(X6,X10))))) = s(fun(X2,fun(X3,fun(X4,X0))),happ(s(fun(X6,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x0000x1111_3'(s(fun(X1,fun(X5,X0)),X7),s(X1,X8),s(fun(X6,fun(X2,fun(X3,fun(X4,X5)))),X9))),s(X6,X10)))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X35,X58,X36,X55,X57,X63,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,X35))),'_dst_x000x000x111_3'(s(fun(X58,fun(X63,X35)),X38),s(X58,X39),s(fun(X36,fun(X55,fun(X57,X63))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X63)))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,X35))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X35)))),'_dst_x0000x0000x1111_3'(s(fun(X58,fun(X63,X35)),X38),s(X58,X39),s(fun(X37,fun(X36,fun(X55,fun(X57,X63)))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000x0000x1111_ax')).
fof(f521,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111x000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X5,X10))),s(X6,X9))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X6,X9))),s(X5,X10)))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111x000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X5,X10))),s(X6,X9))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X6,X9))),s(X5,X10)))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111x000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X5,X10))),s(X6,X9))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111x0000_3'(s(fun(X4,fun(X6,X0)),X7),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X8),s(X6,X9))),s(X5,X10)))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X35,X36,X55,X57,X58,X37,X63] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,X35))),'_dst_x000x111x000_3'(s(fun(X58,fun(X63,X35)),X38),s(fun(X36,fun(X55,fun(X57,X58))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X58)))),X39),s(X37,X40))),s(X63,X56))) = s(fun(X36,fun(X55,fun(X57,X35))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X35)))),'_dst_x0000x1111x0000_3'(s(fun(X58,fun(X63,X35)),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,X58)))),X39),s(X63,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000x1111x0000_ax')).
fof(f520,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X6,X0)),X7),s(X1,happ(s(fun(X2,X1),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1000x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,X1),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X6,X0)),X7),s(X1,happ(s(fun(X2,X1),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1000x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,X1),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))),
  inference(flattening,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9,X10] : s(fun(X3,fun(X4,fun(X5,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X6,X0)),X7),s(X1,happ(s(fun(X2,X1),X8),s(X2,X10))),s(fun(X3,fun(X4,fun(X5,X6))),X9))) = s(fun(X3,fun(X4,fun(X5,X0))),happ(s(fun(X2,fun(X3,fun(X4,fun(X5,X0)))),'_dst_x0000x1000x0111_3'(s(fun(X1,fun(X6,X0)),X7),s(fun(X2,X1),X8),s(fun(X3,fun(X4,fun(X5,X6))),X9))),s(X2,X10)))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X35,X58,X37,X36,X55,X57,X63] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,fun(X57,X35))),'_dst_x000x000x111_3'(s(fun(X58,fun(X63,X35)),X38),s(X58,happ(s(fun(X37,X58),X39),s(X37,X40))),s(fun(X36,fun(X55,fun(X57,X63))),X56))) = s(fun(X36,fun(X55,fun(X57,X35))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X35)))),'_dst_x0000x1000x0111_3'(s(fun(X58,fun(X63,X35)),X38),s(fun(X37,X58),X39),s(fun(X36,fun(X55,fun(X57,X63))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000x1000x0111_ax')).
fof(f519,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8),s(X6,X9))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111_2'(s(fun(X5,X0),X7),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8))),s(X6,X9)))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8),s(X6,X9))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111_2'(s(fun(X5,X0),X7),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8))),s(X6,X9)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ! [X7,X8,X9] : s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),'_dst_x0000x1111_2'(s(fun(X5,X0),X7),s(fun(X1,fun(X2,fun(X3,fun(X4,X5)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8),s(X6,X9))))) = s(fun(X1,fun(X2,fun(X3,fun(X4,X0)))),happ(s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X0))))),'_dst_x00000x11111_2'(s(fun(X5,X0),X7),s(fun(X6,fun(X1,fun(X2,fun(X3,fun(X4,X5))))),X8))),s(X6,X9)))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X35,X36,X55,X57,X58,X63,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),'_dst_x0000x1111_2'(s(fun(X63,X35),X38),s(fun(X36,fun(X55,fun(X57,fun(X58,X63)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X63))))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,fun(X58,X35)))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X35))))),'_dst_x00000x11111_2'(s(fun(X63,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,fun(X58,X63))))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00000x11111_ax')).
fof(f518,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X2,X0)),'_dst_x00x11x00_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X4,X9))),s(X5,X8))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111x000_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X5,X8))),s(X4,X9)))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X2,X0)),'_dst_x00x11x00_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X4,X9))),s(X5,X8))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111x000_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X5,X8))),s(X4,X9)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X1,fun(X2,X0)),'_dst_x00x11x00_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X4,X9))),s(X5,X8))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111x000_3'(s(fun(X3,fun(X5,X0)),X6),s(fun(X4,fun(X1,fun(X2,X3))),X7),s(X5,X8))),s(X4,X9)))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X35,X36,X55,X57,X37,X58] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,X35)),'_dst_x00x11x00_3'(s(fun(X57,fun(X58,X35)),X38),s(fun(X36,fun(X55,X57)),happ(s(fun(X37,fun(X36,fun(X55,X57))),X39),s(X37,X40))),s(X58,X56))) = s(fun(X36,fun(X55,X35)),happ(s(fun(X37,fun(X36,fun(X55,X35))),'_dst_x000x111x000_3'(s(fun(X57,fun(X58,X35)),X38),s(fun(X37,fun(X36,fun(X55,X57))),X39),s(X58,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000x111x000_ax')).
fof(f517,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X2,fun(X3,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X2,fun(X3,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X2,fun(X3,X4)),happ(s(fun(X5,fun(X2,fun(X3,X4))),X8),s(X5,X9))))) = s(fun(X2,fun(X3,X0)),happ(s(fun(X5,fun(X2,fun(X3,X0))),'_dst_x000x000x111_3'(s(fun(X1,fun(X4,X0)),X6),s(X1,X7),s(fun(X5,fun(X2,fun(X3,X4))),X8))),s(X5,X9)))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X35,X57,X36,X55,X58,X37] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,X35)),'_dst_x00x00x11_3'(s(fun(X57,fun(X58,X35)),X38),s(X57,X39),s(fun(X36,fun(X55,X58)),happ(s(fun(X37,fun(X36,fun(X55,X58))),X56),s(X37,X40))))) = s(fun(X36,fun(X55,X35)),happ(s(fun(X37,fun(X36,fun(X55,X35))),'_dst_x000x000x111_3'(s(fun(X57,fun(X58,X35)),X38),s(X57,X39),s(fun(X37,fun(X36,fun(X55,X58))),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000x000x111_ax')).
fof(f516,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X35,X57,X37,X36,X55,X58] : ! [X38,X39,X56,X40] : s(fun(X36,fun(X55,X35)),'_dst_x00x00x11_3'(s(fun(X57,fun(X58,X35)),X38),s(X57,happ(s(fun(X37,X57),X39),s(X37,X40))),s(fun(X36,fun(X55,X58)),X56))) = s(fun(X36,fun(X55,X35)),happ(s(fun(X37,fun(X36,fun(X55,X35))),'_dst_x000x100x011_3'(s(fun(X57,fun(X58,X35)),X38),s(fun(X37,X57),X39),s(fun(X36,fun(X55,X58)),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000x100x011_ax')).
fof(f515,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8] : s(fun(X1,fun(X2,fun(X3,X0))),'_dst_x000x111_2'(s(fun(X4,X0),X6),s(fun(X1,fun(X2,fun(X3,X4))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7),s(X5,X8))))) = s(fun(X1,fun(X2,fun(X3,X0))),happ(s(fun(X5,fun(X1,fun(X2,fun(X3,X0)))),'_dst_x0000x1111_2'(s(fun(X4,X0),X6),s(fun(X5,fun(X1,fun(X2,fun(X3,X4)))),X7))),s(X5,X8)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X35,X36,X55,X57,X58,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,fun(X57,X35))),'_dst_x000x111_2'(s(fun(X58,X35),X38),s(fun(X36,fun(X55,fun(X57,X58))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X58)))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,fun(X57,X35))),happ(s(fun(X37,fun(X36,fun(X55,fun(X57,X35)))),'_dst_x0000x1111_2'(s(fun(X58,X35),X38),s(fun(X37,fun(X36,fun(X55,fun(X57,X58)))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0000x1111_ax')).
fof(f514,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X35,X55,X36,X57,X37] : ! [X38,X39,X56,X40] : s(fun(X36,X35),'_dst_x0x1x1_3'(s(fun(X55,fun(X57,X35)),X38),s(fun(X36,X55),X39),s(fun(X36,X57),happ(s(fun(X37,fun(X36,X57)),X56),s(X37,X40))))) = s(fun(X36,X35),happ(s(fun(X37,fun(X36,X35)),'_dst_x00x01x11_3'(s(fun(X55,fun(X57,X35)),X38),s(fun(X36,X55),X39),s(fun(X37,fun(X36,X57)),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00x01x11_ax')).
fof(f513,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X35,X55,X36,X57,X37] : ! [X38,X39,X56,X40] : s(fun(X36,X35),'_dst_x0x0x1_3'(s(fun(X55,fun(X57,X35)),X38),s(X55,X39),s(fun(X36,X57),happ(s(fun(X37,fun(X36,X57)),X56),s(X37,X40))))) = s(fun(X36,X35),happ(s(fun(X37,fun(X36,X35)),'_dst_x00x00x11_3'(s(fun(X55,fun(X57,X35)),X38),s(X55,X39),s(fun(X37,fun(X36,X57)),X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00x00x11_ax')).
fof(f512,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X35,X36,X55,X37,X57] : ! [X38,X39,X56,X40] : s(fun(X36,X35),'_dst_x0x1x0_3'(s(fun(X55,fun(X57,X35)),X38),s(fun(X36,X55),happ(s(fun(X37,fun(X36,X55)),X39),s(X37,X40))),s(X57,X56))) = s(fun(X36,X35),happ(s(fun(X37,fun(X36,X35)),'_dst_x00x11x00_3'(s(fun(X55,fun(X57,X35)),X38),s(fun(X37,fun(X36,X55)),X39),s(X57,X56))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00x11x00_ax')).
fof(f511,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X35,X36,X55,X57,X37] : ! [X38,X39,X40] : s(fun(X36,fun(X55,X35)),'_dst_x00x11_2'(s(fun(X57,X35),X38),s(fun(X36,fun(X55,X57)),happ(s(fun(X37,fun(X36,fun(X55,X57))),X39),s(X37,X40))))) = s(fun(X36,fun(X55,X35)),happ(s(fun(X37,fun(X36,fun(X55,X35))),'_dst_x000x111_2'(s(fun(X57,X35),X38),s(fun(X37,fun(X36,fun(X55,X57))),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x000x111_ax')).
fof(f554,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(X3,X7))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(forward_demodulation,[],[f510,f502])).
fof(f510,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X35,X36,X55,X37] : ! [X38,X39,X56,X40] : s(X35,happ(s(fun(X37,X35),'_dst_x0x1x1_3'(s(fun(X36,fun(X55,X35)),X38),s(fun(X37,X36),X39),s(fun(X37,X55),X56))),s(X37,X40))) = s(X35,happ(s(fun(X55,X35),happ(s(fun(X36,fun(X55,X35)),X38),s(X36,happ(s(fun(X37,X36),X39),s(X37,X40))))),s(X55,happ(s(fun(X37,X55),X56),s(X37,X40)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0x1x1_ax')).
fof(f553,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),'const.combin.o_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X2,X7))),s(X3,X6)))) )),
  inference(forward_demodulation,[],[f509,f502])).
fof(f509,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X35,X36,X37,X55] : ! [X38,X39,X56,X40] : s(X35,happ(s(fun(X37,X35),'_dst_x0x1x0_3'(s(fun(X36,fun(X55,X35)),X38),s(fun(X37,X36),X39),s(X55,X56))),s(X37,X40))) = s(X35,happ(s(fun(X55,X35),happ(s(fun(X36,fun(X55,X35)),X38),s(X36,happ(s(fun(X37,X36),X39),s(X37,X40))))),s(X55,X56)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0x1x0_ax')).
fof(f507,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X35,X36,X55,X37] : ! [X38,X39,X40] : s(fun(X36,X35),'_dst_x0x1_2'(s(fun(X55,X35),X38),s(fun(X36,X55),happ(s(fun(X37,fun(X36,X55)),X39),s(X37,X40))))) = s(fun(X36,X35),happ(s(fun(X37,fun(X36,X35)),'_dst_x00x11_2'(s(fun(X55,X35),X38),s(fun(X37,fun(X36,X55)),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x00x11_ax')).
fof(f505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4),s('type.bvpSem.state'(X2),X3))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4))),s('type.bvpSem.state'(X2),X3)))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4),s('type.bvpSem.state'(X2),X3))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4))),s('type.bvpSem.state'(X2),X3)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s('type.bvpSem.state'(X1),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4),s('type.bvpSem.state'(X2),X3))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X0)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X5),s(fun('type.ffi.ffi_state'(X2),'type.ffi.ffi_state'(X1)),X4))),s('type.bvpSem.state'(X2),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X23,X24,X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X23),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X24),'type.ffi.ffi_state'(X23)),X10),s('type.bvpSem.state'(X24),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X24)),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X23),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X23)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X24),'type.ffi.ffi_state'(X23)),X10),s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X24)),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|3|48|_prove_ltb')).
fof(f552,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(forward_demodulation,[],[f504,f502])).
fof(f504,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X35,X36,X37] : ! [X38,X39,X40] : s(X35,happ(s(fun(X36,X35),X38),s(X36,happ(s(fun(X37,X36),X39),s(X37,X40))))) = s(X35,happ(s(fun(X37,X35),'_dst_x0x1_2'(s(fun(X36,X35),X38),s(fun(X37,X36),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dst_x0x1_ax')).
fof(f503,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X15,X6] : ! [X11,X12,X17] : s(X14,'const.combin.o_3'(s(fun(X15,X14),X11),s(fun(X6,X15),X12),s(X6,X17))) = s(X14,happ(s(fun(X6,X14),'const.combin.o_2'(s(fun(X15,X14),X11),s(fun(X6,X15),X12))),s(X6,X17)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_2_3_const.combin.o')).
fof(f501,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X35,X36,X37] : ! [X39,X40] : s(fun(fun(X36,X35),X35),'_dstsel_0_x0_1'(s(X36,happ(s(fun(X37,X36),X39),s(X37,X40))))) = s(fun(fun(X36,X35),X35),happ(s(fun(X37,fun(fun(X36,X35),X35)),'_dstsel_1_x10_1'(s(fun(X37,X36),X39))),s(X37,X40)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','_dstsel_1_x10_ax')).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X1),s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X0),s('type.modSem.environment',X2))))) = s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'const.combin.o_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X1),s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X0))),s('type.modSem.environment',X2)))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X1),s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X0),s('type.modSem.environment',X2))))) = s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'const.combin.o_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X1),s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X0))),s('type.modSem.environment',X2)))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X9,X10,X70] : s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X10),s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X9),s('type.modSem.environment',X70))))) = s('type.modSem.environment','const.modSem.environment_c_fupd_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'const.combin.o_2'(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X10),s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),X9))),s('type.modSem.environment',X70)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.modSem.modSem_|49|59_prove_ltb')).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X2))))) = s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X2)))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X2))))) = s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X2)))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X9,X10,X70] : s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X10),s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X9),s('type.modSem.environment',X70))))) = s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X10),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X9))),s('type.modSem.environment',X70)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.modSem.modSem_|49|6|48|_prove_ltb')).
fof(f496,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvpSem.state_CASE_2'(s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X1),X9),s('type.num.num',X10))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))),happ(s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))),happ(s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))),happ(s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))))),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X1),X9))),s('type.num.num',X10)))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(X0,'const.bvpSem.state_CASE_2'(s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X1),X9),s('type.num.num',X10))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))),happ(s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))),happ(s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))),happ(s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))))),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X1),X9))),s('type.num.num',X10)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s(X0,'const.bvpSem.state_CASE_2'(s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X1),X9),s('type.num.num',X10))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))),happ(s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))),happ(s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))),happ(s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0)))))))),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X1),fun('type.num.num',X0))))))))),X11),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X1),X9))),s('type.num.num',X10)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6,X7] : ! [X46,X47,X48,X49,X50,X51,X52,X53,X54,X10] : s(X6,'const.bvpSem.state_CASE_2'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X46),s('type.list.list'('type.bvpSem.stack'),X47),s('type.option.option'('type.num.num'),X48),s('type.num.num',X49),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X50),s('type.num.num',X51),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X52),s('type.ffi.ffi_state'(X7),X53),s('type.num.num',X54))),s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6))))))))),X10))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6)),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6))),happ(s(fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6)))),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6))))),happ(s(fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6)))))),happ(s(fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6))))))),happ(s(fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6)))))))),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),fun('type.list.list'('type.bvpSem.stack'),fun('type.option.option'('type.num.num'),fun('type.num.num',fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),fun('type.num.num',fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),fun('type.ffi.ffi_state'(X7),fun('type.num.num',X6))))))))),X10),s('type.sptree.spt'('type.bvlSem.v'),X46))),s('type.list.list'('type.bvpSem.stack'),X47))),s('type.option.option'('type.num.num'),X48))),s('type.num.num',X49))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X50))),s('type.num.num',X51))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X52))),s('type.ffi.ffi_state'(X7),X53))),s('type.num.num',X54)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_case_def')).
fof(f495,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11)))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.bvpSem.state'(X1),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X1),X10),s('type.num.num',X11))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X3),s('type.list.list'('type.bvpSem.stack'),X4),s('type.option.option'('type.num.num'),X5),s('type.num.num',X6),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X7),s('type.num.num',X8),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X9),s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X2),s('type.ffi.ffi_state'(X1),X10))),s('type.num.num',X11)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X24,X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X24),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X24)),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))))) = s('type.bvpSem.state'(X24),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X24),happ(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X24)),X26),s('type.ffi.ffi_state'(X7),X32))),s('type.num.num',X33)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_ffi_fupd')).
fof(f494,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2))) = s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4))),s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2)))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2))) = s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4))),s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2)))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2))) = s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X4))),s(fun(X0,'type.ffi.ffi_state'(X1)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X1)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X3))),s(fun(X0,'type.ffi.ffi_state'(X1)),X2)))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6,X7] : ! [X8,X9,X32] : s(fun(X6,'type.ffi.ffi_state'(X7)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X7)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X32),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X9))))),s(fun(X6,'type.ffi.ffi_state'(X7)),X8))) = s(fun(X6,'type.ffi.ffi_state'(X7)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X7)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X32))),s(fun(X6,'type.ffi.ffi_state'(X7)),'const.combin.o_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X7)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X9))),s(fun(X6,'type.ffi.ffi_state'(X7)),X8)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ffi.ffi_state_fupdfupds_comp|split|7')).
fof(f493,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) = s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) = s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))))),s(fun(X0,'type.bvpSem.state'(X1)),X2))) = s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X4))),s(fun(X0,'type.bvpSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X1),'type.bvpSem.state'(X1)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3))),s(fun(X0,'type.bvpSem.state'(X1)),X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9))))),s(fun(X6,'type.bvpSem.state'(X7)),X8))) = s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10))),s(fun(X6,'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9))),s(fun(X6,'type.bvpSem.state'(X7)),X8)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_fupdfupds_comp|split|1')).
fof(f492,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f383])).
fof(f383,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3',aHL_EXT)).
fof(f490,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X1),X2))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X1),'type.ffi.ffi_state'(X0)),X3),s('type.bvpSem.state'(X1),X2)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X24,X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X24),'const.bvpSem.state_ffi_fupd_2'(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X24)),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_55_prove_ltb')).
fof(f488,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X14] : ! [X16,X20] : s(X14,X20) = s(X14,'const.pair.SND_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X16),s(X14,X20)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.pair.SND')).
fof(f485,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X14,X6] : ! [X42] : ? [X10] : ! [X43,X44,X45] : s(X14,happ(s(fun('type.ind_type.recspace'(X6),X14),X10),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X43),s(X6,X44),s(fun('type.num.num','type.ind_type.recspace'(X6)),X45))))) = s(X14,happ(s(fun(fun('type.num.num',X14),X14),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X14),X14)))),X42),s('type.num.num',X43))),s(X6,X44))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X45))),s(fun('type.num.num',X14),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X14),X10),s(fun('type.num.num','type.ind_type.recspace'(X6)),X45)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ind_type.CONSTR_REC')).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X14] : ! [X11] : s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),X11))) = s(X6,happ(s(fun('type.pair.prod'(X6,X14),X6),'const.pair.FST_0'),s('type.pair.prod'(X6,X14),X11)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_0_1_const.pair.FST')).
fof(f472,plain,(
  ( ! [X0,X1] : (s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))) = s('type.modSem.environment',happ(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1)))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.modSem.environment',X1))) = s('type.modSem.environment',happ(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.modSem.environment',X1)))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X11,X12] : s('type.modSem.environment','const.modSem.environment_v_fupd_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X11),s('type.modSem.environment',X12))) = s('type.modSem.environment',happ(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X11))),s('type.modSem.environment',X12)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_2_const.modSem.environment_v_fupd')).
fof(f551,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f550,f471])).
fof(f550,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f549,f471])).
fof(f549,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f548,f471])).
fof(f548,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f547,f471])).
fof(f547,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f546,f471])).
fof(f546,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f545,f471])).
fof(f545,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(forward_demodulation,[],[f544,f471])).
fof(f544,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(backward_demodulation,[],[f471,f469])).
fof(f471,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X11,X12] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X11))),s('type.num.num',X12)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_2_const.arithmetic.+')).
fof(f469,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X2))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X3))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X4))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X5))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X7))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ffi.ffi_state'(X0),X9))))),s('type.num.num',X10)))))))))))))))))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X7] : ! [X10,X46,X47,X48,X49,X50,X51,X52,X53,X54] : s('type.num.num','const.bvpSem.state_size_2'(s(fun(X7,'type.num.num'),X10),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X46),s('type.list.list'('type.bvpSem.stack'),X47),s('type.option.option'('type.num.num'),X48),s('type.num.num',X49),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X50),s('type.num.num',X51),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X52),s('type.ffi.ffi_state'(X7),X53),s('type.num.num',X54))))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.sptree.spt'('type.bvlSem.v'),X46))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.list.list_size_2'(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.list.list'('type.bvpSem.stack'),X47))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.option.option'('type.num.num'),X48))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X49))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.finite_map.fmap_size_3'(s(fun('type.num.num','type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'_dst_x0x1_2'(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(fun('type.closSem.ref'('type.bvlSem.v'),'type.num.num'),'const.closSem.ref_size_1'(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'))))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X50))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X51))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.sptree.spt_size_2'(s(fun('type.pair.prod'('type.num.num','type.bvp.prog'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.num.num','type.num.num'),'_i_0_0'),s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X52))))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num','const.ffi.ffi_state_size_2'(s(fun(X7,'type.num.num'),X10),s('type.ffi.ffi_state'(X7),X53))))),s('type.num.num',X54)))))))))))))))))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_size_def')).
fof(f468,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X1),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8))),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))))) = s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),happ(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X26),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31))),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_code_fupd')).
fof(f467,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X1),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6))),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))))) = s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),happ(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10))),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_refs_fupd')).
fof(f466,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.option.option'('type.num.num'),X4))),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),happ(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X26),s('type.option.option'('type.num.num'),X28))),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_global_fupd')).
fof(f465,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.list.list'('type.bvpSem.stack'),X3))),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X26),s('type.list.list'('type.bvpSem.stack'),X27))),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_stack_fupd')).
fof(f464,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.sptree.spt'('type.bvlSem.v'),X2))),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),happ(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X26),s('type.sptree.spt'('type.bvlSem.v'),X13))),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_locals_fupd')).
fof(f463,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X7))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X26),s('type.num.num',X30))),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_clock_fupd')).
fof(f462,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X10))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X26),s('type.num.num',X33))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_space_fupd')).
fof(f461,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X1),s('type.num.num',X5))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X1),s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X2),s('type.list.list'('type.bvpSem.stack'),X3),s('type.option.option'('type.num.num'),X4),s('type.num.num',X5),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X7),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X8),s('type.ffi.ffi_state'(X0),X9),s('type.num.num',X10)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7] : ! [X26,X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X26),s('type.num.num',X29))),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X26),s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_handler_fupd')).
fof(f460,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_code')).
fof(f459,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_refs')).
fof(f458,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.sptree.spt'('type.bvlSem.v'),X1) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.sptree.spt'('type.bvlSem.v'),X13) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_locals')).
fof(f457,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.option.option'('type.num.num'),X3) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.option.option'('type.num.num'),X28) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_global')).
fof(f456,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.list.list'('type.bvpSem.stack'),X2) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.list.list'('type.bvpSem.stack'),X27) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_stack')).
fof(f455,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.ffi.ffi_state'(X0),X8) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.ffi.ffi_state'(X7),X32) = s('type.ffi.ffi_state'(X7),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_ffi')).
fof(f454,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X9) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.num.num',X33) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_space')).
fof(f453,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X6) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.num.num',X30) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_clock')).
fof(f452,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s('type.num.num',X4) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X1),s('type.list.list'('type.bvpSem.stack'),X2),s('type.option.option'('type.num.num'),X3),s('type.num.num',X4),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X6),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X7),s('type.ffi.ffi_state'(X0),X8),s('type.num.num',X9)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X7] : ! [X13,X27,X28,X29,X10,X30,X31,X32,X33] : s('type.num.num',X29) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X7),'const.bvpSem.recordtype|46|state_9'(s('type.sptree.spt'('type.bvlSem.v'),X13),s('type.list.list'('type.bvpSem.stack'),X27),s('type.option.option'('type.num.num'),X28),s('type.num.num',X29),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X10),s('type.num.num',X30),s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),X31),s('type.ffi.ffi_state'(X7),X32),s('type.num.num',X33)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_handler')).
fof(f451,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3))),s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))),s(fun(X0,'type.modSem.environment'),X1))))) = s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))))),s(fun(X0,'type.modSem.environment'),X1)))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3))),s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))),s(fun(X0,'type.modSem.environment'),X1))))) = s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))))),s(fun(X0,'type.modSem.environment'),X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : ! [X1,X2,X3] : s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3))),s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))),s(fun(X0,'type.modSem.environment'),X1))))) = s(fun(X0,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X2))))),s(fun(X0,'type.modSem.environment'),X1)))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X6] : ! [X8,X9,X10] : s(fun(X6,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X10))),s(fun(X6,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X9))),s(fun(X6,'type.modSem.environment'),X8))))) = s(fun(X6,'type.modSem.environment'),'const.combin.o_2'(s(fun('type.modSem.environment','type.modSem.environment'),'const.modSem.environment_v_fupd_1'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X10),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v')),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X9))))),s(fun(X6,'type.modSem.environment'),X8)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.modSem.environment_fupdfupds_comp|split|3')).
fof(f450,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X3),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X10),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|29_prove_ltb')).
fof(f449,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X3),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),'const.combin.o_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|27_prove_ltb')).
fof(f448,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X2),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X7] : ! [X9,X32,X10] : s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X32),s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X9),s('type.ffi.ffi_state'(X7),X10))))) = s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),'const.combin.o_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X32),s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X9))),s('type.ffi.ffi_state'(X7),X10)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ffi.ffi_|49||48||48|_prove_ltb')).
fof(f447,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),'const.combin.o_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),'const.combin.o_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),'const.combin.o_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X2),s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X7] : ! [X9,X32,X10] : s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X32),s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X9),s('type.ffi.ffi_state'(X7),X10))))) = s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_final_event_fupd_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),'const.combin.o_2'(s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X32),s(fun('type.option.option'('type.ffi.final_event'),'type.option.option'('type.ffi.final_event')),X9))),s('type.ffi.ffi_state'(X7),X10)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ffi.ffi_99_prove_ltb')).
fof(f446,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X3),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|23_prove_ltb')).
fof(f445,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.o_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.o_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.o_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X3),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),'const.combin.o_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X10),s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|25_prove_ltb')).
fof(f444,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X3),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),'const.combin.o_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|24_prove_ltb')).
fof(f443,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X11,X12,X17] : s(X6,'const.ind_type.FCONS_3'(s(X6,X11),s(fun('type.num.num',X6),X12),s('type.num.num',X17))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X11))),s(fun('type.num.num',X6),X12))),s('type.num.num',X17)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_0_3_const.ind_type.FCONS')).
fof(f442,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X11,X12,X17] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X11),s(X6,X12),s(fun('type.num.num','type.ind_type.recspace'(X6)),X17))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X11))),s(X6,X12))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X17)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_3_const.ind_type.CONSTR')).
fof(f441,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X10),s(fun('type.num.num','type.num.num'),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|3|49|_prove_ltb')).
fof(f440,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))),s('type.ffi.ffi_state'(X0),X3)))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X2),s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),X1),s('type.ffi.ffi_state'(X0),X3))))) = s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))),s('type.ffi.ffi_state'(X0),X3)))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X7] : ! [X9,X32,X10] : s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X7,X7),X32),s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X7,X7),X9),s('type.ffi.ffi_state'(X7),X10))))) = s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_ffi_state_fupd_2'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X32),s(fun(X7,X7),X9))),s('type.ffi.ffi_state'(X7),X10)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ffi.ffi_98_prove_ltb')).
fof(f439,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X10),s(fun('type.num.num','type.num.num'),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|26_prove_ltb')).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X3),s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1))))) = s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))),s('type.bvpSem.state'(X0),X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X7] : ! [X13,X9,X10] : s('type.bvpSem.state'(X7),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X9),s('type.bvpSem.state'(X7),X13))))) = s('type.bvpSem.state'(X7),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X10),s(fun('type.num.num','type.num.num'),X9))),s('type.bvpSem.state'(X7),X13)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49|28_prove_ltb')).
fof(f437,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X41,X10,X29] : s(X6,happ(s(fun('type.num.num',X6),X10),s('type.num.num',X29))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X41),s(fun('type.num.num',X6),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X29)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ind_type.FCONS|split|1')).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1)))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))))) = s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2))),s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),'const.combin.o_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10),s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9))))) = s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10))),s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X9)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.state_fupdfupds_comp|split|0')).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),happ(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49||49|4_prove_ltb')).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ! [X1,X2] : s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X1))),s('type.bvpSem.state'(X0),X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X11,X12] : s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X11),s('type.bvpSem.state'(X7),X12))) = s('type.bvpSem.state'(X7),happ(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_stack_fupd_1'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X11))),s('type.bvpSem.state'(X7),X12)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_2_const.bvpSem.state_stack_fupd')).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))) = s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))) = s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1),s('type.ffi.ffi_state'(X0),X2))) = s('type.ffi.ffi_state'(X0),happ(s(fun('type.ffi.ffi_state'(X0),'type.ffi.ffi_state'(X0)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X1))),s('type.ffi.ffi_state'(X0),X2)))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X7] : ! [X11,X12] : s('type.ffi.ffi_state'(X7),'const.ffi.ffi_state_io_events_fupd_2'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X11),s('type.ffi.ffi_state'(X7),X12))) = s('type.ffi.ffi_state'(X7),happ(s(fun('type.ffi.ffi_state'(X7),'type.ffi.ffi_state'(X7)),'const.ffi.ffi_state_io_events_fupd_1'(s(fun('type.list.list'('type.ffi.io_event'),'type.list.list'('type.ffi.io_event')),X11))),s('type.ffi.ffi_state'(X7),X12)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_2_const.ffi.ffi_state_io_events_fupd')).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0] : ! [X1,X2] : s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1),s('type.bvpSem.state'(X0),X2))) = s('type.bvpSem.state'(X0),happ(s(fun('type.bvpSem.state'(X0),'type.bvpSem.state'(X0)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X1))),s('type.bvpSem.state'(X0),X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X11,X12] : s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X11),s('type.bvpSem.state'(X7),X12))) = s('type.bvpSem.state'(X7),happ(s(fun('type.bvpSem.state'(X7),'type.bvpSem.state'(X7)),'const.bvpSem.state_locals_fupd_1'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X11))),s('type.bvpSem.state'(X7),X12)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_1_2_const.bvpSem.state_locals_fupd')).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X7] : ! [X13,X10] : s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X7),X13))) = s('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'const.bvpSem.state_code_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_9|48|_prove_ltb')).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_code_fupd_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvp.prog'))),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_54_prove_ltb')).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_refs_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v'))),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_52_prove_ltb')).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0] : ! [X1,X2] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X7] : ! [X13,X10] : s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X7),X13))) = s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),'const.bvpSem.state_refs_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_74_prove_ltb')).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0] : ! [X1,X2] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),X1))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X7] : ! [X13,X10] : s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X7),X13))) = s('type.sptree.spt'('type.bvlSem.v'),'const.bvpSem.state_locals_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_4|49|_prove_ltb')).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_global_fupd_2'(s(fun('type.option.option'('type.num.num'),'type.option.option'('type.num.num')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_5|48|_prove_ltb')).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0] : ! [X1,X2] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),X1))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X7] : ! [X13,X10] : s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X7),X13))) = s('type.option.option'('type.num.num'),'const.bvpSem.state_global_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_58_prove_ltb')).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_locals_fupd_2'(s(fun('type.sptree.spt'('type.bvlSem.v'),'type.sptree.spt'('type.bvlSem.v')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_49_prove_ltb')).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : ! [X1,X2] : s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),X1))) = s('type.ffi.ffi_state'(X0),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X7] : ! [X13,X10] : s('type.ffi.ffi_state'(X7),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X7),X13))) = s('type.ffi.ffi_state'(X7),'const.bvpSem.state_ffi_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_98_prove_ltb')).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_space_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_56_prove_ltb')).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X7] : ! [X13,X10] : s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X7),X13))) = s('type.num.num','const.bvpSem.state_clock_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_82_prove_ltb')).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X7] : ! [X13,X10] : s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X7),X13))) = s('type.num.num','const.bvpSem.state_handler_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_66_prove_ltb')).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_53_prove_ltb')).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),X1))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X7] : ! [X13,X10] : s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X7),X13))) = s('type.num.num','const.bvpSem.state_space_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_stack_fupd_2'(s(fun('type.list.list'('type.bvpSem.stack'),'type.list.list'('type.bvpSem.stack')),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_|49||48|6_prove_ltb')).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),X1))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X0),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.bvpSem.state'(X0),X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X7] : ! [X13,X10] : s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),X13))) = s('type.list.list'('type.bvpSem.stack'),'const.bvpSem.state_stack_1'(s('type.bvpSem.state'(X7),'const.bvpSem.state_handler_fupd_2'(s(fun('type.num.num','type.num.num'),X10),s('type.bvpSem.state'(X7),X13)))))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.bvpSem.bvpSem_5|49|_prove_ltb')).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X41,X10] : s(X6,X41) = s(X6,'const.ind_type.FCONS_3'(s(X6,X41),s(fun('type.num.num',X6),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','thm.ind_type.FCONS|split|0')).
fof(f406,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X0))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X0)))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X11] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v',X11))) = s('type.num.num',happ(s(fun('type.bvlSem.v','type.num.num'),'const.bvlSem.v_size_0'),s('type.bvlSem.v',X11)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_0_1_const.bvlSem.v_size')).
fof(f405,plain,(
  ( ! [X0] : (s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X0))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X0)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X11] : s('type.num.num','const.bvp.prog_size_1'(s('type.bvp.prog',X11))) = s('type.num.num',happ(s(fun('type.bvp.prog','type.num.num'),'const.bvp.prog_size_0'),s('type.bvp.prog',X11)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_0_1_const.bvp.prog_size')).
fof(f404,plain,(
  ( ! [X0] : (s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X0))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X0)))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X11] : s('type.num.num','const.bvpSem.stack_size_1'(s('type.bvpSem.stack',X11))) = s('type.num.num',happ(s(fun('type.bvpSem.stack','type.num.num'),'const.bvpSem.stack_size_0'),s('type.bvpSem.stack',X11)))),
  file('thm.bvpSem.state_fupdfupds_comp|split|3','ARITY_0_1_const.bvpSem.stack_size')).
fof(f18818,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8('_i_0_0',bool))) ) | $spl58),
  inference(cnf_transformation,[],[f18818_D])).
fof(f18818_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8('_i_0_0',bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f18819,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl60),
  inference(cnf_transformation,[],[f18819_D])).
fof(f18819_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f17974,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK8('_i_0_0',bool)) | s(bool,X0) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | ($spl3 | $spl22)),
  inference(equality_factoring,[],[f9095])).
fof(f18665,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,sK8('_i_0_0',bool))) ) | $spl56),
  inference(cnf_transformation,[],[f18665_D])).
fof(f18665_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,sK8('_i_0_0',bool))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f18664,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl8 | $spl22)),
  inference(condensation,[],[f18663])).
fof(f18663,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'F_0') = s(bool,X84) | s(bool,X83) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl8 | $spl22)),
  inference(subsumption_resolution,[],[f17673,f1940])).
fof(f17673,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'F_0') = s(bool,X84) | s(bool,X83) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X82)) ) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f1940,f9095])).
fof(f18266,plain,(
  ( ! [X47,X45] : (s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,X45) = s(bool,X47) | s(bool,X47) = s(bool,sK8('_i_0_0',bool))) ) | $spl54),
  inference(cnf_transformation,[],[f18266_D])).
fof(f18266_D,plain,(
  ( ! [X47,X45] : (s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,X45) = s(bool,X47) | s(bool,X47) = s(bool,sK8('_i_0_0',bool))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f18143,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,sK8('_i_0_0',bool)) | s(bool,X42) = s(bool,X44) | s(bool,X42) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl22)),
  inference(duplicate_literal_removal,[],[f17271])).
fof(f17271,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,sK8('_i_0_0',bool)) | s(bool,X42) = s(bool,X44) | s(bool,X42) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X43)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f9095,f9095])).
fof(f18153,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,sK8('_i_0_0',bool))) ) | ($spl3 | $spl22)),
  inference(duplicate_literal_removal,[],[f17182])).
fof(f17182,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X51)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f9095,f9095])).
fof(f18193,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,X47) = s(bool,sK8('_i_0_0',bool)) | s(bool,X45) = s(bool,X47)) ) | ($spl3 | $spl22)),
  inference(duplicate_literal_removal,[],[f17099])).
fof(f17099,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,sK8('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,sK8('_i_0_0',bool)) | s(bool,X45) = s(bool,X47)) ) | ($spl3 | $spl22)),
  inference(superposition,[],[f9095,f9095])).
fof(f12521,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | $spl52),
  inference(cnf_transformation,[],[f12521_D])).
fof(f12521_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f11659,plain,(
  s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl37),
  inference(cnf_transformation,[],[f11659_D])).
fof(f11659_D,plain,(
  s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f12401,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | $spl40),
  inference(equality_factoring,[],[f11663])).
fof(f12477,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | ($spl3 | $spl8 | $spl40)),
  inference(condensation,[],[f12476])).
fof(f12476,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) ) | ($spl3 | $spl8 | $spl40)),
  inference(subsumption_resolution,[],[f12276,f1940])).
fof(f12276,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | ($spl3 | $spl8 | $spl40)),
  inference(superposition,[],[f1940,f11663])).
fof(f11686,plain,(
  $false | ($spl3 | $spl36 | $spl46)),
  inference(subsumption_resolution,[],[f11685,f563])).
fof(f11685,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl36 | $spl46)),
  inference(backward_demodulation,[],[f11660,f11669])).
fof(f11669,plain,(
  s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl46),
  inference(cnf_transformation,[],[f11669_D])).
fof(f11669_D,plain,(
  s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f11665,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) | $spl42),
  inference(cnf_transformation,[],[f11665_D])).
fof(f11665_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f11673,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | ($spl3 | $spl8 | $spl22)),
  inference(forward_demodulation,[],[f11631,f8998])).
fof(f11631,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,sK8('_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | ($spl3 | $spl8 | $spl22)),
  inference(superposition,[],[f10683,f9019])).
fof(f11660,plain,(
  s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl36),
  inference(cnf_transformation,[],[f11660_D])).
fof(f11660_D,plain,(
  s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f11662,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) | $spl38),
  inference(cnf_transformation,[],[f11662_D])).
fof(f11662_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f11672,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | ($spl3 | $spl24)),
  inference(forward_demodulation,[],[f11630,f8583])).
fof(f11630,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,sK7('_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | ($spl3 | $spl24)),
  inference(superposition,[],[f10683,f8653])).
fof(f11670,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X5)) ) | $spl48),
  inference(cnf_transformation,[],[f11670_D])).
fof(f11670_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11629,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X8) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f10683,f1556])).
fof(f11671,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X6)))) | s(bool,'F_0') = s(bool,X6)) ) | $spl50),
  inference(cnf_transformation,[],[f11671_D])).
fof(f11671_D,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X6)))) | s(bool,'F_0') = s(bool,X6)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f11628,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X7)))) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) ) | $spl8),
  inference(superposition,[],[f10683,f1556])).
fof(f11627,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X6)))) | s(bool,'F_0') = s(bool,X6) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) ) | $spl8),
  inference(superposition,[],[f10683,f1556])).
fof(f11626,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X5) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f10683,f1556])).
fof(f11667,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl8),
  inference(forward_demodulation,[],[f11625,f829])).
fof(f11625,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | $spl8),
  inference(superposition,[],[f10683,f1534])).
fof(f11663,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(cnf_transformation,[],[f11663_D])).
fof(f11663_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f11624,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X3) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f10683,f555])).
fof(f11666,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f11666_D])).
fof(f11666_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f11623,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f10683,f555])).
fof(f11622,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f10683,f555])).
fof(f11621,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f10683,f555])).
fof(f11620,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK8('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK8(sK8('?_1',fun(bool,bool)),bool))),
  inference(superposition,[],[f10683,f407])).
fof(f9154,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) ) | $spl34),
  inference(cnf_transformation,[],[f9154_D])).
fof(f9154_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9162,plain,(
  ( ! [X59,X58] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X59)))) | p(s(bool,X59)) | s(bool,'F_0') = s(bool,X58)) ) | $spl6),
  inference(forward_demodulation,[],[f9161,f898])).
fof(f9161,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X59)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X57,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X59)))) | s(bool,'F_0') = s(bool,X58)) ) | $spl6),
  inference(subsumption_resolution,[],[f9124,f592])).
fof(f9124,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X57,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X59)))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f481,f1535])).
fof(f9156,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) )),
  inference(condensation,[],[f9155])).
fof(f9155,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f9117,f593])).
fof(f9117,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X42)))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f481,f555])).
fof(f9153,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) )),
  inference(condensation,[],[f9152])).
fof(f9152,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X39)))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f9116,f593])).
fof(f9116,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X39)))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f481,f555])).
fof(f8951,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('_i_0_0',bool))) ) | $spl32),
  inference(cnf_transformation,[],[f8951_D])).
fof(f8951_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7('_i_0_0',bool))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8913,plain,(
  s(bool,'T_0') != s(bool,sK7('_i_0_0',bool)) | $spl31),
  inference(cnf_transformation,[],[f8913_D])).
fof(f8913_D,plain,(
  s(bool,'T_0') != s(bool,sK7('_i_0_0',bool)) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f8843,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK7('_i_0_0',bool)) | s(bool,X0) = s(bool,sK7('_i_0_0',bool))) ) | ($spl3 | $spl24)),
  inference(equality_factoring,[],[f8653])).
fof(f8911,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,sK7('_i_0_0',bool))) ) | ($spl3 | $spl8 | $spl24)),
  inference(condensation,[],[f8910])).
fof(f8910,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,sK7('_i_0_0',bool)) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X52)) ) | ($spl3 | $spl8 | $spl24)),
  inference(subsumption_resolution,[],[f8742,f1940])).
fof(f8742,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,sK7('_i_0_0',bool)) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X51)) ) | ($spl3 | $spl8 | $spl24)),
  inference(superposition,[],[f1940,f8653])).
fof(f8585,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl26),
  inference(cnf_transformation,[],[f8585_D])).
fof(f8585_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8590,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) | $spl29),
  inference(cnf_transformation,[],[f8590_D])).
fof(f8590_D,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f8566,plain,(
  ~p(s(bool,sK7('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f478,f408])).
fof(f8583,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) | $spl24),
  inference(cnf_transformation,[],[f8583_D])).
fof(f8583_D,plain,(
  s(bool,'F_0') = s(bool,sK7('_i_0_0',bool)) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f8581,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7('_i_0_0',bool))),
  inference(subsumption_resolution,[],[f8552,f401])).
fof(f8552,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7('_i_0_0',bool))),
  inference(superposition,[],[f478,f571])).
fof(f6778,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl22),
  inference(cnf_transformation,[],[f6778_D])).
fof(f6778_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6759,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f480,f408])).
fof(f6776,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f6746,f401])).
fof(f6746,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f480,f571])).
fof(f5084,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f5042,f403])).
fof(f5042,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f563,f3621,f473])).
fof(f5083,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f5048,f403])).
fof(f5048,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f563,f3621,f473])).
fof(f5049,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f403,f563,f3621,f473])).
fof(f5047,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f3621,f563,f3621,f473])).
fof(f5043,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f403,f563,f3621,f473])).
fof(f5041,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f3621,f563,f3621,f473])).
fof(f5037,plain,(
  $false | ($spl3 | $spl14)),
  inference(unit_resulting_resolution,[],[f563,f563,f3621,f593])).
fof(f5033,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f401,f3621])).
fof(f5012,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4872,f4902])).
fof(f4902,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f4816,f403])).
fof(f4816,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,X1))) ) | $spl0),
  inference(backward_demodulation,[],[f561,f483])).
fof(f4872,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f561,f470])).
fof(f5011,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4832,f4902])).
fof(f4832,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(backward_demodulation,[],[f561,f3425])).
fof(f4901,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4893,f403])).
fof(f4893,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(condensation,[],[f4815])).
fof(f4815,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl0),
  inference(backward_demodulation,[],[f561,f482])).
fof(f4783,plain,(
  ( ! [X105,X102,X104] : (s(bool,X104) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X105) | s(bool,X104) = s(bool,X105)) ) | $spl20),
  inference(cnf_transformation,[],[f4783_D])).
fof(f4783_D,plain,(
  ( ! [X105,X102,X104] : (s(bool,X104) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X105) | s(bool,X104) = s(bool,X105)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4782,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,X104) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'))) | s(bool,X104) = s(bool,X105) | s(bool,'T_0') = s(bool,X105) | s(bool,'F_0') = s(bool,X103)) ) | ($spl3 | $spl6 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f4745,f592])).
fof(f4745,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X103)) | s(bool,X104) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'))) | s(bool,X104) = s(bool,X105) | s(bool,'T_0') = s(bool,X105) | s(bool,'F_0') = s(bool,X103)) ) | ($spl3 | $spl8 | $spl12)),
  inference(superposition,[],[f1409,f3679])).
fof(f3679,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl3 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f3629,f563])).
fof(f3629,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl8 | $spl12)),
  inference(superposition,[],[f3620,f1556])).
fof(f4778,plain,(
  ( ! [X88,X89,X86] : (s(bool,X88) = s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X89) | s(bool,X88) = s(bool,X89)) ) | $spl18),
  inference(cnf_transformation,[],[f4778_D])).
fof(f4778_D,plain,(
  ( ! [X88,X89,X86] : (s(bool,X88) = s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X89) | s(bool,X88) = s(bool,X89)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4777,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X87)) ) | $spl6),
  inference(subsumption_resolution,[],[f4741,f592])).
fof(f4741,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X87)) | s(bool,X88) = s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f1409,f1535])).
fof(f4773,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X78) | s(bool,X77) = s(bool,X78)) ) | $spl16),
  inference(cnf_transformation,[],[f4773_D])).
fof(f4773_D,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X78) | s(bool,X77) = s(bool,X78)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4772,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X76)) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f4737,f592])).
fof(f4737,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,X76)) | s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X76)) ) | $spl8),
  inference(superposition,[],[f1409,f1534])).
fof(f3620,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0')))) ) | $spl12),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0')))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3615,plain,(
  ( ! [X31,X32] : (~p(s(bool,X32)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X31),bool),'const.pair.FST_0')))) )),
  inference(superposition,[],[f2975,f486])).
fof(f815,plain,(
  ( ! [X3] : (~p(s(bool,sK5(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f434,f408])).
fof(f613,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f598,f401])).
fof(f598,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f591,f403])).
fof(f592,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f595,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f587,f557])).
fof(f587,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f473,f407])).
fof(f590,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f584,f557])).
fof(f584,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f473,f407])).
fof(f563,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f560,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f407])).
% SZS output end Proof for thm.bvpSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 278
% Passive clauses: 2774
% Generated clauses: 46283
% Final active clauses: 157
% Final passive clauses: 1385
% Input formulas: 132
% Initial clauses: 143
% 
% Duplicate literals: 1515
% Trivial inequalities: 193
% Fw subsumption resolutions: 1697
% Fw demodulations: 6968
% Bw demodulations: 968
% Condensations: 2221
% Global subsumptions: 1
% 
% Simple tautologies: 243
% Equational tautologies: 241
% Forward subsumptions: 30326
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 1565
% Factoring: 1
% Forward superposition: 15233
% Backward superposition: 14005
% Self superposition: 372
% Equality factoring: 158
% Equality resolution: 16
% 
% Split clauses: 78
% Split components: 156
% SAT solver clauses: 2965
% SAT solver unit clauses: 558
% SAT solver binary clauses: 1075
% 
% Sat splits: 78
% Sat splitting refutations: 32
% 
% Memory used [KB]: 44647
% Time elapsed: 1.741 s
% ------------------------------
% ------------------------------
% Success in time 3.1 s
