Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modSem
% (20639)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 3032
% Generated clauses: 14856
% Final active clauses: 147
% Final passive clauses: 2718
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 489
% Trivial inequalities: 383
% Fw subsumption resolutions: 1373
% Fw demodulations: 2165
% Bw demodulations: 6
% 
% Simple tautologies: 408
% Equational tautologies: 45
% Forward subsumptions: 6641
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 288
% Factoring: 1
% Forward superposition: 5380
% Backward superposition: 4332
% Self superposition: 115
% Equality factoring: 66
% Equality resolution: 40
% 
% Split clauses: 175
% Split components: 361
% SAT solver clauses: 407
% SAT solver unit clauses: 11
% SAT solver binary clauses: 92
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 175
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 406
% TWLsolver calls for satisfiability: 138
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modSem
% (20666)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 1368
% Generated clauses: 14265
% Final active clauses: 123
% Final passive clauses: 1194
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 443
% Trivial inequalities: 271
% Fw subsumption resolutions: 972
% Fw demodulations: 1823
% Bw demodulations: 43
% Condensations: 731
% Global subsumptions: 8
% 
% Simple tautologies: 239
% Equational tautologies: 43
% Forward subsumptions: 6157
% 
% Binary resolution: 185
% Factoring: 1
% Forward superposition: 4166
% Backward superposition: 3879
% Self superposition: 117
% Equality factoring: 61
% Equality resolution: 24
% 
% Split clauses: 53
% Split components: 106
% SAT solver clauses: 2243
% SAT solver unit clauses: 155
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 53
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2209
% TWLsolver calls for satisfiability: 656
% 
% Memory used [KB]: 18805
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.modSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modSem
% SZS output start Proof for thm.modSem
fof(f30831,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f429,f427,f260,f428,f426,f431,f1384,f1905,f1908,f1910,f1916,f1918,f1921,f1923,f1926,f1928,f1930,f1932,f1934,f1936,f1938,f1954,f1956,f1958,f1960,f1964,f1966,f1968,f1970,f1972,f1974,f1977,f2044,f2045,f2046,f2047,f5078,f5085,f5644,f5645,f5648,f5649,f6041,f6050,f6043,f6555,f6573,f6571,f8232,f8243,f8242,f13178,f13252,f13330,f13331,f13332,f13333,f14071,f14073,f14081,f14083,f14091,f14092,f14093,f14135,f14136,f14137,f14140,f14145,f14148,f14182,f14184,f14213,f14256,f14259,f14260,f14261,f14262,f15897,f15947,f20201,f20274,f18551,f17806,f4759,f1906,f4779,f6042,f20538,f20539,f2265,f1395,f1394,f23586,f23686,f13251,f30705,f30722])).
fof(f30722,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30721,f429])).
fof(f30721,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f30364,f309])).
fof(f309,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.modSem.v_|49||49||split|4','_i0_ax')).
fof(f30364,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f309,f30096])).
fof(f30096,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f29645,f327])).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_case_def|split|4')).
fof(f29645,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X1),s(fun('type.list.list'('type.modSem.v'),X0),X6)))) ) | $spl0),
  inference(superposition,[],[f327,f426])).
fof(f30705,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f30704,f429])).
fof(f30704,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f30339,f309])).
fof(f30339,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f30096,f309])).
fof(f13251,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) | $spl30),
  inference(cnf_transformation,[],[f13251_D])).
fof(f13251_D,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f23686,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) | $spl44),
  inference(cnf_transformation,[],[f23686_D])).
fof(f23686_D,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f23586,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(duplicate_literal_removal,[],[f23410])).
fof(f23410,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f20350,f20869])).
fof(f20869,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(trivial_inequality_removal,[],[f20854])).
fof(f20854,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f20348,f445])).
fof(f445,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.modSem.v_|49||49||split|4',aHL_BOOL_CASES)).
fof(f20348,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26 | $spl34)),
  inference(global_subsumption,[],[f20291,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1497,f1391,f1390,f1389,f1388,f1387,f1533,f1386,f1385,f319,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2228,f2229,f2230,f2231,f2232,f2524,f2730,f2704,f2803,f3321,f3322,f3415,f3633,f4107,f4527,f4528,f4687,f4610,f4622,f4691,f4761,f4847,f4891,f4890,f4884,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5307,f5308,f5318,f5457,f5462,f5473,f5475,f5478,f5495,f5535,f5540,f5551,f5553,f5556,f5631,f5642,f5644,f5646,f5647,f5651,f6041,f6050,f6283,f6500,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6592,f6593,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6622,f6623,f6634,f6635,f6636,f6639,f6640,f6651,f6652,f6653,f6654,f7121,f7122,f7552,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8242,f8270,f8271,f8367,f8368,f8378,f8632,f8798,f8799,f8725,f8731,f8736,f8845,f8848,f8851,f8853,f8854,f8855,f8856,f8857,f9233,f9236,f9239,f9241,f9242,f9243,f9244,f9245,f9734,f9694,f9697,f9699,f9700,f9701,f9702,f9703,f9829,f10037,f9997,f10000,f10002,f10003,f10004,f10005,f10006,f10136,f10270,f10431,f10655,f10983,f11095,f11096,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12464,f12433,f12574,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12745,f12748,f12764,f12766,f12767,f12799,f12802,f12818,f12820,f12821,f12831,f12833,f12834,f12906,f12909,f12912,f12914,f12915,f12916,f12917,f12918,f12919,f13330,f13332,f6915,f7141,f8233,f8238,f13336,f264,f275,f276,f472,f513,f540,f633,f686,f743,f909,f1378,f1488,f2371,f2417,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2529,f2530,f2531,f2689,f2913,f2915,f2917,f2949,f2951,f2965,f3320,f3328,f3337,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3718,f3728,f3729,f4073,f4078,f4081,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4520,f4526,f4534,f4535,f4577,f4578,f4588,f4590,f4591,f4595,f4596,f4606,f4685,f4784,f4844,f4866,f4936,f5306,f5314,f5315,f5423,f5424,f5434,f5436,f5437,f5441,f5442,f5452,f5502,f5503,f5513,f5515,f5516,f5520,f5521,f5531,f5571,f5634,f5641,f5643,f5655,f6047,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6293,f6311,f6319,f6320,f6352,f6356,f6360,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6510,f6590,f6591,f6599,f6600,f6621,f6629,f6630,f6637,f6638,f6646,f6647,f6667,f6681,f6685,f6849,f6857,f6858,f6862,f6870,f6871,f6901,f6905,f6911,f7125,f7144,f14263,f7555,f8617,f8625,f8626,f8664,f8666,f8667,f8671,f8672,f8684,f8757,f9371,f9636,f9738,f9938,f10042,f10248,f10632,f10851,f10852,f10854,f10855,f10857,f10862,f10932,f10933,f10993,f11227,f12470,f3330,f3329,f1376,f6312,f6313,f6365,f1374,f2661,f2663,f5625,f9698,f10001,f12428,f317,f320,f321,f346,f347,f1904,f4609,f4686,f4702,f4781,f5952,f6031,f6032,f6034,f6036,f6037,f6038,f6039,f6044,f6300,f6301,f6310,f6339,f6341,f6342,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6357,f8789,f9367,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9591,f9592,f9602,f9620,f9621,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9632,f9633,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9893,f9894,f9904,f9922,f9923,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9934,f9935,f10244,f12747,f12754,f12755,f12756,f12757,f12801,f12808,f12809,f12810,f12811,f345,f348,f5421,f5422,f5425,f5427,f5428,f5433,f5435,f5496,f5497,f5500,f5501,f5504,f5506,f5507,f5512,f5514,f5621,f6019,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9356,f9357,f9360,f9362,f9363,f9364,f9365,f9368,f9693,f9996,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10233,f10234,f10237,f10239,f10240,f10241,f10242,f10245,f10324,f10485,f10709,f12423,f12753,f12759,f12807,f12813,f5498,f10606,f6889,f10238,f9361,f10203,f9326,f10236,f9359,f10204,f9327,f12812,f12806,f10232,f9355,f12758,f12752,f10235,f9358,f10214,f9337,f5419,f6018,f5942,f5943,f12800,f6030,f6033,f8788,f12746,f6035,f311,f10610,f10612,f10618,f10622,f10624,f10626,f10628,f10629,f420,f2994,f2995,f2997,f2999,f3000,f3001,f3002,f3003,f10588,f4612,f6020,f6051,f6366,f6850,f6851,f6863,f6864,f339,f6891,f6892,f6894,f6896,f6897,f6898,f6899,f6908,f6848,f8647,f8618,f8619,f8648,f8649,f8685,f9696,f9999,f12426,f5536,f5534,f5458,f5456,f12749,f12743,f12797,f8646,f10577,f10616,f10578,f12804,f5461,f8791,f5539,f12805,f12750,f12751,f6838,f6839,f8804,f12776,f12830,f10620,f2998,f6895,f10608,f5138,f5418,f5620,f9692,f9995,f10325,f10486,f10710,f12422,f12803,f6890,f10614,f2996,f6893,f10605,f1889,f398,f349,f350,f354,f355,f356,f4517,f4613,f4614,f4797,f4807,f4841,f7073,f7086,f7503,f7516,f5543,f5545,f7077,f7507,f5465,f5467,f7076,f7506,f12760,f12762,f12769,f12771,f12814,f12816,f12823,f12825,f5542,f5544,f7064,f7494,f5464,f5466,f7063,f7493,f12761,f12763,f12768,f12770,f12815,f12817,f12822,f12824,f1484,f1486,f1909,f2227,f5623,f9695,f9998,f12425,f318,f829,f830,f831,f832,f833,f834,f930,f4538,f4539,f4541,f4543,f4544,f4545,f4546,f4547,f4611,f4689,f4690,f10345,f10506,f10730,f4542,f2699,f5538,f8700,f8696,f5460,f12744,f5116,f5113,f12798,f4540,f1515,f4575,f4576,f4579,f4581,f4582,f8042,f11852,f4620,f6620,f6586,f4615,f4587,f4589,f5630,f10335,f10496,f10720,f4861,f4760,f5619,f5622,f5624,f5626,f5627,f5628,f5629,f14548,f14549,f15708,f15711,f15714,f15716,f15717,f15718,f15719,f15720])).
fof(f15720,plain,(
  ( ! [X132,X133,X131] : (s(bool,X131) = s(bool,sK7('const.pair.FST_0',X133,bool)) | s(bool,'T_0') = s(bool,sK4(X131,X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f836])).
fof(f15719,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,X129) = s(bool,sK7('const.pair.FST_0',X130,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X129)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f342])).
fof(f15718,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) = s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,X125)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f336])).
fof(f15717,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X121) = s(bool,sK7('const.pair.FST_0',X123,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,X121),s(X120,X122)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f341])).
fof(f15716,plain,(
  ( ! [X118,X119,X117] : (s(bool,X118) = s(bool,sK7('const.pair.FST_0',X119,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X117,bool),'const.sum.INR_1'(s(bool,X118)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f337])).
fof(f15714,plain,(
  ( ! [X111,X112,X110] : (s(bool,X110) = s(bool,sK7('const.pair.FST_0',X112,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f316])).
fof(f15711,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,sK7('const.pair.FST_0',X100,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X97),bool),sK18(X98,X99,X97,bool)),s('type.list.list'(X97),'const.list.NIL_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f357])).
fof(f15708,plain,(
  ( ! [X83,X84] : (s(bool,X83) = s(bool,sK7('const.pair.FST_0',X84,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X83)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14548,f309])).
fof(f14549,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14492,f264])).
fof(f14492,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f271,f14263])).
fof(f14548,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14485,f432])).
fof(f14485,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f14263,f445])).
fof(f5629,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X157)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X158),bool),sK18(X159,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))),X158,bool))))) )),
  inference(superposition,[],[f4861,f342])).
fof(f5628,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,X152)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X153),bool),sK18(X154,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X151),'const.sum.INL_1'(s(bool,X152)))),X153,bool))))) )),
  inference(superposition,[],[f4861,f336])).
fof(f5627,plain,(
  ( ! [X146,X149,X147,X150,X148] : (~p(s(bool,X147)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X149),bool),sK18(X150,'const.pair.FST_1'(s('type.pair.prod'(bool,X146),'const.pair.,_2'(s(bool,X147),s(X146,X148)))),X149,bool))))) )),
  inference(superposition,[],[f4861,f341])).
fof(f5626,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.sum.OUTR_1'(s('type.sum.sum'(X142,bool),'const.sum.INR_1'(s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4861,f337])).
fof(f5624,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.|48|_0')),X135,bool))))) )),
  inference(superposition,[],[f4861,f316])).
fof(f5622,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X127),bool),sK18(X128,happ(s(fun('type.list.list'(X124),bool),sK18(X125,X126,X124,bool)),s('type.list.list'(X124),'const.list.NIL_0')),X127,bool))))) )),
  inference(superposition,[],[f4861,f357])).
fof(f5619,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4861,f309])).
fof(f4760,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK18(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4686,f357])).
fof(f4861,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4849,f264])).
fof(f4849,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4760,f270])).
fof(f10720,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2231,f836])).
fof(f10496,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2229,f836])).
fof(f10335,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | $spl10),
  inference(superposition,[],[f2225,f836])).
fof(f5630,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,X160)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X162),bool),sK18(X163,sK4(X160,X161,bool),X162,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161)))) )),
  inference(superposition,[],[f4861,f836])).
fof(f4589,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f494,f836])).
fof(f4587,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f447,f836])).
fof(f4615,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (s(fun(X178,X179),X180) != s(fun(X178,X179),X182) | s(X179,happ(s(fun(X178,X179),X182),s(X178,X183))) = s(X179,happ(s(fun(X178,X179),sK4(X180,X181,fun(X178,X179))),s(X178,X183))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X178,X179),bool),X181)))) )),
  inference(superposition,[],[f352,f836])).
fof(f6586,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f6620,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f353,f836])).
fof(f4620,plain,(
  ( ! [X200,X199,X201] : (s('type.ast.lit',X199) != s('type.ast.lit',X201) | s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X201))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',sK4(X199,X200,'type.ast.lit')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X200)))) )),
  inference(superposition,[],[f335,f836])).
fof(f11852,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X72) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X74) | s('type.list.list'('type.modSem.v'),X75) != s('type.list.list'('type.modSem.v'),X76) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X74),s('type.list.list'('type.modSem.v'),X76))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),sK4(X72,X73,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.list.list'('type.modSem.v'),X75))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),bool),X73)))) )),
  inference(superposition,[],[f396,f836])).
fof(f8042,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X64),s('type.list.list'('type.modSem.v'),X66))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X67),s('type.list.list'('type.modSem.v'),X68))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X67) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),sK4(X64,X65,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),bool),X65)))) )),
  inference(superposition,[],[f394,f836])).
fof(f4582,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f418,f836])).
fof(f4581,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f333,f836])).
fof(f4579,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f330,f836])).
fof(f4576,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f272,f836])).
fof(f4575,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f271,f836])).
fof(f1515,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f318,f338])).
fof(f4540,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f836,f338])).
fof(f12798,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X35,X37),happ(s(fun(X36,fun(X35,X37)),X39),s(X36,X40)))))))) | s(X35,X41) = s(X35,sK4(X41,happ(s(fun(X36,fun(X35,bool)),'_dst_x00x11_2'(s(fun(X37,bool),X38),s(fun(X36,fun(X35,X37)),X39))),s(X36,X40)),X35))) )),
  inference(superposition,[],[f318,f397])).
fof(f5113,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0x1_2'(s(fun(X17,bool),X19),s(fun(X16,X17),X18))))) | s(X16,X20) = s(X16,sK4(X20,happ(s(fun(fun(X17,bool),fun(X16,bool)),'_dstsel_0_x01_1'(s(fun(X16,X17),X18))),s(fun(X17,bool),X19)),X16))) )),
  inference(superposition,[],[f836,f368])).
fof(f5116,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun(X31,X32),X33)))))) | s(X31,X35) = s(X31,sK4(X35,happ(s(fun(fun(X32,bool),fun(X31,bool)),'_dstsel_0_x01_1'(s(fun(X31,X32),X33))),s(fun(X32,bool),X34)),X31))) )),
  inference(superposition,[],[f318,f368])).
fof(f12744,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,'|dollar|exists_1'(s(fun(X43,bool),happ(s(fun(X46,fun(X43,bool)),'_dst_x00x11_2'(s(fun(X44,bool),X45),s(fun(X46,fun(X43,X44)),X47))),s(X46,X48)))))) | s(X43,X49) = s(X43,sK4(X49,'_dst_x0x1_2'(s(fun(X44,bool),X45),s(fun(X43,X44),happ(s(fun(X46,fun(X43,X44)),X47),s(X46,X48)))),X43))) )),
  inference(superposition,[],[f318,f397])).
fof(f5460,plain,(
  ( ! [X344,X346,X341,X343,X345,X347,X342] : (p(s(bool,'|dollar|exists_1'(s(fun(X341,bool),happ(s(fun(X344,fun(X341,bool)),'_dst_x0x1_2'(s(fun(X342,fun(X341,bool)),X343),s(fun(X344,X342),X345))),s(X344,X346)))))) | s(X341,X347) = s(X341,sK4(X347,happ(s(fun(X342,fun(X341,bool)),X343),s(X342,happ(s(fun(X344,X342),X345),s(X344,X346)))),X341))) )),
  inference(superposition,[],[f318,f369])).
fof(f8696,plain,(
  ( ! [X397,X399,X393,X395,X396,X398,X394] : (p(s(bool,'|dollar|exists_1'(s(fun(X393,bool),happ(s(fun(X395,fun(X393,bool)),happ(s(fun(X394,fun(X395,fun(X393,bool))),X396),s(X394,X397))),s(X395,X398)))))) | s(X393,X399) = s(X393,sK4(X399,happ(s(fun('type.pair.prod'(X394,X395),fun(X393,bool)),sK19(X396,X395,X394,fun(X393,bool))),s('type.pair.prod'(X394,X395),'const.pair.,_2'(s(X394,X397),s(X395,X398)))),X393))) )),
  inference(superposition,[],[f318,f367])).
fof(f8700,plain,(
  ( ! [X428,X424,X426,X423,X429,X425,X427] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X423,bool),happ(s(fun(X425,fun(X423,bool)),happ(s(fun(X424,fun(X425,fun(X423,bool))),X426),s(X424,X427))),s(X425,X428))))) | s(X423,X429) = s(X423,sK4(X429,happ(s(fun('type.pair.prod'(X424,X425),fun(X423,bool)),sK19(X426,X425,X424,fun(X423,bool))),s('type.pair.prod'(X424,X425),'const.pair.,_2'(s(X424,X427),s(X425,X428)))),X423))) )),
  inference(superposition,[],[f836,f367])).
fof(f5538,plain,(
  ( ! [X333,X335,X331,X337,X332,X334,X336] : (p(s(bool,'|dollar|exists_1'(s(fun(X331,bool),happ(s(fun(X333,fun(X331,bool)),X334),s(X333,happ(s(fun(X332,X333),X335),s(X332,X336)))))))) | s(X331,X337) = s(X331,sK4(X337,happ(s(fun(X332,fun(X331,bool)),'_dst_x0x1_2'(s(fun(X333,fun(X331,bool)),X334),s(fun(X332,X333),X335))),s(X332,X336)),X331))) )),
  inference(superposition,[],[f318,f369])).
fof(f2699,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,'|dollar|exists_1'(s(fun(X139,bool),happ(s(fun('type.num.num',fun(X139,bool)),X141),s('type.num.num',X142)))))) | s(X139,X143) = s(X139,sK4(X143,'const.ind_type.FCONS_3'(s(fun(X139,bool),X140),s(fun('type.num.num',fun(X139,bool)),X141),s('type.num.num','const.num.SUC_1'(s('type.num.num',X142)))),X139))) )),
  inference(superposition,[],[f318,f324])).
fof(f4542,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))) | s(X16,X20) = s(X16,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))) )),
  inference(superposition,[],[f836,f324])).
fof(f10730,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10690])).
fof(f10690,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2231,f836])).
fof(f10506,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10466])).
fof(f10466,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2229,f836])).
fof(f10345,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10305])).
fof(f10305,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(superposition,[],[f2225,f836])).
fof(f4690,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4638,f262])).
fof(f4638,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f418,f836])).
fof(f4689,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4637,f262])).
fof(f4637,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f333,f836])).
fof(f4611,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f318,f836])).
fof(f4547,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f836,f836])).
fof(f4546,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X37))) | s(X34,X38) = s(X34,sK4(X38,'const.pair.SND_1'(s('type.pair.prod'(X35,fun(X34,bool)),'const.pair.,_2'(s(X35,X36),s(fun(X34,bool),X37)))),X34))) )),
  inference(superposition,[],[f836,f342])).
fof(f4545,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X32))) | s(X30,X33) = s(X30,sK4(X33,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X30,bool),X31),'const.sum.INL_1'(s(fun(X30,bool),X32)))),X30))) )),
  inference(superposition,[],[f836,f336])).
fof(f4544,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f836,f341])).
fof(f4543,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTR_1'(s('type.sum.sum'(X22,fun(X21,bool)),'const.sum.INR_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f836,f337])).
fof(f4541,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f836,f316])).
fof(f4539,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f836,f357])).
fof(f4538,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f836,f309])).
fof(f930,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK18(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f318,f357])).
fof(f834,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X23)))) | s(X20,X24) = s(X20,sK4(X24,'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))) )),
  inference(superposition,[],[f318,f342])).
fof(f833,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | s(X16,X19) = s(X16,sK4(X19,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X16,bool),X17),'const.sum.INL_1'(s(fun(X16,bool),X18)))),X16))) )),
  inference(superposition,[],[f318,f336])).
fof(f832,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f318,f341])).
fof(f831,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X9)))) | s(X7,X10) = s(X7,sK4(X10,'const.sum.OUTR_1'(s('type.sum.sum'(X8,fun(X7,bool)),'const.sum.INR_1'(s(fun(X7,bool),X9)))),X7))) )),
  inference(superposition,[],[f318,f337])).
fof(f830,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X6) = s(X3,sK4(X6,'const.ind_type.FCONS_3'(s(fun(X3,bool),X4),s(fun('type.num.num',fun(X3,bool)),X5),s('type.num.num','const.num.|48|_0')),X3))) )),
  inference(superposition,[],[f318,f316])).
fof(f829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f318,f309])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : ! [X34,X24] : (! [X38] : (s(X9,X24) = s(X9,X38) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X34)))))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.JRH_INDUCT_UTIL')).
fof(f12425,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) ) | $spl10),
  inference(superposition,[],[f10983,f338])).
fof(f9998,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f6500,f338])).
fof(f9695,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f6283,f338])).
fof(f5623,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X131),bool),sK18(X132,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)),X131,bool))))) )),
  inference(superposition,[],[f4861,f338])).
fof(f2227,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(forward_demodulation,[],[f2192,f2070])).
fof(f2192,plain,(
  ( ! [X23,X22] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(superposition,[],[f2168,f338])).
fof(f1909,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28))))) )),
  inference(forward_demodulation,[],[f1890,f338])).
fof(f1890,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X27),'const.min.@_1'(s(fun('type.pair.prod'(bool,X27),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f317,f338])).
fof(f1486,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f338,f445])).
fof(f1484,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f338,f445])).
fof(f12824,plain,(
  ( ! [X269,X271,X265,X267,X268,X270,X264,X266,X272] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X264,X265),'_dst_x0x1_2'(s(fun(X267,X265),X268),s(fun(X264,X267),happ(s(fun(X266,fun(X264,X267)),X269),s(X266,X270))))),s(fun(X265,X264),X271)))) | s(X264,X272) = s(X264,happ(s(fun(X265,X264),X271),s(X265,happ(s(fun(X264,X265),happ(s(fun(X266,fun(X264,X265)),'_dst_x00x11_2'(s(fun(X267,X265),X268),s(fun(X266,fun(X264,X267)),X269))),s(X266,X270))),s(X264,X272)))))) )),
  inference(superposition,[],[f355,f397])).
fof(f12822,plain,(
  ( ! [X253,X251,X249,X246,X254,X252,X250,X248,X247] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X246,X247),'_dst_x0x1_2'(s(fun(X249,X247),X250),s(fun(X246,X249),happ(s(fun(X248,fun(X246,X249)),X251),s(X248,X252))))),s(fun(X247,X246),X253)))) | s(X247,X254) = s(X247,happ(s(fun(X246,X247),happ(s(fun(X248,fun(X246,X247)),'_dst_x00x11_2'(s(fun(X249,X247),X250),s(fun(X248,fun(X246,X249)),X251))),s(X248,X252))),s(X246,happ(s(fun(X247,X246),X253),s(X247,X254)))))) )),
  inference(superposition,[],[f354,f397])).
fof(f12817,plain,(
  ( ! [X210,X208,X206,X204,X202,X211,X209,X207,X205,X203] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X202,X203),'_dst_x0x1_2'(s(fun(X205,X203),X206),s(fun(X202,X205),happ(s(fun(X204,fun(X202,X205)),X207),s(X204,X208))))),s(fun(X203,X202),X209)))) | s(X203,X210) != s(X203,happ(s(fun(X202,X203),happ(s(fun(X204,fun(X202,X203)),'_dst_x00x11_2'(s(fun(X205,X203),X206),s(fun(X204,fun(X202,X205)),X207))),s(X204,X208))),s(X202,X211))) | s(X202,X211) = s(X202,happ(s(fun(X203,X202),X209),s(X203,X210)))) )),
  inference(superposition,[],[f350,f397])).
fof(f12815,plain,(
  ( ! [X191,X189,X187,X185,X182,X190,X188,X186,X184,X183] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X182,X183),'_dst_x0x1_2'(s(fun(X185,X183),X186),s(fun(X182,X185),happ(s(fun(X184,fun(X182,X185)),X187),s(X184,X188))))),s(fun(X183,X182),X189)))) | s(X182,X190) != s(X182,happ(s(fun(X183,X182),X189),s(X183,X191))) | s(X183,X191) = s(X183,happ(s(fun(X182,X183),happ(s(fun(X184,fun(X182,X183)),'_dst_x00x11_2'(s(fun(X185,X183),X186),s(fun(X184,fun(X182,X185)),X187))),s(X184,X188))),s(X182,X190)))) )),
  inference(superposition,[],[f349,f397])).
fof(f12770,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X272,X274] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X272,X273),happ(s(fun(X276,fun(X272,X273)),'_dst_x00x11_2'(s(fun(X274,X273),X275),s(fun(X276,fun(X272,X274)),X277))),s(X276,X278))),s(fun(X273,X272),X279)))) | s(X272,X280) = s(X272,happ(s(fun(X273,X272),X279),s(X273,happ(s(fun(X272,X273),'_dst_x0x1_2'(s(fun(X274,X273),X275),s(fun(X272,X274),happ(s(fun(X276,fun(X272,X274)),X277),s(X276,X278))))),s(X272,X280)))))) )),
  inference(superposition,[],[f355,f397])).
fof(f12768,plain,(
  ( ! [X255,X261,X257,X259,X260,X254,X256,X258,X262] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X254,X255),happ(s(fun(X258,fun(X254,X255)),'_dst_x00x11_2'(s(fun(X256,X255),X257),s(fun(X258,fun(X254,X256)),X259))),s(X258,X260))),s(fun(X255,X254),X261)))) | s(X255,X262) = s(X255,happ(s(fun(X254,X255),'_dst_x0x1_2'(s(fun(X256,X255),X257),s(fun(X254,X256),happ(s(fun(X258,fun(X254,X256)),X259),s(X258,X260))))),s(X254,happ(s(fun(X255,X254),X261),s(X255,X262)))))) )),
  inference(superposition,[],[f354,f397])).
fof(f12763,plain,(
  ( ! [X212,X210,X216,X218,X215,X213,X211,X217,X219,X214] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X210,X211),happ(s(fun(X214,fun(X210,X211)),'_dst_x00x11_2'(s(fun(X212,X211),X213),s(fun(X214,fun(X210,X212)),X215))),s(X214,X216))),s(fun(X211,X210),X217)))) | s(X211,X218) != s(X211,happ(s(fun(X210,X211),'_dst_x0x1_2'(s(fun(X212,X211),X213),s(fun(X210,X212),happ(s(fun(X214,fun(X210,X212)),X215),s(X214,X216))))),s(X210,X219))) | s(X210,X219) = s(X210,happ(s(fun(X211,X210),X217),s(X211,X218)))) )),
  inference(superposition,[],[f350,f397])).
fof(f12761,plain,(
  ( ! [X191,X198,X196,X194,X192,X190,X199,X197,X195,X193] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X190,X191),happ(s(fun(X194,fun(X190,X191)),'_dst_x00x11_2'(s(fun(X192,X191),X193),s(fun(X194,fun(X190,X192)),X195))),s(X194,X196))),s(fun(X191,X190),X197)))) | s(X190,X198) != s(X190,happ(s(fun(X191,X190),X197),s(X191,X199))) | s(X191,X199) = s(X191,happ(s(fun(X190,X191),'_dst_x0x1_2'(s(fun(X192,X191),X193),s(fun(X190,X192),happ(s(fun(X194,fun(X190,X192)),X195),s(X194,X196))))),s(X190,X198)))) )),
  inference(superposition,[],[f349,f397])).
fof(f7493,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f7063,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5466,plain,(
  ( ! [X391,X397,X393,X395,X390,X396,X398,X392,X394] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X390,X391),happ(s(fun(X394,fun(X390,X391)),'_dst_x0x1_2'(s(fun(X392,fun(X390,X391)),X393),s(fun(X394,X392),X395))),s(X394,X396))),s(fun(X391,X390),X397)))) | s(X390,X398) = s(X390,happ(s(fun(X391,X390),X397),s(X391,happ(s(fun(X390,X391),happ(s(fun(X392,fun(X390,X391)),X393),s(X392,happ(s(fun(X394,X392),X395),s(X394,X396))))),s(X390,X398)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5464,plain,(
  ( ! [X377,X379,X372,X374,X380,X376,X378,X373,X375] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X372,X373),happ(s(fun(X376,fun(X372,X373)),'_dst_x0x1_2'(s(fun(X374,fun(X372,X373)),X375),s(fun(X376,X374),X377))),s(X376,X378))),s(fun(X373,X372),X379)))) | s(X373,X380) = s(X373,happ(s(fun(X372,X373),happ(s(fun(X374,fun(X372,X373)),X375),s(X374,happ(s(fun(X376,X374),X377),s(X376,X378))))),s(X372,happ(s(fun(X373,X372),X379),s(X373,X380)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7494,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f7064,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5544,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X380,X381),happ(s(fun(X383,fun(X380,X381)),X384),s(X383,happ(s(fun(X382,X383),X385),s(X382,X386))))),s(fun(X381,X380),X387)))) | s(X380,X388) = s(X380,happ(s(fun(X381,X380),X387),s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5542,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X362,X363),happ(s(fun(X365,fun(X362,X363)),X366),s(X365,happ(s(fun(X364,X365),X367),s(X364,X368))))),s(fun(X363,X362),X369)))) | s(X363,X370) = s(X363,happ(s(fun(X362,X363),happ(s(fun(X364,fun(X362,X363)),'_dst_x0x1_2'(s(fun(X365,fun(X362,X363)),X366),s(fun(X364,X365),X367))),s(X364,X368))),s(X362,happ(s(fun(X363,X362),X369),s(X363,X370)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f12825,plain,(
  ( ! [X280,X277,X279,X273,X275,X281,X276,X278,X274] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X274,X273),X280),s(fun(X273,X274),'_dst_x0x1_2'(s(fun(X276,X274),X277),s(fun(X273,X276),happ(s(fun(X275,fun(X273,X276)),X278),s(X275,X279)))))))) | s(X274,X281) = s(X274,happ(s(fun(X273,X274),happ(s(fun(X275,fun(X273,X274)),'_dst_x00x11_2'(s(fun(X276,X274),X277),s(fun(X275,fun(X273,X276)),X278))),s(X275,X279))),s(X273,happ(s(fun(X274,X273),X280),s(X274,X281)))))) )),
  inference(superposition,[],[f355,f397])).
fof(f12823,plain,(
  ( ! [X255,X261,X257,X259,X263,X260,X262,X256,X258] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X256,X255),X262),s(fun(X255,X256),'_dst_x0x1_2'(s(fun(X258,X256),X259),s(fun(X255,X258),happ(s(fun(X257,fun(X255,X258)),X260),s(X257,X261)))))))) | s(X255,X263) = s(X255,happ(s(fun(X256,X255),X262),s(X256,happ(s(fun(X255,X256),happ(s(fun(X257,fun(X255,X256)),'_dst_x00x11_2'(s(fun(X258,X256),X259),s(fun(X257,fun(X255,X258)),X260))),s(X257,X261))),s(X255,X263)))))) )),
  inference(superposition,[],[f354,f397])).
fof(f12816,plain,(
  ( ! [X198,X196,X194,X192,X200,X199,X197,X195,X193,X201] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X193,X192),X199),s(fun(X192,X193),'_dst_x0x1_2'(s(fun(X195,X193),X196),s(fun(X192,X195),happ(s(fun(X194,fun(X192,X195)),X197),s(X194,X198)))))))) | s(X192,X200) != s(X192,happ(s(fun(X193,X192),X199),s(X193,X201))) | s(X193,X201) = s(X193,happ(s(fun(X192,X193),happ(s(fun(X194,fun(X192,X193)),'_dst_x00x11_2'(s(fun(X195,X193),X196),s(fun(X194,fun(X192,X195)),X197))),s(X194,X198))),s(X192,X200)))) )),
  inference(superposition,[],[f350,f397])).
fof(f12814,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X172,X181,X179] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X173,X172),X179),s(fun(X172,X173),'_dst_x0x1_2'(s(fun(X175,X173),X176),s(fun(X172,X175),happ(s(fun(X174,fun(X172,X175)),X177),s(X174,X178)))))))) | s(X173,X180) != s(X173,happ(s(fun(X172,X173),happ(s(fun(X174,fun(X172,X173)),'_dst_x00x11_2'(s(fun(X175,X173),X176),s(fun(X174,fun(X172,X175)),X177))),s(X174,X178))),s(X172,X181))) | s(X172,X181) = s(X172,happ(s(fun(X173,X172),X179),s(X173,X180)))) )),
  inference(superposition,[],[f349,f397])).
fof(f12771,plain,(
  ( ! [X282,X288,X285,X287,X281,X283,X289,X284,X286] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X282,X281),X288),s(fun(X281,X282),happ(s(fun(X285,fun(X281,X282)),'_dst_x00x11_2'(s(fun(X283,X282),X284),s(fun(X285,fun(X281,X283)),X286))),s(X285,X287)))))) | s(X282,X289) = s(X282,happ(s(fun(X281,X282),'_dst_x0x1_2'(s(fun(X283,X282),X284),s(fun(X281,X283),happ(s(fun(X285,fun(X281,X283)),X286),s(X285,X287))))),s(X281,happ(s(fun(X282,X281),X288),s(X282,X289)))))) )),
  inference(superposition,[],[f355,f397])).
fof(f12769,plain,(
  ( ! [X263,X269,X271,X265,X267,X268,X270,X264,X266] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X264,X263),X270),s(fun(X263,X264),happ(s(fun(X267,fun(X263,X264)),'_dst_x00x11_2'(s(fun(X265,X264),X266),s(fun(X267,fun(X263,X265)),X268))),s(X267,X269)))))) | s(X263,X271) = s(X263,happ(s(fun(X264,X263),X270),s(X264,happ(s(fun(X263,X264),'_dst_x0x1_2'(s(fun(X265,X264),X266),s(fun(X263,X265),happ(s(fun(X267,fun(X263,X265)),X268),s(X267,X269))))),s(X263,X271)))))) )),
  inference(superposition,[],[f354,f397])).
fof(f12762,plain,(
  ( ! [X208,X206,X204,X202,X200,X209,X207,X205,X203,X201] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X201,X200),X207),s(fun(X200,X201),happ(s(fun(X204,fun(X200,X201)),'_dst_x00x11_2'(s(fun(X202,X201),X203),s(fun(X204,fun(X200,X202)),X205))),s(X204,X206)))))) | s(X200,X208) != s(X200,happ(s(fun(X201,X200),X207),s(X201,X209))) | s(X201,X209) = s(X201,happ(s(fun(X200,X201),'_dst_x0x1_2'(s(fun(X202,X201),X203),s(fun(X200,X202),happ(s(fun(X204,fun(X200,X202)),X205),s(X204,X206))))),s(X200,X208)))) )),
  inference(superposition,[],[f350,f397])).
fof(f12760,plain,(
  ( ! [X189,X187,X185,X182,X180,X188,X186,X184,X183,X181] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X181,X180),X187),s(fun(X180,X181),happ(s(fun(X184,fun(X180,X181)),'_dst_x00x11_2'(s(fun(X182,X181),X183),s(fun(X184,fun(X180,X182)),X185))),s(X184,X186)))))) | s(X181,X188) != s(X181,happ(s(fun(X180,X181),'_dst_x0x1_2'(s(fun(X182,X181),X183),s(fun(X180,X182),happ(s(fun(X184,fun(X180,X182)),X185),s(X184,X186))))),s(X180,X189))) | s(X180,X189) = s(X180,happ(s(fun(X181,X180),X187),s(X181,X188)))) )),
  inference(superposition,[],[f349,f397])).
fof(f7506,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f350,f369])).
fof(f7076,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5467,plain,(
  ( ! [X399,X405,X407,X401,X403,X404,X406,X400,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X400,X399),X406),s(fun(X399,X400),happ(s(fun(X403,fun(X399,X400)),'_dst_x0x1_2'(s(fun(X401,fun(X399,X400)),X402),s(fun(X403,X401),X404))),s(X403,X405)))))) | s(X400,X407) = s(X400,happ(s(fun(X399,X400),happ(s(fun(X401,fun(X399,X400)),X402),s(X401,happ(s(fun(X403,X401),X404),s(X403,X405))))),s(X399,happ(s(fun(X400,X399),X406),s(X400,X407)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5465,plain,(
  ( ! [X383,X389,X385,X387,X382,X388,X384,X386,X381] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X382,X381),X388),s(fun(X381,X382),happ(s(fun(X385,fun(X381,X382)),'_dst_x0x1_2'(s(fun(X383,fun(X381,X382)),X384),s(fun(X385,X383),X386))),s(X385,X387)))))) | s(X381,X389) = s(X381,happ(s(fun(X382,X381),X388),s(X382,happ(s(fun(X381,X382),happ(s(fun(X383,fun(X381,X382)),X384),s(X383,happ(s(fun(X385,X383),X386),s(X385,X387))))),s(X381,X389)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7507,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f350,f369])).
fof(f7077,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f349,f369])).
fof(f5545,plain,(
  ( ! [X389,X391,X397,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X390,X389),X396),s(fun(X389,X390),happ(s(fun(X392,fun(X389,X390)),X393),s(X392,happ(s(fun(X391,X392),X394),s(X391,X395)))))))) | s(X390,X397) = s(X390,happ(s(fun(X389,X390),happ(s(fun(X391,fun(X389,X390)),'_dst_x0x1_2'(s(fun(X392,fun(X389,X390)),X393),s(fun(X391,X392),X394))),s(X391,X395))),s(X389,happ(s(fun(X390,X389),X396),s(X390,X397)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f5543,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X372,X371),X378),s(fun(X371,X372),happ(s(fun(X374,fun(X371,X372)),X375),s(X374,happ(s(fun(X373,X374),X376),s(X373,X377)))))))) | s(X371,X379) = s(X371,happ(s(fun(X372,X371),X378),s(X372,happ(s(fun(X371,X372),happ(s(fun(X373,fun(X371,X372)),'_dst_x0x1_2'(s(fun(X374,fun(X371,X372)),X375),s(fun(X373,X374),X376))),s(X373,X377))),s(X371,X379)))))) )),
  inference(superposition,[],[f354,f369])).
fof(f7516,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f7503,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f350,f836])).
fof(f7086,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f7073,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f836])).
fof(f4841,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4810,f262])).
fof(f4810,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f355,f445])).
fof(f4807,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f4797,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f355,f836])).
fof(f4614,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X172,X173),X174),s(fun(X173,X172),X176)))) | s(X173,X177) = s(X173,happ(s(fun(X172,X173),sK4(X174,X175,fun(X172,X173))),s(X172,happ(s(fun(X173,X172),X176),s(X173,X177))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X172,X173),bool),X175)))) )),
  inference(superposition,[],[f354,f836])).
fof(f4613,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X167,X166),X170),s(fun(X166,X167),X168)))) | s(X166,X171) = s(X166,happ(s(fun(X167,X166),X170),s(X167,happ(s(fun(X166,X167),sK4(X168,X169,fun(X166,X167))),s(X166,X171))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X166,X167),bool),X169)))) )),
  inference(superposition,[],[f354,f836])).
fof(f4517,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4486,f262])).
fof(f4486,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f354,f445])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0)))))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X31,X9] : ! [X10,X37] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) <=> (! [X38] : s(X31,X38) = s(X31,happ(s(fun(X9,X31),X10),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))) & ! [X43] : s(X9,X43) = s(X9,happ(s(fun(X31,X9),X37),s(X31,happ(s(fun(X9,X31),X10),s(X9,X43)))))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.ISO')).
fof(f355,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f354,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X31] : ! [X37,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X37)))) => (! [X34] : (! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ! [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X34] : (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) <=> ? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,happ(s(fun(X31,X9),X37),s(X31,X38))))))) & ! [X6,X39] : (s(X9,X6) = s(X9,happ(s(fun(X31,X9),X37),s(X31,X39))) <=> s(X31,X39) = s(X31,happ(s(fun(X9,X31),X10),s(X9,X6))))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.ISO_USAGE')).
fof(f349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f239])).
fof(f398,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7))))))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X31,X44,X9,X45] : ! [X46,X37,X47,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X45),X10),s(fun(X45,X9),X47)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X31,X44),X37),s(fun(X44,X31),X46))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X31),fun(X45,X44)),'_dst_x00x11_2'(s(fun(X31,X44),X37),s(fun(fun(X9,X31),fun(X45,X31)),'_dstsel_0_x01_1'(s(fun(X45,X9),X47))))),s(fun(fun(X45,X44),fun(X9,X31)),'_dst_x00x11_2'(s(fun(X44,X31),X46),s(fun(fun(X45,X44),fun(X9,X44)),'_dstsel_0_x01_1'(s(fun(X9,X45),X10)))))))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.ISO_FUN')).
fof(f1889,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.min.@_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)))))))) )),
  inference(superposition,[],[f317,f357])).
fof(f10605,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | ~p(s(bool,happ(s(fun('type.list.list'(X59),bool),sK18(X60,X61,X59,bool)),s('type.list.list'(X59),'const.list.CONS_2'(s(X59,sK3(sK18(X60,X61,X59,bool),X59)),s('type.list.list'(X59),sK2(sK18(X60,X61,X59,bool),X59))))))) | p(s(bool,happ(s(fun('type.list.list'(X59),bool),sK18(X60,X61,X59,bool)),s('type.list.list'(X59),X62))))) )),
  inference(superposition,[],[f311,f357])).
fof(f6893,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)))))))) )),
  inference(forward_demodulation,[],[f6841,f338])).
fof(f6841,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36))))) )),
  inference(superposition,[],[f339,f338])).
fof(f2996,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f2980,f338])).
fof(f2980,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f338])).
fof(f10614,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30)))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10613,f338])).
fof(f10613,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30))))) )),
  inference(forward_demodulation,[],[f10580,f338])).
fof(f10580,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30))))) )),
  inference(superposition,[],[f311,f338])).
fof(f6890,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),X5)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f6836,f368])).
fof(f6836,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4))),s('type.pair.prod'(X0,X1),X5))))) )),
  inference(superposition,[],[f339,f368])).
fof(f12803,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X72,X74),happ(s(fun(X73,fun(X72,X74)),X76),s(X73,X77))))),s(X72,X78)))) | p(s(bool,'|dollar|exists_1'(s(fun(X72,bool),happ(s(fun(X73,fun(X72,bool)),'_dst_x00x11_2'(s(fun(X74,bool),X75),s(fun(X73,fun(X72,X74)),X76))),s(X73,X77))))))) )),
  inference(superposition,[],[f4686,f397])).
fof(f12422,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X72,bool),X73),s(fun(X74,X72),X75))),s(X74,X76))))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76)))))) ) | $spl10),
  inference(superposition,[],[f10983,f369])).
fof(f10710,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X126),'const.sum.INL_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2231,f369])).
fof(f10486,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2229,f369])).
fof(f10325,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun(X75,bool),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X75,bool),X76),s(fun(X74,X75),X77))),s(X74,X78)))))) ) | $spl10),
  inference(superposition,[],[f2225,f369])).
fof(f9995,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6500,f369])).
fof(f9692,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6283,f369])).
fof(f5620,plain,(
  ( ! [X111,X116,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),'_dst_x0x1_2'(s(fun(X110,bool),X111),s(fun(X112,X110),X113))),s(X112,X114)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X115),bool),sK18(X116,happ(s(fun(X110,bool),X111),s(X110,happ(s(fun(X112,X110),X113),s(X112,X114)))),X115,bool))))) )),
  inference(superposition,[],[f4861,f369])).
fof(f5418,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f4686,f369])).
fof(f5138,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36)))))) )),
  inference(forward_demodulation,[],[f5117,f368])).
fof(f5117,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36))))) | ~p(s(bool,happ(s(fun('type.list.list'(X36),bool),happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39))),s('type.list.list'(X36),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f368])).
fof(f10608,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),X4)))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10607,f368])).
fof(f10607,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),X4))))) )),
  inference(forward_demodulation,[],[f10575,f368])).
fof(f10575,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),X4))))) )),
  inference(superposition,[],[f311,f368])).
fof(f6895,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),X47)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)))))))) )),
  inference(forward_demodulation,[],[f6843,f324])).
fof(f6843,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))),s('type.pair.prod'(X42,X43),X47))))) )),
  inference(superposition,[],[f339,f324])).
fof(f2998,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f2982,f324])).
fof(f2982,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f324])).
fof(f10620,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),X46)))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10619,f324])).
fof(f10619,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),X46))))) )),
  inference(forward_demodulation,[],[f10583,f324])).
fof(f10583,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),X46))))) )),
  inference(superposition,[],[f311,f324])).
fof(f12830,plain,(
  ( ! [X325,X327,X321,X323,X324,X326,X320,X322] : (~p(s(bool,happ(s(fun('type.pair.prod'(X320,X321),bool),'_dst_x0x1_2'(s(fun(X323,bool),X324),s(fun('type.pair.prod'(X320,X321),X323),happ(s(fun(X322,fun('type.pair.prod'(X320,X321),X323)),X325),s(X322,X326))))),s('type.pair.prod'(X320,X321),'const.pair.,_2'(s(X320,sK7(happ(s(fun(X322,fun('type.pair.prod'(X320,X321),bool)),'_dst_x00x11_2'(s(fun(X323,bool),X324),s(fun(X322,fun('type.pair.prod'(X320,X321),X323)),X325))),s(X322,X326)),X321,X320)),s(X321,sK8(happ(s(fun(X322,fun('type.pair.prod'(X320,X321),bool)),'_dst_x00x11_2'(s(fun(X323,bool),X324),s(fun(X322,fun('type.pair.prod'(X320,X321),X323)),X325))),s(X322,X326)),X321,X320))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X320,X321),bool),happ(s(fun(X322,fun('type.pair.prod'(X320,X321),bool)),'_dst_x00x11_2'(s(fun(X323,bool),X324),s(fun(X322,fun('type.pair.prod'(X320,X321),X323)),X325))),s(X322,X326))),s('type.pair.prod'(X320,X321),X327))))) )),
  inference(superposition,[],[f339,f397])).
fof(f12776,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X328,X330] : (~p(s(bool,happ(s(fun('type.pair.prod'(X328,X329),bool),happ(s(fun(X332,fun('type.pair.prod'(X328,X329),bool)),'_dst_x00x11_2'(s(fun(X330,bool),X331),s(fun(X332,fun('type.pair.prod'(X328,X329),X330)),X333))),s(X332,X334))),s('type.pair.prod'(X328,X329),'const.pair.,_2'(s(X328,sK7('_dst_x0x1_2'(s(fun(X330,bool),X331),s(fun('type.pair.prod'(X328,X329),X330),happ(s(fun(X332,fun('type.pair.prod'(X328,X329),X330)),X333),s(X332,X334)))),X329,X328)),s(X329,sK8('_dst_x0x1_2'(s(fun(X330,bool),X331),s(fun('type.pair.prod'(X328,X329),X330),happ(s(fun(X332,fun('type.pair.prod'(X328,X329),X330)),X333),s(X332,X334)))),X329,X328))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X328,X329),bool),'_dst_x0x1_2'(s(fun(X330,bool),X331),s(fun('type.pair.prod'(X328,X329),X330),happ(s(fun(X332,fun('type.pair.prod'(X328,X329),X330)),X333),s(X332,X334))))),s('type.pair.prod'(X328,X329),X335))))) )),
  inference(superposition,[],[f339,f397])).
fof(f8804,plain,(
  ( ! [X596,X598,X592,X594,X597,X593,X595,X591] : (p(s(bool,happ(s(fun('type.pair.prod'(X591,X592),bool),happ(s(fun(X594,fun('type.pair.prod'(X591,X592),bool)),happ(s(fun(X593,fun(X594,fun('type.pair.prod'(X591,X592),bool))),X595),s(X593,X596))),s(X594,X597))),s('type.pair.prod'(X591,X592),X598)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X591,X592),bool),happ(s(fun(X594,fun('type.pair.prod'(X591,X592),bool)),happ(s(fun(X593,fun(X594,fun('type.pair.prod'(X591,X592),bool))),X595),s(X593,X596))),s(X594,X597))),s('type.pair.prod'(X591,X592),'const.pair.,_2'(s(X591,sK7(happ(s(fun('type.pair.prod'(X593,X594),fun('type.pair.prod'(X591,X592),bool)),sK19(X595,X594,X593,fun('type.pair.prod'(X591,X592),bool))),s('type.pair.prod'(X593,X594),'const.pair.,_2'(s(X593,X596),s(X594,X597)))),X592,X591)),s(X592,sK8(happ(s(fun('type.pair.prod'(X593,X594),fun('type.pair.prod'(X591,X592),bool)),sK19(X595,X594,X593,fun('type.pair.prod'(X591,X592),bool))),s('type.pair.prod'(X593,X594),'const.pair.,_2'(s(X593,X596),s(X594,X597)))),X592,X591)))))))) )),
  inference(forward_demodulation,[],[f8719,f367])).
fof(f8719,plain,(
  ( ! [X596,X598,X592,X594,X597,X593,X595,X591] : (~p(s(bool,happ(s(fun('type.pair.prod'(X591,X592),bool),happ(s(fun(X594,fun('type.pair.prod'(X591,X592),bool)),happ(s(fun(X593,fun(X594,fun('type.pair.prod'(X591,X592),bool))),X595),s(X593,X596))),s(X594,X597))),s('type.pair.prod'(X591,X592),'const.pair.,_2'(s(X591,sK7(happ(s(fun('type.pair.prod'(X593,X594),fun('type.pair.prod'(X591,X592),bool)),sK19(X595,X594,X593,fun('type.pair.prod'(X591,X592),bool))),s('type.pair.prod'(X593,X594),'const.pair.,_2'(s(X593,X596),s(X594,X597)))),X592,X591)),s(X592,sK8(happ(s(fun('type.pair.prod'(X593,X594),fun('type.pair.prod'(X591,X592),bool)),sK19(X595,X594,X593,fun('type.pair.prod'(X591,X592),bool))),s('type.pair.prod'(X593,X594),'const.pair.,_2'(s(X593,X596),s(X594,X597)))),X592,X591))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X591,X592),bool),happ(s(fun('type.pair.prod'(X593,X594),fun('type.pair.prod'(X591,X592),bool)),sK19(X595,X594,X593,fun('type.pair.prod'(X591,X592),bool))),s('type.pair.prod'(X593,X594),'const.pair.,_2'(s(X593,X596),s(X594,X597))))),s('type.pair.prod'(X591,X592),X598))))) )),
  inference(superposition,[],[f339,f367])).
fof(f6839,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,sK7(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18)),s(X19,sK8(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24))),s('type.pair.prod'(X18,X19),X25))))) )),
  inference(superposition,[],[f339,f369])).
fof(f6838,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X14,fun('type.pair.prod'(X10,X11),bool)),'_dst_x0x1_2'(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s('type.pair.prod'(X10,X11),'const.pair.,_2'(s(X10,sK7(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10)),s(X11,sK8(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16))))),s('type.pair.prod'(X10,X11),X17))))) )),
  inference(superposition,[],[f339,f369])).
fof(f12751,plain,(
  ( ! [X94,X99,X97,X95,X98,X96] : (p(s(bool,happ(s(fun('type.list.list'(X94),bool),happ(s(fun(X97,fun('type.list.list'(X94),bool)),'_dst_x00x11_2'(s(fun(X95,bool),X96),s(fun(X97,fun('type.list.list'(X94),X95)),X98))),s(X97,X99))),s('type.list.list'(X94),sK2('_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun('type.list.list'(X94),X95),happ(s(fun(X97,fun('type.list.list'(X94),X95)),X98),s(X97,X99)))),X94))))) | ~p(s(bool,happ(s(fun('type.list.list'(X94),bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun('type.list.list'(X94),X95),happ(s(fun(X97,fun('type.list.list'(X94),X95)),X98),s(X97,X99))))),s('type.list.list'(X94),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f397])).
fof(f12750,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,happ(s(fun('type.list.list'(X87),bool),happ(s(fun(X90,fun('type.list.list'(X87),bool)),'_dst_x00x11_2'(s(fun(X88,bool),X89),s(fun(X90,fun('type.list.list'(X87),X88)),X91))),s(X90,X92))),s('type.list.list'(X87),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X87),bool),'_dst_x0x1_2'(s(fun(X88,bool),X89),s(fun('type.list.list'(X87),X88),happ(s(fun(X90,fun('type.list.list'(X87),X88)),X91),s(X90,X92))))),s('type.list.list'(X87),'const.list.CONS_2'(s(X87,sK3('_dst_x0x1_2'(s(fun(X88,bool),X89),s(fun('type.list.list'(X87),X88),happ(s(fun(X90,fun('type.list.list'(X87),X88)),X91),s(X90,X92)))),X87)),s('type.list.list'(X87),sK2('_dst_x0x1_2'(s(fun(X88,bool),X89),s(fun('type.list.list'(X87),X88),happ(s(fun(X90,fun('type.list.list'(X87),X88)),X91),s(X90,X92)))),X87))))))) | p(s(bool,happ(s(fun('type.list.list'(X87),bool),'_dst_x0x1_2'(s(fun(X88,bool),X89),s(fun('type.list.list'(X87),X88),happ(s(fun(X90,fun('type.list.list'(X87),X88)),X91),s(X90,X92))))),s('type.list.list'(X87),X93))))) )),
  inference(superposition,[],[f311,f397])).
fof(f12805,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (p(s(bool,happ(s(fun('type.list.list'(X86),bool),'_dst_x0x1_2'(s(fun(X88,bool),X89),s(fun('type.list.list'(X86),X88),happ(s(fun(X87,fun('type.list.list'(X86),X88)),X90),s(X87,X91))))),s('type.list.list'(X86),sK2(happ(s(fun(X87,fun('type.list.list'(X86),bool)),'_dst_x00x11_2'(s(fun(X88,bool),X89),s(fun(X87,fun('type.list.list'(X86),X88)),X90))),s(X87,X91)),X86))))) | ~p(s(bool,happ(s(fun('type.list.list'(X86),bool),happ(s(fun(X87,fun('type.list.list'(X86),bool)),'_dst_x00x11_2'(s(fun(X88,bool),X89),s(fun(X87,fun('type.list.list'(X86),X88)),X90))),s(X87,X91))),s('type.list.list'(X86),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f397])).
fof(f5539,plain,(
  ( ! [X341,X343,X339,X340,X342,X338] : (p(s(bool,happ(s(fun('type.list.list'(X338),bool),happ(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(X340,happ(s(fun(X339,X340),X342),s(X339,X343))))),s('type.list.list'(X338),sK2(happ(s(fun(X339,fun('type.list.list'(X338),bool)),'_dst_x0x1_2'(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(fun(X339,X340),X342))),s(X339,X343)),X338))))) | ~p(s(bool,happ(s(fun('type.list.list'(X338),bool),happ(s(fun(X339,fun('type.list.list'(X338),bool)),'_dst_x0x1_2'(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(fun(X339,X340),X342))),s(X339,X343))),s('type.list.list'(X338),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f8791,plain,(
  ( ! [X441,X438,X440,X442,X437,X439] : (~p(s(bool,happ(s(fun('type.list.list'(X437),bool),happ(s(fun(X439,fun('type.list.list'(X437),bool)),happ(s(fun(X438,fun(X439,fun('type.list.list'(X437),bool))),X440),s(X438,X441))),s(X439,X442))),s('type.list.list'(X437),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X437),bool),happ(s(fun(X439,fun('type.list.list'(X437),bool)),happ(s(fun(X438,fun(X439,fun('type.list.list'(X437),bool))),X440),s(X438,X441))),s(X439,X442))),s('type.list.list'(X437),sK2(happ(s(fun('type.pair.prod'(X438,X439),fun('type.list.list'(X437),bool)),sK19(X440,X439,X438,fun('type.list.list'(X437),bool))),s('type.pair.prod'(X438,X439),'const.pair.,_2'(s(X438,X441),s(X439,X442)))),X437)))))) )),
  inference(forward_demodulation,[],[f8702,f367])).
fof(f8702,plain,(
  ( ! [X441,X438,X440,X442,X437,X439] : (p(s(bool,happ(s(fun('type.list.list'(X437),bool),happ(s(fun(X439,fun('type.list.list'(X437),bool)),happ(s(fun(X438,fun(X439,fun('type.list.list'(X437),bool))),X440),s(X438,X441))),s(X439,X442))),s('type.list.list'(X437),sK2(happ(s(fun('type.pair.prod'(X438,X439),fun('type.list.list'(X437),bool)),sK19(X440,X439,X438,fun('type.list.list'(X437),bool))),s('type.pair.prod'(X438,X439),'const.pair.,_2'(s(X438,X441),s(X439,X442)))),X437))))) | ~p(s(bool,happ(s(fun('type.list.list'(X437),bool),happ(s(fun('type.pair.prod'(X438,X439),fun('type.list.list'(X437),bool)),sK19(X440,X439,X438,fun('type.list.list'(X437),bool))),s('type.pair.prod'(X438,X439),'const.pair.,_2'(s(X438,X441),s(X439,X442))))),s('type.list.list'(X437),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f367])).
fof(f5461,plain,(
  ( ! [X350,X352,X349,X351,X353,X348] : (p(s(bool,happ(s(fun('type.list.list'(X348),bool),happ(s(fun(X351,fun('type.list.list'(X348),bool)),'_dst_x0x1_2'(s(fun(X349,fun('type.list.list'(X348),bool)),X350),s(fun(X351,X349),X352))),s(X351,X353))),s('type.list.list'(X348),sK2(happ(s(fun(X349,fun('type.list.list'(X348),bool)),X350),s(X349,happ(s(fun(X351,X349),X352),s(X351,X353)))),X348))))) | ~p(s(bool,happ(s(fun('type.list.list'(X348),bool),happ(s(fun(X349,fun('type.list.list'(X348),bool)),X350),s(X349,happ(s(fun(X351,X349),X352),s(X351,X353))))),s('type.list.list'(X348),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f12804,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun('type.list.list'(X79),bool),'_dst_x0x1_2'(s(fun(X81,bool),X82),s(fun('type.list.list'(X79),X81),happ(s(fun(X80,fun('type.list.list'(X79),X81)),X83),s(X80,X84))))),s('type.list.list'(X79),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X79),bool),happ(s(fun(X80,fun('type.list.list'(X79),bool)),'_dst_x00x11_2'(s(fun(X81,bool),X82),s(fun(X80,fun('type.list.list'(X79),X81)),X83))),s(X80,X84))),s('type.list.list'(X79),'const.list.CONS_2'(s(X79,sK3(happ(s(fun(X80,fun('type.list.list'(X79),bool)),'_dst_x00x11_2'(s(fun(X81,bool),X82),s(fun(X80,fun('type.list.list'(X79),X81)),X83))),s(X80,X84)),X79)),s('type.list.list'(X79),sK2(happ(s(fun(X80,fun('type.list.list'(X79),bool)),'_dst_x00x11_2'(s(fun(X81,bool),X82),s(fun(X80,fun('type.list.list'(X79),X81)),X83))),s(X80,X84)),X79))))))) | p(s(bool,happ(s(fun('type.list.list'(X79),bool),happ(s(fun(X80,fun('type.list.list'(X79),bool)),'_dst_x00x11_2'(s(fun(X81,bool),X82),s(fun(X80,fun('type.list.list'(X79),X81)),X83))),s(X80,X84))),s('type.list.list'(X79),X85))))) )),
  inference(superposition,[],[f311,f397])).
fof(f10578,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.list.list'(X15),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20))),s('type.list.list'(X15),'const.list.CONS_2'(s(X15,sK3(happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20)),X15)),s('type.list.list'(X15),sK2(happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20)),X15))))))) | p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20))),s('type.list.list'(X15),X21))))) )),
  inference(superposition,[],[f311,f369])).
fof(f10616,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),X37)))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10615,f367])).
fof(f10615,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),X37))))) )),
  inference(forward_demodulation,[],[f10581,f367])).
fof(f10581,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),X37))))) )),
  inference(superposition,[],[f311,f367])).
fof(f10577,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X11,fun('type.list.list'(X8),bool)),'_dst_x0x1_2'(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(fun(X11,X9),X12))),s(X11,X13))),s('type.list.list'(X8),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13))))),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,sK3(happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))),X8)),s('type.list.list'(X8),sK2(happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))),X8))))))) | p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13))))),s('type.list.list'(X8),X14))))) )),
  inference(superposition,[],[f311,f369])).
fof(f8646,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(X108,bool),happ(s(fun(X107,fun(X108,bool)),X109),s(X107,sK7(sK19(X109,X108,X107,bool),X108,X107)))),s(X108,sK8(sK19(X109,X108,X107,bool),X108,X107))))) | p(s(bool,happ(s(fun('type.pair.prod'(X107,X108),bool),sK19(X109,X108,X107,bool)),s('type.pair.prod'(X107,X108),X110))))) )),
  inference(superposition,[],[f339,f367])).
fof(f12797,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (~p(s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X30,bool),X31),s(fun(X28,X30),happ(s(fun(X29,fun(X28,X30)),X32),s(X29,X33))))),s(X28,X34)))) | p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),'_dst_x00x11_2'(s(fun(X30,bool),X31),s(fun(X29,fun(X28,X30)),X32))),s(X29,X33))),s(X28,'const.min.@_1'(s(fun(X28,bool),happ(s(fun(X29,fun(X28,bool)),'_dst_x00x11_2'(s(fun(X30,bool),X31),s(fun(X29,fun(X28,X30)),X32))),s(X29,X33))))))))) )),
  inference(superposition,[],[f317,f397])).
fof(f12743,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X39,fun(X36,bool)),'_dst_x00x11_2'(s(fun(X37,bool),X38),s(fun(X39,fun(X36,X37)),X40))),s(X39,X41))),s(X36,X42)))) | p(s(bool,happ(s(fun(X36,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X36,X37),happ(s(fun(X39,fun(X36,X37)),X40),s(X39,X41))))),s(X36,'const.min.@_1'(s(fun(X36,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X36,X37),happ(s(fun(X39,fun(X36,X37)),X40),s(X39,X41))))))))))) )),
  inference(superposition,[],[f317,f397])).
fof(f12749,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X80,bool),happ(s(fun(X83,fun(X80,bool)),'_dst_x00x11_2'(s(fun(X81,bool),X82),s(fun(X83,fun(X80,X81)),X84))),s(X83,X85))),s(X80,X86)))) | p(s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0x1_2'(s(fun(X81,bool),X82),s(fun(X80,X81),happ(s(fun(X83,fun(X80,X81)),X84),s(X83,X85))))))))) )),
  inference(superposition,[],[f4686,f397])).
fof(f5456,plain,(
  ( ! [X313,X315,X316,X318,X314,X317,X319] : (~p(s(bool,happ(s(fun(X313,bool),happ(s(fun(X316,fun(X313,bool)),'_dst_x0x1_2'(s(fun(X314,fun(X313,bool)),X315),s(fun(X316,X314),X317))),s(X316,X318))),s(X313,X319)))) | p(s(bool,'|dollar|exists_1'(s(fun(X313,bool),happ(s(fun(X314,fun(X313,bool)),X315),s(X314,happ(s(fun(X316,X314),X317),s(X316,X318))))))))) )),
  inference(superposition,[],[f4686,f369])).
fof(f5458,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(X327,bool),happ(s(fun(X330,fun(X327,bool)),'_dst_x0x1_2'(s(fun(X328,fun(X327,bool)),X329),s(fun(X330,X328),X331))),s(X330,X332))),s(X327,X333)))) | p(s(bool,happ(s(fun(X327,bool),happ(s(fun(X328,fun(X327,bool)),X329),s(X328,happ(s(fun(X330,X328),X331),s(X330,X332))))),s(X327,'const.min.@_1'(s(fun(X327,bool),happ(s(fun(X328,fun(X327,bool)),X329),s(X328,happ(s(fun(X330,X328),X331),s(X330,X332))))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f5534,plain,(
  ( ! [X308,X304,X306,X303,X309,X305,X307] : (~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X305,fun(X303,bool)),X306),s(X305,happ(s(fun(X304,X305),X307),s(X304,X308))))),s(X303,X309)))) | p(s(bool,'|dollar|exists_1'(s(fun(X303,bool),happ(s(fun(X304,fun(X303,bool)),'_dst_x0x1_2'(s(fun(X305,fun(X303,bool)),X306),s(fun(X304,X305),X307))),s(X304,X308))))))) )),
  inference(superposition,[],[f4686,f369])).
fof(f5536,plain,(
  ( ! [X321,X323,X318,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(X317,bool),happ(s(fun(X319,fun(X317,bool)),X320),s(X319,happ(s(fun(X318,X319),X321),s(X318,X322))))),s(X317,X323)))) | p(s(bool,happ(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))),s(X317,'const.min.@_1'(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f12426,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X88,bool),happ(s(fun(X87,fun(X88,bool)),X89),s(X87,X90))),s(X88,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X87,X88),bool),sK19(X89,X88,X87,bool)),s('type.pair.prod'(X87,X88),'const.pair.,_2'(s(X87,X90),s(X88,X91)))))) ) | $spl10),
  inference(superposition,[],[f10983,f367])).
fof(f9999,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,happ(s(fun(X133,bool),happ(s(fun(X132,fun(X133,bool)),X134),s(X132,X135))),s(X133,X136))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X132,X133),bool),sK19(X134,X133,X132,bool)),s('type.pair.prod'(X132,X133),'const.pair.,_2'(s(X132,X135),s(X133,X136)))))) )),
  inference(superposition,[],[f6500,f367])).
fof(f9696,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X133,bool),happ(s(fun(X132,fun(X133,bool)),X134),s(X132,X135))),s(X133,X136))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X132,X133),bool),sK19(X134,X133,X132,bool)),s('type.pair.prod'(X132,X133),'const.pair.,_2'(s(X132,X135),s(X133,X136)))))) )),
  inference(superposition,[],[f6283,f367])).
fof(f8685,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(X328,bool),happ(s(fun(X327,fun(X328,bool)),X329),s(X327,X330))),s(X328,X331)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X332),bool),sK18(X333,happ(s(fun('type.pair.prod'(X327,X328),bool),sK19(X329,X328,X327,bool)),s('type.pair.prod'(X327,X328),'const.pair.,_2'(s(X327,X330),s(X328,X331)))),X332,bool))))) )),
  inference(superposition,[],[f4861,f367])).
fof(f8649,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X122,fun(X123,bool)),X124),s(X122,X125))),s(X123,X126)))) | p(s(bool,happ(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)),s('type.pair.prod'(X122,X123),'const.min.@_1'(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)))))))) )),
  inference(superposition,[],[f317,f367])).
fof(f8648,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X117,fun(X118,bool)),X119),s(X117,X120))),s(X118,X121)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X117,X118),bool),sK19(X119,X118,X117,bool)))))) )),
  inference(superposition,[],[f4686,f367])).
fof(f8619,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK19(X19,X18,X17,bool)),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f367,f445])).
fof(f8618,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,happ(s(fun('type.pair.prod'(X11,X12),bool),sK19(X13,X12,X11,bool)),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f367,f445])).
fof(f8647,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X113),s(X111,X114))),s(X112,X115)))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),X116),s('type.pair.prod'(X111,X112),sK6(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),sK19(X113,X112,X111,bool)),s('type.pair.prod'(X111,X112),sK5(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112))))))) )),
  inference(superposition,[],[f320,f367])).
fof(f6848,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),X72),s('type.pair.prod'(X70,X71),'const.pair.,_2'(s(X70,sK7(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70)),s(X71,sK8(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool))),s('type.pair.prod'(X70,X71),X74)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.pair.prod'(X70,X71),bool),bool),X73)))) )),
  inference(superposition,[],[f339,f836])).
fof(f6908,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f6876,f262])).
fof(f6876,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f339,f445])).
fof(f6899,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),X69)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)))))))) )),
  inference(forward_demodulation,[],[f6847,f342])).
fof(f6847,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),'const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68))))),s('type.pair.prod'(X64,X65),X69))))) )),
  inference(superposition,[],[f339,f342])).
fof(f6898,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),X63)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)))))))) )),
  inference(forward_demodulation,[],[f6846,f336])).
fof(f6846,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62))))),s('type.pair.prod'(X59,X60),X63))))) )),
  inference(superposition,[],[f339,f336])).
fof(f6897,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),X58)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)))))))) )),
  inference(forward_demodulation,[],[f6845,f341])).
fof(f6845,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57))))),s('type.pair.prod'(X53,X54),X58))))) )),
  inference(superposition,[],[f339,f341])).
fof(f6896,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),X52)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)))))))) )),
  inference(forward_demodulation,[],[f6844,f337])).
fof(f6844,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51))))),s('type.pair.prod'(X48,X49),X52))))) )),
  inference(superposition,[],[f339,f337])).
fof(f6894,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),X41)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)))))))) )),
  inference(forward_demodulation,[],[f6842,f316])).
fof(f6842,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'(X37,X38),X41))))) )),
  inference(superposition,[],[f339,f316])).
fof(f6892,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),X31)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f6840,f357])).
fof(f6840,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0'))),s('type.pair.prod'(X26,X27),X31))))) )),
  inference(superposition,[],[f339,f357])).
fof(f6891,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),X9)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)))))))) )),
  inference(forward_demodulation,[],[f6837,f309])).
fof(f6837,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8))),s('type.pair.prod'(X6,X7),X9))))) )),
  inference(superposition,[],[f339,f309])).
fof(f339,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9,X31] : ! [X34] : (! [X59,X60] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X59),s(X31,X60)))))) => ! [X61] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X34),s('type.pair.prod'(X9,X31),X61)))))),
  file('thm.modSem.v_|49||49||split|4','thm.pair.pair_induction')).
fof(f6864,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,sK7(X7,bool,X8)),s(bool,X9)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK8(X7,bool,X8))) )),
  inference(superposition,[],[f339,f445])).
fof(f6863,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,sK7(X3,bool,X4)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),X5)))) | s(bool,X6) = s(bool,sK8(X3,bool,X4)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f6851,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X9),s(X8,sK8(X7,X8,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7(X7,X8,bool))) )),
  inference(superposition,[],[f339,f445])).
fof(f6850,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'T_0'),s(X4,sK8(X3,X4,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),X5)))) | s(bool,X6) = s(bool,sK7(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f339,f445])).
fof(f6366,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74))))) )),
  inference(forward_demodulation,[],[f6340,f338])).
fof(f6340,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72))))))) )),
  inference(superposition,[],[f321,f338])).
fof(f6051,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(forward_demodulation,[],[f6021,f338])).
fof(f6021,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(superposition,[],[f320,f338])).
fof(f6020,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X80)) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),X81),s('type.list.list'(X78),sK6(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78)))))) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),sK18(X79,X80,X78,bool)),s('type.list.list'(X78),sK5(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78))))))) )),
  inference(superposition,[],[f320,f357])).
fof(f4612,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f420,f836])).
fof(f10588,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.list.list'(X65),bool),X66),s('type.list.list'(X65),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X65),bool),sK4(X66,X67,fun('type.list.list'(X65),bool))),s('type.list.list'(X65),'const.list.CONS_2'(s(X65,sK3(sK4(X66,X67,fun('type.list.list'(X65),bool)),X65)),s('type.list.list'(X65),sK2(sK4(X66,X67,fun('type.list.list'(X65),bool)),X65))))))) | p(s(bool,happ(s(fun('type.list.list'(X65),bool),sK4(X66,X67,fun('type.list.list'(X65),bool))),s('type.list.list'(X65),X68)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X65),bool),bool),X67)))) )),
  inference(superposition,[],[f311,f836])).
fof(f3003,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f2987,f264])).
fof(f2987,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f420,f270])).
fof(f3002,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26)))))) )),
  inference(forward_demodulation,[],[f2986,f342])).
fof(f2986,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26))))) | ~p(s(bool,happ(s(fun('type.list.list'(X26),bool),'const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29))))),s('type.list.list'(X26),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f342])).
fof(f3001,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23)))))) )),
  inference(forward_demodulation,[],[f2985,f336])).
fof(f2985,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f336])).
fof(f3000,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f2984,f341])).
fof(f2984,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f341])).
fof(f2999,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16)))))) )),
  inference(forward_demodulation,[],[f2983,f337])).
fof(f2983,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f337])).
fof(f2997,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f2981,f316])).
fof(f2981,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f316])).
fof(f2995,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f2979,f357])).
fof(f2979,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f357])).
fof(f2994,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2978,f309])).
fof(f2978,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f309])).
fof(f420,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X34] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X24] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X24)))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X35),s('type.list.list'(X9),X24)))))))) => ! [X36] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X34),s('type.list.list'(X9),X36)))))),
  file('thm.modSem.v_|49||49||split|4','thm.list.list_INDUCT')).
fof(f10629,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK3(X4,X3)),s('type.list.list'(X3),sK2(X4,X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),X5)))) | s(bool,X6) = s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f10590,f262])).
fof(f10590,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK3(X4,X3)),s('type.list.list'(X3),sK2(X4,X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),X5)))) | s(bool,X6) = s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f311,f445])).
fof(f10628,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),X64)))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10627,f342])).
fof(f10627,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),X64))))) )),
  inference(forward_demodulation,[],[f10587,f342])).
fof(f10587,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),X64))))) )),
  inference(superposition,[],[f311,f342])).
fof(f10626,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),X59)))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10625,f336])).
fof(f10625,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),X59))))) )),
  inference(forward_demodulation,[],[f10586,f336])).
fof(f10586,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),X59))))) )),
  inference(superposition,[],[f311,f336])).
fof(f10624,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),X55)))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10623,f341])).
fof(f10623,plain,(
  ( ! [X54,X52,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),X55))))) )),
  inference(forward_demodulation,[],[f10585,f341])).
fof(f10585,plain,(
  ( ! [X54,X52,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),X55))))) )),
  inference(superposition,[],[f311,f341])).
fof(f10622,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),X50)))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10621,f337])).
fof(f10621,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),X50))))) )),
  inference(forward_demodulation,[],[f10584,f337])).
fof(f10584,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),X50))))) )),
  inference(superposition,[],[f311,f337])).
fof(f10618,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),X41)))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10617,f316])).
fof(f10617,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),X41))))) )),
  inference(forward_demodulation,[],[f10582,f316])).
fof(f10582,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),X41))))) )),
  inference(superposition,[],[f311,f316])).
fof(f10612,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),X26)))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10611,f357])).
fof(f10611,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),X26))))) )),
  inference(forward_demodulation,[],[f10579,f357])).
fof(f10579,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),X26))))) )),
  inference(superposition,[],[f311,f357])).
fof(f10610,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),X7)))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f10609,f309])).
fof(f10609,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),X7))))) )),
  inference(forward_demodulation,[],[f10576,f309])).
fof(f10576,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),X7))))) )),
  inference(superposition,[],[f311,f309])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f6035,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(forward_demodulation,[],[f5947,f324])).
fof(f5947,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(superposition,[],[f320,f324])).
fof(f12746,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (~p(s(bool,happ(s(fun(X57,bool),happ(s(fun(X60,fun(X57,bool)),'_dst_x00x11_2'(s(fun(X58,bool),X59),s(fun(X60,fun(X57,X58)),X61))),s(X60,X62))),s(X57,X63)))) | p(s(bool,happ(s(fun(X57,bool),X64),s(X57,sK6('_dst_x0x1_2'(s(fun(X58,bool),X59),s(fun(X57,X58),happ(s(fun(X60,fun(X57,X58)),X61),s(X60,X62)))),X64,X57))))) | p(s(bool,happ(s(fun(X57,bool),'_dst_x0x1_2'(s(fun(X58,bool),X59),s(fun(X57,X58),happ(s(fun(X60,fun(X57,X58)),X61),s(X60,X62))))),s(X57,sK5('_dst_x0x1_2'(s(fun(X58,bool),X59),s(fun(X57,X58),happ(s(fun(X60,fun(X57,X58)),X61),s(X60,X62)))),X64,X57)))))) )),
  inference(superposition,[],[f320,f397])).
fof(f8788,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X409,X411] : (p(s(bool,happ(s(fun(X407,bool),happ(s(fun(X409,fun(X407,bool)),happ(s(fun(X408,fun(X409,fun(X407,bool))),X410),s(X408,X411))),s(X409,X412))),s(X407,sK5(happ(s(fun('type.pair.prod'(X408,X409),fun(X407,bool)),sK19(X410,X409,X408,fun(X407,bool))),s('type.pair.prod'(X408,X409),'const.pair.,_2'(s(X408,X411),s(X409,X412)))),X414,X407))))) | ~p(s(bool,happ(s(fun(X407,bool),happ(s(fun(X409,fun(X407,bool)),happ(s(fun(X408,fun(X409,fun(X407,bool))),X410),s(X408,X411))),s(X409,X412))),s(X407,X413)))) | p(s(bool,happ(s(fun(X407,bool),X414),s(X407,sK6(happ(s(fun('type.pair.prod'(X408,X409),fun(X407,bool)),sK19(X410,X409,X408,fun(X407,bool))),s('type.pair.prod'(X408,X409),'const.pair.,_2'(s(X408,X411),s(X409,X412)))),X414,X407)))))) )),
  inference(forward_demodulation,[],[f8698,f367])).
fof(f8698,plain,(
  ( ! [X412,X414,X408,X410,X407,X413,X409,X411] : (~p(s(bool,happ(s(fun(X407,bool),happ(s(fun(X409,fun(X407,bool)),happ(s(fun(X408,fun(X409,fun(X407,bool))),X410),s(X408,X411))),s(X409,X412))),s(X407,X413)))) | p(s(bool,happ(s(fun(X407,bool),X414),s(X407,sK6(happ(s(fun('type.pair.prod'(X408,X409),fun(X407,bool)),sK19(X410,X409,X408,fun(X407,bool))),s('type.pair.prod'(X408,X409),'const.pair.,_2'(s(X408,X411),s(X409,X412)))),X414,X407))))) | p(s(bool,happ(s(fun(X407,bool),happ(s(fun('type.pair.prod'(X408,X409),fun(X407,bool)),sK19(X410,X409,X408,fun(X407,bool))),s('type.pair.prod'(X408,X409),'const.pair.,_2'(s(X408,X411),s(X409,X412))))),s(X407,sK5(happ(s(fun('type.pair.prod'(X408,X409),fun(X407,bool)),sK19(X410,X409,X408,fun(X407,bool))),s('type.pair.prod'(X408,X409),'const.pair.,_2'(s(X408,X411),s(X409,X412)))),X414,X407)))))) )),
  inference(superposition,[],[f320,f367])).
fof(f6033,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(forward_demodulation,[],[f5945,f338])).
fof(f5945,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(superposition,[],[f320,f338])).
fof(f6030,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f5940,f368])).
fof(f5940,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(superposition,[],[f320,f368])).
fof(f12800,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X49,bool),'_dst_x0x1_2'(s(fun(X51,bool),X52),s(fun(X49,X51),happ(s(fun(X50,fun(X49,X51)),X53),s(X50,X54))))),s(X49,X55)))) | p(s(bool,happ(s(fun(X49,bool),X56),s(X49,sK6(happ(s(fun(X50,fun(X49,bool)),'_dst_x00x11_2'(s(fun(X51,bool),X52),s(fun(X50,fun(X49,X51)),X53))),s(X50,X54)),X56,X49))))) | p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X50,fun(X49,bool)),'_dst_x00x11_2'(s(fun(X51,bool),X52),s(fun(X50,fun(X49,X51)),X53))),s(X50,X54))),s(X49,sK5(happ(s(fun(X50,fun(X49,bool)),'_dst_x00x11_2'(s(fun(X51,bool),X52),s(fun(X50,fun(X49,X51)),X53))),s(X50,X54)),X56,X49)))))) )),
  inference(superposition,[],[f320,f397])).
fof(f5943,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,X24)))) | p(s(bool,happ(s(fun(X18,bool),X25),s(X18,sK6(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18))))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK5(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f5942,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,X16)))) | p(s(bool,happ(s(fun(X10,bool),X17),s(X10,sK6(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10))))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK5(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6018,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70)))) | p(s(bool,happ(s(fun(X66,bool),X71),s(X66,sK6(X67,X71,X66))))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X67,X71,X66)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f5419,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun(X97,X95),X98))),s(X97,X99)))) | p(s(bool,happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f9337,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK10(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f10214,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK13(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f348,f836])).
fof(f9358,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f9329,f338])).
fof(f9329,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f10235,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK13(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f10206,f338])).
fof(f10206,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK13(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(superposition,[],[f348,f338])).
fof(f12752,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X101,X100),X107),s(fun(X100,X101),happ(s(fun(X104,fun(X100,X101)),'_dst_x00x11_2'(s(fun(X102,X101),X103),s(fun(X104,fun(X100,X102)),X105))),s(X104,X106)))))) | ~p(s(bool,happ(s(fun(X101,bool),X108),s(X101,sK10(X108,X107,'_dst_x0x1_2'(s(fun(X102,X101),X103),s(fun(X100,X102),happ(s(fun(X104,fun(X100,X102)),X105),s(X104,X106)))),X100,X101))))) | p(s(bool,happ(s(fun(X101,bool),X108),s(X101,happ(s(fun(X100,X101),'_dst_x0x1_2'(s(fun(X102,X101),X103),s(fun(X100,X102),happ(s(fun(X104,fun(X100,X102)),X105),s(X104,X106))))),s(X100,X109))))))) )),
  inference(superposition,[],[f345,f397])).
fof(f12758,plain,(
  ( ! [X167,X165,X163,X161,X169,X166,X164,X162,X160,X168] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X161,X160),X167),s(fun(X160,X161),happ(s(fun(X164,fun(X160,X161)),'_dst_x00x11_2'(s(fun(X162,X161),X163),s(fun(X164,fun(X160,X162)),X165))),s(X164,X166)))))) | ~p(s(bool,happ(s(fun(X161,bool),X168),s(X161,happ(s(fun(X160,X161),'_dst_x0x1_2'(s(fun(X162,X161),X163),s(fun(X160,X162),happ(s(fun(X164,fun(X160,X162)),X165),s(X164,X166))))),s(X160,X169)))))) | p(s(bool,happ(s(fun(X161,bool),X168),s(X161,sK13(X168,X167,'_dst_x0x1_2'(s(fun(X162,X161),X163),s(fun(X160,X162),happ(s(fun(X164,fun(X160,X162)),X165),s(X164,X166)))),X160,X161)))))) )),
  inference(superposition,[],[f348,f397])).
fof(f9355,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f9324,f368])).
fof(f9324,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f10232,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK13(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f10201,f368])).
fof(f10201,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK13(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(superposition,[],[f348,f368])).
fof(f12806,plain,(
  ( ! [X94,X92,X101,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X93,X92),X99),s(fun(X92,X93),'_dst_x0x1_2'(s(fun(X95,X93),X96),s(fun(X92,X95),happ(s(fun(X94,fun(X92,X95)),X97),s(X94,X98)))))))) | ~p(s(bool,happ(s(fun(X93,bool),X100),s(X93,sK10(X100,X99,happ(s(fun(X94,fun(X92,X93)),'_dst_x00x11_2'(s(fun(X95,X93),X96),s(fun(X94,fun(X92,X95)),X97))),s(X94,X98)),X92,X93))))) | p(s(bool,happ(s(fun(X93,bool),X100),s(X93,happ(s(fun(X92,X93),happ(s(fun(X94,fun(X92,X93)),'_dst_x00x11_2'(s(fun(X95,X93),X96),s(fun(X94,fun(X92,X95)),X97))),s(X94,X98))),s(X92,X101))))))) )),
  inference(superposition,[],[f345,f397])).
fof(f12812,plain,(
  ( ! [X158,X156,X154,X152,X161,X159,X157,X155,X153,X160] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X153,X152),X159),s(fun(X152,X153),'_dst_x0x1_2'(s(fun(X155,X153),X156),s(fun(X152,X155),happ(s(fun(X154,fun(X152,X155)),X157),s(X154,X158)))))))) | ~p(s(bool,happ(s(fun(X153,bool),X160),s(X153,happ(s(fun(X152,X153),happ(s(fun(X154,fun(X152,X153)),'_dst_x00x11_2'(s(fun(X155,X153),X156),s(fun(X154,fun(X152,X155)),X157))),s(X154,X158))),s(X152,X161)))))) | p(s(bool,happ(s(fun(X153,bool),X160),s(X153,sK13(X160,X159,happ(s(fun(X154,fun(X152,X153)),'_dst_x00x11_2'(s(fun(X155,X153),X156),s(fun(X154,fun(X152,X155)),X157))),s(X154,X158)),X152,X153)))))) )),
  inference(superposition,[],[f348,f397])).
fof(f9327,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK10(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f10204,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33)))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK13(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9359,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f9330,f367])).
fof(f9330,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,X58))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f10236,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK13(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f10207,f367])).
fof(f10207,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,X58)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK13(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(superposition,[],[f348,f367])).
fof(f9326,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK10(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f10203,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23)))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK13(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9361,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f9332,f324])).
fof(f9332,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f10238,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK13(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f10209,f324])).
fof(f10209,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK13(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(superposition,[],[f348,f324])).
fof(f6889,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X66,bool),X67),s(X66,happ(s(fun('type.pair.prod'(X64,X65),X66),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun('type.pair.prod'(X64,X65),X66),X68)),X65,X64)),s(X65,sK8('_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun('type.pair.prod'(X64,X65),X66),X68)),X65,X64))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun('type.pair.prod'(X64,X65),X66),X68))),s('type.pair.prod'(X64,X65),X69))))) )),
  inference(superposition,[],[f339,f369])).
fof(f10606,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,happ(s(fun(X64,bool),X65),s(X64,happ(s(fun('type.list.list'(X63),X64),X66),s('type.list.list'(X63),'const.list.NIL_0')))))) | ~p(s(bool,happ(s(fun('type.list.list'(X63),bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66))),s('type.list.list'(X63),'const.list.CONS_2'(s(X63,sK3('_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66)),X63)),s('type.list.list'(X63),sK2('_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66)),X63))))))) | p(s(bool,happ(s(fun('type.list.list'(X63),bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66))),s('type.list.list'(X63),X67))))) )),
  inference(superposition,[],[f311,f369])).
fof(f5498,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,happ(s(fun('type.list.list'(X96),X97),X99),s('type.list.list'(X96),sK2('_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99)),X96))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X96),bool),'_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99))),s('type.list.list'(X96),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f420,f369])).
fof(f12813,plain,(
  ( ! [X167,X165,X163,X171,X169,X166,X164,X162,X170,X168] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X162,X163),'_dst_x0x1_2'(s(fun(X165,X163),X166),s(fun(X162,X165),happ(s(fun(X164,fun(X162,X165)),X167),s(X164,X168))))),s(fun(X163,X162),X169)))) | ~p(s(bool,happ(s(fun(X162,bool),X170),s(X162,happ(s(fun(X163,X162),X169),s(X163,X171)))))) | p(s(bool,happ(s(fun(X162,bool),X170),s(X162,sK13(X170,happ(s(fun(X164,fun(X162,X163)),'_dst_x00x11_2'(s(fun(X165,X163),X166),s(fun(X164,fun(X162,X165)),X167))),s(X164,X168)),X169,X163,X162)))))) )),
  inference(superposition,[],[f348,f397])).
fof(f12807,plain,(
  ( ! [X103,X111,X109,X107,X105,X102,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X102,X103),'_dst_x0x1_2'(s(fun(X105,X103),X106),s(fun(X102,X105),happ(s(fun(X104,fun(X102,X105)),X107),s(X104,X108))))),s(fun(X103,X102),X109)))) | ~p(s(bool,happ(s(fun(X102,bool),X110),s(X102,sK10(X110,happ(s(fun(X104,fun(X102,X103)),'_dst_x00x11_2'(s(fun(X105,X103),X106),s(fun(X104,fun(X102,X105)),X107))),s(X104,X108)),X109,X103,X102))))) | p(s(bool,happ(s(fun(X102,bool),X110),s(X102,happ(s(fun(X103,X102),X109),s(X103,X111))))))) )),
  inference(superposition,[],[f345,f397])).
fof(f12759,plain,(
  ( ! [X177,X175,X173,X171,X178,X176,X174,X172,X170,X179] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X170,X171),happ(s(fun(X174,fun(X170,X171)),'_dst_x00x11_2'(s(fun(X172,X171),X173),s(fun(X174,fun(X170,X172)),X175))),s(X174,X176))),s(fun(X171,X170),X177)))) | ~p(s(bool,happ(s(fun(X170,bool),X178),s(X170,happ(s(fun(X171,X170),X177),s(X171,X179)))))) | p(s(bool,happ(s(fun(X170,bool),X178),s(X170,sK13(X178,'_dst_x0x1_2'(s(fun(X172,X171),X173),s(fun(X170,X172),happ(s(fun(X174,fun(X170,X172)),X175),s(X174,X176)))),X177,X171,X170)))))) )),
  inference(superposition,[],[f348,f397])).
fof(f12753,plain,(
  ( ! [X111,X118,X116,X114,X112,X110,X119,X117,X115,X113] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X110,X111),happ(s(fun(X114,fun(X110,X111)),'_dst_x00x11_2'(s(fun(X112,X111),X113),s(fun(X114,fun(X110,X112)),X115))),s(X114,X116))),s(fun(X111,X110),X117)))) | ~p(s(bool,happ(s(fun(X110,bool),X118),s(X110,sK10(X118,'_dst_x0x1_2'(s(fun(X112,X111),X113),s(fun(X110,X112),happ(s(fun(X114,fun(X110,X112)),X115),s(X114,X116)))),X117,X111,X110))))) | p(s(bool,happ(s(fun(X110,bool),X118),s(X110,happ(s(fun(X111,X110),X117),s(X111,X119))))))) )),
  inference(superposition,[],[f345,f397])).
fof(f12423,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X78,bool),X79),s(X78,happ(s(fun(X77,X78),X80),s(X77,X81))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X77,bool),'_dst_x0x1_2'(s(fun(X78,bool),X79),s(fun(X77,X78),X80))),s(X77,X81)))) ) | $spl10),
  inference(superposition,[],[f10983,f369])).
fof(f10709,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X120),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2231,f369])).
fof(f10485,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X120,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2229,f369])).
fof(f10324,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(X71,bool),'_dst_x0x1_2'(s(fun(X69,bool),X70),s(fun(X71,X69),X72))),s(X71,X73))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X69,bool),X70),s(X69,happ(s(fun(X71,X69),X72),s(X71,X73)))))))) ) | $spl10),
  inference(superposition,[],[f2225,f369])).
fof(f10245,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK13(X13,X11,X12,X10,X9))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f10217,f262])).
fof(f10217,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK13(X13,X11,X12,X10,X9))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f348,f445])).
fof(f10242,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK13(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f10213,f342])).
fof(f10213,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103)))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK13(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(superposition,[],[f348,f342])).
fof(f10241,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK13(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f10212,f336])).
fof(f10212,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95)))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK13(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(superposition,[],[f348,f336])).
fof(f10240,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK13(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f10211,f341])).
fof(f10211,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88)))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK13(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(superposition,[],[f348,f341])).
fof(f10239,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK13(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f10210,f337])).
fof(f10210,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80)))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK13(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(superposition,[],[f348,f337])).
fof(f10237,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK13(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f10208,f316])).
fof(f10208,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65)))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK13(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(superposition,[],[f348,f316])).
fof(f10234,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK13(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f10205,f357])).
fof(f10205,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41)))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK13(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(superposition,[],[f348,f357])).
fof(f10233,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK13(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f10202,f309])).
fof(f10202,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13)))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK13(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(superposition,[],[f348,f309])).
fof(f10200,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK13(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f348,f836])).
fof(f10199,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103)))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK13(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))) )),
  inference(superposition,[],[f348,f342])).
fof(f10198,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95)))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK13(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))) )),
  inference(superposition,[],[f348,f336])).
fof(f10197,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88)))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK13(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))) )),
  inference(superposition,[],[f348,f341])).
fof(f10196,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80)))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK13(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))) )),
  inference(superposition,[],[f348,f337])).
fof(f10195,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73)))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK13(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))) )),
  inference(superposition,[],[f348,f324])).
fof(f10194,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65)))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK13(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))) )),
  inference(superposition,[],[f348,f316])).
fof(f10193,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,X58)))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,sK13(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49)))))) )),
  inference(superposition,[],[f348,f367])).
fof(f10192,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48)))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK13(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))) )),
  inference(superposition,[],[f348,f338])).
fof(f10191,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41)))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK13(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))) )),
  inference(superposition,[],[f348,f357])).
fof(f10190,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33)))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK13(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))) )),
  inference(superposition,[],[f348,f369])).
fof(f10189,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23)))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK13(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))) )),
  inference(superposition,[],[f348,f369])).
fof(f10188,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13)))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK13(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))) )),
  inference(superposition,[],[f348,f309])).
fof(f10187,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9996,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6500,f369])).
fof(f9693,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X124,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6283,f369])).
fof(f9368,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9340,f262])).
fof(f9340,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f345,f445])).
fof(f9365,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f9336,f342])).
fof(f9336,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f9364,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f9335,f336])).
fof(f9335,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f9363,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f9334,f341])).
fof(f9334,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f9362,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f9333,f337])).
fof(f9333,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f9360,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f9331,f316])).
fof(f9331,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f9357,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f9328,f357])).
fof(f9328,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f9356,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f9325,f309])).
fof(f9325,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f9323,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK10(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f345,f836])).
fof(f9322,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK10(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103))))))) )),
  inference(superposition,[],[f345,f342])).
fof(f9321,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK10(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95))))))) )),
  inference(superposition,[],[f345,f336])).
fof(f9320,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK10(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88))))))) )),
  inference(superposition,[],[f345,f341])).
fof(f9319,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK10(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80))))))) )),
  inference(superposition,[],[f345,f337])).
fof(f9318,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK10(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73))))))) )),
  inference(superposition,[],[f345,f324])).
fof(f9317,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK10(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65))))))) )),
  inference(superposition,[],[f345,f316])).
fof(f9316,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,sK10(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,X58))))))) )),
  inference(superposition,[],[f345,f367])).
fof(f9315,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK10(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48))))))) )),
  inference(superposition,[],[f345,f338])).
fof(f9314,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK10(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41))))))) )),
  inference(superposition,[],[f345,f357])).
fof(f9313,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK10(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f9312,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK10(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))))))) )),
  inference(superposition,[],[f345,f369])).
fof(f9311,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK10(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13))))))) )),
  inference(superposition,[],[f345,f309])).
fof(f9310,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK10(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) )),
  inference(superposition,[],[f345,f368])).
fof(f6019,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | p(s(bool,happ(s(fun(X72,bool),X77),s(X72,sK6('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72))))) | p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,sK5('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f5621,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X122),bool),sK18(X123,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X117,X118),X120))),s(X117,X121)),X122,bool))))) )),
  inference(superposition,[],[f4861,f369])).
fof(f5514,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X184,X185),X187),s(X184,X188)))))) | s(bool,X189) = s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X184,X185),X187))),s(X184,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f494,f369])).
fof(f5512,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X173,X174),X176),s(X173,X177)))))) | s(bool,X178) != s(bool,happ(s(fun(X173,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X173,X174),X176))),s(X173,X177))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f447,f369])).
fof(f5507,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X144,bool),X145),s(X144,happ(s(fun(X143,X144),X146),s(X143,X147)))))) | s(bool,X148) != s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X144,bool),X145),s(fun(X143,X144),X146))),s(X143,X147))) | p(s(bool,X148))) )),
  inference(superposition,[],[f418,f369])).
fof(f5506,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | ~p(s(bool,X142)) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f333,f369])).
fof(f5504,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129)))))) | p(s(bool,X130)) | s(bool,X130) = s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f330,f369])).
fof(f5501,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114)))) )),
  inference(superposition,[],[f272,f369])).
fof(f5500,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f271,f369])).
fof(f5497,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun(X91,X92),X94),s(X91,X95)))))) | p(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))),s(X91,'const.min.@_1'(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))))))))) )),
  inference(superposition,[],[f317,f369])).
fof(f5496,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))) )),
  inference(superposition,[],[f4686,f369])).
fof(f5435,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (p(s(bool,happ(s(fun(X189,bool),'_dst_x0x1_2'(s(fun(X187,bool),X188),s(fun(X189,X187),X190))),s(X189,X191)))) | s(bool,X192) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f494,f369])).
fof(f5433,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X178,bool),'_dst_x0x1_2'(s(fun(X176,bool),X177),s(fun(X178,X176),X179))),s(X178,X180)))) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),X177),s(X176,happ(s(fun(X178,X176),X179),s(X178,X180))))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f447,f369])).
fof(f5428,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,X151) != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150))))) | p(s(bool,X151))) )),
  inference(superposition,[],[f418,f369])).
fof(f5427,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | ~p(s(bool,X145)) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) )),
  inference(superposition,[],[f333,f369])).
fof(f5425,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X130,X128),X131))),s(X130,X132)))) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f330,f369])).
fof(f5422,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117)))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f272,f369])).
fof(f5421,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X108,bool),X109),s(fun(X110,X108),X111))),s(X110,X112)))) | s(bool,'T_0') != s(bool,happ(s(fun(X108,bool),X109),s(X108,happ(s(fun(X110,X108),X111),s(X110,X112)))))) )),
  inference(superposition,[],[f271,f369])).
fof(f348,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f345,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f12811,plain,(
  ( ! [X146,X144,X142,X151,X149,X147,X145,X143,X150,X148] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X142,X143),'_dst_x0x1_2'(s(fun(X145,X143),X146),s(fun(X142,X145),happ(s(fun(X144,fun(X142,X145)),X147),s(X144,X148))))),s(fun(X143,X142),X149)))) | ~p(s(bool,happ(s(fun(X142,bool),X150),s(X142,X151)))) | p(s(bool,happ(s(fun(X142,bool),X150),s(X142,happ(s(fun(X143,X142),X149),s(X143,sK12(X150,happ(s(fun(X144,fun(X142,X143)),'_dst_x00x11_2'(s(fun(X145,X143),X146),s(fun(X144,fun(X142,X145)),X147))),s(X144,X148)),X149,X143,X142)))))))) )),
  inference(superposition,[],[f347,f397])).
fof(f12810,plain,(
  ( ! [X134,X132,X140,X138,X136,X135,X133,X141,X139,X137] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X133,X132),X139),s(fun(X132,X133),'_dst_x0x1_2'(s(fun(X135,X133),X136),s(fun(X132,X135),happ(s(fun(X134,fun(X132,X135)),X137),s(X134,X138)))))))) | ~p(s(bool,happ(s(fun(X133,bool),X140),s(X133,X141)))) | p(s(bool,happ(s(fun(X133,bool),X140),s(X133,happ(s(fun(X132,X133),happ(s(fun(X134,fun(X132,X133)),'_dst_x00x11_2'(s(fun(X135,X133),X136),s(fun(X134,fun(X132,X135)),X137))),s(X134,X138))),s(X132,sK12(X140,X139,happ(s(fun(X134,fun(X132,X133)),'_dst_x00x11_2'(s(fun(X135,X133),X136),s(fun(X134,fun(X132,X135)),X137))),s(X134,X138)),X132,X133)))))))) )),
  inference(superposition,[],[f347,f397])).
fof(f12809,plain,(
  ( ! [X127,X125,X123,X130,X128,X126,X124,X122,X131,X129] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X122,X123),'_dst_x0x1_2'(s(fun(X125,X123),X126),s(fun(X122,X125),happ(s(fun(X124,fun(X122,X125)),X127),s(X124,X128))))),s(fun(X123,X122),X129)))) | ~p(s(bool,happ(s(fun(X122,bool),X130),s(X122,happ(s(fun(X123,X122),X129),s(X123,sK11(X130,happ(s(fun(X124,fun(X122,X123)),'_dst_x00x11_2'(s(fun(X125,X123),X126),s(fun(X124,fun(X122,X125)),X127))),s(X124,X128)),X129,X123,X122))))))) | p(s(bool,happ(s(fun(X122,bool),X130),s(X122,X131))))) )),
  inference(superposition,[],[f346,f397])).
fof(f12808,plain,(
  ( ! [X121,X118,X116,X114,X112,X120,X119,X117,X115,X113] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X113,X112),X119),s(fun(X112,X113),'_dst_x0x1_2'(s(fun(X115,X113),X116),s(fun(X112,X115),happ(s(fun(X114,fun(X112,X115)),X117),s(X114,X118)))))))) | ~p(s(bool,happ(s(fun(X113,bool),X120),s(X113,happ(s(fun(X112,X113),happ(s(fun(X114,fun(X112,X113)),'_dst_x00x11_2'(s(fun(X115,X113),X116),s(fun(X114,fun(X112,X115)),X117))),s(X114,X118))),s(X112,sK11(X120,X119,happ(s(fun(X114,fun(X112,X113)),'_dst_x00x11_2'(s(fun(X115,X113),X116),s(fun(X114,fun(X112,X115)),X117))),s(X114,X118)),X112,X113))))))) | p(s(bool,happ(s(fun(X113,bool),X120),s(X113,X121))))) )),
  inference(superposition,[],[f346,f397])).
fof(f12801,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X63] : (~p(s(bool,happ(s(fun(X57,bool),'_dst_x0x1_2'(s(fun(X59,bool),X60),s(fun(X57,X59),happ(s(fun(X58,fun(X57,X59)),X61),s(X58,X62))))),s(X57,sK5(X63,happ(s(fun(X58,fun(X57,bool)),'_dst_x00x11_2'(s(fun(X59,bool),X60),s(fun(X58,fun(X57,X59)),X61))),s(X58,X62)),X57))))) | ~p(s(bool,happ(s(fun(X57,bool),X63),s(X57,X64)))) | p(s(bool,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'_dst_x00x11_2'(s(fun(X59,bool),X60),s(fun(X58,fun(X57,X59)),X61))),s(X58,X62))),s(X57,sK6(X63,happ(s(fun(X58,fun(X57,bool)),'_dst_x00x11_2'(s(fun(X59,bool),X60),s(fun(X58,fun(X57,X59)),X61))),s(X58,X62)),X57)))))) )),
  inference(superposition,[],[f321,f397])).
fof(f12757,plain,(
  ( ! [X158,X156,X154,X152,X151,X159,X157,X155,X153,X150] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X150,X151),happ(s(fun(X154,fun(X150,X151)),'_dst_x00x11_2'(s(fun(X152,X151),X153),s(fun(X154,fun(X150,X152)),X155))),s(X154,X156))),s(fun(X151,X150),X157)))) | ~p(s(bool,happ(s(fun(X150,bool),X158),s(X150,X159)))) | p(s(bool,happ(s(fun(X150,bool),X158),s(X150,happ(s(fun(X151,X150),X157),s(X151,sK12(X158,'_dst_x0x1_2'(s(fun(X152,X151),X153),s(fun(X150,X152),happ(s(fun(X154,fun(X150,X152)),X155),s(X154,X156)))),X157,X151,X150)))))))) )),
  inference(superposition,[],[f347,f397])).
fof(f12756,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X148] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X141,X140),X147),s(fun(X140,X141),happ(s(fun(X144,fun(X140,X141)),'_dst_x00x11_2'(s(fun(X142,X141),X143),s(fun(X144,fun(X140,X142)),X145))),s(X144,X146)))))) | ~p(s(bool,happ(s(fun(X141,bool),X148),s(X141,X149)))) | p(s(bool,happ(s(fun(X141,bool),X148),s(X141,happ(s(fun(X140,X141),'_dst_x0x1_2'(s(fun(X142,X141),X143),s(fun(X140,X142),happ(s(fun(X144,fun(X140,X142)),X145),s(X144,X146))))),s(X140,sK12(X148,X147,'_dst_x0x1_2'(s(fun(X142,X141),X143),s(fun(X140,X142),happ(s(fun(X144,fun(X140,X142)),X145),s(X144,X146)))),X140,X141)))))))) )),
  inference(superposition,[],[f347,f397])).
fof(f12755,plain,(
  ( ! [X134,X132,X130,X138,X136,X135,X133,X131,X139,X137] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X130,X131),happ(s(fun(X134,fun(X130,X131)),'_dst_x00x11_2'(s(fun(X132,X131),X133),s(fun(X134,fun(X130,X132)),X135))),s(X134,X136))),s(fun(X131,X130),X137)))) | ~p(s(bool,happ(s(fun(X130,bool),X138),s(X130,happ(s(fun(X131,X130),X137),s(X131,sK11(X138,'_dst_x0x1_2'(s(fun(X132,X131),X133),s(fun(X130,X132),happ(s(fun(X134,fun(X130,X132)),X135),s(X134,X136)))),X137,X131,X130))))))) | p(s(bool,happ(s(fun(X130,bool),X138),s(X130,X139))))) )),
  inference(superposition,[],[f346,f397])).
fof(f12754,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120,X129] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X121,X120),X127),s(fun(X120,X121),happ(s(fun(X124,fun(X120,X121)),'_dst_x00x11_2'(s(fun(X122,X121),X123),s(fun(X124,fun(X120,X122)),X125))),s(X124,X126)))))) | ~p(s(bool,happ(s(fun(X121,bool),X128),s(X121,happ(s(fun(X120,X121),'_dst_x0x1_2'(s(fun(X122,X121),X123),s(fun(X120,X122),happ(s(fun(X124,fun(X120,X122)),X125),s(X124,X126))))),s(X120,sK11(X128,X127,'_dst_x0x1_2'(s(fun(X122,X121),X123),s(fun(X120,X122),happ(s(fun(X124,fun(X120,X122)),X125),s(X124,X126)))),X120,X121))))))) | p(s(bool,happ(s(fun(X121,bool),X128),s(X121,X129))))) )),
  inference(superposition,[],[f346,f397])).
fof(f12747,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),happ(s(fun(X68,fun(X65,bool)),'_dst_x00x11_2'(s(fun(X66,bool),X67),s(fun(X68,fun(X65,X66)),X69))),s(X68,X70))),s(X65,sK5(X71,'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X65,X66),happ(s(fun(X68,fun(X65,X66)),X69),s(X68,X70)))),X65))))) | ~p(s(bool,happ(s(fun(X65,bool),X71),s(X65,X72)))) | p(s(bool,happ(s(fun(X65,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X65,X66),happ(s(fun(X68,fun(X65,X66)),X69),s(X68,X70))))),s(X65,sK6(X71,'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X65,X66),happ(s(fun(X68,fun(X65,X66)),X69),s(X68,X70)))),X65)))))) )),
  inference(superposition,[],[f321,f397])).
fof(f10244,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f10243,f309])).
fof(f10243,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(subsumption_resolution,[],[f10215,f262])).
fof(f10215,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(superposition,[],[f348,f440])).
fof(f9935,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9907,f262])).
fof(f9907,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f347,f445])).
fof(f9934,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f9933,f309])).
fof(f9933,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(subsumption_resolution,[],[f9905,f262])).
fof(f9905,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(superposition,[],[f347,f440])).
fof(f9932,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f9903,f342])).
fof(f9903,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f9931,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f9902,f336])).
fof(f9902,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f9930,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f9901,f341])).
fof(f9901,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f9929,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f9900,f337])).
fof(f9900,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f9928,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f9899,f324])).
fof(f9899,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73)))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f9927,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f9898,f316])).
fof(f9898,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f9926,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f9897,f367])).
fof(f9897,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58)))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f9925,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f9896,f338])).
fof(f9896,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f9924,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f9895,f357])).
fof(f9895,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f9923,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f9892,f309])).
fof(f9892,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f9922,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f9891,f368])).
fof(f9891,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9904,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK12(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f9894,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33)))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK12(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9893,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23)))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK12(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9890,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK12(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f836])).
fof(f9889,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103)))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK12(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))))) )),
  inference(superposition,[],[f347,f342])).
fof(f9888,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95)))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK12(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))))) )),
  inference(superposition,[],[f347,f336])).
fof(f9887,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88)))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK12(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f9886,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80)))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK12(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))))) )),
  inference(superposition,[],[f347,f337])).
fof(f9885,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73)))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK12(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))))) )),
  inference(superposition,[],[f347,f324])).
fof(f9884,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65)))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK12(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))))) )),
  inference(superposition,[],[f347,f316])).
fof(f9883,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58)))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK12(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49)))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f9882,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK12(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f9881,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41)))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK12(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))))) )),
  inference(superposition,[],[f347,f357])).
fof(f9880,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33)))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK12(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9879,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23)))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK12(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9878,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK12(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))))) )),
  inference(superposition,[],[f347,f309])).
fof(f9877,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9633,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9605,f262])).
fof(f9605,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f346,f445])).
fof(f9632,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f9631,f309])).
fof(f9631,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(subsumption_resolution,[],[f9603,f262])).
fof(f9603,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(superposition,[],[f346,f440])).
fof(f9630,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f9601,f342])).
fof(f9601,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f9629,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f9600,f336])).
fof(f9600,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f9628,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f9599,f341])).
fof(f9599,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f9627,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f9598,f337])).
fof(f9598,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f9626,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f9597,f324])).
fof(f9597,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f9625,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f9596,f316])).
fof(f9596,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f9624,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f9595,f367])).
fof(f9595,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(superposition,[],[f346,f367])).
fof(f9623,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f9594,f338])).
fof(f9594,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(superposition,[],[f346,f338])).
fof(f9622,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f9593,f357])).
fof(f9593,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(superposition,[],[f346,f357])).
fof(f9621,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f9590,f309])).
fof(f9590,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f9620,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f9589,f368])).
fof(f9589,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f9602,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK11(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f9592,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK11(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f9591,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK11(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f9588,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK11(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f346,f836])).
fof(f9587,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK11(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103))))) )),
  inference(superposition,[],[f346,f342])).
fof(f9586,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK11(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95))))) )),
  inference(superposition,[],[f346,f336])).
fof(f9585,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK11(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88))))) )),
  inference(superposition,[],[f346,f341])).
fof(f9584,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK11(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80))))) )),
  inference(superposition,[],[f346,f337])).
fof(f9583,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK11(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73))))) )),
  inference(superposition,[],[f346,f324])).
fof(f9582,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK11(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65))))) )),
  inference(superposition,[],[f346,f316])).
fof(f9581,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK11(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58))))) )),
  inference(superposition,[],[f346,f367])).
fof(f9580,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK11(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48))))) )),
  inference(superposition,[],[f346,f338])).
fof(f9579,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK11(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41))))) )),
  inference(superposition,[],[f346,f357])).
fof(f9578,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK11(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33))))) )),
  inference(superposition,[],[f346,f369])).
fof(f9577,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK11(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23))))) )),
  inference(superposition,[],[f346,f369])).
fof(f9576,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK11(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13))))) )),
  inference(superposition,[],[f346,f309])).
fof(f9575,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK11(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7))))) )),
  inference(superposition,[],[f346,f368])).
fof(f9367,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f9366,f309])).
fof(f9366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f9338,f262])).
fof(f9338,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f345,f440])).
fof(f8789,plain,(
  ( ! [X420,X422,X416,X418,X415,X421,X417,X419] : (p(s(bool,happ(s(fun(X415,bool),happ(s(fun(X417,fun(X415,bool)),happ(s(fun(X416,fun(X417,fun(X415,bool))),X418),s(X416,X419))),s(X417,X420))),s(X415,sK6(X421,happ(s(fun('type.pair.prod'(X416,X417),fun(X415,bool)),sK19(X418,X417,X416,fun(X415,bool))),s('type.pair.prod'(X416,X417),'const.pair.,_2'(s(X416,X419),s(X417,X420)))),X415))))) | ~p(s(bool,happ(s(fun(X415,bool),happ(s(fun(X417,fun(X415,bool)),happ(s(fun(X416,fun(X417,fun(X415,bool))),X418),s(X416,X419))),s(X417,X420))),s(X415,sK5(X421,happ(s(fun('type.pair.prod'(X416,X417),fun(X415,bool)),sK19(X418,X417,X416,fun(X415,bool))),s('type.pair.prod'(X416,X417),'const.pair.,_2'(s(X416,X419),s(X417,X420)))),X415))))) | ~p(s(bool,happ(s(fun(X415,bool),X421),s(X415,X422))))) )),
  inference(forward_demodulation,[],[f8699,f367])).
fof(f8699,plain,(
  ( ! [X420,X422,X416,X418,X415,X421,X417,X419] : (~p(s(bool,happ(s(fun(X415,bool),happ(s(fun(X417,fun(X415,bool)),happ(s(fun(X416,fun(X417,fun(X415,bool))),X418),s(X416,X419))),s(X417,X420))),s(X415,sK5(X421,happ(s(fun('type.pair.prod'(X416,X417),fun(X415,bool)),sK19(X418,X417,X416,fun(X415,bool))),s('type.pair.prod'(X416,X417),'const.pair.,_2'(s(X416,X419),s(X417,X420)))),X415))))) | ~p(s(bool,happ(s(fun(X415,bool),X421),s(X415,X422)))) | p(s(bool,happ(s(fun(X415,bool),happ(s(fun('type.pair.prod'(X416,X417),fun(X415,bool)),sK19(X418,X417,X416,fun(X415,bool))),s('type.pair.prod'(X416,X417),'const.pair.,_2'(s(X416,X419),s(X417,X420))))),s(X415,sK6(X421,happ(s(fun('type.pair.prod'(X416,X417),fun(X415,bool)),sK19(X418,X417,X416,fun(X415,bool))),s('type.pair.prod'(X416,X417),'const.pair.,_2'(s(X416,X419),s(X417,X420)))),X415)))))) )),
  inference(superposition,[],[f321,f367])).
fof(f6357,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6326,f262])).
fof(f6326,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f321,f445])).
fof(f6350,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69))))) )),
  inference(forward_demodulation,[],[f6309,f342])).
fof(f6309,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69)))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64)))))) )),
  inference(superposition,[],[f321,f342])).
fof(f6349,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63))))) )),
  inference(forward_demodulation,[],[f6308,f336])).
fof(f6308,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63)))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59)))))) )),
  inference(superposition,[],[f321,f336])).
fof(f6348,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58))))) )),
  inference(forward_demodulation,[],[f6307,f341])).
fof(f6307,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58)))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53)))))) )),
  inference(superposition,[],[f321,f341])).
fof(f6347,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52))))) )),
  inference(forward_demodulation,[],[f6306,f337])).
fof(f6306,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48)))))) )),
  inference(superposition,[],[f321,f337])).
fof(f6346,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47))))) )),
  inference(forward_demodulation,[],[f6305,f324])).
fof(f6305,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47)))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)))))) )),
  inference(superposition,[],[f321,f324])).
fof(f6345,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41))))) )),
  inference(forward_demodulation,[],[f6304,f316])).
fof(f6304,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37)))))) )),
  inference(superposition,[],[f321,f316])).
fof(f6344,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36))))) )),
  inference(forward_demodulation,[],[f6303,f338])).
fof(f6303,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36)))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32)))))) )),
  inference(superposition,[],[f321,f338])).
fof(f6343,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31))))) )),
  inference(forward_demodulation,[],[f6302,f357])).
fof(f6302,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31)))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26)))))) )),
  inference(superposition,[],[f321,f357])).
fof(f6342,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f6299,f309])).
fof(f6299,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f6341,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f6298,f368])).
fof(f6298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0)))))) )),
  inference(superposition,[],[f321,f368])).
fof(f6339,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,happ(s(fun(X66,X67),X69),s(X66,sK5(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66))))))) | ~p(s(bool,happ(s(fun(X66,bool),X70),s(X66,X71)))) | p(s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69))),s(X66,sK6(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6310,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | ~p(s(bool,happ(s(fun(X70,bool),X73),s(X70,X74)))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK6(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f321,f836])).
fof(f6301,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,sK5(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18))))) | ~p(s(bool,happ(s(fun(X18,bool),X24),s(X18,X25)))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK6(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6300,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,sK5(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10))))) | ~p(s(bool,happ(s(fun(X10,bool),X16),s(X10,X17)))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK6(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10)))))) )),
  inference(superposition,[],[f321,f369])).
fof(f6044,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6005,f262])).
fof(f6005,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f320,f445])).
fof(f6039,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(forward_demodulation,[],[f5951,f342])).
fof(f5951,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(superposition,[],[f320,f342])).
fof(f6038,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(forward_demodulation,[],[f5950,f336])).
fof(f5950,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(superposition,[],[f320,f336])).
fof(f6037,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(forward_demodulation,[],[f5949,f341])).
fof(f5949,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(superposition,[],[f320,f341])).
fof(f6036,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(forward_demodulation,[],[f5948,f337])).
fof(f5948,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(superposition,[],[f320,f337])).
fof(f6034,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(forward_demodulation,[],[f5946,f316])).
fof(f5946,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(superposition,[],[f320,f316])).
fof(f6032,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(forward_demodulation,[],[f5944,f357])).
fof(f5944,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(superposition,[],[f320,f357])).
fof(f6031,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(forward_demodulation,[],[f5941,f309])).
fof(f5941,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f5952,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,X73)))) | p(s(bool,happ(s(fun(X70,bool),X74),s(X70,sK6(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK5(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f320,f836])).
fof(f4781,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4748,f262])).
fof(f4748,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4686,f445])).
fof(f4702,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4686,f836])).
fof(f4686,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4559,f271])).
fof(f4559,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f319,f836])).
fof(f4609,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f317,f836])).
fof(f1904,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1883,f262])).
fof(f1883,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f317,f445])).
fof(f347,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f346,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f321,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X34] : (! [X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))) => (? [X38] : p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => ? [X38] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X38))))))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.MONO_EXISTS')).
fof(f320,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9] : ! [X34,X38] : (p(s(bool,happ(s(fun(X9,bool),X34),s(X9,X38)))) => p(s(bool,happ(s(fun(X9,bool),X34),s(X9,'const.min.@_1'(s(fun(X9,bool),X34)))))))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.SELECT_AX')).
fof(f12428,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X95),s('type.num.num',X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X94),s(fun('type.num.num',bool),X95),s('type.num.num','const.num.SUC_1'(s('type.num.num',X96)))))) ) | $spl10),
  inference(superposition,[],[f10983,f324])).
fof(f10001,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X142),s('type.num.num',X143))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X141),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.SUC_1'(s('type.num.num',X143)))))) )),
  inference(superposition,[],[f6500,f324])).
fof(f9698,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X144,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X142),s('type.num.num',X143))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X141),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.SUC_1'(s('type.num.num',X143)))))) )),
  inference(superposition,[],[f6283,f324])).
fof(f5625,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun('type.num.num',bool),X138),s('type.num.num',X139)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK18(X141,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))),X140,bool))))) )),
  inference(superposition,[],[f4861,f324])).
fof(f2663,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f324,f445])).
fof(f2661,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f324,f445])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f319,f445])).
fof(f6365,plain,(
  ( ! [X64,X65] : (p(s(bool,sK6(X64,'_i_0_0',bool))) | ~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))))) )),
  inference(forward_demodulation,[],[f6338,f309])).
fof(f6338,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X64,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f321,f309])).
fof(f6313,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X7,X8,bool))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool))) )),
  inference(superposition,[],[f321,f445])).
fof(f6312,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,sK6(X3,X4,bool))))) | s(bool,X6) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f321,f445])).
fof(f1376,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f319,f445])).
fof(f3329,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK14(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK14(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f3330,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK14(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK14(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f351,f445])).
fof(f12470,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(subsumption_resolution,[],[f12391,f432])).
fof(f12391,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(superposition,[],[f10983,f2803])).
fof(f11227,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f11135,f432])).
fof(f11135,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5646,f2803])).
fof(f10993,plain,(
  ( ! [X54,X53] : (s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(subsumption_resolution,[],[f10911,f264])).
fof(f10911,plain,(
  ( ! [X54,X53] : (p(s(bool,'F_0')) | s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(superposition,[],[f2426,f2913])).
fof(f10933,plain,(
  ( ! [X98] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X98),s(bool,sK15('_i_0_0',X98,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X98) | s(bool,'T_0') = s(bool,sK15('_i_0_0',X98,bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2913])).
fof(f10932,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X97),s(bool,sK15(X97,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X97) | s(bool,'T_0') = s(bool,sK15(X97,'_i_0_0',bool,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2913])).
fof(f10862,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103))) | s(bool,'T_0') = s(bool,sK4(X103,X104,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X104)))) ) | $spl10),
  inference(superposition,[],[f2913,f836])).
fof(f10857,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) ) | $spl10),
  inference(superposition,[],[f2913,f324])).
fof(f10855,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X84,bool),happ(s(fun(X83,fun(X84,bool)),X85),s(X83,X86))),s(X84,X87))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X83,X84),bool),sK19(X85,X84,X83,bool)),s('type.pair.prod'(X83,X84),'const.pair.,_2'(s(X83,X86),s(X84,X87)))))) ) | $spl10),
  inference(superposition,[],[f2913,f367])).
fof(f10854,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),X82))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X81),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X81),X82)))) ) | $spl10),
  inference(superposition,[],[f2913,f338])).
fof(f10852,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X73,X74),X76),s(X73,X77))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X73,X74),X76))),s(X73,X77)))) ) | $spl10),
  inference(superposition,[],[f2913,f369])).
fof(f10851,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X70,bool),'_dst_x0x1_2'(s(fun(X68,bool),X69),s(fun(X70,X68),X71))),s(X70,X72))))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,happ(s(fun(X70,X68),X71),s(X70,X72)))))) ) | $spl10),
  inference(superposition,[],[f2913,f369])).
fof(f10632,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f10631,f262])).
fof(f10631,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f10595,f2070])).
fof(f10595,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f311,f2531])).
fof(f10248,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f10247,f262])).
fof(f10247,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f10222,f2070])).
fof(f10222,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f348,f2531])).
fof(f10042,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9967,f432])).
fof(f9967,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6500,f2803])).
fof(f9938,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9937,f262])).
fof(f9937,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9912,f2070])).
fof(f9912,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f347,f2531])).
fof(f9738,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9664,f432])).
fof(f9664,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6283,f2803])).
fof(f9636,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9635,f262])).
fof(f9635,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9610,f2070])).
fof(f9610,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f346,f2531])).
fof(f9371,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9370,f262])).
fof(f9370,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9345,f2070])).
fof(f9345,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f345,f2531])).
fof(f8757,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(forward_demodulation,[],[f8623,f2070])).
fof(f8623,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(superposition,[],[f367,f2531])).
fof(f8684,plain,(
  ( ! [X325,X321,X323,X324,X326,X322] : (s(bool,'T_0') != s(bool,happ(s(fun(X322,bool),happ(s(fun(X321,fun(X322,bool)),X323),s(X321,X324))),s(X322,X325))) | s(bool,X326) = s(bool,happ(s(fun('type.pair.prod'(X321,X322),bool),sK19(X323,X322,X321,bool)),s('type.pair.prod'(X321,X322),'const.pair.,_2'(s(X321,X324),s(X322,X325))))) | s(bool,'F_0') = s(bool,X326)) ) | $spl10),
  inference(superposition,[],[f3633,f367])).
fof(f8672,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X256),'const.sum.INL_1'(s(bool,happ(s(fun(X252,bool),happ(s(fun(X251,fun(X252,bool)),X253),s(X251,X254))),s(X252,X255))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X251,X252),bool),sK19(X253,X252,X251,bool)),s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X254),s(X252,X255)))))) )),
  inference(superposition,[],[f2421,f367])).
fof(f8671,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X250,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X246,bool),happ(s(fun(X245,fun(X246,bool)),X247),s(X245,X248))),s(X246,X249))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X245,X246),bool),sK19(X247,X246,X245,bool)),s('type.pair.prod'(X245,X246),'const.pair.,_2'(s(X245,X248),s(X246,X249)))))) )),
  inference(superposition,[],[f2419,f367])).
fof(f8667,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X228,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X224,bool),happ(s(fun(X223,fun(X224,bool)),X225),s(X223,X226))),s(X224,X227))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X223,X224),bool),sK19(X225,X224,X223,bool)),s('type.pair.prod'(X223,X224),'const.pair.,_2'(s(X223,X226),s(X224,X227)))))) )),
  inference(superposition,[],[f540,f367])).
fof(f8666,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X222),'const.sum.INL_1'(s(bool,happ(s(fun(X218,bool),happ(s(fun(X217,fun(X218,bool)),X219),s(X217,X220))),s(X218,X221))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X217,X218),bool),sK19(X219,X218,X217,bool)),s('type.pair.prod'(X217,X218),'const.pair.,_2'(s(X217,X220),s(X218,X221)))))) )),
  inference(superposition,[],[f513,f367])).
fof(f8664,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X207,bool),happ(s(fun(X206,fun(X207,bool)),X208),s(X206,X209))),s(X207,X210))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X206,X207),bool),sK19(X208,X207,X206,bool)),s('type.pair.prod'(X206,X207),'const.pair.,_2'(s(X206,X209),s(X207,X210)))))) )),
  inference(superposition,[],[f472,f367])).
fof(f8626,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X58,bool),happ(s(fun(X57,fun(X58,bool)),X59),s(X57,X60))),s(X58,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X57,X58),bool),sK19(X59,X58,X57,bool)),s('type.pair.prod'(X57,X58),'const.pair.,_2'(s(X57,X60),s(X58,X61)))))) ) | $spl10),
  inference(superposition,[],[f367,f2803])).
fof(f8625,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,X54))),s(X52,X55))) | s(bool,X56) = s(bool,happ(s(fun('type.pair.prod'(X51,X52),bool),sK19(X53,X52,X51,bool)),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl10),
  inference(superposition,[],[f367,f2803])).
fof(f8617,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),sK19(X8,X7,X6,bool)),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f367,f270])).
fof(f7555,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f7554,f262])).
fof(f7554,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f7524,f2070])).
fof(f7524,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f350,f2531])).
fof(f14263,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f7144,f261,f260,f262,f265,f266,f267,f268,f269,f270,f272,f271,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f494,f803,f836,f837,f838,f839,f840,f841,f842,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2524,f2661,f2663,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3321,f3322,f3329,f3330,f4517,f4527,f4528,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4579,f4581,f4582,f4587,f4589,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4760,f4761,f4797,f4807,f4841,f4847,f4861,f4891,f4890,f4884,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5307,f5308,f5318,f5418,f5419,f5421,f5422,f5425,f5427,f5428,f5433,f5435,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5504,f5506,f5507,f5512,f5514,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5642,f5647,f5651,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6050,f6018,f6019,f6020,f6051,f6283,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6312,f6313,f6357,f6365,f6339,f6366,f6500,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6592,f6593,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6622,f6623,f6634,f6635,f6636,f6639,f6640,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6850,f6851,f6863,f6864,f6908,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8242,f8270,f8271,f8367,f8368,f8378,f8617,f8618,f8619,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f8845,f8848,f8851,f8853,f8854,f8855,f8856,f8857,f9233,f9236,f9239,f9241,f9242,f9243,f9244,f9245,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9734,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9829,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f10037,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10136,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10605,f10606,f11095,f11096,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f12906,f12909,f12912,f12914,f12915,f12916,f12917,f12918,f6915,f7141,f8042,f11852,f4620,f6620,f6586,f4615,f264,f275,f276,f472,f513,f540,f633,f686,f743,f909,f1378,f1488,f2371,f2417,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2529,f2530,f2689,f3320,f3328,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4526,f4577,f4578,f4588,f4590,f4591,f4595,f4596,f5306,f5423,f5424,f5434,f5436,f5437,f5441,f5442,f5502,f5503,f5513,f5515,f5516,f5520,f5521,f5641,f5643,f5655,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6293,f6311,f6356,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6510,f6590,f6591,f6621,f6637,f6638,f6849,f6862])).
fof(f7144,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(subsumption_resolution,[],[f7143,f262])).
fof(f7143,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f7134,f2070])).
fof(f7134,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(superposition,[],[f6915,f2531])).
fof(f7125,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f7124,f262])).
fof(f7124,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f7094,f2070])).
fof(f7094,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f349,f2531])).
fof(f6911,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6910,f262])).
fof(f6910,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(forward_demodulation,[],[f6881,f2070])).
fof(f6881,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f6905,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(forward_demodulation,[],[f6868,f2070])).
fof(f6868,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f6901,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'T_0'),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f6855,f2070])).
fof(f6855,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2531])).
fof(f6871,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,sK7(X33,bool,X34)),s(bool,X35)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK8(X33,bool,X34))) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f6870,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,sK7(X29,bool,X30)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),X31)))) | s(bool,X32) = s(bool,sK8(X29,bool,X30)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f6862,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK7(X0,bool,X1)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK8(X0,bool,X1))) )),
  inference(superposition,[],[f339,f270])).
fof(f6858,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,sK8(X33,X34,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK7(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f6857,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,sK8(X29,X30,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),X31)))) | s(bool,X32) = s(bool,sK7(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f339,f2803])).
fof(f6849,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK8(X0,X1,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool))) )),
  inference(superposition,[],[f339,f270])).
fof(f6685,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6644,f2070])).
fof(f6644,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6681,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'T_0'))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(forward_demodulation,[],[f6627,f2070])).
fof(f6627,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6667,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6597,f2070])).
fof(f6597,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2531])).
fof(f6647,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X46,X47,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X46,X47,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6646,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X42,X43,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X42,X43,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6638,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X13,X14,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,X14,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f6637,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15('_i_0_0',X11,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15('_i_0_0',X11,bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f6630,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,happ(s(fun(bool,X35),X33),s(bool,X36))) != s(X35,happ(s(fun(bool,X35),X34),s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK15(X33,X34,bool,X35))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6629,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,happ(s(fun(bool,X31),X29),s(bool,'F_0'))) != s(X31,happ(s(fun(bool,X31),X30),s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK15(X29,X30,bool,X31)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6621,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,X2))) )),
  inference(superposition,[],[f353,f270])).
fof(f6600,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X47,X46,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X47,X46,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6599,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X43,X42,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X43,X42,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f353,f2803])).
fof(f6591,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X14,X13,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X14,X13,X12,bool))))) )),
  inference(superposition,[],[f353,f270])).
fof(f6590,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15(X11,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15(X11,'_i_0_0',bool,bool))) )),
  inference(superposition,[],[f353,f472])).
fof(f6510,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6450,f264])).
fof(f6450,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'F_0')) | s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2426,f2421])).
fof(f6414,plain,(
  ( ! [X130,X131,X129] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f2421,f836])).
fof(f6413,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X125,bool),'const.pair.,_2'(s(X125,X126),s(bool,X127)))))) )),
  inference(superposition,[],[f2421,f342])).
fof(f6411,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X121),'const.sum.INL_1'(s(bool,X119))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120)))))) )),
  inference(superposition,[],[f2421,f341])).
fof(f6409,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X112),s('type.num.num',X113))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X111),s(fun('type.num.num',bool),X112),s('type.num.num','const.num.SUC_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f2421,f324])).
fof(f6408,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X110),'const.sum.INL_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2421,f316])).
fof(f6407,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f2421,f338])).
fof(f6406,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X104),'const.sum.INL_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK18(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2421,f357])).
fof(f6405,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f2421,f369])).
fof(f6404,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f2421,f369])).
fof(f6360,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6359,f262])).
fof(f6359,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(forward_demodulation,[],[f6331,f2070])).
fof(f6331,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(superposition,[],[f321,f2531])).
fof(f6356,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6(X0,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(forward_demodulation,[],[f6355,f309])).
fof(f6355,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f6324,f262])).
fof(f6324,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f321,f472])).
fof(f6352,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f6317,f2070])).
fof(f6317,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2531])).
fof(f6320,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,sK6(X33,X34,bool))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK5(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f321,f2803])).
fof(f6319,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X30),s(bool,sK6(X29,X30,bool))))) | s(bool,X32) = s(bool,sK5(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f321,f2803])).
fof(f6311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f321,f270])).
fof(f6293,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f6233,f264])).
fof(f6233,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f2426,f2419])).
fof(f6198,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X126))))) | s(bool,'T_0') = s(bool,sK4(X126,X127,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X127)))) )),
  inference(superposition,[],[f2419,f836])).
fof(f6197,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X125,bool),'const.sum.INR_1'(s(bool,X124))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X122,bool),'const.pair.,_2'(s(X122,X123),s(bool,X124)))))) )),
  inference(superposition,[],[f2419,f342])).
fof(f6195,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,X116))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),'const.pair.,_2'(s(bool,X116),s(X115,X117)))))) )),
  inference(superposition,[],[f2419,f341])).
fof(f6193,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X109),s('type.num.num',X110))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.SUC_1'(s('type.num.num',X110)))))) )),
  inference(superposition,[],[f2419,f324])).
fof(f6192,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,X105))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X105),s(fun('type.num.num',bool),X106),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2419,f316])).
fof(f6191,plain,(
  ( ! [X103,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X102),X103)))) )),
  inference(superposition,[],[f2419,f338])).
fof(f6190,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X100))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X98),bool),sK18(X99,X100,X98,bool)),s('type.list.list'(X98),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2419,f357])).
fof(f6189,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,happ(s(fun(X92,X93),X95),s(X92,X96))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X93,bool),X94),s(fun(X92,X93),X95))),s(X92,X96)))) )),
  inference(superposition,[],[f2419,f369])).
fof(f6188,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X88,bool),'_dst_x0x1_2'(s(fun(X86,bool),X87),s(fun(X88,X86),X89))),s(X88,X90))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90)))))) )),
  inference(superposition,[],[f2419,f369])).
fof(f6047,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(subsumption_resolution,[],[f6046,f262])).
fof(f6046,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(forward_demodulation,[],[f6010,f2070])).
fof(f6010,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(superposition,[],[f320,f2531])).
fof(f5655,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5616,f262])).
fof(f5616,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4861,f513])).
fof(f5643,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f5608,f262])).
fof(f5608,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f4861,f540])).
fof(f5641,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5606,f262])).
fof(f5606,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4861,f472])).
fof(f5634,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f5633,f262])).
fof(f5633,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f5595,f2070])).
fof(f5595,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f4861,f2531])).
fof(f5571,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(forward_demodulation,[],[f5312,f2070])).
fof(f5312,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(superposition,[],[f369,f2531])).
fof(f5531,plain,(
  ( ! [X282,X285,X287,X283,X284,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X283,bool),X284),s(X283,happ(s(fun(X282,X283),X285),s(X282,X286))))) | s(bool,X287) = s(bool,happ(s(fun(X282,bool),'_dst_x0x1_2'(s(fun(X283,bool),X284),s(fun(X282,X283),X285))),s(X282,X286))) | s(bool,'F_0') = s(bool,X287)) ) | $spl10),
  inference(superposition,[],[f3633,f369])).
fof(f5521,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (p(s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X224,X225),X227),s(X224,X228)))))) | s(bool,X229) != s(bool,happ(s(fun(X224,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X224,X225),X227))),s(X224,X228))) | s(bool,'F_0') = s(bool,X229)) )),
  inference(superposition,[],[f2426,f369])).
fof(f5520,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (~p(s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X218,X219),X221),s(X218,X222)))))) | s(bool,X223) = s(bool,happ(s(fun(X218,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X218,X219),X221))),s(X218,X222))) | s(bool,'F_0') = s(bool,X223)) )),
  inference(superposition,[],[f2424,f369])).
fof(f5516,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X197,bool),X198),s(X197,happ(s(fun(X196,X197),X199),s(X196,X200))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X197,bool),X198),s(fun(X196,X197),X199))),s(X196,X200)))) )),
  inference(superposition,[],[f540,f369])).
fof(f5515,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X195),'const.sum.INL_1'(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f513,f369])).
fof(f5513,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183)))) )),
  inference(superposition,[],[f472,f369])).
fof(f5503,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f276,f369])).
fof(f5502,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) )),
  inference(superposition,[],[f275,f369])).
fof(f5452,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X287,bool),'_dst_x0x1_2'(s(fun(X285,bool),X286),s(fun(X287,X285),X288))),s(X287,X289))) | s(bool,X290) = s(bool,happ(s(fun(X285,bool),X286),s(X285,happ(s(fun(X287,X285),X288),s(X287,X289))))) | s(bool,'F_0') = s(bool,X290)) ) | $spl10),
  inference(superposition,[],[f3633,f369])).
fof(f5442,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (p(s(bool,happ(s(fun(X229,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X229,X227),X230))),s(X229,X231)))) | s(bool,X232) != s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X229,X227),X230),s(X229,X231))))) | s(bool,'F_0') = s(bool,X232)) )),
  inference(superposition,[],[f2426,f369])).
fof(f5441,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (~p(s(bool,happ(s(fun(X223,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X223,X221),X224))),s(X223,X225)))) | s(bool,X226) = s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X223,X221),X224),s(X223,X225))))) | s(bool,'F_0') = s(bool,X226)) )),
  inference(superposition,[],[f2424,f369])).
fof(f5437,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),X200),s(fun(X201,X199),X202))),s(X201,X203))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X199,bool),X200),s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))))) )),
  inference(superposition,[],[f540,f369])).
fof(f5436,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X198),'const.sum.INL_1'(s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X193,bool),X194),s(fun(X195,X193),X196))),s(X195,X197))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,happ(s(fun(X195,X193),X196),s(X195,X197)))))) )),
  inference(superposition,[],[f513,f369])).
fof(f5434,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X184,X182),X185))),s(X184,X186))))) | s(bool,'F_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X184,X182),X185),s(X184,X186)))))) )),
  inference(superposition,[],[f472,f369])).
fof(f5424,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X123,bool),X124),s(fun(X125,X123),X126))),s(X125,X127)))) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,happ(s(fun(X125,X123),X126),s(X125,X127)))))) )),
  inference(superposition,[],[f276,f369])).
fof(f5423,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'F_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) )),
  inference(superposition,[],[f275,f369])).
fof(f5315,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,X82) = s(bool,happ(s(fun(X79,bool),'_dst_x0x1_2'(s(fun(X77,bool),X78),s(fun(X79,X77),X80))),s(X79,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,happ(s(fun(X77,bool),X78),s(X77,happ(s(fun(X79,X77),X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f369,f2803])).
fof(f5314,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X71,bool),X72),s(fun(X73,X71),X74))),s(X73,X75))) | s(bool,X76) = s(bool,happ(s(fun(X71,bool),X72),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f369,f2803])).
fof(f5306,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X28,X26),X29))),s(X28,X30))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30)))))) )),
  inference(superposition,[],[f369,f270])).
fof(f4936,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4901,f264])).
fof(f4901,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4847,f2803])).
fof(f4866,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4857,f264])).
fof(f4857,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4760,f2803])).
fof(f4844,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4843,f262])).
fof(f4843,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4815,f2070])).
fof(f4815,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f355,f2531])).
fof(f4784,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4783,f262])).
fof(f4783,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f4753,f2070])).
fof(f4753,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f4686,f2531])).
fof(f4685,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4532,f2070])).
fof(f4532,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2531])).
fof(f4606,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3633,f836])).
fof(f4596,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | s(bool,X113) != s(bool,sK4(X111,X112,bool)) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) )),
  inference(superposition,[],[f2426,f836])).
fof(f4595,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | s(bool,X110) = s(bool,sK4(X108,X109,bool)) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X109)))) )),
  inference(superposition,[],[f2424,f836])).
fof(f4591,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f540,f836])).
fof(f4590,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f513,f836])).
fof(f4588,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f472,f836])).
fof(f4578,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f276,f836])).
fof(f4577,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f275,f836])).
fof(f4535,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f836,f2803])).
fof(f4534,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f836,f2803])).
fof(f4526,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f836,f270])).
fof(f4520,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4519,f262])).
fof(f4519,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4491,f2070])).
fof(f4491,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f354,f2531])).
fof(f4463,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4406,f262])).
fof(f4406,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f540])).
fof(f4383,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94)))))) )),
  inference(superposition,[],[f540,f342])).
fof(f4382,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90)))))) )),
  inference(superposition,[],[f540,f336])).
fof(f4381,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X88,bool),'const.sum.INR_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f540,f341])).
fof(f4380,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83)))))) )),
  inference(superposition,[],[f540,f337])).
fof(f4379,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f540,f324])).
fof(f4378,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f540,f316])).
fof(f4377,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f540,f338])).
fof(f4376,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK18(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f540,f357])).
fof(f4344,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4288,f262])).
fof(f4288,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f447,f513])).
fof(f4265,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f513,f342])).
fof(f4264,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X86),'const.sum.INL_1'(s(bool,X87)))))) )),
  inference(superposition,[],[f513,f336])).
fof(f4263,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f513,f341])).
fof(f4262,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,X80)))))) )),
  inference(superposition,[],[f513,f337])).
fof(f4261,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f513,f324])).
fof(f4260,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f513,f316])).
fof(f4259,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f513,f338])).
fof(f4258,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK18(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f513,f357])).
fof(f4148,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3633,f342])).
fof(f4147,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,X87)) ) | $spl10),
  inference(superposition,[],[f3633,f336])).
fof(f4146,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3633,f341])).
fof(f4145,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X80) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3633,f337])).
fof(f4144,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X75),s('type.num.num',X76))) | s(bool,X77) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl10),
  inference(superposition,[],[f3633,f324])).
fof(f4143,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3633,f316])).
fof(f4142,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3633,f338])).
fof(f4141,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK18(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3633,f357])).
fof(f4140,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3633,f309])).
fof(f4081,plain,(
  ( ! [X101,X99,X100,X98] : (s(X98,X100) = s(X98,sK4(X100,X99,X98)) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(bool,'F_0') = s(bool,X101)) ) | $spl10),
  inference(subsumption_resolution,[],[f3704,f264])).
fof(f3704,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,'F_0')) | s(X98,X100) = s(X98,sK4(X100,X99,X98)) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99))) | s(bool,'F_0') = s(bool,X101)) ) | $spl10),
  inference(superposition,[],[f318,f2803])).
fof(f4078,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(subsumption_resolution,[],[f3696,f264])).
fof(f3696,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(superposition,[],[f420,f2803])).
fof(f4073,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3668,f264])).
fof(f3668,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f330,f2803])).
fof(f3729,plain,(
  ( ! [X187,X185,X188,X186] : (s(X187,f(s(bool,'F_0'))) != s(X187,g(s(bool,'F_0'))) | s(fun(bool,X187),X185) = s(fun(bool,X187),X186) | s(bool,X188) = s(bool,sK14(X185,X186,X187,bool)) | s(bool,'F_0') = s(bool,X188)) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3728,plain,(
  ( ! [X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,X184) = s(bool,sK4(X182,X183,bool)) | s(bool,'F_0') = s(bool,X184)) ) | $spl10),
  inference(superposition,[],[f319,f2803])).
fof(f3718,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'F_0') != s(bool,f(s(X148,sK14(X149,X150,bool,X148)))) | s(fun(X148,bool),X149) = s(fun(X148,bool),X150) | s(bool,X151) = s(bool,g(s(X148,sK14(X149,X150,bool,X148)))) | s(bool,'F_0') = s(bool,X151)) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3632,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f275,f2803])).
fof(f3559,plain,(
  ( ! [X198,X200,X199,X201] : (s(X200,f(s(bool,X201))) != s(X200,g(s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,'F_0') = s(bool,X201) | s(bool,'F_0') = s(bool,sK14(X198,X199,X200,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3558,plain,(
  ( ! [X196,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X196),s(bool,X197)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X196)))) | s(bool,'F_0') = s(bool,X197) | s(bool,'F_0') = s(bool,sK4(X195,X196,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2803])).
fof(f3548,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X164) != s(bool,f(s(X161,sK14(X162,X163,bool,X161)))) | s(fun(X161,bool),X162) = s(fun(X161,bool),X163) | s(bool,'F_0') = s(bool,X164) | s(bool,'F_0') = s(bool,g(s(X161,sK14(X162,X163,bool,X161))))) ) | $spl10),
  inference(superposition,[],[f351,f2803])).
fof(f3451,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2803,f342])).
fof(f3450,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') = s(bool,X71) | s(bool,X72) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X70),'const.sum.INL_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f2803,f336])).
fof(f3449,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2803,f341])).
fof(f3448,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X63,bool),'const.sum.INR_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2803,f337])).
fof(f3447,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61))) | s(bool,X62) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl10),
  inference(superposition,[],[f2803,f324])).
fof(f3446,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2803,f316])).
fof(f3445,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2803,f338])).
fof(f3444,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK18(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2803,f357])).
fof(f3380,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2803,f342])).
fof(f3379,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X83),'const.sum.INL_1'(s(bool,X84)))))) ) | $spl10),
  inference(superposition,[],[f2803,f336])).
fof(f3378,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2803,f341])).
fof(f3377,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X76,bool),'const.sum.INR_1'(s(bool,X77)))))) ) | $spl10),
  inference(superposition,[],[f2803,f337])).
fof(f3376,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl10),
  inference(superposition,[],[f2803,f324])).
fof(f3375,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2803,f316])).
fof(f3374,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2803,f338])).
fof(f3373,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK18(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2803,f357])).
fof(f3372,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2803,f309])).
fof(f3340,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f3334,f2070])).
fof(f3334,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f351,f2531])).
fof(f3337,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f3326,f2070])).
fof(f3326,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f351,f2531])).
fof(f3328,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK14(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK14(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f351,f270])).
fof(f3320,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,X2,bool))) )),
  inference(superposition,[],[f351,f270])).
fof(f2965,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2866,f2070])).
fof(f2866,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f319,f2531])).
fof(f2951,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2950,f262])).
fof(f2950,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2835,f2070])).
fof(f2835,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f317,f2531])).
fof(f2949,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2948,f262])).
fof(f2948,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2834,f2070])).
fof(f2834,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f319,f2531])).
fof(f2917,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2798,f2070])).
fof(f2798,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(superposition,[],[f2531,f324])).
fof(f2915,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2796,f2070])).
fof(f2796,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2531,f338])).
fof(f2913,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2794,f2070])).
fof(f2794,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2531,f309])).
fof(f2689,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f472,f324])).
fof(f2531,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2509,f262])).
fof(f2509,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f2424,f2070])).
fof(f2530,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2508,f262])).
fof(f2508,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2424,f440])).
fof(f2529,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2507,f262])).
fof(f2507,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2424,f472])).
fof(f2427,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2394,f262])).
fof(f2394,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f447,f472])).
fof(f2426,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2425,f309])).
fof(f2425,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2391,f262])).
fof(f2391,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f418,f472])).
fof(f2424,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2423,f309])).
fof(f2423,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2390,f262])).
fof(f2390,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f333,f472])).
fof(f2422,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(forward_demodulation,[],[f2376,f309])).
fof(f2376,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f472,f342])).
fof(f2421,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(forward_demodulation,[],[f2375,f309])).
fof(f2375,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X35))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35)))))) )),
  inference(superposition,[],[f472,f336])).
fof(f2420,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2374,f309])).
fof(f2374,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f472,f341])).
fof(f2419,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f2373,f309])).
fof(f2373,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30)))))) )),
  inference(superposition,[],[f472,f337])).
fof(f2418,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2372,f309])).
fof(f2372,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f472,f316])).
fof(f2417,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2370,f309])).
fof(f2370,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f472,f357])).
fof(f2371,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f472,f338])).
fof(f1488,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X12),X13)))) )),
  inference(superposition,[],[f338,f270])).
fof(f1378,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f319,f270])).
fof(f909,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,X2,X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f357,f270])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f342,f270])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f341,f270])).
fof(f633,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f270])).
fof(f540,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f337,f270])).
fof(f513,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f336,f270])).
fof(f472,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f309,f270])).
fof(f276,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|3')).
fof(f275,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f264,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.modSem.v_|49||49||split|4',aHL_FALSITY)).
fof(f13336,plain,(
  ( ! [X23,X21,X22] : (s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool))) ) | $spl10),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4615,f4620,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6586,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6620,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8042,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10335,f10431,f10506,f10485,f10486,f10496,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10720,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11852,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238])).
fof(f8238,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f8225,f264])).
fof(f8225,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f7144,f2803])).
fof(f8233,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8217,f264])).
fof(f8217,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(superposition,[],[f7144,f270])).
fof(f7141,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7129,f262])).
fof(f7129,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6915,f445])).
fof(f6915,plain,(
  ( ! [X62,X63] : (~p(s(bool,sK7('const.pair.FST_0',X62,bool))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X62),X63))))) )),
  inference(forward_demodulation,[],[f6914,f338])).
fof(f6914,plain,(
  ( ! [X62,X63] : (~p(s(bool,sK7('const.pair.FST_0',X62,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X62),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X62),X63))))) )),
  inference(forward_demodulation,[],[f6888,f341])).
fof(f6888,plain,(
  ( ! [X62,X63] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X62),'const.pair.,_2'(s(bool,sK7('const.pair.FST_0',X62,bool)),s(X62,sK8('const.pair.FST_0',X62,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X62),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X62),X63))))) )),
  inference(superposition,[],[f339,f338])).
fof(f12919,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f9829,f2803])).
fof(f12918,plain,(
  ( ! [X175,X173,X174] : (s(bool,X173) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X175,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X173,X174,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X174)))) ) | $spl26),
  inference(superposition,[],[f9829,f836])).
fof(f12917,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,X171) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X172,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X169,bool),'const.pair.,_2'(s(X169,X170),s(bool,X171)))))) ) | $spl26),
  inference(superposition,[],[f9829,f342])).
fof(f12916,plain,(
  ( ! [X167,X166,X168] : (s(bool,X167) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X168,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X166),'const.sum.INL_1'(s(bool,X167)))))) ) | $spl26),
  inference(superposition,[],[f9829,f336])).
fof(f12915,plain,(
  ( ! [X165,X163,X164,X162] : (s(bool,X163) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X165,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) ) | $spl26),
  inference(superposition,[],[f9829,f341])).
fof(f12914,plain,(
  ( ! [X161,X159,X160] : (s(bool,X160) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X161,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X159,bool),'const.sum.INR_1'(s(bool,X160)))))) ) | $spl26),
  inference(superposition,[],[f9829,f337])).
fof(f12912,plain,(
  ( ! [X154,X152,X153] : (s(bool,X152) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X154,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X152),s(fun('type.num.num',bool),X153),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f9829,f316])).
fof(f12909,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,X141) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X142,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X139),bool),sK18(X140,X141,X139,bool)),s('type.list.list'(X139),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f9829,f357])).
fof(f12906,plain,(
  ( ! [X125,X126] : (s(bool,X125) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X125)))) ) | $spl26),
  inference(superposition,[],[f9829,f309])).
fof(f12834,plain,(
  ( ! [X356,X358,X354,X357,X359,X353,X355] : (s(fun(X353,X354),happ(s(fun(X355,fun(X353,X354)),X359),s(X355,sK15('_dst_x00x11_2'(s(fun(X356,X354),X357),s(fun(X355,fun(X353,X356)),X358)),X359,X355,fun(X353,X354))))) != s(fun(X353,X354),'_dst_x0x1_2'(s(fun(X356,X354),X357),s(fun(X353,X356),happ(s(fun(X355,fun(X353,X356)),X358),s(X355,sK15('_dst_x00x11_2'(s(fun(X356,X354),X357),s(fun(X355,fun(X353,X356)),X358)),X359,X355,fun(X353,X354))))))) | s(fun(X355,fun(X353,X354)),X359) = s(fun(X355,fun(X353,X354)),'_dst_x00x11_2'(s(fun(X356,X354),X357),s(fun(X355,fun(X353,X356)),X358)))) )),
  inference(superposition,[],[f353,f397])).
fof(f12833,plain,(
  ( ! [X350,X346,X352,X349,X351,X347,X348] : (s(fun(X346,X347),happ(s(fun(X348,fun(X346,X347)),X352),s(X348,sK15(X352,'_dst_x00x11_2'(s(fun(X349,X347),X350),s(fun(X348,fun(X346,X349)),X351)),X348,fun(X346,X347))))) != s(fun(X346,X347),'_dst_x0x1_2'(s(fun(X349,X347),X350),s(fun(X346,X349),happ(s(fun(X348,fun(X346,X349)),X351),s(X348,sK15(X352,'_dst_x00x11_2'(s(fun(X349,X347),X350),s(fun(X348,fun(X346,X349)),X351)),X348,fun(X346,X347))))))) | s(fun(X348,fun(X346,X347)),X352) = s(fun(X348,fun(X346,X347)),'_dst_x00x11_2'(s(fun(X349,X347),X350),s(fun(X348,fun(X346,X349)),X351)))) )),
  inference(superposition,[],[f353,f397])).
fof(f12831,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X328,X330,X336] : (s(X335,happ(s(fun(fun(X328,X329),X335),X336),s(fun(X328,X329),'_dst_x0x1_2'(s(fun(X331,X329),X332),s(fun(X328,X331),happ(s(fun(X330,fun(X328,X331)),X333),s(X330,X334))))))) = s(X335,happ(s(fun(X330,X335),'_dst_x0x1_2'(s(fun(fun(X328,X329),X335),X336),s(fun(X330,fun(X328,X329)),'_dst_x00x11_2'(s(fun(X331,X329),X332),s(fun(X330,fun(X328,X331)),X333))))),s(X330,X334)))) )),
  inference(superposition,[],[f369,f397])).
fof(f12821,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X238,X245] : (s(X239,happ(s(fun(X238,X239),X245),s(X238,sK15(happ(s(fun(X240,fun(X238,X239)),'_dst_x00x11_2'(s(fun(X241,X239),X242),s(fun(X240,fun(X238,X241)),X243))),s(X240,X244)),X245,X238,X239)))) != s(X239,happ(s(fun(X238,X239),'_dst_x0x1_2'(s(fun(X241,X239),X242),s(fun(X238,X241),happ(s(fun(X240,fun(X238,X241)),X243),s(X240,X244))))),s(X238,sK15(happ(s(fun(X240,fun(X238,X239)),'_dst_x00x11_2'(s(fun(X241,X239),X242),s(fun(X240,fun(X238,X241)),X243))),s(X240,X244)),X245,X238,X239)))) | s(fun(X238,X239),X245) = s(fun(X238,X239),happ(s(fun(X240,fun(X238,X239)),'_dst_x00x11_2'(s(fun(X241,X239),X242),s(fun(X240,fun(X238,X241)),X243))),s(X240,X244)))) )),
  inference(superposition,[],[f353,f397])).
fof(f12820,plain,(
  ( ! [X231,X237,X235,X233,X230,X236,X234,X232] : (s(X231,happ(s(fun(X230,X231),X237),s(X230,sK15(X237,happ(s(fun(X232,fun(X230,X231)),'_dst_x00x11_2'(s(fun(X233,X231),X234),s(fun(X232,fun(X230,X233)),X235))),s(X232,X236)),X230,X231)))) != s(X231,happ(s(fun(X230,X231),'_dst_x0x1_2'(s(fun(X233,X231),X234),s(fun(X230,X233),happ(s(fun(X232,fun(X230,X233)),X235),s(X232,X236))))),s(X230,sK15(X237,happ(s(fun(X232,fun(X230,X231)),'_dst_x00x11_2'(s(fun(X233,X231),X234),s(fun(X232,fun(X230,X233)),X235))),s(X232,X236)),X230,X231)))) | s(fun(X230,X231),X237) = s(fun(X230,X231),happ(s(fun(X232,fun(X230,X231)),'_dst_x00x11_2'(s(fun(X233,X231),X234),s(fun(X232,fun(X230,X233)),X235))),s(X232,X236)))) )),
  inference(superposition,[],[f353,f397])).
fof(f12818,plain,(
  ( ! [X212,X218,X216,X220,X215,X213,X219,X217,X214] : (s(fun(X212,X213),X219) != s(fun(X212,X213),'_dst_x0x1_2'(s(fun(X215,X213),X216),s(fun(X212,X215),happ(s(fun(X214,fun(X212,X215)),X217),s(X214,X218))))) | s(X213,happ(s(fun(X212,X213),X219),s(X212,X220))) = s(X213,happ(s(fun(X212,X213),happ(s(fun(X214,fun(X212,X213)),'_dst_x00x11_2'(s(fun(X215,X213),X216),s(fun(X214,fun(X212,X215)),X217))),s(X214,X218))),s(X212,X220)))) )),
  inference(superposition,[],[f352,f397])).
fof(f12802,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X65,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X65,X67),happ(s(fun(X66,fun(X65,X67)),X69),s(X66,X70))))))) | s(X65,X71) = s(X65,sK4(X71,happ(s(fun(X66,fun(X65,bool)),'_dst_x00x11_2'(s(fun(X67,bool),X68),s(fun(X66,fun(X65,X67)),X69))),s(X66,X70)),X65))) )),
  inference(superposition,[],[f836,f397])).
fof(f12799,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),'_dst_x0x1_2'(s(fun(X44,bool),X45),s(fun(X42,X44),happ(s(fun(X43,fun(X42,X44)),X46),s(X43,X47))))),s(X42,sK4(X48,happ(s(fun(X43,fun(X42,bool)),'_dst_x00x11_2'(s(fun(X44,bool),X45),s(fun(X43,fun(X42,X44)),X46))),s(X43,X47)),X42))))) | p(s(bool,'|dollar|exists_1'(s(fun(X42,bool),happ(s(fun(X43,fun(X42,bool)),'_dst_x00x11_2'(s(fun(X44,bool),X45),s(fun(X43,fun(X42,X44)),X46))),s(X43,X47))))))) )),
  inference(superposition,[],[f319,f397])).
fof(f12767,plain,(
  ( ! [X253,X251,X249,X246,X252,X250,X248,X247] : (s(X247,happ(s(fun(X246,X247),X253),s(X246,sK15('_dst_x0x1_2'(s(fun(X248,X247),X249),s(fun(X246,X248),happ(s(fun(X250,fun(X246,X248)),X251),s(X250,X252)))),X253,X246,X247)))) != s(X247,happ(s(fun(X246,X247),happ(s(fun(X250,fun(X246,X247)),'_dst_x00x11_2'(s(fun(X248,X247),X249),s(fun(X250,fun(X246,X248)),X251))),s(X250,X252))),s(X246,sK15('_dst_x0x1_2'(s(fun(X248,X247),X249),s(fun(X246,X248),happ(s(fun(X250,fun(X246,X248)),X251),s(X250,X252)))),X253,X246,X247)))) | s(fun(X246,X247),X253) = s(fun(X246,X247),'_dst_x0x1_2'(s(fun(X248,X247),X249),s(fun(X246,X248),happ(s(fun(X250,fun(X246,X248)),X251),s(X250,X252)))))) )),
  inference(superposition,[],[f353,f397])).
fof(f12766,plain,(
  ( ! [X243,X241,X239,X244,X242,X240,X238,X245] : (s(X239,happ(s(fun(X238,X239),X245),s(X238,sK15(X245,'_dst_x0x1_2'(s(fun(X240,X239),X241),s(fun(X238,X240),happ(s(fun(X242,fun(X238,X240)),X243),s(X242,X244)))),X238,X239)))) != s(X239,happ(s(fun(X238,X239),happ(s(fun(X242,fun(X238,X239)),'_dst_x00x11_2'(s(fun(X240,X239),X241),s(fun(X242,fun(X238,X240)),X243))),s(X242,X244))),s(X238,sK15(X245,'_dst_x0x1_2'(s(fun(X240,X239),X241),s(fun(X238,X240),happ(s(fun(X242,fun(X238,X240)),X243),s(X242,X244)))),X238,X239)))) | s(fun(X238,X239),X245) = s(fun(X238,X239),'_dst_x0x1_2'(s(fun(X240,X239),X241),s(fun(X238,X240),happ(s(fun(X242,fun(X238,X240)),X243),s(X242,X244)))))) )),
  inference(superposition,[],[f353,f397])).
fof(f12764,plain,(
  ( ! [X222,X220,X227,X225,X223,X221,X228,X226,X224] : (s(fun(X220,X221),X227) != s(fun(X220,X221),happ(s(fun(X224,fun(X220,X221)),'_dst_x00x11_2'(s(fun(X222,X221),X223),s(fun(X224,fun(X220,X222)),X225))),s(X224,X226))) | s(X221,happ(s(fun(X220,X221),X227),s(X220,X228))) = s(X221,happ(s(fun(X220,X221),'_dst_x0x1_2'(s(fun(X222,X221),X223),s(fun(X220,X222),happ(s(fun(X224,fun(X220,X222)),X225),s(X224,X226))))),s(X220,X228)))) )),
  inference(superposition,[],[f352,f397])).
fof(f12748,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X73,bool),happ(s(fun(X76,fun(X73,bool)),'_dst_x00x11_2'(s(fun(X74,bool),X75),s(fun(X76,fun(X73,X74)),X77))),s(X76,X78))))) | s(X73,X79) = s(X73,sK4(X79,'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X73,X74),happ(s(fun(X76,fun(X73,X74)),X77),s(X76,X78)))),X73))) )),
  inference(superposition,[],[f836,f397])).
fof(f12745,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51] : (~p(s(bool,happ(s(fun(X50,bool),happ(s(fun(X53,fun(X50,bool)),'_dst_x00x11_2'(s(fun(X51,bool),X52),s(fun(X53,fun(X50,X51)),X54))),s(X53,X55))),s(X50,sK4(X56,'_dst_x0x1_2'(s(fun(X51,bool),X52),s(fun(X50,X51),happ(s(fun(X53,fun(X50,X51)),X54),s(X53,X55)))),X50))))) | p(s(bool,'|dollar|exists_1'(s(fun(X50,bool),'_dst_x0x1_2'(s(fun(X51,bool),X52),s(fun(X50,X51),happ(s(fun(X53,fun(X50,X51)),X54),s(X53,X55))))))))) )),
  inference(superposition,[],[f319,f397])).
fof(f12738,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X3),s(X2,happ(s(fun(X0,X2),happ(s(fun(X4,fun(X0,X2)),X5),s(X4,X6))),s(X0,X7))))) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X4,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X2,X1),X3),s(fun(X4,fun(X0,X2)),X5))),s(X4,X6))),s(X0,X7)))) )),
  inference(superposition,[],[f369,f397])).
fof(f12737,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(fun(X7,X8),'_dst_x0x1_2'(s(fun(X10,X8),X11),s(fun(X7,X10),happ(s(fun(X9,fun(X7,X10)),X12),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15))))))) = s(fun(X7,X8),happ(s(fun(X13,fun(X7,X8)),'_dst_x0x1_2'(s(fun(X9,fun(X7,X8)),'_dst_x00x11_2'(s(fun(X10,X8),X11),s(fun(X9,fun(X7,X10)),X12))),s(fun(X13,X9),X14))),s(X13,X15)))) )),
  inference(superposition,[],[f397,f369])).
fof(f12633,plain,(
  ( ! [X59,X57,X54,X52,X60,X58,X56,X55,X53] : (s(fun(X52,X59),'_dst_x0x1_2'(s(fun(X53,X59),X60),s(fun(X52,X53),happ(s(fun(X55,fun(X52,X53)),happ(s(fun(X54,fun(X55,fun(X52,X53))),X56),s(X54,X57))),s(X55,X58))))) = s(fun(X52,X59),happ(s(fun('type.pair.prod'(X54,X55),fun(X52,X59)),'_dst_x00x11_2'(s(fun(X53,X59),X60),s(fun('type.pair.prod'(X54,X55),fun(X52,X53)),sK19(X56,X55,X54,fun(X52,X53))))),s('type.pair.prod'(X54,X55),'const.pair.,_2'(s(X54,X57),s(X55,X58)))))) )),
  inference(superposition,[],[f397,f367])).
fof(f12632,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(fun(X39,X44),'_dst_x0x1_2'(s(fun(X40,X44),X45),s(fun(X39,X40),X43))) = s(fun(X39,X44),happ(s(fun('type.list.list'(X41),fun(X39,X44)),'_dst_x00x11_2'(s(fun(X40,X44),X45),s(fun('type.list.list'(X41),fun(X39,X40)),sK18(X42,X43,X41,fun(X39,X40))))),s('type.list.list'(X41),'const.list.NIL_0')))) )),
  inference(superposition,[],[f397,f357])).
fof(f12631,plain,(
  ( ! [X30,X37,X35,X33,X31,X38,X36,X34,X32] : (s(fun(X30,X37),'_dst_x0x1_2'(s(fun(X31,X37),X38),s(fun(X30,X31),happ(s(fun(X33,fun(X30,X31)),X34),s(X33,happ(s(fun(X32,X33),X35),s(X32,X36))))))) = s(fun(X30,X37),happ(s(fun(X32,fun(X30,X37)),'_dst_x00x11_2'(s(fun(X31,X37),X38),s(fun(X32,fun(X30,X31)),'_dst_x0x1_2'(s(fun(X33,fun(X30,X31)),X34),s(fun(X32,X33),X35))))),s(X32,X36)))) )),
  inference(superposition,[],[f397,f369])).
fof(f12630,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22] : (s(fun(X21,X28),happ(s(fun(X23,fun(X21,X28)),'_dst_x00x11_2'(s(fun(X22,X28),X29),s(fun(X23,fun(X21,X22)),X24))),s(X23,happ(s(fun(X25,X23),X26),s(X25,X27))))) = s(fun(X21,X28),'_dst_x0x1_2'(s(fun(X22,X28),X29),s(fun(X21,X22),happ(s(fun(X25,fun(X21,X22)),'_dst_x0x1_2'(s(fun(X23,fun(X21,X22)),X24),s(fun(X25,X23),X26))),s(X25,X27)))))) )),
  inference(superposition,[],[f397,f369])).
fof(f12629,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(fun(X7,X14),'_dst_x0x1_2'(s(fun(X8,X14),X15),s(fun(X7,X8),'_dst_x0x1_2'(s(fun(X10,X8),X11),s(fun(X7,X10),happ(s(fun(X9,fun(X7,X10)),X12),s(X9,X13))))))) = s(fun(X7,X14),happ(s(fun(X9,fun(X7,X14)),'_dst_x00x11_2'(s(fun(X8,X14),X15),s(fun(X9,fun(X7,X8)),'_dst_x00x11_2'(s(fun(X10,X8),X11),s(fun(X9,fun(X7,X10)),X12))))),s(X9,X13)))) )),
  inference(superposition,[],[f397,f397])).
fof(f12628,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),happ(s(fun(X4,fun(X1,X2)),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X3),s(fun(X4,fun(X1,X2)),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f397,f352])).
fof(f12574,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f12557])).
fof(f12557,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f12464,f445])).
fof(f12433,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107))) | s(bool,'T_0') = s(bool,sK4(X107,X108,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X108)))) ) | $spl10),
  inference(superposition,[],[f10983,f836])).
fof(f12464,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f12371,f8242])).
fof(f12371,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('const.pair.FST_0',X0,bool))))) ) | ($spl10 | $spl26)),
  inference(unit_resulting_resolution,[],[f8368,f10983])).
fof(f11866,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.modSem.v'),X0) != s('type.list.list'('type.modSem.v'),X1) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X1)))) )),
  inference(equality_resolution,[],[f396])).
fof(f11842,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X17) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(X13,happ(s(fun(X12,X13),X15),s(X12,X16))))) | s('type.list.list'('type.modSem.v'),X18) != s('type.list.list'('type.modSem.v'),X19) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.modSem.v'),X19))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X12,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(fun(X12,X13),X15))),s(X12,X16))),s('type.list.list'('type.modSem.v'),X18)))) )),
  inference(superposition,[],[f396,f369])).
fof(f11841,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X9) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X6,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(fun(X6,X4),X7))),s(X6,X8))) | s('type.list.list'('type.modSem.v'),X10) != s('type.list.list'('type.modSem.v'),X11) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X9),s('type.list.list'('type.modSem.v'),X11))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(X4,happ(s(fun(X6,X4),X7),s(X6,X8))))),s('type.list.list'('type.modSem.v'),X10)))) )),
  inference(superposition,[],[f396,f369])).
fof(f11954,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),happ(s(fun(X1,'type.list.list'('type.modSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),happ(s(fun(X3,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.modSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(forward_demodulation,[],[f11750,f342])).
fof(f11750,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),happ(s(fun(X1,'type.list.list'('type.modSem.v')),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.SND_1'(s('type.pair.prod'(X6,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.,_2'(s(X6,X7),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X3,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X1,'type.list.list'('type.modSem.v')),X2),s(fun(X3,X1),X4))),s(X3,X5)))))) )),
  inference(unit_resulting_resolution,[],[f369,f342,f396])).
fof(f11932,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X4,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f11772,f338])).
fof(f11772,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.FST_0'),s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X4,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f369,f338,f396])).
fof(f11910,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X4,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(forward_demodulation,[],[f11794,f324])).
fof(f11794,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num',X1))),s('type.list.list'('type.modSem.v'),happ(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X7),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X4,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X2,'type.list.list'('type.modSem.v')),X3),s(fun(X4,X2),X5))),s(X4,X6)))))) )),
  inference(unit_resulting_resolution,[],[f369,f324,f396])).
fof(f11805,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X6,'type.list.list'('type.modSem.v')),X7),s(fun(X5,X6),X8))),s(X5,X9))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X6,'type.list.list'('type.modSem.v')),X7),s(X6,happ(s(fun(X5,X6),X8),s(X5,X9)))))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f11892,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),X5))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),X5)))) )),
  inference(forward_demodulation,[],[f11813,f342])).
fof(f11813,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),X5))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X6,'type.list.list'('type.modSem.v')),'const.pair.,_2'(s(X6,X7),s('type.list.list'('type.modSem.v'),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f342,f369,f396])).
fof(f11891,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X5),X6))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X5),X6)))))) )),
  inference(forward_demodulation,[],[f11814,f338])).
fof(f11814,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X5),X6))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.modSem.v'),X5),'type.list.list'('type.modSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.modSem.v'),X5),X6)))))) )),
  inference(unit_resulting_resolution,[],[f338,f369,f396])).
fof(f11890,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun('type.num.num','type.list.list'('type.modSem.v')),X5),s('type.num.num',X6))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun('type.num.num','type.list.list'('type.modSem.v')),X5),s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f11815,f324])).
fof(f11815,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun('type.num.num','type.list.list'('type.modSem.v')),X5),s('type.num.num',X6))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.modSem.v'),X7),s(fun('type.num.num','type.list.list'('type.modSem.v')),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(unit_resulting_resolution,[],[f324,f369,f396])).
fof(f11816,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X7,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f369,f396])).
fof(f11889,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.modSem.v'))),X7),s(X6,X8))),s(X5,X9))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.modSem.v'))),X7),s(X6,X8))),s(X5,X9)))))) )),
  inference(forward_demodulation,[],[f11817,f367])).
fof(f11817,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),happ(s(fun(X6,fun(X5,'type.list.list'('type.modSem.v'))),X7),s(X6,X8))),s(X5,X9))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X2,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'(X6,X5),'type.list.list'('type.modSem.v')),sK19(X7,X5,X6,'type.list.list'('type.modSem.v'))),s('type.pair.prod'(X6,X5),'const.pair.,_2'(s(X6,X8),s(X5,X9)))))))) )),
  inference(unit_resulting_resolution,[],[f367,f369,f396])).
fof(f11868,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X7,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(forward_demodulation,[],[f11838,f367])).
fof(f11838,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X1,fun(X0,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),X2),s(X1,X3))),s(X0,X4))),s('type.list.list'('type.modSem.v'),happ(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9))))))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'(X1,X0),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),sK19(X2,X0,X1,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X3),s(X0,X4))))),s('type.list.list'('type.modSem.v'),happ(s(fun(X7,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X5,'type.list.list'('type.modSem.v')),X6),s(fun(X7,X5),X8))),s(X7,X9)))))) )),
  inference(unit_resulting_resolution,[],[f369,f367,f396])).
fof(f11237,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f11182,f262])).
fof(f11182,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(superposition,[],[f4861,f5646])).
fof(f11232,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f11156,f262])).
fof(f11156,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f447,f5646])).
fof(f11096,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.ind_type.recspace'(X16),'const.ind_type.CONSTR_3'(s('type.num.num',X17),s(X16,X18),s(fun('type.num.num','type.ind_type.recspace'(X16)),sK15(happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)),X19,fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))))) != s('type.ind_type.recspace'(X16),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),X19),s(fun('type.num.num','type.ind_type.recspace'(X16)),sK15(happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)),X19,fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))))) | s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),X19) = s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)))) )),
  inference(superposition,[],[f353,f325])).
fof(f11095,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.ind_type.recspace'(X12),'const.ind_type.CONSTR_3'(s('type.num.num',X13),s(X12,X14),s(fun('type.num.num','type.ind_type.recspace'(X12)),sK15(X15,happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)),fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))))) != s('type.ind_type.recspace'(X12),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),X15),s(fun('type.num.num','type.ind_type.recspace'(X12)),sK15(X15,happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)),fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))))) | s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),X15) = s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)))) )),
  inference(superposition,[],[f353,f325])).
fof(f10983,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f10885,f264])).
fof(f10885,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f271,f2913])).
fof(f10655,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2231,f445])).
fof(f10431,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2229,f445])).
fof(f10270,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f2225,f445])).
fof(f10136,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f10119])).
fof(f10119,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f10037,f445])).
fof(f10006,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X161),'const.sum.INL_1'(s(bool,X159))))) | s(bool,'T_0') = s(bool,sK4(X159,X160,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X160)))) )),
  inference(superposition,[],[f6500,f836])).
fof(f10005,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X158),'const.sum.INL_1'(s(bool,X157))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))))) )),
  inference(superposition,[],[f6500,f342])).
fof(f10004,plain,(
  ( ! [X154,X152,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X154),'const.sum.INL_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X153)))))) )),
  inference(superposition,[],[f6500,f336])).
fof(f10003,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X151),'const.sum.INL_1'(s(bool,X149))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X149),s(X148,X150)))))) )),
  inference(superposition,[],[f6500,f341])).
fof(f10002,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X147),'const.sum.INL_1'(s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X146)))))) )),
  inference(superposition,[],[f6500,f337])).
fof(f10000,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X140),'const.sum.INL_1'(s(bool,X138))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6500,f316])).
fof(f9997,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X125),bool),sK18(X126,X127,X125,bool)),s('type.list.list'(X125),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6500,f357])).
fof(f10037,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9947,f8242])).
fof(f9947,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8368,f6500])).
fof(f9829,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f9812])).
fof(f9812,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f9734,f445])).
fof(f9703,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X161,bool),'const.sum.INR_1'(s(bool,X159))))) | s(bool,'T_0') = s(bool,sK4(X159,X160,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X160)))) )),
  inference(superposition,[],[f6283,f836])).
fof(f9702,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X158,bool),'const.sum.INR_1'(s(bool,X157))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))))) )),
  inference(superposition,[],[f6283,f342])).
fof(f9701,plain,(
  ( ! [X154,X152,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X154,bool),'const.sum.INR_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X153)))))) )),
  inference(superposition,[],[f6283,f336])).
fof(f9700,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X151,bool),'const.sum.INR_1'(s(bool,X149))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X149),s(X148,X150)))))) )),
  inference(superposition,[],[f6283,f341])).
fof(f9699,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X147,bool),'const.sum.INR_1'(s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X146)))))) )),
  inference(superposition,[],[f6283,f337])).
fof(f9697,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X140,bool),'const.sum.INR_1'(s(bool,X138))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6283,f316])).
fof(f9694,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X125),bool),sK18(X126,X127,X125,bool)),s('type.list.list'(X125),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6283,f357])).
fof(f9734,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9644,f8242])).
fof(f9644,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8368,f6283])).
fof(f9245,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) != s(bool,sK7('const.pair.FST_0',X153,bool)) | s(bool,'F_0') = s(bool,sK4(X151,X152,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X152)))) ) | $spl26),
  inference(superposition,[],[f8378,f836])).
fof(f9244,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) != s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X147,bool),'const.pair.,_2'(s(X147,X148),s(bool,X149)))))) ) | $spl26),
  inference(superposition,[],[f8378,f342])).
fof(f9243,plain,(
  ( ! [X146,X144,X145] : (s(bool,X145) != s(bool,sK7('const.pair.FST_0',X146,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,X145)))))) ) | $spl26),
  inference(superposition,[],[f8378,f336])).
fof(f9242,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X141) != s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,X141),s(X140,X142)))))) ) | $spl26),
  inference(superposition,[],[f8378,f341])).
fof(f9241,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) != s(bool,sK7('const.pair.FST_0',X139,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,X138)))))) ) | $spl26),
  inference(superposition,[],[f8378,f337])).
fof(f9239,plain,(
  ( ! [X132,X130,X131] : (s(bool,X130) != s(bool,sK7('const.pair.FST_0',X132,bool)) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X130),s(fun('type.num.num',bool),X131),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8378,f316])).
fof(f9236,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X119) != s(bool,sK7('const.pair.FST_0',X120,bool)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X117),bool),sK18(X118,X119,X117,bool)),s('type.list.list'(X117),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8378,f357])).
fof(f9233,plain,(
  ( ! [X103,X104] : (s(bool,X103) != s(bool,sK7('const.pair.FST_0',X104,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103)))) ) | $spl26),
  inference(superposition,[],[f8378,f309])).
fof(f8857,plain,(
  ( ! [X152,X151,X153] : (s(bool,X151) = s(bool,sK7('const.pair.FST_0',X153,bool)) | s(bool,'T_0') = s(bool,sK4(X151,X152,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X152)))) ) | $spl26),
  inference(superposition,[],[f8367,f836])).
fof(f8856,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,X149) = s(bool,sK7('const.pair.FST_0',X150,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X147,bool),'const.pair.,_2'(s(X147,X148),s(bool,X149)))))) ) | $spl26),
  inference(superposition,[],[f8367,f342])).
fof(f8855,plain,(
  ( ! [X146,X144,X145] : (s(bool,X145) = s(bool,sK7('const.pair.FST_0',X146,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,X145)))))) ) | $spl26),
  inference(superposition,[],[f8367,f336])).
fof(f8854,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X141) = s(bool,sK7('const.pair.FST_0',X143,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X140),'const.pair.,_2'(s(bool,X141),s(X140,X142)))))) ) | $spl26),
  inference(superposition,[],[f8367,f341])).
fof(f8853,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) = s(bool,sK7('const.pair.FST_0',X139,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,X138)))))) ) | $spl26),
  inference(superposition,[],[f8367,f337])).
fof(f8851,plain,(
  ( ! [X132,X130,X131] : (s(bool,X130) = s(bool,sK7('const.pair.FST_0',X132,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X130),s(fun('type.num.num',bool),X131),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8367,f316])).
fof(f8848,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X119) = s(bool,sK7('const.pair.FST_0',X120,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X117),bool),sK18(X118,X119,X117,bool)),s('type.list.list'(X117),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8367,f357])).
fof(f8845,plain,(
  ( ! [X103,X104] : (s(bool,X103) = s(bool,sK7('const.pair.FST_0',X104,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103)))) ) | $spl26),
  inference(superposition,[],[f8367,f309])).
fof(f8736,plain,(
  ( ! [X707,X705,X708,X702,X704,X706,X709,X703] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X707),s('type.list.list'('type.modSem.v'),X708))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X703,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X702,fun(X703,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),X704),s(X702,X705))),s(X703,X706))),s('type.list.list'('type.modSem.v'),X709))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X707) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'(X702,X703),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),sK19(X704,X703,X702,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.pair.prod'(X702,X703),'const.pair.,_2'(s(X702,X705),s(X703,X706)))))) )),
  inference(superposition,[],[f394,f367])).
fof(f8731,plain,(
  ( ! [X668,X670,X673,X675,X669,X671,X672,X674] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X673),s('type.list.list'('type.modSem.v'),X674))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X675),s('type.list.list'('type.modSem.v'),happ(s(fun(X669,'type.list.list'('type.modSem.v')),happ(s(fun(X668,fun(X669,'type.list.list'('type.modSem.v'))),X670),s(X668,X671))),s(X669,X672))))) | s('type.list.list'('type.modSem.v'),X674) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'(X668,X669),'type.list.list'('type.modSem.v')),sK19(X670,X669,X668,'type.list.list'('type.modSem.v'))),s('type.pair.prod'(X668,X669),'const.pair.,_2'(s(X668,X671),s(X669,X672)))))) )),
  inference(superposition,[],[f395,f367])).
fof(f8725,plain,(
  ( ! [X633,X629,X631,X632,X634,X630] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X634))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X630,'type.ast.lit'),happ(s(fun(X629,fun(X630,'type.ast.lit')),X631),s(X629,X632))),s(X630,X633))))) | s('type.ast.lit',X634) = s('type.ast.lit',happ(s(fun('type.pair.prod'(X629,X630),'type.ast.lit'),sK19(X631,X630,X629,'type.ast.lit')),s('type.pair.prod'(X629,X630),'const.pair.,_2'(s(X629,X632),s(X630,X633)))))) )),
  inference(superposition,[],[f334,f367])).
fof(f8799,plain,(
  ( ! [X509,X511,X516,X512,X514,X510,X513,X515] : (s(fun(X509,X510),X516) = s(fun(X509,X510),happ(s(fun(X512,fun(X509,X510)),happ(s(fun(X511,fun(X512,fun(X509,X510))),X513),s(X511,X514))),s(X512,X515))) | s(X510,happ(s(fun(X509,X510),X516),s(X509,sK15(happ(s(fun('type.pair.prod'(X511,X512),fun(X509,X510)),sK19(X513,X512,X511,fun(X509,X510))),s('type.pair.prod'(X511,X512),'const.pair.,_2'(s(X511,X514),s(X512,X515)))),X516,X509,X510)))) != s(X510,happ(s(fun(X509,X510),happ(s(fun(X512,fun(X509,X510)),happ(s(fun(X511,fun(X512,fun(X509,X510))),X513),s(X511,X514))),s(X512,X515))),s(X509,sK15(happ(s(fun('type.pair.prod'(X511,X512),fun(X509,X510)),sK19(X513,X512,X511,fun(X509,X510))),s('type.pair.prod'(X511,X512),'const.pair.,_2'(s(X511,X514),s(X512,X515)))),X516,X509,X510))))) )),
  inference(forward_demodulation,[],[f8710,f367])).
fof(f8710,plain,(
  ( ! [X509,X511,X516,X512,X514,X510,X513,X515] : (s(X510,happ(s(fun(X509,X510),X516),s(X509,sK15(happ(s(fun('type.pair.prod'(X511,X512),fun(X509,X510)),sK19(X513,X512,X511,fun(X509,X510))),s('type.pair.prod'(X511,X512),'const.pair.,_2'(s(X511,X514),s(X512,X515)))),X516,X509,X510)))) != s(X510,happ(s(fun(X509,X510),happ(s(fun(X512,fun(X509,X510)),happ(s(fun(X511,fun(X512,fun(X509,X510))),X513),s(X511,X514))),s(X512,X515))),s(X509,sK15(happ(s(fun('type.pair.prod'(X511,X512),fun(X509,X510)),sK19(X513,X512,X511,fun(X509,X510))),s('type.pair.prod'(X511,X512),'const.pair.,_2'(s(X511,X514),s(X512,X515)))),X516,X509,X510)))) | s(fun(X509,X510),X516) = s(fun(X509,X510),happ(s(fun('type.pair.prod'(X511,X512),fun(X509,X510)),sK19(X513,X512,X511,fun(X509,X510))),s('type.pair.prod'(X511,X512),'const.pair.,_2'(s(X511,X514),s(X512,X515)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f8798,plain,(
  ( ! [X505,X507,X502,X508,X504,X506,X501,X503] : (s(fun(X501,X502),X508) = s(fun(X501,X502),happ(s(fun(X504,fun(X501,X502)),happ(s(fun(X503,fun(X504,fun(X501,X502))),X505),s(X503,X506))),s(X504,X507))) | s(X502,happ(s(fun(X501,X502),X508),s(X501,sK15(X508,happ(s(fun('type.pair.prod'(X503,X504),fun(X501,X502)),sK19(X505,X504,X503,fun(X501,X502))),s('type.pair.prod'(X503,X504),'const.pair.,_2'(s(X503,X506),s(X504,X507)))),X501,X502)))) != s(X502,happ(s(fun(X501,X502),happ(s(fun(X504,fun(X501,X502)),happ(s(fun(X503,fun(X504,fun(X501,X502))),X505),s(X503,X506))),s(X504,X507))),s(X501,sK15(X508,happ(s(fun('type.pair.prod'(X503,X504),fun(X501,X502)),sK19(X505,X504,X503,fun(X501,X502))),s('type.pair.prod'(X503,X504),'const.pair.,_2'(s(X503,X506),s(X504,X507)))),X501,X502))))) )),
  inference(forward_demodulation,[],[f8709,f367])).
fof(f8709,plain,(
  ( ! [X505,X507,X502,X508,X504,X506,X501,X503] : (s(X502,happ(s(fun(X501,X502),X508),s(X501,sK15(X508,happ(s(fun('type.pair.prod'(X503,X504),fun(X501,X502)),sK19(X505,X504,X503,fun(X501,X502))),s('type.pair.prod'(X503,X504),'const.pair.,_2'(s(X503,X506),s(X504,X507)))),X501,X502)))) != s(X502,happ(s(fun(X501,X502),happ(s(fun(X504,fun(X501,X502)),happ(s(fun(X503,fun(X504,fun(X501,X502))),X505),s(X503,X506))),s(X504,X507))),s(X501,sK15(X508,happ(s(fun('type.pair.prod'(X503,X504),fun(X501,X502)),sK19(X505,X504,X503,fun(X501,X502))),s('type.pair.prod'(X503,X504),'const.pair.,_2'(s(X503,X506),s(X504,X507)))),X501,X502)))) | s(fun(X501,X502),X508) = s(fun(X501,X502),happ(s(fun('type.pair.prod'(X503,X504),fun(X501,X502)),sK19(X505,X504,X503,fun(X501,X502))),s('type.pair.prod'(X503,X504),'const.pair.,_2'(s(X503,X506),s(X504,X507)))))) )),
  inference(superposition,[],[f353,f367])).
fof(f8632,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))) = s(X6,happ(s(fun('type.pair.prod'(X1,X2),X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(superposition,[],[f369,f367])).
fof(f8378,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(subsumption_resolution,[],[f8325,f264])).
fof(f8325,plain,(
  ( ! [X41,X40] : (p(s(bool,'F_0')) | s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(superposition,[],[f2426,f8242])).
fof(f8368,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(subsumption_resolution,[],[f8302,f264])).
fof(f8302,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(superposition,[],[f271,f8242])).
fof(f8367,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(subsumption_resolution,[],[f8290,f432])).
fof(f8290,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f8242,f445])).
fof(f8271,plain,(
  ( ! [X10,X11] : (s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X11),s('type.list.list'('type.modSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10)),X11,'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10),s('type.list.list'('type.modSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10)),X11,'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X11) = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10)))) )),
  inference(superposition,[],[f353,f328])).
fof(f8270,plain,(
  ( ! [X8,X9] : (s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X9),s('type.list.list'('type.modSem.v'),sK15(X9,happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8)),'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.modSem.v'),sK15(X9,happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8)),'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X9) = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8)))) )),
  inference(superposition,[],[f353,f328])).
fof(f8041,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X60),s('type.list.list'('type.modSem.v'),X61))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X62),s('type.list.list'('type.modSem.v'),X63))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X62) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.SND_1'(s('type.pair.prod'(X58,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.,_2'(s(X58,X59),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X60)))))) )),
  inference(superposition,[],[f394,f342])).
fof(f8040,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X54),s('type.list.list'('type.modSem.v'),X55))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X56),s('type.list.list'('type.modSem.v'),X57))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X56) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.sum.OUTL_1'(s('type.sum.sum'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X53),'const.sum.INL_1'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X54)))))) )),
  inference(superposition,[],[f394,f336])).
fof(f8039,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X48),s('type.list.list'('type.modSem.v'),X50))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51),s('type.list.list'('type.modSem.v'),X52))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X47),'const.pair.,_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X48),s(X47,X49)))))) )),
  inference(superposition,[],[f394,f341])).
fof(f8038,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X43),s('type.list.list'('type.modSem.v'),X44))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45),s('type.list.list'('type.modSem.v'),X46))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.sum.OUTR_1'(s('type.sum.sum'(X42,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.sum.INR_1'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X43)))))) )),
  inference(superposition,[],[f394,f337])).
fof(f8037,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.modSem.v'),X41))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X37),s('type.num.num',X38))),s('type.list.list'('type.modSem.v'),X39))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X36),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X37),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38)))))) )),
  inference(superposition,[],[f394,f324])).
fof(f8036,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X31),s('type.list.list'('type.modSem.v'),X33))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34),s('type.list.list'('type.modSem.v'),X35))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X31),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f394,f316])).
fof(f8035,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.modSem.v'),X30))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),X27))),s('type.list.list'('type.modSem.v'),X28))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.FST_0'),s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),X27)))) )),
  inference(superposition,[],[f394,f338])).
fof(f8034,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X22),s('type.list.list'('type.modSem.v'),X23))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.modSem.v'),X25))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.list.list'(X20),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),sK18(X21,X22,X20,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.list.list'(X20),'const.list.NIL_0')))) )),
  inference(superposition,[],[f394,f357])).
fof(f8033,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.modSem.v'),X19))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(X13,happ(s(fun(X12,X13),X15),s(X12,X16))))),s('type.list.list'('type.modSem.v'),X17))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X12,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(fun(X12,X13),X15))),s(X12,X16)))) )),
  inference(superposition,[],[f394,f369])).
fof(f8032,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10),s('type.list.list'('type.modSem.v'),X11))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X6,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(fun(X6,X4),X7))),s(X6,X8))),s('type.list.list'('type.modSem.v'),X9))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(X4,happ(s(fun(X6,X4),X7),s(X6,X8)))))) )),
  inference(superposition,[],[f394,f369])).
fof(f8031,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_i_0_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0)))) )),
  inference(superposition,[],[f394,f309])).
fof(f7552,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f7519,f262])).
fof(f7519,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f350,f445])).
fof(f7122,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f7089,f262])).
fof(f7089,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f349,f445])).
fof(f7121,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(subsumption_resolution,[],[f7087,f262])).
fof(f7087,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f349,f440])).
fof(f6654,plain,(
  ( ! [X68] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK15('const.modSem.Vectorv_0',X68,'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X68),s('type.list.list'('type.modSem.v'),sK15('const.modSem.Vectorv_0',X68,'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0') = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f6653,plain,(
  ( ! [X67] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',sK15('const.modSem.Litv_0',X67,'type.ast.lit','type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.ast.lit','type.modSem.v'),X67),s('type.ast.lit',sK15('const.modSem.Litv_0',X67,'type.ast.lit','type.modSem.v')))) | s(fun('type.ast.lit','type.modSem.v'),'const.modSem.Litv_0') = s(fun('type.ast.lit','type.modSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f6652,plain,(
  ( ! [X66] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK15('const.modSem.Loc_0',X66,'type.num.num','type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.num.num','type.modSem.v'),X66),s('type.num.num',sK15('const.modSem.Loc_0',X66,'type.num.num','type.modSem.v')))) | s(fun('type.num.num','type.modSem.v'),'const.modSem.Loc_0') = s(fun('type.num.num','type.modSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f6651,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6640,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X20,X21,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,X21,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f6639,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f6636,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f6635,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6634,plain,(
  ( ! [X0,X1] : (s(X0,sK15('_i_0_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15('_i_0_0',X1,X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f6623,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,X9))) )),
  inference(superposition,[],[f353,f445])).
fof(f6622,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK15(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f353,f445])).
fof(f6679,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65))))) )),
  inference(forward_demodulation,[],[f6619,f342])).
fof(f6619,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f6678,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60))))) )),
  inference(forward_demodulation,[],[f6618,f336])).
fof(f6618,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f6677,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54))))) )),
  inference(forward_demodulation,[],[f6617,f341])).
fof(f6617,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f6676,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49))))) )),
  inference(forward_demodulation,[],[f6616,f337])).
fof(f6616,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f6675,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43))))) )),
  inference(forward_demodulation,[],[f6615,f324])).
fof(f6615,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f6674,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38))))) )),
  inference(forward_demodulation,[],[f6614,f316])).
fof(f6614,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f6673,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33))))) )),
  inference(forward_demodulation,[],[f6613,f338])).
fof(f6613,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f353,f338])).
fof(f6672,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27))))) )),
  inference(forward_demodulation,[],[f6612,f357])).
fof(f6612,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f6611,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6610,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f6671,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7))))) )),
  inference(forward_demodulation,[],[f6609,f309])).
fof(f6609,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f6670,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1))))) )),
  inference(forward_demodulation,[],[f6608,f368])).
fof(f6608,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6607,plain,(
  ( ! [X68] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK15(X68,'const.modSem.Vectorv_0','type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X68),s('type.list.list'('type.modSem.v'),sK15(X68,'const.modSem.Vectorv_0','type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0') = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X68)) )),
  inference(superposition,[],[f353,f268])).
fof(f6606,plain,(
  ( ! [X67] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',sK15(X67,'const.modSem.Litv_0','type.ast.lit','type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.ast.lit','type.modSem.v'),X67),s('type.ast.lit',sK15(X67,'const.modSem.Litv_0','type.ast.lit','type.modSem.v')))) | s(fun('type.ast.lit','type.modSem.v'),'const.modSem.Litv_0') = s(fun('type.ast.lit','type.modSem.v'),X67)) )),
  inference(superposition,[],[f353,f267])).
fof(f6605,plain,(
  ( ! [X66] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK15(X66,'const.modSem.Loc_0','type.num.num','type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.num.num','type.modSem.v'),X66),s('type.num.num',sK15(X66,'const.modSem.Loc_0','type.num.num','type.modSem.v')))) | s(fun('type.num.num','type.modSem.v'),'const.modSem.Loc_0') = s(fun('type.num.num','type.modSem.v'),X66)) )),
  inference(superposition,[],[f353,f266])).
fof(f6604,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6593,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X21,X20,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X21,X20,X19,bool))))) )),
  inference(superposition,[],[f353,f445])).
fof(f6592,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f353,f445])).
fof(f6589,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f353,f338])).
fof(f6588,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6587,plain,(
  ( ! [X0,X1] : (s(X0,sK15(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f353,f309])).
fof(f6665,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65))))) )),
  inference(forward_demodulation,[],[f6585,f342])).
fof(f6585,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f353,f342])).
fof(f6664,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60))))) )),
  inference(forward_demodulation,[],[f6584,f336])).
fof(f6584,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f353,f336])).
fof(f6663,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54))))) )),
  inference(forward_demodulation,[],[f6583,f341])).
fof(f6583,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f353,f341])).
fof(f6662,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49))))) )),
  inference(forward_demodulation,[],[f6582,f337])).
fof(f6582,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f353,f337])).
fof(f6661,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43))))) )),
  inference(forward_demodulation,[],[f6581,f324])).
fof(f6581,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f353,f324])).
fof(f6660,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38))))) )),
  inference(forward_demodulation,[],[f6580,f316])).
fof(f6580,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f353,f316])).
fof(f6659,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33))))) )),
  inference(forward_demodulation,[],[f6579,f338])).
fof(f6579,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f353,f338])).
fof(f6658,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27))))) )),
  inference(forward_demodulation,[],[f6578,f357])).
fof(f6578,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f357])).
fof(f6577,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f353,f369])).
fof(f6576,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f353,f369])).
fof(f6657,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7))))) )),
  inference(forward_demodulation,[],[f6575,f309])).
fof(f6575,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f353,f309])).
fof(f6656,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) )),
  inference(forward_demodulation,[],[f6574,f368])).
fof(f6574,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f353,f368])).
fof(f6500,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6428,f264])).
fof(f6428,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2421])).
fof(f6283,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6212,f264])).
fof(f6212,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f271,f2419])).
fof(f5651,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f5614,f262])).
fof(f5614,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f4861,f4890])).
fof(f5647,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f5611,f262])).
fof(f5611,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4861,f836])).
fof(f5646,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5610,f262])).
fof(f5610,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4861,f2070])).
fof(f5642,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(subsumption_resolution,[],[f5607,f262])).
fof(f5607,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(superposition,[],[f4861,f440])).
fof(f5631,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5590,f262])).
fof(f5590,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4861,f445])).
fof(f5556,plain,(
  ( ! [X466,X463,X469,X465,X467,X468,X470,X464] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X468),s('type.list.list'('type.modSem.v'),X469))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X470),s('type.list.list'('type.modSem.v'),happ(s(fun(X464,'type.list.list'('type.modSem.v')),X465),s(X464,happ(s(fun(X463,X464),X466),s(X463,X467))))))) | s('type.list.list'('type.modSem.v'),X469) = s('type.list.list'('type.modSem.v'),happ(s(fun(X463,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X464,'type.list.list'('type.modSem.v')),X465),s(fun(X463,X464),X466))),s(X463,X467)))) )),
  inference(superposition,[],[f395,f369])).
fof(f5553,plain,(
  ( ! [X447,X449,X451,X446,X448,X450] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X451))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X447,'type.ast.lit'),X448),s(X447,happ(s(fun(X446,X447),X449),s(X446,X450))))))) | s('type.ast.lit',X451) = s('type.ast.lit',happ(s(fun(X446,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X447,'type.ast.lit'),X448),s(fun(X446,X447),X449))),s(X446,X450)))) )),
  inference(superposition,[],[f334,f369])).
fof(f5551,plain,(
  ( ! [X436,X438,X434,X437,X439,X435] : (s('type.ast.lit',X439) != s('type.ast.lit',happ(s(fun(X435,'type.ast.lit'),X436),s(X435,happ(s(fun(X434,X435),X437),s(X434,X438))))) | s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X439))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X434,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X435,'type.ast.lit'),X436),s(fun(X434,X435),X437))),s(X434,X438)))))) )),
  inference(superposition,[],[f335,f369])).
fof(f5540,plain,(
  ( ! [X350,X344,X346,X352,X349,X351,X345,X347,X348] : (s(fun(X344,X345),X351) != s(fun(X344,X345),happ(s(fun(X347,fun(X344,X345)),X348),s(X347,happ(s(fun(X346,X347),X349),s(X346,X350))))) | s(X345,happ(s(fun(X344,X345),X351),s(X344,X352))) = s(X345,happ(s(fun(X344,X345),happ(s(fun(X346,fun(X344,X345)),'_dst_x0x1_2'(s(fun(X347,fun(X344,X345)),X348),s(fun(X346,X347),X349))),s(X346,X350))),s(X344,X352)))) )),
  inference(superposition,[],[f352,f369])).
fof(f5535,plain,(
  ( ! [X313,X315,X310,X316,X312,X314,X311] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X310,bool),happ(s(fun(X312,fun(X310,bool)),X313),s(X312,happ(s(fun(X311,X312),X314),s(X311,X315))))))) | s(X310,X316) = s(X310,sK4(X316,happ(s(fun(X311,fun(X310,bool)),'_dst_x0x1_2'(s(fun(X312,fun(X310,bool)),X313),s(fun(X311,X312),X314))),s(X311,X315)),X310))) )),
  inference(superposition,[],[f836,f369])).
fof(f5495,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK4(X85,'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84)),X81))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f319,f369])).
fof(f5478,plain,(
  ( ! [X484,X486,X482,X488,X485,X487,X483,X489] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X487),s('type.list.list'('type.modSem.v'),X488))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X489),s('type.list.list'('type.modSem.v'),happ(s(fun(X484,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X482,'type.list.list'('type.modSem.v')),X483),s(fun(X484,X482),X485))),s(X484,X486))))) | s('type.list.list'('type.modSem.v'),X488) = s('type.list.list'('type.modSem.v'),happ(s(fun(X482,'type.list.list'('type.modSem.v')),X483),s(X482,happ(s(fun(X484,X482),X485),s(X484,X486)))))) )),
  inference(superposition,[],[f395,f369])).
fof(f5475,plain,(
  ( ! [X466,X469,X465,X467,X468,X470] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X470))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X467,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X465,'type.ast.lit'),X466),s(fun(X467,X465),X468))),s(X467,X469))))) | s('type.ast.lit',X470) = s('type.ast.lit',happ(s(fun(X465,'type.ast.lit'),X466),s(X465,happ(s(fun(X467,X465),X468),s(X467,X469)))))) )),
  inference(superposition,[],[f334,f369])).
fof(f5473,plain,(
  ( ! [X453,X455,X457,X454,X456,X458] : (s('type.ast.lit',X458) != s('type.ast.lit',happ(s(fun(X455,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X453,'type.ast.lit'),X454),s(fun(X455,X453),X456))),s(X455,X457))) | s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X458))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X453,'type.ast.lit'),X454),s(X453,happ(s(fun(X455,X453),X456),s(X455,X457)))))))) )),
  inference(superposition,[],[f335,f369])).
fof(f5462,plain,(
  ( ! [X356,X358,X354,X360,X362,X357,X359,X355,X361] : (s(fun(X354,X355),X361) != s(fun(X354,X355),happ(s(fun(X358,fun(X354,X355)),'_dst_x0x1_2'(s(fun(X356,fun(X354,X355)),X357),s(fun(X358,X356),X359))),s(X358,X360))) | s(X355,happ(s(fun(X354,X355),X361),s(X354,X362))) = s(X355,happ(s(fun(X354,X355),happ(s(fun(X356,fun(X354,X355)),X357),s(X356,happ(s(fun(X358,X356),X359),s(X358,X360))))),s(X354,X362)))) )),
  inference(superposition,[],[f352,f369])).
fof(f5457,plain,(
  ( ! [X325,X321,X323,X324,X326,X320,X322] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X320,bool),happ(s(fun(X323,fun(X320,bool)),'_dst_x0x1_2'(s(fun(X321,fun(X320,bool)),X322),s(fun(X323,X321),X324))),s(X323,X325))))) | s(X320,X326) = s(X320,sK4(X326,happ(s(fun(X321,fun(X320,bool)),X322),s(X321,happ(s(fun(X323,X321),X324),s(X323,X325)))),X320))) )),
  inference(superposition,[],[f836,f369])).
fof(f5318,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s(fun(X95,X96),'_dst_x0x1_2'(s(fun(X97,X96),happ(s(fun(X99,fun(X97,X96)),X100),s(X99,X101))),s(fun(X95,X97),X98))) = s(fun(X95,X96),happ(s(fun(X99,fun(X95,X96)),'_dst_x0x1_2'(s(fun(fun(X97,X96),fun(X95,X96)),'_dstsel_0_x01_1'(s(fun(X95,X97),X98))),s(fun(X99,fun(X97,X96)),X100))),s(X99,X101)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5308,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X39,X37),X40))),s(X39,X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X39,X37),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f369,f445])).
fof(f5307,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X33,X31),X34))),s(X33,X35))) | s(bool,X36) = s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X33,X31),X34),s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f369,f445])).
fof(f5305,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f369,f338])).
fof(f5304,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5303,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(X104,happ(s(fun(fun(X99,X100),X104),X105),s(fun(X99,X100),'_dst_x0x1_2'(s(fun(X101,X100),X103),s(fun(X99,X101),X102))))) = s(X104,happ(s(fun(fun(X101,X100),X104),'_dst_x0x1_2'(s(fun(fun(X99,X100),X104),X105),s(fun(fun(X101,X100),fun(X99,X100)),'_dstsel_0_x01_1'(s(fun(X99,X101),X102))))),s(fun(X101,X100),X103)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5302,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK18(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f369,f357])).
fof(f5301,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5300,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f369,f369])).
fof(f5298,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f369,f352])).
fof(f5299,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X0,'type.ast.lit'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun(X2,'type.ast.lit'),'_dst_x0x1_2'(s(fun(X0,'type.ast.lit'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f369,f335])).
fof(f5082,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f5081,f262])).
fof(f5081,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f5072,f2070])).
fof(f5072,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f4761,f2531])).
fof(f5079,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5067,f262])).
fof(f5067,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4761,f445])).
fof(f5054,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f5007,f262])).
fof(f5007,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f447,f4890])).
fof(f4956,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X50),s('type.list.list'('type.modSem.v'),X48))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51),s('type.list.list'('type.modSem.v'),X52))) | s('type.list.list'('type.modSem.v'),X52) = s('type.list.list'('type.modSem.v'),sK4(X48,X49,'type.list.list'('type.modSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.modSem.v'),bool),X49)))) )),
  inference(superposition,[],[f395,f836])).
fof(f4955,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45),s('type.list.list'('type.modSem.v'),X44))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.modSem.v'),X47))) | s('type.list.list'('type.modSem.v'),X47) = s('type.list.list'('type.modSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X42,'type.list.list'('type.modSem.v')),'const.pair.,_2'(s(X42,X43),s('type.list.list'('type.modSem.v'),X44)))))) )),
  inference(superposition,[],[f395,f342])).
fof(f4954,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X39),s('type.list.list'('type.modSem.v'),X38))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.modSem.v'),X41))) | s('type.list.list'('type.modSem.v'),X41) = s('type.list.list'('type.modSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.modSem.v'),X37),'const.sum.INL_1'(s('type.list.list'('type.modSem.v'),X38)))))) )),
  inference(superposition,[],[f395,f336])).
fof(f4953,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34),s('type.list.list'('type.modSem.v'),X32))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X35),s('type.list.list'('type.modSem.v'),X36))) | s('type.list.list'('type.modSem.v'),X36) = s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.modSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f395,f341])).
fof(f4952,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X28),s('type.list.list'('type.modSem.v'),X27))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.modSem.v'),X30))) | s('type.list.list'('type.modSem.v'),X30) = s('type.list.list'('type.modSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.list.list'('type.modSem.v')),'const.sum.INR_1'(s('type.list.list'('type.modSem.v'),X27)))))) )),
  inference(superposition,[],[f395,f337])).
fof(f4951,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.modSem.v'),X25))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X23),s('type.list.list'('type.modSem.v'),happ(s(fun('type.num.num','type.list.list'('type.modSem.v')),X21),s('type.num.num',X22))))) | s('type.list.list'('type.modSem.v'),X25) = s('type.list.list'('type.modSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.modSem.v'),X20),s(fun('type.num.num','type.list.list'('type.modSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f395,f324])).
fof(f4950,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.modSem.v'),X15))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.modSem.v'),X19))) | s('type.list.list'('type.modSem.v'),X19) = s('type.list.list'('type.modSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.modSem.v'),X15),s(fun('type.num.num','type.list.list'('type.modSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f395,f316])).
fof(f4949,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.modSem.v'),X14))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X12),s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X10),X11))))) | s('type.list.list'('type.modSem.v'),X14) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.modSem.v'),X10),'type.list.list'('type.modSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.modSem.v'),X10),X11)))) )),
  inference(superposition,[],[f395,f338])).
fof(f4948,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X7),s('type.list.list'('type.modSem.v'),X6))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.modSem.v'),X9))) | s('type.list.list'('type.modSem.v'),X9) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.modSem.v')),sK18(X5,X6,X4,'type.list.list'('type.modSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f395,f357])).
fof(f4947,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),X0))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.list.list'('type.modSem.v'),X3) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'('type.modSem.v'),'type.list.list'('type.modSem.v')),'_i_0_0'),s('type.list.list'('type.modSem.v'),X0)))) )),
  inference(superposition,[],[f395,f309])).
fof(f4884,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f357,f4847,f447])).
fof(f4890,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f432,f4847,f2424])).
fof(f4891,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4847,f4760])).
fof(f4847,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f262,f4760])).
fof(f4761,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4686,f338])).
fof(f4691,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4641,f262])).
fof(f4641,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f447,f836])).
fof(f4622,plain,(
  ( ! [X206,X207,X205] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X205))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X207))) | s('type.ast.lit',X207) = s('type.ast.lit',sK4(X205,X206,'type.ast.lit')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X206)))) )),
  inference(superposition,[],[f334,f836])).
fof(f4610,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f319,f836])).
fof(f4687,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(forward_demodulation,[],[f4594,f2070])).
fof(f4594,plain,(
  ( ! [X107,X106] : (s(bool,X106) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) ) | $spl10),
  inference(superposition,[],[f2168,f836])).
fof(f4528,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f836,f445])).
fof(f4527,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f836,f445])).
fof(f4107,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3633,f445])).
fof(f3633,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f432,f2803])).
fof(f3415,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2803,f2803])).
fof(f3322,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK14(X7,X8,X9,bool))) )),
  inference(superposition,[],[f351,f445])).
fof(f3321,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK14(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f351,f445])).
fof(f2803,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2531,f2531])).
fof(f2704,plain,(
  ( ! [X158,X161,X159,X160] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X161))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',happ(s(fun('type.num.num','type.ast.lit'),X159),s('type.num.num',X160))))) | s('type.ast.lit',X161) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X158),s(fun('type.num.num','type.ast.lit'),X159),s('type.num.num','const.num.SUC_1'(s('type.num.num',X160)))))) )),
  inference(superposition,[],[f334,f324])).
fof(f2730,plain,(
  ( ! [X134,X138,X136,X135,X137] : (p(s(bool,'|dollar|exists_1'(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137)))))) | ~p(s(bool,happ(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137))),s(X134,sK4(X138,'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))),X134)))))) )),
  inference(forward_demodulation,[],[f2698,f324])).
fof(f2698,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,happ(s(fun(X134,bool),happ(s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num',X137))),s(X134,sK4(X138,'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137)))),X134))))) | p(s(bool,'|dollar|exists_1'(s(fun(X134,bool),'const.ind_type.FCONS_3'(s(fun(X134,bool),X135),s(fun('type.num.num',fun(X134,bool)),X136),s('type.num.num','const.num.SUC_1'(s('type.num.num',X137))))))))) )),
  inference(superposition,[],[f319,f324])).
fof(f2524,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2499,f262])).
fof(f2499,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2424,f445])).
fof(f2232,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2197,f2070])).
fof(f2197,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(superposition,[],[f2168,f342])).
fof(f2231,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2196,f2070])).
fof(f2196,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(superposition,[],[f2168,f336])).
fof(f2230,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2195,f2070])).
fof(f2195,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(superposition,[],[f2168,f341])).
fof(f2229,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2194,f2070])).
fof(f2194,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(superposition,[],[f2168,f337])).
fof(f2228,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X24) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2193,f2070])).
fof(f2193,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2168,f316])).
fof(f2226,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2191,f2070])).
fof(f2191,plain,(
  ( ! [X21,X19,X20] : (s(bool,X21) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2168,f357])).
fof(f2225,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f2190,f2070])).
fof(f2190,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(superposition,[],[f2168,f309])).
fof(f2168,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2143,f262])).
fof(f2143,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f447,f2070])).
fof(f2051,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f357,f1908,f447])).
fof(f2070,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1908,f272])).
fof(f1392,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1379,f262])).
fof(f1379,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f319,f445])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1367,f309])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f319,f309])).
fof(f1386,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1368,f357])).
fof(f1368,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f319,f357])).
fof(f1533,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1514,f338])).
fof(f1514,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f319,f338])).
fof(f1387,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1369,f316])).
fof(f1369,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f319,f316])).
fof(f1388,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12)))))) )),
  inference(forward_demodulation,[],[f1370,f337])).
fof(f1370,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14))))))))) )),
  inference(superposition,[],[f319,f337])).
fof(f1389,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1371,f341])).
fof(f1371,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f319,f341])).
fof(f1390,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | ~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21)))))) )),
  inference(forward_demodulation,[],[f1372,f336])).
fof(f1372,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23))))))))) )),
  inference(superposition,[],[f319,f336])).
fof(f1391,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X28)))) | ~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25)))))) )),
  inference(forward_demodulation,[],[f1373,f342])).
fof(f1373,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25))))) | p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28))))))))) )),
  inference(superposition,[],[f319,f342])).
fof(f1497,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f319,f338])).
fof(f1519,plain,(
  ( ! [X99,X97,X98] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X99))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X97),X98))))) | s('type.ast.lit',X99) = s('type.ast.lit',happ(s(fun('type.pair.prod'('type.ast.lit',X97),'type.ast.lit'),'const.pair.FST_0'),s('type.pair.prod'('type.ast.lit',X97),X98)))) )),
  inference(superposition,[],[f334,f338])).
fof(f1360,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1117,f262])).
fof(f1117,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f447,f445])).
fof(f1359,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1116,f262])).
fof(f1116,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f418,f445])).
fof(f1358,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1115,f262])).
fof(f1115,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f333,f445])).
fof(f1006,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f445,f342])).
fof(f1005,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f445,f336])).
fof(f1004,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f445,f341])).
fof(f1003,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f445,f337])).
fof(f1002,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f445,f316])).
fof(f1001,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK18(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f445,f357])).
fof(f1000,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f445,f309])).
fof(f990,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f445,f445])).
fof(f973,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) )),
  inference(superposition,[],[f445,f342])).
fof(f972,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X39)))))) )),
  inference(superposition,[],[f445,f336])).
fof(f971,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f445,f341])).
fof(f970,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,X32)))))) )),
  inference(superposition,[],[f445,f337])).
fof(f969,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f445,f316])).
fof(f968,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f445,f357])).
fof(f967,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f445,f309])).
fof(f934,plain,(
  ( ! [X103,X105,X106,X104] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X105))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X106))) | s('type.ast.lit',X106) = s('type.ast.lit',happ(s(fun('type.list.list'(X103),'type.ast.lit'),sK18(X104,X105,X103,'type.ast.lit')),s('type.list.list'(X103),'const.list.NIL_0')))) )),
  inference(superposition,[],[f334,f357])).
fof(f842,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X17))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X18))) | s('type.ast.lit',X18) = s('type.ast.lit','const.pair.SND_1'(s('type.pair.prod'(X15,'type.ast.lit'),'const.pair.,_2'(s(X15,X16),s('type.ast.lit',X17)))))) )),
  inference(superposition,[],[f334,f342])).
fof(f841,plain,(
  ( ! [X14,X12,X13] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X13))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X14))) | s('type.ast.lit',X14) = s('type.ast.lit','const.sum.OUTL_1'(s('type.sum.sum'('type.ast.lit',X12),'const.sum.INL_1'(s('type.ast.lit',X13)))))) )),
  inference(superposition,[],[f334,f336])).
fof(f840,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X9))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X11))) | s('type.ast.lit',X11) = s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X8),'const.pair.,_2'(s('type.ast.lit',X9),s(X8,X10)))))) )),
  inference(superposition,[],[f334,f341])).
fof(f839,plain,(
  ( ! [X6,X7,X5] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X6))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X7))) | s('type.ast.lit',X7) = s('type.ast.lit','const.sum.OUTR_1'(s('type.sum.sum'(X5,'type.ast.lit'),'const.sum.INR_1'(s('type.ast.lit',X6)))))) )),
  inference(superposition,[],[f334,f337])).
fof(f838,plain,(
  ( ! [X4,X2,X3] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X2))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X4))) | s('type.ast.lit',X4) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X2),s(fun('type.num.num','type.ast.lit'),X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f334,f316])).
fof(f837,plain,(
  ( ! [X0,X1] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X1) = s('type.ast.lit',happ(s(fun('type.ast.lit','type.ast.lit'),'_i_0_0'),s('type.ast.lit',X0)))) )),
  inference(superposition,[],[f334,f309])).
fof(f836,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f835,f264])).
fof(f835,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f318,f270])).
fof(f803,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f790,f262])).
fof(f790,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f447,f440])).
fof(f494,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f489,f264])).
fof(f489,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f330,f270])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f275,f270])).
fof(f440,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f265,f272])).
fof(f432,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f264,f271])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5))) & ((s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5)) | s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) <=> (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) & s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16,X17,X18,X19,X20,X21] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X17),s('type.list.list'('type.string.char'),X18))) = s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X19),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X20),s('type.list.list'('type.string.char'),X21))) <=> (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X19) & s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X17) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X20) & s('type.list.list'('type.string.char'),X18) = s('type.list.list'('type.string.char'),X21)))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_|49||49||split|3')).
fof(f404,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) = s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4)) )),
  inference(cnf_transformation,[],[f259])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) != s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) = s('type.list.list'('type.string.char'),X5)) )),
  inference(cnf_transformation,[],[f259])).
fof(f406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4),s('type.list.list'('type.string.char'),X5))) | s('type.list.list'('type.string.char'),X2) != s('type.list.list'('type.string.char'),X5) | s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1) != s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X4) | s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3)) )),
  inference(cnf_transformation,[],[f259])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) != s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5)))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) != s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) | (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.modLang.exp',X2) = s('type.modLang.exp',X5))) & (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.modLang.exp',X2) != s('type.modLang.exp',X5) | s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5)))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) != s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) | (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.modLang.exp',X2) = s('type.modLang.exp',X5))) & ((s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.modLang.exp',X2) != s('type.modLang.exp',X5)) | s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5)))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) <=> (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3) & s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4) & s('type.modLang.exp',X2) = s('type.modLang.exp',X5)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19,X20,X21] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16),s('type.list.list'('type.string.char'),X17),s('type.modLang.exp',X18))) = s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X19),s('type.list.list'('type.string.char'),X20),s('type.modLang.exp',X21))) <=> (s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16) = s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X19) & s('type.list.list'('type.string.char'),X17) = s('type.list.list'('type.string.char'),X20) & s('type.modLang.exp',X18) = s('type.modLang.exp',X21)))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_|49||49||split|2')).
fof(f400,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) != s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) | s('type.list.list'('type.string.char'),X1) = s('type.list.list'('type.string.char'),X4)) )),
  inference(cnf_transformation,[],[f257])).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) != s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) | s('type.modLang.exp',X2) = s('type.modLang.exp',X5)) )),
  inference(cnf_transformation,[],[f257])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3),s('type.list.list'('type.string.char'),X4),s('type.modLang.exp',X5))) | s('type.modLang.exp',X2) != s('type.modLang.exp',X5) | s('type.list.list'('type.string.char'),X1) != s('type.list.list'('type.string.char'),X4) | s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0) != s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f397,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X25,X26,X52,X27] : ! [X28,X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X52,X25),X28),s(fun(X26,X52),happ(s(fun(X27,fun(X26,X52)),X29),s(X27,X30))))) = s(fun(X26,X25),happ(s(fun(X27,fun(X26,X25)),'_dst_x00x11_2'(s(fun(X52,X25),X28),s(fun(X27,fun(X26,X52)),X29))),s(X27,X30)))),
  file('thm.modSem.v_|49||49||split|4','_dst_x00x11_ax')).
fof(f394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : ((s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3))) & (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : ((s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3))) & ((s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3)) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) <=> (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X16,X17,X19,X20] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.modSem.v'),X17))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X19),s('type.list.list'('type.modSem.v'),X20))) <=> (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X19) & s('type.list.list'('type.modSem.v'),X17) = s('type.list.list'('type.modSem.v'),X20)))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_|49||49||split|1')).
fof(f395,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f396,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))) )),
  inference(cnf_transformation,[],[f255])).
fof(f418,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X38,X55,X43,X56] : ((s(bool,X38) = s(bool,X55) & (p(s(bool,X55)) => s(bool,X43) = s(bool,X56))) => ((p(s(bool,X38)) => p(s(bool,X43))) <=> (p(s(bool,X55)) => p(s(bool,X56)))))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.IMP_CONG')).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.modSem.v_|49||49||split|4','_dst_x0x1_ax')).
fof(f368,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X25,X26,X27] : ! [X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X30),s(fun(X26,X27),X29))) = s(fun(X26,X25),happ(s(fun(fun(X27,X25),fun(X26,X25)),'_dstsel_0_x01_1'(s(fun(X26,X27),X29))),s(fun(X27,X25),X30)))),
  file('thm.modSem.v_|49||49||split|4','_dstsel_0_x01_ax')).
fof(f367,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ! [X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X45,X9,X31] : ! [X10] : ? [X33] : ! [X38,X43] : s(X45,happ(s(fun(X31,X45),happ(s(fun(X9,fun(X31,X45)),X10),s(X9,X38))),s(X31,X43))) = s(X45,happ(s(fun('type.pair.prod'(X9,X31),X45),X33),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.modSem.v_|49||49||split|4','thm.pair.pair_Axiom')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.modSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v'))),'const.modSem.Recclosure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1),s('type.list.list'('type.string.char'),X2))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.modSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v'))),'const.modSem.Recclosure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X1))),s('type.list.list'('type.string.char'),X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X22,X23] : s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X8),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X22),s('type.list.list'('type.string.char'),X23))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.string.char'),'type.modSem.v'),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),'type.modSem.v'))),'const.modSem.Recclosure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X8))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X22))),s('type.list.list'('type.string.char'),X23)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_3_const.modSem.Recclosure')).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v',happ(s(fun('type.modLang.exp','type.modSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v'))),'const.modSem.Closure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.list.list'('type.string.char'),X1))),s('type.modLang.exp',X2)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0),s('type.list.list'('type.string.char'),X1),s('type.modLang.exp',X2))) = s('type.modSem.v',happ(s(fun('type.modLang.exp','type.modSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v'))),'const.modSem.Closure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X0))),s('type.list.list'('type.string.char'),X1))),s('type.modLang.exp',X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X22,X23] : s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X8),s('type.list.list'('type.string.char'),X22),s('type.modLang.exp',X23))) = s('type.modSem.v',happ(s(fun('type.modLang.exp','type.modSem.v'),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v')),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp','type.modSem.v'))),'const.modSem.Closure_0'),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X8))),s('type.list.list'('type.string.char'),X22))),s('type.modLang.exp',X23)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_3_const.modSem.Closure')).
fof(f357,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X31,X9] : ! [X32,X11] : ? [X33] : (s(X31,X32) = s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X31,happ(s(fun(X31,X31),happ(s(fun('type.list.list'(X9),fun(X31,X31)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X31,X31))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X31,happ(s(fun('type.list.list'(X9),X31),X33),s('type.list.list'(X9),X17))))))),
  file('thm.modSem.v_|49||49||split|4','thm.list.list_Axiom')).
fof(f358,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f352,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X31,X9] : ! [X10,X37] : (s(fun(X9,X31),X10) = s(fun(X9,X31),X37) <=> ! [X38] : s(X31,happ(s(fun(X9,X31),X10),s(X9,X38))) = s(X31,happ(s(fun(X9,X31),X37),s(X9,X38))))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.FUN_EQ_THM')).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f351,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.modSem.v_|49||49||split|4',aHL_EXT)).
fof(f342,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X31] : ! [X38,X43] : s(X31,X43) = s(X31,'const.pair.SND_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.modSem.v_|49||49||split|4','thm.pair.SND')).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31,X9] : ! [X38,X43] : s(X9,X38) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X38),s(X31,X43)))))),
  file('thm.modSem.v_|49||49||split|4','thm.pair.FST')).
fof(f340,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X31,X9] : ! [X48] : ? [X10] : ! [X49,X50,X51] : s(X31,happ(s(fun('type.ind_type.recspace'(X9),X31),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X49),s(X9,X50),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))))) = s(X31,happ(s(fun(fun('type.num.num',X31),X31),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)))),X48),s('type.num.num',X49))),s(X9,X50))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51))),s(fun('type.num.num',X31),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X31),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X51)))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.CONSTR_REC')).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9,X31] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X31),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X31),X8)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_1_const.pair.FST')).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X31] : ! [X38] : s(X31,X38) = s(X31,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X31),'const.sum.INR_1'(s(X31,X38)))))),
  file('thm.modSem.v_|49||49||split|4','thm.sum.OUTR')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X31,X9] : ! [X38] : s(X9,X38) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X31),'const.sum.INL_1'(s(X9,X38)))))),
  file('thm.modSem.v_|49||49||split|4','thm.sum.OUTL')).
fof(f334,plain,(
  ( ! [X0,X1] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) != s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X6))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X7))) <=> s('type.ast.lit',X6) = s('type.ast.lit',X7))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_|49||49||split|0')).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) = s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f233])).
fof(f330,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X40,X41] : ((p(s(bool,X40)) => p(s(bool,X41))) => ((p(s(bool,X41)) => p(s(bool,X40))) => s(bool,X40) = s(bool,X41)))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f329,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X16,X17] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.modSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X17)))))))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_size_def|split|1')).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.modSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.modSem.v'),X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X22] : s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.modSem.v'),X22))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8))),s('type.list.list'('type.modSem.v'),X22)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_2_const.modSem.Conv')).
fof(f326,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X8,X22,X23] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X22),s('type.num.num',X23))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X22))),s('type.num.num',X23)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X22,X23] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X22),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X22))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_1_3_const.ind_type.CONSTR')).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X10,X42] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X42))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.FCONS|split|1')).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.modSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f272,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f269,plain,(
  ( ! [X0] : (s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_size_def|split|4')).
fof(f268,plain,(
  ( ! [X0] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X8))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X8)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_1_const.modSem.Vectorv')).
fof(f267,plain,(
  ( ! [X0] : (s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) = s('type.modSem.v',happ(s(fun('type.ast.lit','type.modSem.v'),'const.modSem.Litv_0'),s('type.ast.lit',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X0))) = s('type.modSem.v',happ(s(fun('type.ast.lit','type.modSem.v'),'const.modSem.Litv_0'),s('type.ast.lit',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.modSem.v','const.modSem.Litv_1'(s('type.ast.lit',X8))) = s('type.modSem.v',happ(s(fun('type.ast.lit','type.modSem.v'),'const.modSem.Litv_0'),s('type.ast.lit',X8)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_1_const.modSem.Litv')).
fof(f266,plain,(
  ( ! [X0] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) = s('type.modSem.v',happ(s(fun('type.num.num','type.modSem.v'),'const.modSem.Loc_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) = s('type.modSem.v',happ(s(fun('type.num.num','type.modSem.v'),'const.modSem.Loc_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X8))) = s('type.modSem.v',happ(s(fun('type.num.num','type.modSem.v'),'const.modSem.Loc_0'),s('type.num.num',X8)))),
  file('thm.modSem.v_|49||49||split|4','ARITY_0_1_const.modSem.Loc')).
fof(f265,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.modSem.v_|49||49||split|4','thm.ind_type.ISO_REFL')).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.modSem.v_|49||49||split|4',aHL_TRUTH)).
fof(f20291,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(forward_demodulation,[],[f20194,f336])).
fof(f20194,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(duplicate_literal_removal,[],[f20028])).
fof(f20028,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(superposition,[],[f10983,f14799])).
fof(f14799,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(trivial_inequality_removal,[],[f14788])).
fof(f14788,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(superposition,[],[f13331,f445])).
fof(f20350,plain,(
  ( ! [X74,X72] : (s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(subsumption_resolution,[],[f20349,f264])).
fof(f20349,plain,(
  ( ! [X74,X72] : (p(s(bool,'F_0')) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(forward_demodulation,[],[f20007,f336])).
fof(f20007,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X73),'const.sum.INL_1'(s(bool,'F_0')))))) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(superposition,[],[f2426,f14799])).
fof(f1394,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1395,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2265,plain,(
  ( ! [X23] : (~p(s(bool,sK4(X23,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f20539,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl42),
  inference(cnf_transformation,[],[f20539_D])).
fof(f20539_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f20538,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2382,f262])).
fof(f2382,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f319,f472])).
fof(f6042,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f6042_D])).
fof(f6042_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4779,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4746,f262])).
fof(f4746,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4686,f472])).
fof(f1906,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) | $spl8),
  inference(cnf_transformation,[],[f1906_D])).
fof(f1906_D,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4759,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4686,f309])).
fof(f17806,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4686,f309])).
fof(f18551,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f18445,f262])).
fof(f18445,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f4686,f472])).
fof(f20274,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl40),
  inference(cnf_transformation,[],[f20274_D])).
fof(f20274_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f20201,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(duplicate_literal_removal,[],[f20010])).
fof(f20010,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(superposition,[],[f2594,f14799])).
fof(f2594,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2593,f264])).
fof(f2593,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2560,f1409])).
fof(f1409,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f276])).
fof(f2560,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2426,f1433])).
fof(f1433,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1411,f262])).
fof(f1411,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1395,f445])).
fof(f15947,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(cnf_transformation,[],[f15947_D])).
fof(f15947_D,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f15897,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f15756])).
fof(f15756,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(superposition,[],[f2594,f14548])).
fof(f14262,plain,(
  $false | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f13331,f14110])).
fof(f14110,plain,(
  ( ! [X89] : (s(bool,'T_0') = s(bool,X89)) ) | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4526,f4527,f4528,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4595,f4596,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4760,f4761,f4797,f4807,f4841,f4847,f4861,f4891,f4890,f4884,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5641,f5642,f5643,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6356,f6357,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6862,f6863,f6864,f6908,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8270,f8271,f8617,f8618,f8619,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10605,f10606,f11095,f11096,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f6915,f7141,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f14109])).
fof(f14109,plain,(
  ( ! [X88,X87,X89] : (s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4587,f8243])).
fof(f14105,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4861,f8243])).
fof(f14104,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5619,f8243])).
fof(f14103,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X127),bool),sK18(X128,happ(s(fun('type.list.list'(X124),bool),sK18(X125,X126,X124,bool)),s('type.list.list'(X124),'const.list.NIL_0')),X127,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5622,f8243])).
fof(f14102,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.ind_type.FCONS_3'(s(bool,X133),s(fun('type.num.num',bool),X134),s('type.num.num','const.num.|48|_0')),X135,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5624,f8243])).
fof(f14101,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.sum.OUTR_1'(s('type.sum.sum'(X142,bool),'const.sum.INR_1'(s(bool,X143)))),X144,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5626,f8243])).
fof(f14100,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X149),bool),sK18(X150,'const.pair.FST_1'(s('type.pair.prod'(bool,X146),'const.pair.,_2'(s(bool,X147),s(X146,X148)))),X149,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5627,f8243])).
fof(f14099,plain,(
  ( ! [X154,X152,X151,X153] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X153),bool),sK18(X154,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X151),'const.sum.INL_1'(s(bool,X152)))),X153,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5628,f8243])).
fof(f14098,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X158),bool),sK18(X159,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))),X158,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5629,f8243])).
fof(f14261,plain,(
  $false | ($spl28 | $spl36)),
  inference(subsumption_resolution,[],[f13333,f14110])).
fof(f14260,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f264,f8243])).
fof(f14259,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4526,f4527,f4528,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4595,f4596,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4760,f4761,f4797,f4807,f4841,f4847,f4861,f4891,f4890,f4884,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5641,f5642,f5643,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6356,f6357,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6862,f6863,f6864,f6908,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8270,f8271,f8617,f8618,f8619,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10605,f10606,f11095,f11096,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f6915,f7141,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14255,f14258])).
fof(f14258,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl28),
  inference(subsumption_resolution,[],[f275,f8243])).
fof(f14255,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f14254])).
fof(f14254,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) ) | $spl28),
  inference(subsumption_resolution,[],[f2424,f8243])).
fof(f14133,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f14132,f8243])).
fof(f14132,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5427,f8243])).
fof(f14131,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f14130,f8243])).
fof(f14130,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f5506,f8243])).
fof(f14129,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl28),
  inference(subsumption_resolution,[],[f350,f8243])).
fof(f14128,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f355,f8243])).
fof(f14127,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (s(X363,X370) = s(X363,happ(s(fun(X362,X363),happ(s(fun(X364,fun(X362,X363)),'_dst_x0x1_2'(s(fun(X365,fun(X362,X363)),X366),s(fun(X364,X365),X367))),s(X364,X368))),s(X362,happ(s(fun(X363,X362),X369),s(X363,X370)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5542,f8243])).
fof(f14126,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (s(X380,X388) = s(X380,happ(s(fun(X381,X380),X387),s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5544,f8243])).
fof(f14125,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7064,f8243])).
fof(f14124,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7494,f8243])).
fof(f14123,plain,(
  ( ! [X377,X379,X372,X374,X380,X376,X378,X373,X375] : (s(X373,X380) = s(X373,happ(s(fun(X372,X373),happ(s(fun(X374,fun(X372,X373)),X375),s(X374,happ(s(fun(X376,X374),X377),s(X376,X378))))),s(X372,happ(s(fun(X373,X372),X379),s(X373,X380)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5464,f8243])).
fof(f14122,plain,(
  ( ! [X391,X397,X393,X395,X390,X396,X398,X392,X394] : (s(X390,X398) = s(X390,happ(s(fun(X391,X390),X397),s(X391,happ(s(fun(X390,X391),happ(s(fun(X392,fun(X390,X391)),X393),s(X392,happ(s(fun(X394,X392),X395),s(X394,X396))))),s(X390,X398)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5466,f8243])).
fof(f14121,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7063,f8243])).
fof(f14120,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7493,f8243])).
fof(f14119,plain,(
  ( ! [X191,X198,X196,X194,X192,X190,X199,X197,X195,X193] : (s(X190,X198) != s(X190,happ(s(fun(X191,X190),X197),s(X191,X199))) | s(X191,X199) = s(X191,happ(s(fun(X190,X191),'_dst_x0x1_2'(s(fun(X192,X191),X193),s(fun(X190,X192),happ(s(fun(X194,fun(X190,X192)),X195),s(X194,X196))))),s(X190,X198)))) ) | $spl28),
  inference(subsumption_resolution,[],[f12761,f8243])).
fof(f14118,plain,(
  ( ! [X212,X218,X210,X216,X215,X213,X211,X219,X217,X214] : (s(X211,X218) != s(X211,happ(s(fun(X210,X211),'_dst_x0x1_2'(s(fun(X212,X211),X213),s(fun(X210,X212),happ(s(fun(X214,fun(X210,X212)),X215),s(X214,X216))))),s(X210,X219))) | s(X210,X219) = s(X210,happ(s(fun(X211,X210),X217),s(X211,X218)))) ) | $spl28),
  inference(subsumption_resolution,[],[f12763,f8243])).
fof(f14117,plain,(
  ( ! [X255,X261,X257,X259,X260,X262,X254,X256,X258] : (s(X255,X262) = s(X255,happ(s(fun(X254,X255),'_dst_x0x1_2'(s(fun(X256,X255),X257),s(fun(X254,X256),happ(s(fun(X258,fun(X254,X256)),X259),s(X258,X260))))),s(X254,happ(s(fun(X255,X254),X261),s(X255,X262)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f12768,f8243])).
fof(f14116,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X272,X274] : (s(X272,X280) = s(X272,happ(s(fun(X273,X272),X279),s(X273,happ(s(fun(X272,X273),'_dst_x0x1_2'(s(fun(X274,X273),X275),s(fun(X272,X274),happ(s(fun(X276,fun(X272,X274)),X277),s(X276,X278))))),s(X272,X280)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f12770,f8243])).
fof(f14115,plain,(
  ( ! [X191,X189,X187,X185,X182,X190,X188,X186,X184,X183] : (s(X182,X190) != s(X182,happ(s(fun(X183,X182),X189),s(X183,X191))) | s(X183,X191) = s(X183,happ(s(fun(X182,X183),happ(s(fun(X184,fun(X182,X183)),'_dst_x00x11_2'(s(fun(X185,X183),X186),s(fun(X184,fun(X182,X185)),X187))),s(X184,X188))),s(X182,X190)))) ) | $spl28),
  inference(subsumption_resolution,[],[f12815,f8243])).
fof(f14114,plain,(
  ( ! [X210,X208,X206,X204,X202,X211,X209,X207,X205,X203] : (s(X203,X210) != s(X203,happ(s(fun(X202,X203),happ(s(fun(X204,fun(X202,X203)),'_dst_x00x11_2'(s(fun(X205,X203),X206),s(fun(X204,fun(X202,X205)),X207))),s(X204,X208))),s(X202,X211))) | s(X202,X211) = s(X202,happ(s(fun(X203,X202),X209),s(X203,X210)))) ) | $spl28),
  inference(subsumption_resolution,[],[f12817,f8243])).
fof(f14113,plain,(
  ( ! [X253,X251,X249,X246,X254,X252,X250,X248,X247] : (s(X247,X254) = s(X247,happ(s(fun(X246,X247),happ(s(fun(X248,fun(X246,X247)),'_dst_x00x11_2'(s(fun(X249,X247),X250),s(fun(X248,fun(X246,X249)),X251))),s(X248,X252))),s(X246,happ(s(fun(X247,X246),X253),s(X247,X254)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f12822,f8243])).
fof(f14112,plain,(
  ( ! [X269,X271,X265,X267,X268,X270,X264,X266,X272] : (s(X264,X272) = s(X264,happ(s(fun(X265,X264),X271),s(X265,happ(s(fun(X264,X265),happ(s(fun(X266,fun(X264,X265)),'_dst_x00x11_2'(s(fun(X267,X265),X268),s(fun(X266,fun(X264,X267)),X269))),s(X266,X270))),s(X264,X272)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f12824,f8243])).
fof(f14111,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) ) | $spl28),
  inference(subsumption_resolution,[],[f4689,f8243])).
fof(f14256,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2661,f2663,f2689,f2730,f2699,f2704,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3328,f3329,f3330,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4526,f4527,f4528,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4595,f4596,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4760,f4761,f4797,f4807,f4841,f4847,f4861,f4891,f4890,f4884,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5641,f5642,f5643,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6356,f6357,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6862,f6863,f6864,f6908,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8270,f8271,f8617,f8618,f8619,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10605,f10606,f11095,f11096,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f6915,f7141,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14255])).
fof(f14213,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496,f14107,f10335,f14108,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14139,f14144,f14147,f14212])).
fof(f14212,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f3632,f8243])).
fof(f14147,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f14146])).
fof(f14146,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f9738,f14110])).
fof(f14144,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f14143])).
fof(f14143,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f10042,f14110])).
fof(f14139,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f14138])).
fof(f14138,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f12470,f14110])).
fof(f14108,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool))))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496,f14107,f10335])).
fof(f14107,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool))))))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10335,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496])).
fof(f14106,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool))))))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10335,f10431,f10506,f10485,f10486,f10496,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720])).
fof(f14184,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f14183,f8243])).
fof(f14183,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5423,f14131])).
fof(f14182,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f14181,f8243])).
fof(f14181,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5502,f14133])).
fof(f14148,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496,f14107,f10335,f14108,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14139,f14144,f14147])).
fof(f14145,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496,f14107,f10335,f14108,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14139,f14144])).
fof(f14140,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f261,f260,f262,f264,f265,f266,f267,f268,f269,f270,f272,f271,f276,f275,f309,f316,f324,f325,f326,f327,f328,f329,f333,f330,f335,f334,f336,f337,f338,f340,f341,f342,f351,f353,f352,f358,f357,f359,f360,f367,f368,f369,f418,f396,f395,f394,f397,f402,f401,f400,f399,f406,f405,f404,f403,f432,f440,f445,f447,f472,f494,f513,f540,f633,f686,f743,f803,f836,f837,f838,f839,f840,f841,f842,f909,f934,f967,f968,f969,f970,f971,f972,f973,f990,f1000,f1001,f1002,f1003,f1004,f1005,f1006,f1358,f1359,f1360,f1384,f1519,f1905,f1484,f1486,f1488,f1909,f398,f356,f355,f354,f350,f349,f1889,f339,f348,f345,f1904,f347,f346,f321,f320,f317,f420,f311,f1497,f1374,f1376,f1378,f834,f1391,f833,f1390,f832,f1389,f831,f1388,f830,f1387,f1533,f1515,f930,f1386,f829,f1385,f319,f318,f1392,f1908,f2070,f2051,f2168,f2225,f2226,f2227,f2228,f2229,f2230,f2231,f2232,f2417,f2371,f2418,f2419,f2420,f2421,f2422,f2424,f2426,f2427,f2524,f2529,f2530,f2531,f2661,f2663,f2689,f2730,f2699,f2704,f2913,f2915,f2917,f2803,f2949,f2951,f2965,f2994,f2995,f2996,f2997,f2998,f2999,f3000,f3001,f3002,f3003,f3320,f3321,f3322,f3337,f3328,f3329,f3330,f3340,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3415,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3548,f3558,f3559,f3632,f3633,f4073,f4078,f4081,f3718,f3728,f3729,f4107,f4140,f4141,f4142,f4143,f4144,f4145,f4146,f4147,f4148,f4258,f4259,f4260,f4261,f4262,f4263,f4264,f4265,f4344,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4463,f4517,f4520,f4526,f4527,f4528,f4685,f4534,f4535,f4538,f4539,f4540,f4541,f4542,f4543,f4544,f4545,f4546,f4547,f4686,f4575,f4576,f4577,f4578,f4579,f4581,f4582,f4587,f4588,f4589,f4590,f4591,f4687,f4595,f4596,f4606,f4609,f4610,f4611,f4612,f4613,f4614,f4622,f4689,f4690,f4691,f4702,f4781,f4784,f4760,f4761,f4797,f4807,f4841,f4844,f4847,f4861,f4866,f4891,f4890,f4884,f4936,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f5054,f5079,f5082,f5078,f5113,f5116,f5138,f5299,f5298,f5300,f5301,f5302,f5303,f5304,f5305,f5306,f5307,f5308,f5571,f5314,f5315,f5318,f5418,f5419,f5421,f5422,f5423,f5424,f5425,f5427,f5428,f5433,f5434,f5435,f5436,f5437,f5441,f5442,f5452,f5456,f5457,f5458,f5460,f5461,f5462,f5464,f5465,f5466,f5467,f5473,f5475,f5478,f5495,f5496,f5497,f5498,f5500,f5501,f5502,f5503,f5504,f5506,f5507,f5512,f5513,f5514,f5515,f5516,f5520,f5521,f5531,f5534,f5535,f5536,f5538,f5539,f5540,f5542,f5543,f5544,f5545,f5551,f5553,f5556,f5631,f5634,f5641,f5642,f5643,f5644,f5646,f5647,f5651,f5655,f5619,f5620,f5621,f5622,f5623,f5624,f5625,f5626,f5627,f5628,f5629,f5630,f6030,f6031,f5942,f5943,f6032,f6033,f6034,f6035,f6036,f6037,f6038,f6039,f5952,f6041,f6044,f6047,f6050,f6018,f6019,f6020,f6051,f6188,f6189,f6190,f6191,f6192,f6193,f6195,f6197,f6198,f6283,f6293,f6341,f6342,f6300,f6301,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6310,f6311,f6312,f6313,f6352,f6319,f6320,f6356,f6357,f6360,f6365,f6339,f6366,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6413,f6414,f6500,f6510,f6656,f6657,f6576,f6577,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6667,f6599,f6600,f6604,f6605,f6606,f6607,f6670,f6671,f6610,f6611,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6621,f6622,f6623,f6681,f6629,f6630,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6685,f6646,f6647,f6651,f6652,f6653,f6654,f6890,f6891,f6838,f6839,f6892,f6893,f6894,f6895,f6896,f6897,f6898,f6899,f6848,f6849,f6850,f6851,f6901,f6857,f6858,f6862,f6863,f6864,f6905,f6870,f6871,f6908,f6911,f6889,f7063,f7064,f7073,f7076,f7077,f7086,f7121,f7122,f7125,f7493,f7494,f7503,f7506,f7507,f7516,f7552,f7555,f8031,f8032,f8033,f8034,f8035,f8036,f8037,f8038,f8039,f8040,f8041,f8232,f8270,f8271,f8617,f8618,f8619,f8757,f8625,f8626,f8632,f8646,f8647,f8648,f8649,f8664,f8666,f8667,f8671,f8672,f8684,f8685,f8696,f8788,f8789,f8700,f8791,f8798,f8799,f8804,f8725,f8731,f8736,f9310,f9311,f9312,f9313,f9314,f9315,f9316,f9317,f9318,f9319,f9320,f9321,f9322,f9323,f9355,f9356,f9326,f9327,f9357,f9358,f9359,f9360,f9361,f9362,f9363,f9364,f9365,f9337,f9367,f9368,f9371,f9575,f9576,f9577,f9578,f9579,f9580,f9581,f9582,f9583,f9584,f9585,f9586,f9587,f9588,f9620,f9621,f9591,f9592,f9622,f9623,f9624,f9625,f9626,f9627,f9628,f9629,f9630,f9602,f9632,f9633,f9636,f9738,f9692,f9693,f9694,f9695,f9696,f9697,f9698,f9699,f9700,f9701,f9702,f9703,f9877,f9878,f9879,f9880,f9881,f9882,f9883,f9884,f9885,f9886,f9887,f9888,f9889,f9890,f9922,f9923,f9893,f9894,f9924,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9904,f9934,f9935,f9938,f10042,f9995,f9996,f9997,f9998,f9999,f10000,f10001,f10002,f10003,f10004,f10005,f10006,f10187,f10188,f10189,f10190,f10191,f10192,f10193,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10232,f10233,f10203,f10204,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10214,f10244,f10245,f10248,f10270,f10345,f10324,f10325,f10431,f10506,f10485,f10486,f10608,f10610,f10577,f10578,f10612,f10614,f10616,f10618,f10620,f10622,f10624,f10626,f10628,f10588,f10629,f10632,f10605,f10606,f10655,f10730,f10709,f10710,f10851,f10852,f10854,f10855,f10857,f10862,f10983,f10993,f10932,f10933,f11095,f11096,f11227,f11232,f11237,f11868,f11889,f11816,f11890,f11891,f11892,f11805,f11910,f11932,f11954,f11841,f11842,f11866,f12470,f12422,f12423,f12425,f12426,f12428,f12433,f12628,f12629,f12630,f12631,f12632,f12633,f12737,f12738,f12743,f12744,f12745,f12746,f12747,f12748,f12749,f12750,f12751,f12752,f12753,f12754,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12766,f12767,f12768,f12769,f12770,f12771,f12776,f12797,f12798,f12799,f12800,f12801,f12802,f12803,f12804,f12805,f12806,f12807,f12808,f12809,f12810,f12811,f12812,f12813,f12814,f12815,f12816,f12817,f12818,f12820,f12821,f12822,f12823,f12824,f12825,f12830,f12831,f12833,f12834,f13330,f13332,f6915,f7141,f7144,f8233,f8238,f13336,f8243,f8042,f11852,f14098,f14099,f14100,f14101,f14102,f14103,f14104,f14105,f4620,f6620,f6586,f4615,f10720,f14106,f10496,f14107,f10335,f14108,f14109,f14110,f14111,f14112,f14113,f14114,f14115,f14116,f14117,f14118,f14119,f14120,f14121,f14122,f14123,f14124,f14125,f14126,f14127,f14128,f14129,f14131,f14133,f14139])).
fof(f14137,plain,(
  $false | ($spl6 | $spl10 | $spl28)),
  inference(subsumption_resolution,[],[f13329,f14110])).
fof(f13329,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f12414,f1410])).
fof(f1410,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1395,f271])).
fof(f12414,plain,(
  ( ! [X65] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X65,'_i_0_0',bool))) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f10983,f1409])).
fof(f14136,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f13334,f14110])).
fof(f13334,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9987,f1410])).
fof(f9987,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6500,f1409])).
fof(f14135,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f13335,f14110])).
fof(f13335,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9684,f1410])).
fof(f9684,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6283,f1409])).
fof(f14093,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f13616,f13251])).
fof(f13616,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f13251,f13335])).
fof(f14092,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f13615,f13251])).
fof(f13615,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f13251,f13334])).
fof(f14091,plain,(
  $false | ($spl6 | $spl10 | $spl30)),
  inference(subsumption_resolution,[],[f13614,f13251])).
fof(f13614,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl6 | $spl10 | $spl30)),
  inference(backward_demodulation,[],[f13251,f13329])).
fof(f14083,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f14082,f8243])).
fof(f14082,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13494,f13251])).
fof(f13494,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(backward_demodulation,[],[f13251,f5502])).
fof(f14081,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f14080,f8243])).
fof(f14080,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f13484,f13251])).
fof(f13484,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(backward_demodulation,[],[f13251,f5423])).
fof(f14073,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f14072,f8243])).
fof(f14072,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f13338,f13251])).
fof(f13338,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl30),
  inference(backward_demodulation,[],[f13251,f275])).
fof(f14071,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f13337,f262])).
fof(f13337,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(backward_demodulation,[],[f13251,f264])).
fof(f13333,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl36),
  inference(cnf_transformation,[],[f13333_D])).
fof(f13333_D,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13332,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(subsumption_resolution,[],[f10923,f10983])).
fof(f10923,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X78))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(superposition,[],[f6283,f2913])).
fof(f13331,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl34),
  inference(cnf_transformation,[],[f13331_D])).
fof(f13331_D,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13330,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f10924,f10983])).
fof(f10924,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f6500,f2913])).
fof(f13252,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl32),
  inference(cnf_transformation,[],[f13252_D])).
fof(f13252_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13178,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(duplicate_literal_removal,[],[f12982])).
fof(f12982,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(superposition,[],[f2594,f9829])).
fof(f8242,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl26),
  inference(cnf_transformation,[],[f8242_D])).
fof(f8242_D,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8243,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl28),
  inference(cnf_transformation,[],[f8243_D])).
fof(f8243_D,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8232,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl10),
  inference(superposition,[],[f7144,f341])).
fof(f6571,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(cnf_transformation,[],[f6571_D])).
fof(f6571_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6573,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl24),
  inference(cnf_transformation,[],[f6573_D])).
fof(f6573_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6555,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f6043,f309])).
fof(f6043,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl20),
  inference(cnf_transformation,[],[f6043_D])).
fof(f6043_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6050,plain,(
  ( ! [X64,X65] : (p(s(bool,sK5('_i_0_0',X65,bool))) | ~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool)))))) )),
  inference(forward_demodulation,[],[f6017,f309])).
fof(f6017,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X65,bool)))))) )),
  inference(superposition,[],[f320,f309])).
fof(f6041,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f6040,f309])).
fof(f6040,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6003,f262])).
fof(f6003,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f320,f472])).
fof(f5649,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5649_D])).
fof(f5649_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5648,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5612,f275])).
fof(f5612,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4861,f2532])).
fof(f2532,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2510,f262])).
fof(f2510,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f2424,f1610])).
fof(f1610,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1557,f272])).
fof(f1557,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1410,f318])).
fof(f5645,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl14),
  inference(cnf_transformation,[],[f5645_D])).
fof(f5645_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5644,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5609,f275])).
fof(f5609,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4861,f2531])).
fof(f5085,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl12),
  inference(cnf_transformation,[],[f5085_D])).
fof(f5085_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5078,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4761,f341])).
fof(f2047,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1957,f264])).
fof(f1957,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f330])).
fof(f2046,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1959,f1906])).
fof(f1959,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f330])).
fof(f2045,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1967,f264])).
fof(f1967,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f330])).
fof(f2044,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1969,f1906])).
fof(f1969,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f330])).
fof(f1977,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1906,f318])).
fof(f1974,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f440,f1906,f271])).
fof(f1972,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1906,f276])).
fof(f1970,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1968,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f1906,f330])).
fof(f1966,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1964,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1906,f432,f1906,f330])).
fof(f1960,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1958,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f432,f1906,f330])).
fof(f1956,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1906,f1410,f1906,f330])).
fof(f1954,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1906,f432,f1906,f330])).
fof(f1938,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f342,f1906,f418])).
fof(f1936,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f336,f1906,f418])).
fof(f1934,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f341,f1906,f418])).
fof(f1932,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f337,f1906,f418])).
fof(f1930,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f316,f1906,f418])).
fof(f1928,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f357,f1906,f418])).
fof(f1926,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f309,f1906,f418])).
fof(f1923,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f262,f1610,f1906,f418])).
fof(f1921,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f440,f1906,f418])).
fof(f1918,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1410,f1410,f1906,f494])).
fof(f1916,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f432,f1906,f494])).
fof(f1910,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f262,f1906])).
fof(f1908,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f1908_D])).
fof(f1908_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1905,plain,(
  ( ! [X23] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f1888,f309])).
fof(f1888,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f317,f309])).
fof(f1384,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f319,f309])).
fof(f431,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f430])).
fof(f430,plain,(
  s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f427,f428])).
fof(f426,plain,(
  s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f428,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f198])).
fof(f198,plain,(
  ? [X0,X1] : ((s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ? [X0,X1] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X0))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X6))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X6))) = s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.modSem.v_|49||49||split|4','thm.modSem.v_|49||49||split|4')).
fof(f427,plain,(
  s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f429,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f261,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK0))) != s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f199])).
% SZS output end Proof for thm.modSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 216
% Passive clauses: 2614
% Generated clauses: 32188
% Final active clauses: 143
% Final passive clauses: 1478
% Input formulas: 69
% Initial clauses: 147
% 
% Duplicate literals: 1025
% Trivial inequalities: 239
% Fw subsumption resolutions: 1540
% Fw demodulations: 5539
% Bw demodulations: 761
% Condensations: 1433
% Global subsumptions: 14
% 
% Simple tautologies: 355
% Equational tautologies: 193
% Forward subsumptions: 18179
% Fw demodulations to eq. taut.: 983
% Bw demodulations to eq. taut.: 64
% 
% Unit resulting resolution: 2670
% Factoring: 1
% Forward superposition: 9108
% Backward superposition: 7602
% Self superposition: 174
% Equality factoring: 90
% Equality resolution: 37
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 2727
% SAT solver unit clauses: 148
% SAT solver binary clauses: 931
% 
% Sat splits: 23
% Sat splitting refutations: 51
% 
% Memory used [KB]: 33389
% Time elapsed: 1.406 s
% ------------------------------
% ------------------------------
% Success in time 2.731 s
