Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.labSem
% (90253)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 145
% Passive clauses: 2718
% Generated clauses: 21130
% Final active clauses: 113
% Final passive clauses: 2360
% Input formulas: 132
% Initial clauses: 141
% 
% Duplicate literals: 849
% Trivial inequalities: 214
% Fw subsumption resolutions: 2613
% Fw demodulations: 3831
% Bw demodulations: 27
% 
% Simple tautologies: 278
% Equational tautologies: 77
% Forward subsumptions: 9983
% 
% Binary resolution: 268
% Factoring: 1
% Forward superposition: 6409
% Backward superposition: 6240
% Self superposition: 254
% Equality factoring: 168
% Equality resolution: 21
% 
% Split clauses: 274
% Split components: 548
% SAT solver clauses: 743
% SAT solver unit clauses: 11
% SAT solver binary clauses: 78
% 
% Sat splits: 274
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 743
% TWLsolver calls for satisfiability: 131
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.labSem
% (90274)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 821
% Generated clauses: 13959
% Final active clauses: 91
% Final passive clauses: 578
% Input formulas: 132
% Initial clauses: 141
% 
% Duplicate literals: 674
% Trivial inequalities: 144
% Fw subsumption resolutions: 999
% Fw demodulations: 1710
% Bw demodulations: 141
% Condensations: 1378
% 
% Simple tautologies: 184
% Equational tautologies: 59
% Forward subsumptions: 6604
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 229
% Factoring: 1
% Forward superposition: 3670
% Backward superposition: 3656
% Self superposition: 214
% Equality factoring: 126
% Equality resolution: 14
% 
% Split clauses: 82
% Split components: 164
% SAT solver clauses: 1426
% SAT solver unit clauses: 258
% SAT solver binary clauses: 53
% 
% Sat splits: 82
% Sat splitting refutations: 28
% 
% TWLsolver clauses: 1426
% TWLsolver calls for satisfiability: 306
% 
% Memory used [KB]: 14711
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.labSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.labSem
% SZS output start Proof for thm.labSem
fof(f19293,plain,(
  $false),
  inference(subsumption_resolution,[],[f19292,f405])).
fof(f405,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f59])).
fof(f59,axiom,(
  p(s(bool,'T_0'))),
  file('thm.labSem.state_fupdfupds_comp|split|21','thm.bool.TRUTH')).
fof(f19292,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f19290,f2845])).
fof(f2845,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2768,f551])).
fof(f551,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f407,f408])).
fof(f408,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.labSem.state_fupdfupds_comp|split|21',aHL_BOOL_CASES)).
fof(f407,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.labSem.state_fupdfupds_comp|split|21',aHL_FALSITY)).
fof(f2768,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f2689,f1264])).
fof(f1264,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X0,X0),'_i_0_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1219,f476])).
fof(f476,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7,X15] : ! [X11,X10,X16] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X11),s(fun(X15,X7),X10),s(X15,X16))) = s(X6,happ(s(fun(X7,X6),X11),s(X7,happ(s(fun(X15,X7),X10),s(X15,X16)))))),
  file('thm.labSem.state_fupdfupds_comp|split|21','thm.combin.o_THM')).
fof(f1219,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f409,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.labSem.state_fupdfupds_comp|split|21',arity_0_HL_EQ)).
fof(f409,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X43] : ! [X45] : s(X43,X45) = s(X43,happ(s(fun(X43,X43),'_i_0_0'),s(X43,X45)))),
  file('thm.labSem.state_fupdfupds_comp|split|21','_i0_ax')).
fof(f2689,plain,(
  ( ! [X76,X74,X75,X73] : (s(X75,happ(s(fun(X73,X75),X76),s(X73,X74))) = s(X75,'const.combin.o_3'(s(fun(X73,X75),X76),s(fun(X73,X73),'_i_0_0'),s(X73,X74)))) )),
  inference(superposition,[],[f476,f409])).
fof(f19290,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))))),
  inference(backward_demodulation,[],[f19252,f4778])).
fof(f4778,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.labSem.state'(sK1,sK2)),fun(fun(sK0,'type.labSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))))))))),
  inference(unit_resulting_resolution,[],[f404,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f403])).
fof(f404,plain,(
  s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3)))))),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),sK5),s(fun('type.num.num','type.num.num'),sK4))))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK5))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.labSem.state'(sK1,sK2),'type.labSem.state'(sK1,sK2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),sK4))),s(fun(sK0,'type.labSem.state'(sK1,sK2)),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f386])).
fof(f386,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) != s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) = s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X5),s(fun('type.num.num','type.num.num'),X4))))),s(fun(X0,'type.labSem.state'(X1,X2)),X3))) = s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X5))),s(fun(X0,'type.labSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.labSem.state'(X1,X2),'type.labSem.state'(X1,X2)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X4))),s(fun(X0,'type.labSem.state'(X1,X2)),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))),s(fun(X6,'type.labSem.state'(X7,X8)),X9))) = s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X10))),s(fun(X6,'type.labSem.state'(X7,X8)),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))),s(fun(X6,'type.labSem.state'(X7,X8)),X9))) = s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun(X6,'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X10))),s(fun(X6,'type.labSem.state'(X7,X8)),X9)))))),
  file('thm.labSem.state_fupdfupds_comp|split|21','thm.labSem.state_fupdfupds_comp|split|21')).
fof(f19252,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X2),s(fun('type.num.num','type.num.num'),X3))))),s(fun(X4,'type.labSem.state'(X0,X1)),X5))) = s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X2))),s(fun(X4,'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun(X4,'type.labSem.state'(X0,X1)),X5)))))) )),
  inference(superposition,[],[f490,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X3),s(fun('type.num.num','type.num.num'),X2))))) = s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X3))),s(fun('type.labSem.state'(X0,X1),'type.labSem.state'(X0,X1)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : ! [X10,X11] : s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),'const.combin.o_2'(s(fun('type.num.num','type.num.num'),X11),s(fun('type.num.num','type.num.num'),X10))))) = s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X11))),s(fun('type.labSem.state'(X7,X8),'type.labSem.state'(X7,X8)),'const.labSem.state_ptr_reg_fupd_1'(s(fun('type.num.num','type.num.num'),X10)))))),
  file('thm.labSem.state_fupdfupds_comp|split|21','thm.labSem.state_fupdfupds_comp|split|20')).
fof(f490,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7,X22,X15] : ! [X11,X10,X9] : s(fun(X22,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X22,X7),'const.combin.o_2'(s(fun(X15,X7),X10),s(fun(X22,X15),X9))))) = s(fun(X22,X6),'const.combin.o_2'(s(fun(X15,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X15,X7),X10))),s(fun(X22,X15),X9)))),
  file('thm.labSem.state_fupdfupds_comp|split|21','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.labSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 167
% Passive clauses: 1841
% Generated clauses: 19652
% Final active clauses: 86
% Final passive clauses: 807
% Input formulas: 132
% Initial clauses: 141
% 
% Duplicate literals: 862
% Trivial inequalities: 104
% Fw subsumption resolutions: 793
% Fw demodulations: 2768
% Bw demodulations: 776
% Condensations: 1204
% Global subsumptions: 2
% 
% Simple tautologies: 141
% Equational tautologies: 197
% Forward subsumptions: 11595
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 793
% Factoring: 1
% Forward superposition: 5484
% Backward superposition: 5384
% Self superposition: 244
% Equality factoring: 132
% Equality resolution: 9
% 
% Split clauses: 45
% Split components: 90
% SAT solver clauses: 2006
% SAT solver unit clauses: 403
% SAT solver binary clauses: 737
% 
% Sat splits: 45
% Sat splitting refutations: 71
% 
% Memory used [KB]: 18933
% Time elapsed: 0.942 s
% ------------------------------
% ------------------------------
% Success in time 2.31 s
