Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.infer_eComplete
% (90303)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 693
% Passive clauses: 9972
% Generated clauses: 21453
% Final active clauses: 690
% Final passive clauses: 8946
% Input formulas: 133
% Initial clauses: 311
% 
% Duplicate literals: 233
% Trivial inequalities: 9
% Fw subsumption resolutions: 1155
% Fw demodulations: 1166
% Bw demodulations: 2
% 
% Simple tautologies: 219
% Equational tautologies: 59
% Forward subsumptions: 8555
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 10780
% Factoring: 1202
% Forward superposition: 3574
% Backward superposition: 2918
% Self superposition: 38
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 182
% Split components: 366
% SAT solver clauses: 421
% SAT solver unit clauses: 11
% SAT solver binary clauses: 80
% 
% Sat splits: 182
% 
% TWLsolver clauses: 421
% TWLsolver calls for satisfiability: 358
% 
% Memory used [KB]: 13816
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.infer_eComplete
% (90336)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 388
% Passive clauses: 2251
% Generated clauses: 10869
% Final active clauses: 385
% Final passive clauses: 1854
% Input formulas: 133
% Initial clauses: 311
% 
% Duplicate literals: 102
% Trivial inequalities: 27
% Fw subsumption resolutions: 493
% Fw demodulations: 1124
% Bw demodulations: 11
% Condensations: 57
% Global subsumptions: 19
% 
% Simple tautologies: 199
% Equational tautologies: 19
% Forward subsumptions: 3805
% Fw demodulations to eq. taut.: 16
% 
% Binary resolution: 1642
% Factoring: 57
% Forward superposition: 2082
% Backward superposition: 2604
% Self superposition: 49
% Equality factoring: 4
% Equality resolution: 21
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 3016
% SAT solver unit clauses: 346
% SAT solver binary clauses: 120
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 18
% 
% Sat splits: 2
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2940
% TWLsolver calls for satisfiability: 2416
% 
% Memory used [KB]: 17782
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.infer_eComplete
% (90379)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1204
% Passive clauses: 12118
% Generated clauses: 141137
% Final active clauses: 982
% Final passive clauses: 2467
% Input formulas: 133
% Initial clauses: 311
% Discarded non-redundant clauses: 13901
% 
% Duplicate literals: 1712
% Trivial inequalities: 337
% Fw subsumption resolutions: 7133
% Fw demodulations: 27890
% Bw demodulations: 4435
% Condensations: 5070
% Global subsumptions: 1345
% 
% Simple tautologies: 335
% Equational tautologies: 316
% Forward subsumptions: 76937
% Fw demodulations to eq. taut.: 3899
% Bw demodulations to eq. taut.: 248
% 
% Unit resulting resolution: 16835
% Factoring: 1
% Forward superposition: 36032
% Backward superposition: 34902
% Self superposition: 280
% Equality factoring: 111
% Equality resolution: 55
% 
% Split clauses: 337
% Split components: 674
% SAT solver clauses: 15977
% SAT solver unit clauses: 4899
% SAT solver binary clauses: 4153
% 
% Sat splits: 337
% Sat splitting refutations: 1838
% 
% Memory used [KB]: 177182
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.infer_eComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.infer_eComplete
% SZS output start Proof for thm.infer_eComplete
fof(f45295,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f45216])).
fof(f45216,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',sK11))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',sK11))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num','const.num.|48|_0'))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f759,f4867])).
fof(f4867,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',X0)))) )),
  inference(forward_demodulation,[],[f4841,f449])).
fof(f449,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t',X0) = s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num','const.num.|48|_0'),s('type.infer_t.infer_t',X0)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : s('type.infer_t.infer_t',X0) = s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num','const.num.|48|_0'),s('type.infer_t.infer_t',X0)))),
  inference(rectify,[],[f157])).
fof(f157,plain,(
  ! [X1] : s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num','const.num.|48|_0'),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X9,X7] : s('type.infer_t.infer_t',X7) = s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num','const.num.|48|_0'),s('type.infer_t.infer_t',X7)))),
  file('thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb','thm.inferProps.infer_deBruijn_inc|48|')).
fof(f4841,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num','const.num.|48|_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',X0))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',X0)))) )),
  inference(superposition,[],[f739,f449])).
fof(f739,plain,(
  ( ! [X4,X5] : (s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X4),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',X5))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X4),s('type.infer_t.infer_t',X5)))))) )),
  inference(forward_demodulation,[],[f738,f448])).
fof(f448,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X16,X17] : s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X16),s('type.infer_t.infer_t',X17))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X16))),s('type.infer_t.infer_t',X17)))),
  file('thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb','ARITY_1_2_const.infer.infer_deBruijn_inc')).
fof(f738,plain,(
  ( ! [X4,X5] : (s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',X5))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t','const.infer.infer_deBruijn_inc_2'(s('type.num.num',X4),s('type.infer_t.infer_t',X5)))))) )),
  inference(forward_demodulation,[],[f630,f448])).
fof(f630,plain,(
  ( ! [X4,X5] : (s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',X5))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X4))),s('type.infer_t.infer_t',X5)))))) )),
  inference(resolution,[],[f396,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t',X2))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t',X2))))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) => s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.infer_t.infer_t',X2))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X6,X7] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) => s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X8))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.infer_t.infer_t',X7))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X8))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X8))),s('type.infer_t.infer_t',X7))))))),
  file('thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb','thm.inferProps.walkstar_inc')).
fof(f396,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))))),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10)))) & s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',sK11))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',sK11)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f235])).
fof(f235,plain,(
  ? [X0] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & ? [X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) => ! [X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) => ! [X7] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.infer_t.infer_t',X7))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6))),s('type.infer_t.infer_t',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) => ! [X7] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.infer_t.infer_t',X7))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6))),s('type.infer_t.infer_t',X7))))),
  file('thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb','thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb')).
fof(f759,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',sK11))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num',X0))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',sK11))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num','const.num.|48|_0'))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f397,f619])).
fof(f619,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : ((s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X26,X27] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X26))) = s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X27))) <=> s('type.num.num',X26) = s('type.num.num',X27))),
  file('thm.infer_eComplete.infer_eComplete_|49|4_prove_ltb','thm.infer_t.infer_t_|49||49||split|2')).
fof(f397,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10),s('type.infer_t.infer_t',sK11))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.finite_map.o_f_2'(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.infer.infer_deBruijn_inc_1'(s('type.num.num','const.num.|48|_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK10))),s('type.infer_t.infer_t',sK11)))),
  inference(cnf_transformation,[],[f298])).
% SZS output end Proof for thm.infer_eComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1073
% Passive clauses: 14561
% Generated clauses: 45283
% Final active clauses: 856
% Final passive clauses: 11128
% Input formulas: 133
% Initial clauses: 226
% 
% Selected by SInE selection: 75
% SInE iterations: 1
% 
% Duplicate literals: 1983
% Trivial inequalities: 262
% Fw subsumption resolutions: 4890
% Fw demodulations: 2644
% Bw demodulations: 17
% 
% Simple tautologies: 3244
% Equational tautologies: 175
% Forward subsumptions: 16122
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 22860
% Factoring: 22
% Forward superposition: 9547
% Backward superposition: 2170
% Self superposition: 4
% Equality resolution: 37
% 
% Split clauses: 2054
% Split components: 4537
% SAT solver clauses: 2136
% SAT solver unit clauses: 24
% SAT solver binary clauses: 560
% 
% Sat splits: 2054
% Sat splitting refutations: 26
% 
% Memory used [KB]: 15991
% Time elapsed: 0.935 s
% ------------------------------
% ------------------------------
% Success in time 12.711 s
