Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.exh_to_patProof
% (113871)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 1547
% Generated clauses: 15312
% Final active clauses: 100
% Final passive clauses: 1366
% Input formulas: 14
% Initial clauses: 15
% 
% Duplicate literals: 1181
% Trivial inequalities: 95
% Fw subsumption resolutions: 1591
% Fw demodulations: 2329
% Bw demodulations: 2
% 
% Simple tautologies: 207
% Equational tautologies: 22
% Forward subsumptions: 7172
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 10
% Forward superposition: 6172
% Backward superposition: 3539
% Self superposition: 140
% Equality factoring: 50
% Equality resolution: 11
% 
% Split clauses: 517
% Split components: 1045
% SAT solver clauses: 944
% SAT solver unit clauses: 12
% SAT solver binary clauses: 115
% 
% Sat splits: 517
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 941
% TWLsolver calls for satisfiability: 102
% 
% Memory used [KB]: 10106
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exh_to_patProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exh_to_patProof
% SZS output start Proof for thm.exh_to_patProof
fof(f2596,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2594,f37,f54,f1649,f1650,f1348,f1356,f42,f61,f48,f65,f1643,f1644,f40,f58,f49,f66,f50,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) | $spl20),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f50,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11,X12] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X11),s(fun('type.num.num',fun('type.num.num',bool)),X12))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X11))),s(fun('type.num.num',fun('type.num.num',bool)),X12)))),
  file('thm.exh_to_patProof.bind_bindn|split|0','ARITY_0_2_const.exh_to_patProof.bindn')).
fof(f66,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),'const.arithmetic.FUNPOW_1'(s(fun(X0,X0),X1))),s('type.num.num',X2))),s(X0,X3)))) ) | $spl18),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),'const.arithmetic.FUNPOW_1'(s(fun(X0,X0),X1))),s('type.num.num',X2))),s(X0,X3)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),'const.arithmetic.FUNPOW_1'(s(fun(X0,X0),X1))),s('type.num.num',X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),'const.arithmetic.FUNPOW_1'(s(fun(X0,X0),X1))),s('type.num.num',X2))),s(X0,X3)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),'const.arithmetic.FUNPOW_1'(s(fun(X0,X0),X1))),s('type.num.num',X2))),s(X0,X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ! [X11,X12,X13] : s(X8,'const.arithmetic.FUNPOW_3'(s(fun(X8,X8),X11),s('type.num.num',X12),s(X8,X13))) = s(X8,happ(s(fun(X8,X8),happ(s(fun('type.num.num',fun(X8,X8)),'const.arithmetic.FUNPOW_1'(s(fun(X8,X8),X11))),s('type.num.num',X12))),s(X8,X13)))),
  file('thm.exh_to_patProof.bind_bindn|split|0','ARITY_1_3_const.arithmetic.FUNPOW')).
fof(f58,plain,(
  s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0') = s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.arithmetic.FUNPOW_1'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'))) | $spl4),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0') = s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.arithmetic.FUNPOW_1'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f40,plain,(
  s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0') = s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.arithmetic.FUNPOW_1'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0')))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0') = s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.arithmetic.FUNPOW_1'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0')))),
  file('thm.exh_to_patProof.bind_bindn|split|0','thm.exh_to_patProof.bindn_def')).
fof(f1644,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) | $spl252),
  inference(cnf_transformation,[],[f1644_D])).
fof(f1644_D,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f1643,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) | ($spl4 | $spl18 | $spl20)),
  inference(forward_demodulation,[],[f1637,f67])).
fof(f1637,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),happ(s(fun('type.num.num',fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool)))),'const.exh_to_patProof.bindn_0'),s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1)))) ) | ($spl4 | $spl18)),
  inference(superposition,[],[f66,f58])).
fof(f65,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),X1),s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3)))))) ) | $spl16),
  inference(cnf_transformation,[],[f65_D])).
fof(f65_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),X1),s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f48,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),X1),s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),X1),s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),X1),s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num',X2),s(X0,X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ! [X9,X6,X10] : s(X8,'const.arithmetic.FUNPOW_3'(s(fun(X8,X8),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s(X8,X10))) = s(X8,happ(s(fun(X8,X8),X9),s(X8,'const.arithmetic.FUNPOW_3'(s(fun(X8,X8),X9),s('type.num.num',X6),s(X8,X10)))))),
  file('thm.exh_to_patProof.bind_bindn|split|0','thm.arithmetic.FUNPOW_SUC')).
fof(f61,plain,(
  ( ! [X0] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),X0))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s(fun('type.num.num',fun('type.num.num',bool)),X0)))) ) | $spl8),
  inference(cnf_transformation,[],[f61_D])).
fof(f61_D,plain,(
  ( ! [X0] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),X0))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s(fun('type.num.num',fun('type.num.num',bool)),X0)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f42,plain,(
  ( ! [X0] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),X0))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s(fun('type.num.num',fun('type.num.num',bool)),X0)))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),X0))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s(fun('type.num.num',fun('type.num.num',bool)),X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),X11))) = s(fun('type.num.num',fun('type.num.num',bool)),happ(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s(fun('type.num.num',fun('type.num.num',bool)),X11)))),
  file('thm.exh_to_patProof.bind_bindn|split|0','ARITY_0_1_const.exh_to_patProof.bind')).
fof(f1356,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) | $spl230),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f1348,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) | ($spl8 | $spl16)),
  inference(superposition,[],[f61,f65])).
fof(f1650,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) | $spl256),
  inference(cnf_transformation,[],[f1650_D])).
fof(f1650_D,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f1649,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2)))))) ) | ($spl230 | $spl252)),
  inference(forward_demodulation,[],[f1646,f1644])).
fof(f1646,plain,(
  ( ! [X2,X1] : (s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X1),s(fun('type.num.num',fun('type.num.num',bool)),X2))))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.FUNPOW_3'(s(fun(fun('type.num.num',fun('type.num.num',bool)),fun('type.num.num',fun('type.num.num',bool))),'const.exh_to_patProof.bind_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s(fun('type.num.num',fun('type.num.num',bool)),X2)))) ) | ($spl230 | $spl252)),
  inference(backward_demodulation,[],[f1644,f1356])).
fof(f54,plain,(
  s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',sK0),s(fun('type.num.num',fun('type.num.num',bool)),sK1))))) | $spl1),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',sK0),s(fun('type.num.num',fun('type.num.num',bool)),sK1))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f37,plain,(
  s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',sK0),s(fun('type.num.num',fun('type.num.num',bool)),sK1)))))),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',sK0),s(fun('type.num.num',fun('type.num.num',bool)),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f32])).
fof(f32,plain,(
  ? [X0,X1] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1)))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s(fun('type.num.num',fun('type.num.num',bool)),X1))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X0),s(fun('type.num.num',fun('type.num.num',bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s(fun('type.num.num',fun('type.num.num',bool)),X7))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X6),s(fun('type.num.num',fun('type.num.num',bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s(fun('type.num.num',fun('type.num.num',bool)),X7))) = s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num',X6),s(fun('type.num.num',fun('type.num.num',bool)),X7)))))),
  file('thm.exh_to_patProof.bind_bindn|split|0','thm.exh_to_patProof.bind_bindn|split|0')).
fof(f2594,plain,(
  $false | ($spl1 | $spl256)),
  inference(trivial_inequality_removal,[],[f2584])).
fof(f2584,plain,(
  s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) != s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s(fun('type.num.num',fun('type.num.num',bool)),sK1))) | ($spl1 | $spl256)),
  inference(superposition,[],[f54,f1650])).
% SZS output end Proof for thm.exh_to_patProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 41
% Passive clauses: 191
% Generated clauses: 2495
% Final active clauses: 37
% Final passive clauses: 129
% Input formulas: 14
% Initial clauses: 15
% 
% Duplicate literals: 218
% Trivial inequalities: 42
% Fw subsumption resolutions: 100
% Fw demodulations: 128
% Bw demodulations: 21
% Condensations: 311
% 
% Simple tautologies: 66
% Equational tautologies: 20
% Forward subsumptions: 1088
% 
% Binary resolution: 4
% Forward superposition: 813
% Backward superposition: 523
% Self superposition: 76
% Equality factoring: 46
% Equality resolution: 5
% 
% Split clauses: 38
% Split components: 84
% SAT solver clauses: 346
% SAT solver unit clauses: 26
% SAT solver binary clauses: 24
% SAT solver learnt clauses: 1
% 
% Sat splits: 38
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 345
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 3326
% Time elapsed: 0.097 s
% ------------------------------
% ------------------------------
% Success in time 0.697 s
