Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stack_removeProof
% (63141)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 574
% Passive clauses: 5857
% Generated clauses: 15165
% Final active clauses: 572
% Final passive clauses: 5015
% Input formulas: 68
% Initial clauses: 211
% 
% Duplicate literals: 273
% Trivial inequalities: 72
% Fw subsumption resolutions: 872
% Fw demodulations: 838
% Bw demodulations: 2
% 
% Simple tautologies: 221
% Equational tautologies: 61
% Forward subsumptions: 5821
% 
% Binary resolution: 6760
% Factoring: 1104
% Forward superposition: 1916
% Backward superposition: 2940
% Self superposition: 59
% Equality factoring: 26
% Equality resolution: 26
% 
% Split clauses: 224
% Split components: 448
% SAT solver clauses: 422
% SAT solver unit clauses: 13
% SAT solver binary clauses: 66
% 
% Sat splits: 224
% 
% TWLsolver clauses: 422
% TWLsolver calls for satisfiability: 278
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stack_removeProof
% (63162)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 414
% Passive clauses: 2706
% Generated clauses: 12129
% Final active clauses: 384
% Final passive clauses: 2102
% Input formulas: 68
% Initial clauses: 211
% 
% Duplicate literals: 125
% Trivial inequalities: 12
% Fw subsumption resolutions: 659
% Fw demodulations: 907
% Bw demodulations: 199
% Condensations: 160
% Global subsumptions: 18
% 
% Simple tautologies: 232
% Equational tautologies: 17
% Forward subsumptions: 4309
% 
% Binary resolution: 3685
% Factoring: 353
% Forward superposition: 1683
% Backward superposition: 1262
% Self superposition: 51
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 3104
% SAT solver unit clauses: 213
% SAT solver binary clauses: 99
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 11
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 3032
% TWLsolver calls for satisfiability: 1612
% 
% Memory used [KB]: 15863
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stack_removeProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stack_removeProof
% SZS output start Proof for thm.stack_removeProof
fof(f5251,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f5248])).
fof(f5248,plain,(
  s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11)),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)),sK10)),s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11))))) != s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11)),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)),sK10)),s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11)))))),
  inference(backward_demodulation,[],[f5246,f2833])).
fof(f2833,plain,(
  s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11)),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)),sK10)),s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11))))) != s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11)),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)),sK10)),s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)))))),
  inference(unit_resulting_resolution,[],[f320,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X8,X11] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X11) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb','thm.pred_set.EXTENSION')).
fof(f320,plain,(
  s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11))) != s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)))),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),sK12),s(sK10,sK11))) != s(fun(sK10,bool),'const.pred_set.DELETE_2'(s(fun(sK10,bool),'const.pred_set.INSERT_2'(s(sK10,sK11),s(fun(sK10,bool),sK12))),s(sK10,sK11)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f178])).
fof(f178,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))) != s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(X0,X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(X0,X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))),s(X6,X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))),s(X6,X7)))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb','thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb')).
fof(f5246,plain,(
  ( ! [X2,X0,X1] : (s(fun(X2,bool),'const.pred_set.DELETE_2'(s(fun(X2,bool),X0),s(X2,X1))) = s(fun(X2,bool),'const.pred_set.DELETE_2'(s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X1),s(fun(X2,bool),X0))),s(X2,X1)))) )),
  inference(forward_demodulation,[],[f5007,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X12,X13] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X12),s(X6,X13)))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f5007,plain,(
  ( ! [X2,X0,X1] : (s(fun(X2,bool),'const.pred_set.DELETE_2'(s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X1),s(fun(X2,bool),X0))),s(X2,X1))) = s(fun(X2,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X2,bool),'const.pred_set.DELETE_2'(s(fun(X2,bool),X0),s(X2,X1))),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X1),s(fun(X2,bool),'const.pred_set.DELETE_2'(s(fun(X2,bool),X0),s(X2,X1)))))))) )),
  inference(superposition,[],[f408,f792])).
fof(f792,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK14(X0,X1,X1,X2))) )),
  inference(unit_resulting_resolution,[],[f321,f758,f412])).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f758,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK14(X0,X1,X1,X2)))) )),
  inference(equality_resolution,[],[f407])).
fof(f407,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,sK14(X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK14(X3,X2,X1,X0))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,sK14(X3,X2,X1,X0)))) & s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,sK14(X3,X2,X1,X0)),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,X4))) & s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,X4),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,X4)))) & s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,X4),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(X0,X1) = s(X0,X2)) & s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,X4),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s(X0,X1) = s(X0,X2)) & s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,X4),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X9,X8] : ? [X10] : ((p(s(bool,X10)) <=> s(X6,X7) = s(X6,X9)) & s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))),s(X6,X9))) = s(fun(X6,bool),'const.bool.COND_3'(s(bool,X10),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X9))),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X9))))))))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb','thm.pred_set.DELETE_INSERT')).
fof(f321,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stack_removeProof.stack_removeProof_|49||48|6_prove_ltb',aHL_TRUTH)).
fof(f408,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))),s(X0,X2))) = s(fun(X0,bool),'const.bool.COND_3'(s(bool,sK14(X3,X2,X1,X0)),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2)))))))) )),
  inference(cnf_transformation,[],[f262])).
% SZS output end Proof for thm.stack_removeProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 269
% Passive clauses: 662
% Generated clauses: 4929
% Final active clauses: 265
% Final passive clauses: 387
% Input formulas: 68
% Initial clauses: 211
% 
% Duplicate literals: 6
% Trivial inequalities: 6
% Fw subsumption resolutions: 26
% Fw demodulations: 1438
% Bw demodulations: 10
% Condensations: 1
% 
% Simple tautologies: 74
% Equational tautologies: 18
% Forward subsumptions: 1947
% Fw demodulations to eq. taut.: 618
% 
% Unit resulting resolution: 1756
% Factoring: 1
% Forward superposition: 611
% Backward superposition: 818
% Self superposition: 35
% Equality factoring: 2
% Equality resolution: 9
% 
% SAT solver clauses: 664
% SAT solver unit clauses: 310
% SAT solver binary clauses: 191
% 
% Memory used [KB]: 6780
% Time elapsed: 0.185 s
% ------------------------------
% ------------------------------
% Success in time 1.51 s
