Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modSem
% (65376)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 165
% Passive clauses: 3600
% Generated clauses: 18829
% Final active clauses: 156
% Final passive clauses: 3275
% Input formulas: 69
% Initial clauses: 140
% 
% Duplicate literals: 631
% Trivial inequalities: 483
% Fw subsumption resolutions: 1945
% Fw demodulations: 2976
% Bw demodulations: 6
% 
% Simple tautologies: 497
% Equational tautologies: 49
% Forward subsumptions: 8342
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 307
% Factoring: 1
% Forward superposition: 6272
% Backward superposition: 5968
% Self superposition: 146
% Equality factoring: 88
% Equality resolution: 42
% 
% Split clauses: 201
% Split components: 414
% SAT solver clauses: 467
% SAT solver unit clauses: 11
% SAT solver binary clauses: 95
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 201
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 466
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modSem
% (65403)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1297
% Generated clauses: 14296
% Final active clauses: 124
% Final passive clauses: 1123
% Input formulas: 69
% Initial clauses: 140
% 
% Duplicate literals: 445
% Trivial inequalities: 306
% Fw subsumption resolutions: 984
% Fw demodulations: 1818
% Bw demodulations: 43
% Condensations: 738
% Global subsumptions: 8
% 
% Simple tautologies: 253
% Equational tautologies: 41
% Forward subsumptions: 6285
% 
% Binary resolution: 186
% Factoring: 1
% Forward superposition: 4225
% Backward superposition: 3887
% Self superposition: 116
% Equality factoring: 61
% Equality resolution: 23
% 
% Split clauses: 53
% Split components: 106
% SAT solver clauses: 2236
% SAT solver unit clauses: 150
% SAT solver binary clauses: 54
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 1
% 
% Sat splits: 53
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2203
% TWLsolver calls for satisfiability: 573
% 
% Memory used [KB]: 18166
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.modSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modSem
% SZS output start Proof for thm.modSem
fof(f29529,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f260,f421,f419,f259,f420,f418,f423,f1296,f1597,f1600,f1602,f1608,f1610,f1613,f1616,f1618,f1620,f1622,f1624,f1626,f1628,f1643,f1645,f1647,f1649,f1653,f1655,f1657,f1659,f1661,f1663,f1666,f1730,f1731,f1732,f1733,f4930,f4937,f5266,f5267,f5270,f5271,f6053,f6062,f6055,f6467,f6485,f6483,f8134,f8145,f8144,f11887,f11961,f12039,f12040,f12041,f12042,f12728,f12730,f12738,f12740,f12748,f12749,f12750,f12784,f12785,f12786,f12789,f12794,f12797,f12831,f12833,f12870,f12905,f12908,f12909,f12910,f12911,f14533,f14583,f18785,f18858,f17231,f16420,f4593,f1598,f4613,f6054,f19122,f19123,f1912,f1307,f1306,f23352,f23452,f11960,f29410,f29426,f29431,f29521])).
fof(f29521,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f29520,f4939])).
fof(f4939,plain,(
  ( ! [X0,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),sK1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),sK0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f421,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3)) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3] : ((s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3))) & (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3] : ((s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3))) & ((s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3)) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) <=> (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) & s('type.list.list'('type.modSem.v'),X1) = s('type.list.list'('type.modSem.v'),X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.modSem.v'),X17))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.modSem.v'),X19))) <=> (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18) & s('type.list.list'('type.modSem.v'),X17) = s('type.list.list'('type.modSem.v'),X19)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_|49||49||split|1')).
fof(f29520,plain,(
  ( ! [X569] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X569),s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X569),s('type.list.list'('type.modSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f29242,f332])).
fof(f332,plain,(
  ( ! [X0,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.modSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.list.list'('type.modSem.v'),X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8,X20] : s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.modSem.v'),X20))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8))),s('type.list.list'('type.modSem.v'),X20)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_0_2_const.modSem.Conv')).
fof(f29242,plain,(
  ( ! [X569] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X569),s('type.list.list'('type.modSem.v'),sK1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X569))),s('type.list.list'('type.modSem.v'),sK0)))) ) | $spl0),
  inference(superposition,[],[f332,f28808])).
fof(f28808,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X1),s('type.list.list'('type.modSem.v'),sK0))) = s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X1),s('type.list.list'('type.modSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f28369,f327])).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X7),s('type.list.list'('type.modSem.v'),X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X7),s('type.list.list'('type.modSem.v'),X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X7),s('type.list.list'('type.modSem.v'),X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.modSem.v'),X9),X15),s('type.list.list'('type.modSem.v'),X6))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_case_def|split|5')).
fof(f28369,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),X1),s('type.list.list'('type.modSem.v'),sK1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X1)))) ) | $spl0),
  inference(superposition,[],[f327,f418])).
fof(f29431,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f29430,f421])).
fof(f29430,plain,(
  s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1) | $spl0),
  inference(forward_demodulation,[],[f29086,f307])).
fof(f307,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X23] : ! [X25] : s(X23,X25) = s(X23,happ(s(fun(X23,X23),'_i_0_0'),s(X23,X25)))),
  file('thm.modSem.v_|49||49||split|5','_i0_ax')).
fof(f29086,plain,(
  s('type.list.list'('type.modSem.v'),sK1) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'('type.modSem.v'),'type.list.list'('type.modSem.v')),'_i_0_0'),s('type.list.list'('type.modSem.v'),sK0))) | $spl0),
  inference(superposition,[],[f307,f28808])).
fof(f29426,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f29425,f4939])).
fof(f29425,plain,(
  ( ! [X46] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.modSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f29083,f332])).
fof(f29083,plain,(
  ( ! [X46] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.modSem.v'),sK1))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46))),s('type.list.list'('type.modSem.v'),sK0)))) ) | $spl0),
  inference(superposition,[],[f28808,f332])).
fof(f29410,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f29409,f421])).
fof(f29409,plain,(
  s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1) | $spl0),
  inference(forward_demodulation,[],[f29059,f307])).
fof(f29059,plain,(
  s('type.list.list'('type.modSem.v'),sK1) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'('type.modSem.v'),'type.list.list'('type.modSem.v')),'_i_0_0'),s('type.list.list'('type.modSem.v'),sK0))) | $spl0),
  inference(superposition,[],[f28808,f307])).
fof(f11960,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) | $spl30),
  inference(cnf_transformation,[],[f11960_D])).
fof(f11960_D,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,X79)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f23452,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) | $spl44),
  inference(cnf_transformation,[],[f23452_D])).
fof(f23452_D,plain,(
  ( ! [X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f23352,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(duplicate_literal_removal,[],[f23176])).
fof(f23176,plain,(
  ( ! [X78,X79] : (s(bool,X79) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f18934,f19509])).
fof(f19509,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(trivial_inequality_removal,[],[f19494])).
fof(f19494,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl34)),
  inference(superposition,[],[f18932,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f268,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.modSem.v_|49||49||split|5',aHL_BOOL_CASES)).
fof(f18932,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26 | $spl34)),
  inference(global_subsumption,[],[f18875,f260,f259,f261,f263,f265,f266,f267,f268,f270,f269,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f487,f753,f783,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1407,f1303,f1302,f1301,f1300,f1299,f1437,f1298,f1297,f318,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2127,f2128,f2129,f2130,f2131,f2458,f2491,f2589,f2671,f3290,f3504,f3974,f4043,f4044,f4367,f4368,f4521,f4452,f4525,f4595,f4681,f4725,f4724,f4718,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5253,f5264,f5266,f5268,f5269,f5273,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5287,f5288,f5298,f5299,f5557,f5438,f5447,f5453,f5476,f5563,f5516,f5525,f5531,f6053,f6062,f6195,f6343,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6572,f6573,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6601,f6602,f6613,f6614,f6615,f6618,f6619,f6630,f6631,f6632,f6840,f7559,f7560,f8101,f8134,f8144,f8230,f8231,f8241,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8441,f8444,f8446,f8448,f8449,f8450,f8451,f8452,f8803,f8804,f8861,f8864,f8866,f8868,f8869,f8870,f8871,f8872,f9181,f9346,f9347,f9279,f9285,f9550,f9510,f9513,f9515,f9516,f9517,f9518,f9519,f9712,f9853,f9813,f9816,f9818,f9819,f9820,f9821,f9822,f10019,f10086,f10314,f10475,f10870,f10998,f11003,f11281,f11250,f11399,f11400,f11517,f11617,f11620,f11623,f11625,f11626,f11627,f11628,f11629,f11630,f12039,f12041,f7353,f7368,f8135,f8140,f12045,f264,f273,f274,f465,f507,f535,f615,f664,f717,f829,f1290,f1398,f2291,f2339,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2428,f2496,f2497,f2498,f2781,f2783,f2785,f2817,f2819,f2833,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3598,f3940,f3945,f3947,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4050,f4051,f4054,f4062,f4063,f4067,f4070,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4374,f4375,f4417,f4418,f4428,f4430,f4431,f4435,f4436,f4446,f4519,f4618,f4678,f4700,f4770,f4827,f5256,f5263,f5265,f5277,f5286,f5294,f5295,f5404,f5405,f5415,f5417,f5418,f5422,f5423,f5433,f5483,f5484,f5494,f5496,f5497,f5501,f5502,f5512,f5551,f6059,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6353,f6371,f6379,f6380,f6412,f6416,f6420,f6570,f6571,f6579,f6580,f6600,f6608,f6609,f6616,f6617,f6625,f6626,f6645,f6659,f6663,f7282,f7290,f7291,f7296,f7304,f7305,f7337,f7342,f7349,f7371,f12912,f7563,f8104,f9165,f9173,f9174,f9213,f9215,f9216,f9220,f9221,f9233,f9308,f9554,f9621,f9858,f9925,f10291,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10819,f10820,f10880,f10993,f11061,f11287,f4056,f4055,f1288,f6372,f6373,f6425,f1286,f2400,f2402,f2455,f5247,f9514,f9817,f11245,f316,f319,f320,f348,f349,f1596,f4451,f4520,f4536,f4615,f5964,f6043,f6044,f6046,f6048,f6049,f6050,f6051,f6056,f6360,f6361,f6370,f6399,f6401,f6402,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6417,f9355,f9617,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9880,f9881,f9891,f9909,f9910,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9921,f9922,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10246,f10247,f10257,f10275,f10276,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10287,f10288,f10676,f347,f350,f5402,f5403,f5406,f5408,f5409,f5414,f5416,f5477,f5478,f5481,f5482,f5485,f5487,f5488,f5493,f5495,f5513,f6031,f9509,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9606,f9607,f9610,f9612,f9613,f9614,f9615,f9618,f9812,f10140,f10368,f10529,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10665,f10666,f10669,f10671,f10672,f10673,f10674,f10677,f11240,f5479,f11035,f7325,f10670,f9611,f10635,f9576,f10636,f9577,f10668,f9609,f10664,f9605,f10667,f9608,f10646,f9587,f5400,f6030,f5954,f5955,f6042,f6045,f9354,f6047,f309,f11039,f11041,f11047,f11051,f11053,f11055,f11057,f11058,f412,f2862,f2863,f2865,f2867,f2868,f2869,f2870,f2871,f11017,f4454,f6032,f6063,f6426,f7283,f7284,f7297,f7298,f341,f7327,f7328,f7330,f7332,f7333,f7334,f7335,f7346,f7281,f9196,f9166,f9167,f9197,f9198,f9234,f9325,f9512,f9815,f11243,f5448,f5446,f5526,f5524,f9195,f11006,f11007,f11045,f5451,f5529,f9357,f7271,f7272,f9358,f11049,f2866,f7331,f11037,f5191,f5399,f5434,f9508,f9811,f10141,f10369,f10530,f11239,f7326,f11043,f2864,f7329,f11034,f1582,f398,f351,f352,f356,f357,f358,f4631,f4641,f4675,f4780,f4790,f4824,f7510,f7523,f8051,f8064,f5441,f5443,f7513,f8054,f5519,f5521,f7514,f8055,f5440,f5442,f7500,f8041,f5518,f5520,f7501,f8042,f1394,f1396,f1601,f2126,f5245,f9511,f9814,f11242,f317,f746,f747,f748,f749,f750,f751,f850,f4378,f4379,f4381,f4383,f4384,f4385,f4386,f4387,f4453,f4523,f4524,f10161,f10389,f10550,f4382,f2436,f5528,f5450,f9267,f9263,f5168,f5165,f4380,f1425,f4415,f4416,f4419,f4421,f4422,f8261,f6599,f6566,f4449,f4427,f4429,f5252,f10151,f10379,f10540,f4695,f4594,f5243,f5244,f5246,f5248,f5249,f5250,f5251,f13097,f13098,f14346,f14349,f14352,f14354,f14355,f14356,f14357,f14358])).
fof(f14358,plain,(
  ( ! [X132,X133,X131] : (s(bool,X131) = s(bool,sK7('const.pair.FST_0',X133,bool)) | s(bool,'T_0') = s(bool,sK4(X131,X132,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X132)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f753])).
fof(f14357,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,X129) = s(bool,sK7('const.pair.FST_0',X130,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,X129)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f344])).
fof(f14356,plain,(
  ( ! [X125,X126,X124] : (s(bool,X125) = s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,X125)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f338])).
fof(f14355,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X121) = s(bool,sK7('const.pair.FST_0',X123,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,X121),s(X120,X122)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f343])).
fof(f14354,plain,(
  ( ! [X118,X119,X117] : (s(bool,X118) = s(bool,sK7('const.pair.FST_0',X119,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X117,bool),'const.sum.INR_1'(s(bool,X118)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f339])).
fof(f14352,plain,(
  ( ! [X111,X112,X110] : (s(bool,X110) = s(bool,sK7('const.pair.FST_0',X112,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.|48|_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f314])).
fof(f14349,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,sK7('const.pair.FST_0',X100,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X97),bool),sK18(X98,X99,X97,bool)),s('type.list.list'(X97),'const.list.NIL_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f359])).
fof(f14346,plain,(
  ( ! [X83,X84] : (s(bool,X83) = s(bool,sK7('const.pair.FST_0',X84,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X83)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f13097,f307])).
fof(f13098,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f13045,f264])).
fof(f13045,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f269,f12912])).
fof(f13097,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f13038,f424])).
fof(f13038,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f12912,f437])).
fof(f5251,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4695,f344])).
fof(f5250,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,X138)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X139),bool),sK18(X140,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X138)))),X139,bool))))) )),
  inference(superposition,[],[f4695,f338])).
fof(f5249,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4695,f343])).
fof(f5248,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,X129)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK18(X131,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X129)))),X130,bool))))) )),
  inference(superposition,[],[f4695,f339])).
fof(f5246,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK18(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) )),
  inference(superposition,[],[f4695,f314])).
fof(f5244,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK18(X114,happ(s(fun('type.list.list'(X110),bool),sK18(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4695,f359])).
fof(f5243,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4695,f307])).
fof(f4594,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK18(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4520,f359])).
fof(f4695,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4683,f264])).
fof(f4683,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4594,f268])).
fof(f10540,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2130,f753])).
fof(f10379,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') != s(bool,X161) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | $spl10),
  inference(superposition,[],[f2128,f753])).
fof(f10151,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | $spl10),
  inference(superposition,[],[f2124,f753])).
fof(f5252,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK18(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4695,f753])).
fof(f4429,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f487,f753])).
fof(f4427,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f439,f753])).
fof(f4449,plain,(
  ( ! [X156,X154,X152,X151,X155,X153] : (s(fun(X151,X152),X153) != s(fun(X151,X152),X155) | s(X152,happ(s(fun(X151,X152),X155),s(X151,X156))) = s(X152,happ(s(fun(X151,X152),sK4(X153,X154,fun(X151,X152))),s(X151,X156))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,X152),bool),X154)))) )),
  inference(superposition,[],[f354,f753])).
fof(f6566,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f355,f753])).
fof(f6599,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK15(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f355,f753])).
fof(f8261,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X64),s('type.list.list'('type.modSem.v'),X66))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X67),s('type.list.list'('type.modSem.v'),X68))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X67) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),sK4(X64,X65,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),bool),X65)))) )),
  inference(superposition,[],[f394,f753])).
fof(f4422,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f410,f753])).
fof(f4421,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f337,f753])).
fof(f4419,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f334,f753])).
fof(f4416,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f270,f753])).
fof(f4415,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f269,f753])).
fof(f1425,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f317,f340])).
fof(f4380,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f753,f340])).
fof(f5165,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X58,bool),'_dst_x0x1_2'(s(fun(X59,bool),X61),s(fun(X58,X59),X60))))) | s(X58,X62) = s(X58,sK4(X62,happ(s(fun(fun(X59,bool),fun(X58,bool)),'_dstsel_0_x01_1'(s(fun(X58,X59),X60))),s(fun(X59,bool),X61)),X58))) )),
  inference(superposition,[],[f753,f368])).
fof(f5168,plain,(
  ( ! [X76,X74,X77,X75,X73] : (p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X74,bool),X76),s(fun(X73,X74),X75)))))) | s(X73,X77) = s(X73,sK4(X77,happ(s(fun(fun(X74,bool),fun(X73,bool)),'_dstsel_0_x01_1'(s(fun(X73,X74),X75))),s(fun(X74,bool),X76)),X73))) )),
  inference(superposition,[],[f317,f368])).
fof(f9263,plain,(
  ( ! [X557,X559,X553,X555,X556,X558,X554] : (p(s(bool,'|dollar|exists_1'(s(fun(X553,bool),happ(s(fun(X555,fun(X553,bool)),happ(s(fun(X554,fun(X555,fun(X553,bool))),X556),s(X554,X557))),s(X555,X558)))))) | s(X553,X559) = s(X553,sK4(X559,happ(s(fun('type.pair.prod'(X554,X555),fun(X553,bool)),sK19(X556,X555,X554,fun(X553,bool))),s('type.pair.prod'(X554,X555),'const.pair.,_2'(s(X554,X557),s(X555,X558)))),X553))) )),
  inference(superposition,[],[f317,f367])).
fof(f9267,plain,(
  ( ! [X588,X584,X586,X589,X585,X587,X583] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X583,bool),happ(s(fun(X585,fun(X583,bool)),happ(s(fun(X584,fun(X585,fun(X583,bool))),X586),s(X584,X587))),s(X585,X588))))) | s(X583,X589) = s(X583,sK4(X589,happ(s(fun('type.pair.prod'(X584,X585),fun(X583,bool)),sK19(X586,X585,X584,fun(X583,bool))),s('type.pair.prod'(X584,X585),'const.pair.,_2'(s(X584,X587),s(X585,X588)))),X583))) )),
  inference(superposition,[],[f753,f367])).
fof(f5450,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X425] : (p(s(bool,'|dollar|exists_1'(s(fun(X420,bool),happ(s(fun(X423,fun(X420,bool)),'_dst_x0x1_2'(s(fun(X421,fun(X420,bool)),X422),s(fun(X423,X421),X424))),s(X423,X425)))))) | s(X420,X426) = s(X420,sK4(X426,happ(s(fun(X421,fun(X420,bool)),X422),s(X421,happ(s(fun(X423,X421),X424),s(X423,X425)))),X420))) )),
  inference(superposition,[],[f317,f369])).
fof(f5528,plain,(
  ( ! [X412,X414,X410,X416,X413,X415,X411] : (p(s(bool,'|dollar|exists_1'(s(fun(X410,bool),happ(s(fun(X412,fun(X410,bool)),X413),s(X412,happ(s(fun(X411,X412),X414),s(X411,X415)))))))) | s(X410,X416) = s(X410,sK4(X416,happ(s(fun(X411,fun(X410,bool)),'_dst_x0x1_2'(s(fun(X412,fun(X410,bool)),X413),s(fun(X411,X412),X414))),s(X411,X415)),X410))) )),
  inference(superposition,[],[f317,f369])).
fof(f2436,plain,(
  ( ! [X134,X132,X135,X133,X131] : (p(s(bool,'|dollar|exists_1'(s(fun(X131,bool),happ(s(fun('type.num.num',fun(X131,bool)),X133),s('type.num.num',X134)))))) | s(X131,X135) = s(X131,sK4(X135,'const.ind_type.FCONS_3'(s(fun(X131,bool),X132),s(fun('type.num.num',fun(X131,bool)),X133),s('type.num.num','const.num.SUC_1'(s('type.num.num',X134)))),X131))) )),
  inference(superposition,[],[f317,f323])).
fof(f4382,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))) | s(X16,X20) = s(X16,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))) )),
  inference(superposition,[],[f753,f323])).
fof(f10550,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10510])).
fof(f10510,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2130,f753])).
fof(f10389,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10349])).
fof(f10349,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X70,bool),X71))))))) | s(X70,X73) = s(X70,sK4(X73,X71,X70))) ) | $spl10),
  inference(superposition,[],[f2128,f753])).
fof(f10161,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10121])).
fof(f10121,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44))))) | s(X43,X45) = s(X43,sK4(X45,X44,X43))) ) | $spl10),
  inference(superposition,[],[f2124,f753])).
fof(f4524,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4472,f261])).
fof(f4472,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f410,f753])).
fof(f4523,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4471,f261])).
fof(f4471,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f337,f753])).
fof(f4453,plain,(
  ( ! [X173,X171,X174,X172] : (p(s(bool,'|dollar|exists_1'(s(fun(X171,bool),X172)))) | s(X171,X174) = s(X171,sK4(X174,sK4(X172,X173,fun(X171,bool)),X171)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X171,bool),bool),X173)))) )),
  inference(superposition,[],[f317,f753])).
fof(f4387,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f753,f753])).
fof(f4386,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X37))) | s(X34,X38) = s(X34,sK4(X38,'const.pair.SND_1'(s('type.pair.prod'(X35,fun(X34,bool)),'const.pair.,_2'(s(X35,X36),s(fun(X34,bool),X37)))),X34))) )),
  inference(superposition,[],[f753,f344])).
fof(f4385,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X32))) | s(X30,X33) = s(X30,sK4(X33,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X30,bool),X31),'const.sum.INL_1'(s(fun(X30,bool),X32)))),X30))) )),
  inference(superposition,[],[f753,f338])).
fof(f4384,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f753,f343])).
fof(f4383,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTR_1'(s('type.sum.sum'(X22,fun(X21,bool)),'const.sum.INR_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f753,f339])).
fof(f4381,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f753,f314])).
fof(f4379,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f753,f359])).
fof(f4378,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f753,f307])).
fof(f850,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK18(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f317,f359])).
fof(f751,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X23)))) | s(X20,X24) = s(X20,sK4(X24,'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))) )),
  inference(superposition,[],[f317,f344])).
fof(f750,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | s(X16,X19) = s(X16,sK4(X19,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X16,bool),X17),'const.sum.INL_1'(s(fun(X16,bool),X18)))),X16))) )),
  inference(superposition,[],[f317,f338])).
fof(f749,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f317,f343])).
fof(f748,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X9)))) | s(X7,X10) = s(X7,sK4(X10,'const.sum.OUTR_1'(s('type.sum.sum'(X8,fun(X7,bool)),'const.sum.INR_1'(s(fun(X7,bool),X9)))),X7))) )),
  inference(superposition,[],[f317,f339])).
fof(f747,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X6) = s(X3,sK4(X6,'const.ind_type.FCONS_3'(s(fun(X3,bool),X4),s(fun('type.num.num',fun(X3,bool)),X5),s('type.num.num','const.num.|48|_0')),X3))) )),
  inference(superposition,[],[f317,f314])).
fof(f746,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f317,f307])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : ! [X27,X28] : (! [X33] : (s(X9,X28) = s(X9,X33) => p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X27)))))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.JRH_INDUCT_UTIL')).
fof(f11242,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) ) | $spl10),
  inference(superposition,[],[f10870,f340])).
fof(f9814,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f6343,f340])).
fof(f9511,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)))) )),
  inference(superposition,[],[f6195,f340])).
fof(f5245,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK18(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4695,f340])).
fof(f2126,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(forward_demodulation,[],[f2091,f1756])).
fof(f2091,plain,(
  ( ! [X23,X22] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X22),X23)))) ) | $spl10),
  inference(superposition,[],[f1849,f340])).
fof(f1601,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.min.@_1'(s(fun('type.pair.prod'(bool,X23),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),X24))))) )),
  inference(forward_demodulation,[],[f1583,f340])).
fof(f1583,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),X24)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X23),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X23),'const.min.@_1'(s(fun('type.pair.prod'(bool,X23),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f316,f340])).
fof(f1396,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f340,f437])).
fof(f1394,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f340,f437])).
fof(f8042,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7501,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f351,f369])).
fof(f5520,plain,(
  ( ! [X350,X346,X352,X354,X349,X351,X347,X353,X348] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X346,X347),happ(s(fun(X349,fun(X346,X347)),X350),s(X349,happ(s(fun(X348,X349),X351),s(X348,X352))))),s(fun(X347,X346),X353)))) | s(X346,X354) = s(X346,happ(s(fun(X347,X346),X353),s(X347,happ(s(fun(X346,X347),happ(s(fun(X348,fun(X346,X347)),'_dst_x0x1_2'(s(fun(X349,fun(X346,X347)),X350),s(fun(X348,X349),X351))),s(X348,X352))),s(X346,X354)))))) )),
  inference(superposition,[],[f357,f369])).
fof(f5518,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X328,X330,X336] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X328,X329),happ(s(fun(X331,fun(X328,X329)),X332),s(X331,happ(s(fun(X330,X331),X333),s(X330,X334))))),s(fun(X329,X328),X335)))) | s(X329,X336) = s(X329,happ(s(fun(X328,X329),happ(s(fun(X330,fun(X328,X329)),'_dst_x0x1_2'(s(fun(X331,fun(X328,X329)),X332),s(fun(X330,X331),X333))),s(X330,X334))),s(X328,happ(s(fun(X329,X328),X335),s(X329,X336)))))) )),
  inference(superposition,[],[f356,f369])).
fof(f8041,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7500,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f351,f369])).
fof(f5442,plain,(
  ( ! [X356,X358,X364,X360,X362,X357,X359,X361,X363] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X356,X357),happ(s(fun(X360,fun(X356,X357)),'_dst_x0x1_2'(s(fun(X358,fun(X356,X357)),X359),s(fun(X360,X358),X361))),s(X360,X362))),s(fun(X357,X356),X363)))) | s(X356,X364) = s(X356,happ(s(fun(X357,X356),X363),s(X357,happ(s(fun(X356,X357),happ(s(fun(X358,fun(X356,X357)),X359),s(X358,happ(s(fun(X360,X358),X361),s(X360,X362))))),s(X356,X364)))))) )),
  inference(superposition,[],[f357,f369])).
fof(f5440,plain,(
  ( ! [X344,X346,X341,X343,X339,X345,X340,X342,X338] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X338,X339),happ(s(fun(X342,fun(X338,X339)),'_dst_x0x1_2'(s(fun(X340,fun(X338,X339)),X341),s(fun(X342,X340),X343))),s(X342,X344))),s(fun(X339,X338),X345)))) | s(X339,X346) = s(X339,happ(s(fun(X338,X339),happ(s(fun(X340,fun(X338,X339)),X341),s(X340,happ(s(fun(X342,X340),X343),s(X342,X344))))),s(X338,happ(s(fun(X339,X338),X345),s(X339,X346)))))) )),
  inference(superposition,[],[f356,f369])).
fof(f8055,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7514,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f351,f369])).
fof(f5521,plain,(
  ( ! [X356,X358,X360,X362,X357,X359,X355,X361,X363] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X356,X355),X362),s(fun(X355,X356),happ(s(fun(X358,fun(X355,X356)),X359),s(X358,happ(s(fun(X357,X358),X360),s(X357,X361)))))))) | s(X356,X363) = s(X356,happ(s(fun(X355,X356),happ(s(fun(X357,fun(X355,X356)),'_dst_x0x1_2'(s(fun(X358,fun(X355,X356)),X359),s(fun(X357,X358),X360))),s(X357,X361))),s(X355,happ(s(fun(X356,X355),X362),s(X356,X363)))))) )),
  inference(superposition,[],[f357,f369])).
fof(f5519,plain,(
  ( ! [X344,X341,X343,X337,X339,X345,X340,X342,X338] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X338,X337),X344),s(fun(X337,X338),happ(s(fun(X340,fun(X337,X338)),X341),s(X340,happ(s(fun(X339,X340),X342),s(X339,X343)))))))) | s(X337,X345) = s(X337,happ(s(fun(X338,X337),X344),s(X338,happ(s(fun(X337,X338),happ(s(fun(X339,fun(X337,X338)),'_dst_x0x1_2'(s(fun(X340,fun(X337,X338)),X341),s(fun(X339,X340),X342))),s(X339,X343))),s(X337,X345)))))) )),
  inference(superposition,[],[f356,f369])).
fof(f8054,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f352,f369])).
fof(f7513,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f351,f369])).
fof(f5443,plain,(
  ( ! [X366,X372,X368,X370,X365,X367,X373,X369,X371] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X366,X365),X372),s(fun(X365,X366),happ(s(fun(X369,fun(X365,X366)),'_dst_x0x1_2'(s(fun(X367,fun(X365,X366)),X368),s(fun(X369,X367),X370))),s(X369,X371)))))) | s(X366,X373) = s(X366,happ(s(fun(X365,X366),happ(s(fun(X367,fun(X365,X366)),X368),s(X367,happ(s(fun(X369,X367),X370),s(X369,X371))))),s(X365,happ(s(fun(X366,X365),X372),s(X366,X373)))))) )),
  inference(superposition,[],[f357,f369])).
fof(f5441,plain,(
  ( ! [X350,X352,X354,X349,X351,X347,X353,X355,X348] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X348,X347),X354),s(fun(X347,X348),happ(s(fun(X351,fun(X347,X348)),'_dst_x0x1_2'(s(fun(X349,fun(X347,X348)),X350),s(fun(X351,X349),X352))),s(X351,X353)))))) | s(X347,X355) = s(X347,happ(s(fun(X348,X347),X354),s(X348,happ(s(fun(X347,X348),happ(s(fun(X349,fun(X347,X348)),X350),s(X349,happ(s(fun(X351,X349),X352),s(X351,X353))))),s(X347,X355)))))) )),
  inference(superposition,[],[f356,f369])).
fof(f8064,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f352,f753])).
fof(f8051,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f352,f753])).
fof(f7523,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f351,f753])).
fof(f7510,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f351,f753])).
fof(f4824,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4793,f261])).
fof(f4793,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f357,f437])).
fof(f4790,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f357,f753])).
fof(f4780,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f357,f753])).
fof(f4675,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4644,f261])).
fof(f4644,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f356,f437])).
fof(f4641,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f356,f753])).
fof(f4631,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f356,f753])).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK16(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK16(X3,X2,X1,X0)))))) | s(X1,sK17(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK17(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X31,X9] : ! [X10,X32] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X32)))) <=> (! [X33] : s(X31,X33) = s(X31,happ(s(fun(X9,X31),X10),s(X9,happ(s(fun(X31,X9),X32),s(X31,X33))))) & ! [X45] : s(X9,X45) = s(X9,happ(s(fun(X31,X9),X32),s(X31,happ(s(fun(X9,X31),X10),s(X9,X45)))))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.ISO')).
fof(f357,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f356,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f352,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9,X31] : ! [X32,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X31),X10),s(fun(X31,X9),X32)))) => (! [X27] : (! [X33] : p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33)))) <=> ! [X33] : p(s(bool,happ(s(fun(X9,bool),X27),s(X9,happ(s(fun(X31,X9),X32),s(X31,X33))))))) & ! [X27] : (? [X33] : p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33)))) <=> ? [X33] : p(s(bool,happ(s(fun(X9,bool),X27),s(X9,happ(s(fun(X31,X9),X32),s(X31,X33))))))) & ! [X6,X34] : (s(X9,X6) = s(X9,happ(s(fun(X31,X9),X32),s(X31,X34))) <=> s(X31,X34) = s(X31,happ(s(fun(X9,X31),X10),s(X9,X6))))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.ISO_USAGE')).
fof(f351,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f242])).
fof(f398,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7))))))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X31,X46,X9,X47] : ! [X48,X32,X49,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X47),X10),s(fun(X47,X9),X49)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X31,X46),X32),s(fun(X46,X31),X48))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X31),fun(X47,X46)),'_dst_x00x11_2'(s(fun(X31,X46),X32),s(fun(fun(X9,X31),fun(X47,X31)),'_dstsel_0_x01_1'(s(fun(X47,X9),X49))))),s(fun(fun(X47,X46),fun(X9,X31)),'_dst_x00x11_2'(s(fun(X46,X31),X48),s(fun(fun(X47,X46),fun(X9,X46)),'_dstsel_0_x01_1'(s(fun(X9,X47),X10)))))))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.ISO_FUN')).
fof(f1582,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun('type.list.list'(X20),bool),sK18(X21,X22,X20,bool)),s('type.list.list'(X20),'const.min.@_1'(s(fun('type.list.list'(X20),bool),sK18(X21,X22,X20,bool)))))))) )),
  inference(superposition,[],[f316,f359])).
fof(f11034,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | ~p(s(bool,happ(s(fun('type.list.list'(X59),bool),sK18(X60,X61,X59,bool)),s('type.list.list'(X59),'const.list.CONS_2'(s(X59,sK3(sK18(X60,X61,X59,bool),X59)),s('type.list.list'(X59),sK2(sK18(X60,X61,X59,bool),X59))))))) | p(s(bool,happ(s(fun('type.list.list'(X59),bool),sK18(X60,X61,X59,bool)),s('type.list.list'(X59),X62))))) )),
  inference(superposition,[],[f309,f359])).
fof(f7329,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)))))))) )),
  inference(forward_demodulation,[],[f7274,f340])).
fof(f7274,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK7(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32)),s(X33,sK8(happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35)),X33,X32))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),happ(s(fun('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),fun('type.pair.prod'(X32,X33),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.pair.prod'(X32,X33),bool),X34),X35))),s('type.pair.prod'(X32,X33),X36))))) )),
  inference(superposition,[],[f341,f340])).
fof(f2864,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f2848,f340])).
fof(f2848,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f340])).
fof(f11043,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30)))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11042,f340])).
fof(f11042,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30))))) )),
  inference(forward_demodulation,[],[f11009,f340])).
fof(f11009,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),'const.list.CONS_2'(s(X27,sK3(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27)),s('type.list.list'(X27),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29)),X27))))))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X27),bool),X28),fun('type.list.list'(X27),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X27),bool),X28),X29))),s('type.list.list'(X27),X30))))) )),
  inference(superposition,[],[f309,f340])).
fof(f7326,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),X5)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f7269,f368])).
fof(f7269,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0x1_2'(s(fun(X2,bool),X4),s(fun('type.pair.prod'(X0,X1),X2),X3))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0)),s(X1,sK8(happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4)),X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun(X2,bool),fun('type.pair.prod'(X0,X1),bool)),'_dstsel_0_x01_1'(s(fun('type.pair.prod'(X0,X1),X2),X3))),s(fun(X2,bool),X4))),s('type.pair.prod'(X0,X1),X5))))) )),
  inference(superposition,[],[f341,f368])).
fof(f11239,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X72,bool),X73),s(fun(X74,X72),X75))),s(X74,X76))))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76)))))) ) | $spl10),
  inference(superposition,[],[f10870,f369])).
fof(f10530,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X126),'const.sum.INL_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2130,f369])).
fof(f10369,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (s(bool,'T_0') != s(bool,happ(s(fun(X122,bool),X123),s(X122,happ(s(fun(X121,X122),X124),s(X121,X125))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X121,bool),'_dst_x0x1_2'(s(fun(X122,bool),X123),s(fun(X121,X122),X124))),s(X121,X125)))))))) ) | $spl10),
  inference(superposition,[],[f2128,f369])).
fof(f10141,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun(X75,bool),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X75,bool),X76),s(fun(X74,X75),X77))),s(X74,X78)))))) ) | $spl10),
  inference(superposition,[],[f2124,f369])).
fof(f9811,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6343,f369])).
fof(f9508,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f6195,f369])).
fof(f5434,plain,(
  ( ! [X292,X294,X296,X293,X295,X291,X297] : (~p(s(bool,happ(s(fun(X293,bool),'_dst_x0x1_2'(s(fun(X291,bool),X292),s(fun(X293,X291),X294))),s(X293,X295)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X296),bool),sK18(X297,happ(s(fun(X291,bool),X292),s(X291,happ(s(fun(X293,X291),X294),s(X293,X295)))),X296,bool))))) )),
  inference(superposition,[],[f4695,f369])).
fof(f5399,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f4520,f369])).
fof(f5191,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,happ(s(fun('type.list.list'(X78),bool),'_dst_x0x1_2'(s(fun(X79,bool),X81),s(fun('type.list.list'(X78),X79),X80))),s('type.list.list'(X78),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),'_dst_x0x1_2'(s(fun(X79,bool),X81),s(fun('type.list.list'(X78),X79),X80))),s('type.list.list'(X78),sK2(happ(s(fun(fun(X79,bool),fun('type.list.list'(X78),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X78),X79),X80))),s(fun(X79,bool),X81)),X78)))))) )),
  inference(forward_demodulation,[],[f5169,f368])).
fof(f5169,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,happ(s(fun('type.list.list'(X78),bool),'_dst_x0x1_2'(s(fun(X79,bool),X81),s(fun('type.list.list'(X78),X79),X80))),s('type.list.list'(X78),sK2(happ(s(fun(fun(X79,bool),fun('type.list.list'(X78),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X78),X79),X80))),s(fun(X79,bool),X81)),X78))))) | ~p(s(bool,happ(s(fun('type.list.list'(X78),bool),happ(s(fun(fun(X79,bool),fun('type.list.list'(X78),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X78),X79),X80))),s(fun(X79,bool),X81))),s('type.list.list'(X78),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f368])).
fof(f11037,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),X4)))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11036,f368])).
fof(f11036,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),X4))))) )),
  inference(forward_demodulation,[],[f11004,f368])).
fof(f11004,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun('type.list.list'(X0),X1),X2))),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0)),s('type.list.list'(X0),sK2(happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3)),X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun(X1,bool),fun('type.list.list'(X0),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X0),X1),X2))),s(fun(X1,bool),X3))),s('type.list.list'(X0),X4))))) )),
  inference(superposition,[],[f309,f368])).
fof(f7331,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),X47)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)))))))) )),
  inference(forward_demodulation,[],[f7276,f323])).
fof(f7276,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num',X46))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42)),s(X43,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X43,X42))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X42,X43),bool),X44),s(fun('type.num.num',fun('type.pair.prod'(X42,X43),bool)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))),s('type.pair.prod'(X42,X43),X47))))) )),
  inference(superposition,[],[f341,f323])).
fof(f2866,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f2850,f323])).
fof(f2850,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f323])).
fof(f11049,plain,(
  ( ! [X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),X46)))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11048,f323])).
fof(f11048,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),X46))))) )),
  inference(forward_demodulation,[],[f11012,f323])).
fof(f11012,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun('type.list.list'(X42),bool),happ(s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num',X45))),s('type.list.list'(X42),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),'const.list.CONS_2'(s(X42,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)),s('type.list.list'(X42),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))))) | p(s(bool,happ(s(fun('type.list.list'(X42),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X42),bool),X43),s(fun('type.num.num',fun('type.list.list'(X42),bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.list.list'(X42),X46))))) )),
  inference(superposition,[],[f309,f323])).
fof(f9358,plain,(
  ( ! [X606,X610,X609,X605,X607,X603,X608,X604] : (p(s(bool,happ(s(fun('type.pair.prod'(X603,X604),bool),happ(s(fun(X606,fun('type.pair.prod'(X603,X604),bool)),happ(s(fun(X605,fun(X606,fun('type.pair.prod'(X603,X604),bool))),X607),s(X605,X608))),s(X606,X609))),s('type.pair.prod'(X603,X604),X610)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X603,X604),bool),happ(s(fun(X606,fun('type.pair.prod'(X603,X604),bool)),happ(s(fun(X605,fun(X606,fun('type.pair.prod'(X603,X604),bool))),X607),s(X605,X608))),s(X606,X609))),s('type.pair.prod'(X603,X604),'const.pair.,_2'(s(X603,sK7(happ(s(fun('type.pair.prod'(X605,X606),fun('type.pair.prod'(X603,X604),bool)),sK19(X607,X606,X605,fun('type.pair.prod'(X603,X604),bool))),s('type.pair.prod'(X605,X606),'const.pair.,_2'(s(X605,X608),s(X606,X609)))),X604,X603)),s(X604,sK8(happ(s(fun('type.pair.prod'(X605,X606),fun('type.pair.prod'(X603,X604),bool)),sK19(X607,X606,X605,fun('type.pair.prod'(X603,X604),bool))),s('type.pair.prod'(X605,X606),'const.pair.,_2'(s(X605,X608),s(X606,X609)))),X604,X603)))))))) )),
  inference(forward_demodulation,[],[f9270,f367])).
fof(f9270,plain,(
  ( ! [X606,X610,X609,X605,X607,X603,X608,X604] : (~p(s(bool,happ(s(fun('type.pair.prod'(X603,X604),bool),happ(s(fun(X606,fun('type.pair.prod'(X603,X604),bool)),happ(s(fun(X605,fun(X606,fun('type.pair.prod'(X603,X604),bool))),X607),s(X605,X608))),s(X606,X609))),s('type.pair.prod'(X603,X604),'const.pair.,_2'(s(X603,sK7(happ(s(fun('type.pair.prod'(X605,X606),fun('type.pair.prod'(X603,X604),bool)),sK19(X607,X606,X605,fun('type.pair.prod'(X603,X604),bool))),s('type.pair.prod'(X605,X606),'const.pair.,_2'(s(X605,X608),s(X606,X609)))),X604,X603)),s(X604,sK8(happ(s(fun('type.pair.prod'(X605,X606),fun('type.pair.prod'(X603,X604),bool)),sK19(X607,X606,X605,fun('type.pair.prod'(X603,X604),bool))),s('type.pair.prod'(X605,X606),'const.pair.,_2'(s(X605,X608),s(X606,X609)))),X604,X603))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X603,X604),bool),happ(s(fun('type.pair.prod'(X605,X606),fun('type.pair.prod'(X603,X604),bool)),sK19(X607,X606,X605,fun('type.pair.prod'(X603,X604),bool))),s('type.pair.prod'(X605,X606),'const.pair.,_2'(s(X605,X608),s(X606,X609))))),s('type.pair.prod'(X603,X604),X610))))) )),
  inference(superposition,[],[f341,f367])).
fof(f7272,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,sK7(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18)),s(X19,sK8(happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X18,X19),bool),happ(s(fun(X20,fun('type.pair.prod'(X18,X19),bool)),'_dst_x0x1_2'(s(fun(X21,fun('type.pair.prod'(X18,X19),bool)),X22),s(fun(X20,X21),X23))),s(X20,X24))),s('type.pair.prod'(X18,X19),X25))))) )),
  inference(superposition,[],[f341,f369])).
fof(f7271,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X14,fun('type.pair.prod'(X10,X11),bool)),'_dst_x0x1_2'(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s('type.pair.prod'(X10,X11),'const.pair.,_2'(s(X10,sK7(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10)),s(X11,sK8(happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X11,X10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X10,X11),bool),happ(s(fun(X12,fun('type.pair.prod'(X10,X11),bool)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16))))),s('type.pair.prod'(X10,X11),X17))))) )),
  inference(superposition,[],[f341,f369])).
fof(f9357,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (~p(s(bool,happ(s(fun('type.list.list'(X597),bool),happ(s(fun(X599,fun('type.list.list'(X597),bool)),happ(s(fun(X598,fun(X599,fun('type.list.list'(X597),bool))),X600),s(X598,X601))),s(X599,X602))),s('type.list.list'(X597),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X597),bool),happ(s(fun(X599,fun('type.list.list'(X597),bool)),happ(s(fun(X598,fun(X599,fun('type.list.list'(X597),bool))),X600),s(X598,X601))),s(X599,X602))),s('type.list.list'(X597),sK2(happ(s(fun('type.pair.prod'(X598,X599),fun('type.list.list'(X597),bool)),sK19(X600,X599,X598,fun('type.list.list'(X597),bool))),s('type.pair.prod'(X598,X599),'const.pair.,_2'(s(X598,X601),s(X599,X602)))),X597)))))) )),
  inference(forward_demodulation,[],[f9269,f367])).
fof(f9269,plain,(
  ( ! [X600,X602,X598,X601,X597,X599] : (p(s(bool,happ(s(fun('type.list.list'(X597),bool),happ(s(fun(X599,fun('type.list.list'(X597),bool)),happ(s(fun(X598,fun(X599,fun('type.list.list'(X597),bool))),X600),s(X598,X601))),s(X599,X602))),s('type.list.list'(X597),sK2(happ(s(fun('type.pair.prod'(X598,X599),fun('type.list.list'(X597),bool)),sK19(X600,X599,X598,fun('type.list.list'(X597),bool))),s('type.pair.prod'(X598,X599),'const.pair.,_2'(s(X598,X601),s(X599,X602)))),X597))))) | ~p(s(bool,happ(s(fun('type.list.list'(X597),bool),happ(s(fun('type.pair.prod'(X598,X599),fun('type.list.list'(X597),bool)),sK19(X600,X599,X598,fun('type.list.list'(X597),bool))),s('type.pair.prod'(X598,X599),'const.pair.,_2'(s(X598,X601),s(X599,X602))))),s('type.list.list'(X597),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f367])).
fof(f5529,plain,(
  ( ! [X420,X422,X418,X421,X417,X419] : (p(s(bool,happ(s(fun('type.list.list'(X417),bool),happ(s(fun(X419,fun('type.list.list'(X417),bool)),X420),s(X419,happ(s(fun(X418,X419),X421),s(X418,X422))))),s('type.list.list'(X417),sK2(happ(s(fun(X418,fun('type.list.list'(X417),bool)),'_dst_x0x1_2'(s(fun(X419,fun('type.list.list'(X417),bool)),X420),s(fun(X418,X419),X421))),s(X418,X422)),X417))))) | ~p(s(bool,happ(s(fun('type.list.list'(X417),bool),happ(s(fun(X418,fun('type.list.list'(X417),bool)),'_dst_x0x1_2'(s(fun(X419,fun('type.list.list'(X417),bool)),X420),s(fun(X418,X419),X421))),s(X418,X422))),s('type.list.list'(X417),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f369])).
fof(f5451,plain,(
  ( ! [X428,X430,X432,X429,X431,X427] : (p(s(bool,happ(s(fun('type.list.list'(X427),bool),happ(s(fun(X430,fun('type.list.list'(X427),bool)),'_dst_x0x1_2'(s(fun(X428,fun('type.list.list'(X427),bool)),X429),s(fun(X430,X428),X431))),s(X430,X432))),s('type.list.list'(X427),sK2(happ(s(fun(X428,fun('type.list.list'(X427),bool)),X429),s(X428,happ(s(fun(X430,X428),X431),s(X430,X432)))),X427))))) | ~p(s(bool,happ(s(fun('type.list.list'(X427),bool),happ(s(fun(X428,fun('type.list.list'(X427),bool)),X429),s(X428,happ(s(fun(X430,X428),X431),s(X430,X432))))),s('type.list.list'(X427),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f369])).
fof(f11045,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),X37)))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11044,f367])).
fof(f11044,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),X37))))) )),
  inference(forward_demodulation,[],[f11010,f367])).
fof(f11010,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun(X33,fun('type.list.list'(X31),bool)),happ(s(fun(X32,fun(X33,fun('type.list.list'(X31),bool))),X34),s(X32,X35))),s(X33,X36))),s('type.list.list'(X31),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),'const.list.CONS_2'(s(X31,sK3(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31)),s('type.list.list'(X31),sK2(happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36)))),X31))))))) | p(s(bool,happ(s(fun('type.list.list'(X31),bool),happ(s(fun('type.pair.prod'(X32,X33),fun('type.list.list'(X31),bool)),sK19(X34,X33,X32,fun('type.list.list'(X31),bool))),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,X35),s(X33,X36))))),s('type.list.list'(X31),X37))))) )),
  inference(superposition,[],[f309,f367])).
fof(f11007,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(X17,happ(s(fun(X16,X17),X19),s(X16,X20))))),s('type.list.list'(X15),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20))),s('type.list.list'(X15),'const.list.CONS_2'(s(X15,sK3(happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20)),X15)),s('type.list.list'(X15),sK2(happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20)),X15))))))) | p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun(X16,fun('type.list.list'(X15),bool)),'_dst_x0x1_2'(s(fun(X17,fun('type.list.list'(X15),bool)),X18),s(fun(X16,X17),X19))),s(X16,X20))),s('type.list.list'(X15),X21))))) )),
  inference(superposition,[],[f309,f369])).
fof(f11006,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X11,fun('type.list.list'(X8),bool)),'_dst_x0x1_2'(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(fun(X11,X9),X12))),s(X11,X13))),s('type.list.list'(X8),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13))))),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,sK3(happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))),X8)),s('type.list.list'(X8),sK2(happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))),X8))))))) | p(s(bool,happ(s(fun('type.list.list'(X8),bool),happ(s(fun(X9,fun('type.list.list'(X8),bool)),X10),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13))))),s('type.list.list'(X8),X14))))) )),
  inference(superposition,[],[f309,f369])).
fof(f9195,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,happ(s(fun(X108,bool),happ(s(fun(X107,fun(X108,bool)),X109),s(X107,sK7(sK19(X109,X108,X107,bool),X108,X107)))),s(X108,sK8(sK19(X109,X108,X107,bool),X108,X107))))) | p(s(bool,happ(s(fun('type.pair.prod'(X107,X108),bool),sK19(X109,X108,X107,bool)),s('type.pair.prod'(X107,X108),X110))))) )),
  inference(superposition,[],[f341,f367])).
fof(f5524,plain,(
  ( ! [X383,X385,X387,X382,X388,X384,X386] : (~p(s(bool,happ(s(fun(X382,bool),happ(s(fun(X384,fun(X382,bool)),X385),s(X384,happ(s(fun(X383,X384),X386),s(X383,X387))))),s(X382,X388)))) | p(s(bool,'|dollar|exists_1'(s(fun(X382,bool),happ(s(fun(X383,fun(X382,bool)),'_dst_x0x1_2'(s(fun(X384,fun(X382,bool)),X385),s(fun(X383,X384),X386))),s(X383,X387))))))) )),
  inference(superposition,[],[f4520,f369])).
fof(f5526,plain,(
  ( ! [X397,X399,X401,X396,X398,X400,X402] : (~p(s(bool,happ(s(fun(X396,bool),happ(s(fun(X398,fun(X396,bool)),X399),s(X398,happ(s(fun(X397,X398),X400),s(X397,X401))))),s(X396,X402)))) | p(s(bool,happ(s(fun(X396,bool),happ(s(fun(X397,fun(X396,bool)),'_dst_x0x1_2'(s(fun(X398,fun(X396,bool)),X399),s(fun(X397,X398),X400))),s(X397,X401))),s(X396,'const.min.@_1'(s(fun(X396,bool),happ(s(fun(X397,fun(X396,bool)),'_dst_x0x1_2'(s(fun(X398,fun(X396,bool)),X399),s(fun(X397,X398),X400))),s(X397,X401))))))))) )),
  inference(superposition,[],[f316,f369])).
fof(f5446,plain,(
  ( ! [X397,X393,X395,X396,X398,X392,X394] : (~p(s(bool,happ(s(fun(X392,bool),happ(s(fun(X395,fun(X392,bool)),'_dst_x0x1_2'(s(fun(X393,fun(X392,bool)),X394),s(fun(X395,X393),X396))),s(X395,X397))),s(X392,X398)))) | p(s(bool,'|dollar|exists_1'(s(fun(X392,bool),happ(s(fun(X393,fun(X392,bool)),X394),s(X393,happ(s(fun(X395,X393),X396),s(X395,X397))))))))) )),
  inference(superposition,[],[f4520,f369])).
fof(f5448,plain,(
  ( ! [X412,X408,X410,X407,X409,X411,X406] : (~p(s(bool,happ(s(fun(X406,bool),happ(s(fun(X409,fun(X406,bool)),'_dst_x0x1_2'(s(fun(X407,fun(X406,bool)),X408),s(fun(X409,X407),X410))),s(X409,X411))),s(X406,X412)))) | p(s(bool,happ(s(fun(X406,bool),happ(s(fun(X407,fun(X406,bool)),X408),s(X407,happ(s(fun(X409,X407),X410),s(X409,X411))))),s(X406,'const.min.@_1'(s(fun(X406,bool),happ(s(fun(X407,fun(X406,bool)),X408),s(X407,happ(s(fun(X409,X407),X410),s(X409,X411))))))))))) )),
  inference(superposition,[],[f316,f369])).
fof(f11243,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X88,bool),happ(s(fun(X87,fun(X88,bool)),X89),s(X87,X90))),s(X88,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X87,X88),bool),sK19(X89,X88,X87,bool)),s('type.pair.prod'(X87,X88),'const.pair.,_2'(s(X87,X90),s(X88,X91)))))) ) | $spl10),
  inference(superposition,[],[f10870,f367])).
fof(f9815,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,happ(s(fun(X133,bool),happ(s(fun(X132,fun(X133,bool)),X134),s(X132,X135))),s(X133,X136))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X132,X133),bool),sK19(X134,X133,X132,bool)),s('type.pair.prod'(X132,X133),'const.pair.,_2'(s(X132,X135),s(X133,X136)))))) )),
  inference(superposition,[],[f6343,f367])).
fof(f9512,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X137,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X133,bool),happ(s(fun(X132,fun(X133,bool)),X134),s(X132,X135))),s(X133,X136))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X132,X133),bool),sK19(X134,X133,X132,bool)),s('type.pair.prod'(X132,X133),'const.pair.,_2'(s(X132,X135),s(X133,X136)))))) )),
  inference(superposition,[],[f6195,f367])).
fof(f9325,plain,(
  ( ! [X239,X237,X235,X238,X236] : (s(bool,'T_0') != s(bool,happ(s(fun(X236,bool),happ(s(fun(X235,fun(X236,bool)),X237),s(X235,X238))),s(X236,X239))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X235,X236),bool),sK19(X237,X236,X235,bool)),s('type.pair.prod'(X235,X236),'const.pair.,_2'(s(X235,X238),s(X236,X239)))))) ) | $spl10),
  inference(forward_demodulation,[],[f9218,f1756])).
fof(f9218,plain,(
  ( ! [X239,X237,X235,X238,X236] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X236,bool),happ(s(fun(X235,fun(X236,bool)),X237),s(X235,X238))),s(X236,X239))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X235,X236),bool),sK19(X237,X236,X235,bool)),s('type.pair.prod'(X235,X236),'const.pair.,_2'(s(X235,X238),s(X236,X239)))))) ) | $spl10),
  inference(superposition,[],[f1849,f367])).
fof(f9234,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(X328,bool),happ(s(fun(X327,fun(X328,bool)),X329),s(X327,X330))),s(X328,X331)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X332),bool),sK18(X333,happ(s(fun('type.pair.prod'(X327,X328),bool),sK19(X329,X328,X327,bool)),s('type.pair.prod'(X327,X328),'const.pair.,_2'(s(X327,X330),s(X328,X331)))),X332,bool))))) )),
  inference(superposition,[],[f4695,f367])).
fof(f9198,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun(X122,fun(X123,bool)),X124),s(X122,X125))),s(X123,X126)))) | p(s(bool,happ(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)),s('type.pair.prod'(X122,X123),'const.min.@_1'(s(fun('type.pair.prod'(X122,X123),bool),sK19(X124,X123,X122,bool)))))))) )),
  inference(superposition,[],[f316,f367])).
fof(f9197,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X117,fun(X118,bool)),X119),s(X117,X120))),s(X118,X121)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(X117,X118),bool),sK19(X119,X118,X117,bool)))))) )),
  inference(superposition,[],[f4520,f367])).
fof(f9167,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X22) = s(bool,happ(s(fun(X18,bool),happ(s(fun(X17,fun(X18,bool)),X19),s(X17,X20))),s(X18,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X17,X18),bool),sK19(X19,X18,X17,bool)),s('type.pair.prod'(X17,X18),'const.pair.,_2'(s(X17,X20),s(X18,X21)))))) )),
  inference(superposition,[],[f367,f437])).
fof(f9166,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X14))),s(X12,X15))) | s(bool,X16) = s(bool,happ(s(fun('type.pair.prod'(X11,X12),bool),sK19(X13,X12,X11,bool)),s('type.pair.prod'(X11,X12),'const.pair.,_2'(s(X11,X14),s(X12,X15))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f367,f437])).
fof(f9196,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),happ(s(fun(X111,fun(X112,bool)),X113),s(X111,X114))),s(X112,X115)))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),X116),s('type.pair.prod'(X111,X112),sK6(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X111,X112),bool),sK19(X113,X112,X111,bool)),s('type.pair.prod'(X111,X112),sK5(sK19(X113,X112,X111,bool),X116,'type.pair.prod'(X111,X112))))))) )),
  inference(superposition,[],[f319,f367])).
fof(f7281,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),X72),s('type.pair.prod'(X70,X71),'const.pair.,_2'(s(X70,sK7(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70)),s(X71,sK8(sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool)),X71,X70))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X70,X71),bool),sK4(X72,X73,fun('type.pair.prod'(X70,X71),bool))),s('type.pair.prod'(X70,X71),X74)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.pair.prod'(X70,X71),bool),bool),X73)))) )),
  inference(superposition,[],[f341,f753])).
fof(f7346,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f7311,f261])).
fof(f7311,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),X7)))) | s(bool,X8) = s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),X6),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK7(X6,X5,X4)),s(X5,sK8(X6,X5,X4)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f341,f437])).
fof(f7335,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),X69)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)))))))) )),
  inference(forward_demodulation,[],[f7280,f344])).
fof(f7280,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),X68),s('type.pair.prod'(X64,X65),'const.pair.,_2'(s(X64,sK7('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64)),s(X65,sK8('const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68)))),X65,X64))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X64,X65),bool),'const.pair.SND_1'(s('type.pair.prod'(X66,fun('type.pair.prod'(X64,X65),bool)),'const.pair.,_2'(s(X66,X67),s(fun('type.pair.prod'(X64,X65),bool),X68))))),s('type.pair.prod'(X64,X65),X69))))) )),
  inference(superposition,[],[f341,f344])).
fof(f7334,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),X63)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)))))))) )),
  inference(forward_demodulation,[],[f7279,f338])).
fof(f7279,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),X62),s('type.pair.prod'(X59,X60),'const.pair.,_2'(s(X59,sK7('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59)),s(X60,sK8('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62)))),X60,X59))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X59,X60),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.pair.prod'(X59,X60),bool),X61),'const.sum.INL_1'(s(fun('type.pair.prod'(X59,X60),bool),X62))))),s('type.pair.prod'(X59,X60),X63))))) )),
  inference(superposition,[],[f341,f338])).
fof(f7333,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),X58)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)))))))) )),
  inference(forward_demodulation,[],[f7278,f343])).
fof(f7278,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),X56),s('type.pair.prod'(X53,X54),'const.pair.,_2'(s(X53,sK7('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53)),s(X54,sK8('const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57)))),X54,X53))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X53,X54),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.pair.prod'(X53,X54),bool),X55),'const.pair.,_2'(s(fun('type.pair.prod'(X53,X54),bool),X56),s(X55,X57))))),s('type.pair.prod'(X53,X54),X58))))) )),
  inference(superposition,[],[f341,f343])).
fof(f7332,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),X52)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)))))))) )),
  inference(forward_demodulation,[],[f7277,f339])).
fof(f7277,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),X51),s('type.pair.prod'(X48,X49),'const.pair.,_2'(s(X48,sK7('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48)),s(X49,sK8('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51)))),X49,X48))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X48,X49),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun('type.pair.prod'(X48,X49),bool)),'const.sum.INR_1'(s(fun('type.pair.prod'(X48,X49),bool),X51))))),s('type.pair.prod'(X48,X49),X52))))) )),
  inference(superposition,[],[f341,f339])).
fof(f7330,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),X41)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)))))))) )),
  inference(forward_demodulation,[],[f7275,f314])).
fof(f7275,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,sK7('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37)),s(X38,sK8('const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38,X37))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X37,X38),bool),'const.ind_type.FCONS_3'(s(fun('type.pair.prod'(X37,X38),bool),X39),s(fun('type.num.num',fun('type.pair.prod'(X37,X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'(X37,X38),X41))))) )),
  inference(superposition,[],[f341,f314])).
fof(f7328,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),X31)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f7273,f359])).
fof(f7273,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),X30),s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,sK7(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26)),s(X27,sK8(happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0')),X27,X26))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X26,X27),bool),happ(s(fun('type.list.list'(X28),fun('type.pair.prod'(X26,X27),bool)),sK18(X29,X30,X28,fun('type.pair.prod'(X26,X27),bool))),s('type.list.list'(X28),'const.list.NIL_0'))),s('type.pair.prod'(X26,X27),X31))))) )),
  inference(superposition,[],[f341,f359])).
fof(f7327,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),X9)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)))))))) )),
  inference(forward_demodulation,[],[f7270,f307])).
fof(f7270,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,sK7(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6)),s(X7,sK8(happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8)),X7,X6))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),happ(s(fun(fun('type.pair.prod'(X6,X7),bool),fun('type.pair.prod'(X6,X7),bool)),'_i_0_0'),s(fun('type.pair.prod'(X6,X7),bool),X8))),s('type.pair.prod'(X6,X7),X9))))) )),
  inference(superposition,[],[f341,f307])).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (? [X3,X4] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) => ! [X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X5)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9,X31] : ! [X27] : (! [X57,X58] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X27),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X57),s(X31,X58)))))) => ! [X59] : p(s(bool,happ(s(fun('type.pair.prod'(X9,X31),bool),X27),s('type.pair.prod'(X9,X31),X59)))))),
  file('thm.modSem.v_|49||49||split|5','thm.pair.pair_induction')).
fof(f7298,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,sK7(X7,bool,X8)),s(bool,X9)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X8,bool),bool),X7),s('type.pair.prod'(X8,bool),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK8(X7,bool,X8))) )),
  inference(superposition,[],[f341,f437])).
fof(f7297,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,sK7(X3,bool,X4)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,bool),bool),X3),s('type.pair.prod'(X4,bool),X5)))) | s(bool,X6) = s(bool,sK8(X3,bool,X4)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f341,f437])).
fof(f7284,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X9),s(X8,sK8(X7,X8,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X8),bool),X7),s('type.pair.prod'(bool,X8),X10)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK7(X7,X8,bool))) )),
  inference(superposition,[],[f341,f437])).
fof(f7283,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'T_0'),s(X4,sK8(X3,X4,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X4),bool),X3),s('type.pair.prod'(bool,X4),X5)))) | s(bool,X6) = s(bool,sK7(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f341,f437])).
fof(f6426,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74))))) )),
  inference(forward_demodulation,[],[f6400,f340])).
fof(f6400,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72))))))) )),
  inference(superposition,[],[f320,f340])).
fof(f6063,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(forward_demodulation,[],[f6033,f340])).
fof(f6033,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(superposition,[],[f319,f340])).
fof(f6032,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X80)) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),X81),s('type.list.list'(X78),sK6(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78)))))) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),sK18(X79,X80,X78,bool)),s('type.list.list'(X78),sK5(sK18(X79,X80,X78,bool),X81,'type.list.list'(X78))))))) )),
  inference(superposition,[],[f319,f359])).
fof(f4454,plain,(
  ( ! [X177,X175,X176] : (p(s(bool,happ(s(fun('type.list.list'(X175),bool),X176),s('type.list.list'(X175),sK2(sK4(X176,X177,fun('type.list.list'(X175),bool)),X175))))) | ~p(s(bool,happ(s(fun('type.list.list'(X175),bool),sK4(X176,X177,fun('type.list.list'(X175),bool))),s('type.list.list'(X175),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X175),bool),bool),X177)))) )),
  inference(superposition,[],[f412,f753])).
fof(f11017,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.list.list'(X65),bool),X66),s('type.list.list'(X65),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X65),bool),sK4(X66,X67,fun('type.list.list'(X65),bool))),s('type.list.list'(X65),'const.list.CONS_2'(s(X65,sK3(sK4(X66,X67,fun('type.list.list'(X65),bool)),X65)),s('type.list.list'(X65),sK2(sK4(X66,X67,fun('type.list.list'(X65),bool)),X65))))))) | p(s(bool,happ(s(fun('type.list.list'(X65),bool),sK4(X66,X67,fun('type.list.list'(X65),bool))),s('type.list.list'(X65),X68)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X65),bool),bool),X67)))) )),
  inference(superposition,[],[f309,f753])).
fof(f2871,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f2855,f264])).
fof(f2855,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f412,f268])).
fof(f2870,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26)))))) )),
  inference(forward_demodulation,[],[f2854,f344])).
fof(f2854,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X26),bool),X29),s('type.list.list'(X26),sK2('const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29)))),X26))))) | ~p(s(bool,happ(s(fun('type.list.list'(X26),bool),'const.pair.SND_1'(s('type.pair.prod'(X27,fun('type.list.list'(X26),bool)),'const.pair.,_2'(s(X27,X28),s(fun('type.list.list'(X26),bool),X29))))),s('type.list.list'(X26),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f344])).
fof(f2869,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23)))))) )),
  inference(forward_demodulation,[],[f2853,f338])).
fof(f2853,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X25),s('type.list.list'(X23),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X23),bool),X24),'const.sum.INL_1'(s(fun('type.list.list'(X23),bool),X25))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f338])).
fof(f2868,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f2852,f343])).
fof(f2852,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f343])).
fof(f2867,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16)))))) )),
  inference(forward_demodulation,[],[f2851,f339])).
fof(f2851,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X17,fun('type.list.list'(X16),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X16),bool),X18))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f339])).
fof(f2865,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f2849,f314])).
fof(f2849,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f314])).
fof(f2863,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f2847,f359])).
fof(f2847,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK18(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f359])).
fof(f2862,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2846,f307])).
fof(f2846,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f307])).
fof(f412,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f228])).
fof(f228,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ! [X27] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X27),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X28] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X27),s('type.list.list'(X9),X28)))) => ! [X29] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X27),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X29),s('type.list.list'(X9),X28)))))))) => ! [X30] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X27),s('type.list.list'(X9),X30)))))),
  file('thm.modSem.v_|49||49||split|5','thm.list.list_INDUCT')).
fof(f11058,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK3(X4,X3)),s('type.list.list'(X3),sK2(X4,X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),X5)))) | s(bool,X6) = s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f11019,f261])).
fof(f11019,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,sK3(X4,X3)),s('type.list.list'(X3),sK2(X4,X3))))))) | p(s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),X5)))) | s(bool,X6) = s(bool,happ(s(fun('type.list.list'(X3),bool),X4),s('type.list.list'(X3),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f309,f437])).
fof(f11057,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),X64)))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11056,f344])).
fof(f11056,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),X64))))) )),
  inference(forward_demodulation,[],[f11016,f344])).
fof(f11016,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun('type.list.list'(X60),bool),X63),s('type.list.list'(X60),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),'const.list.CONS_2'(s(X60,sK3('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60)),s('type.list.list'(X60),sK2('const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63)))),X60))))))) | p(s(bool,happ(s(fun('type.list.list'(X60),bool),'const.pair.SND_1'(s('type.pair.prod'(X61,fun('type.list.list'(X60),bool)),'const.pair.,_2'(s(X61,X62),s(fun('type.list.list'(X60),bool),X63))))),s('type.list.list'(X60),X64))))) )),
  inference(superposition,[],[f309,f344])).
fof(f11055,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),X59)))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11054,f338])).
fof(f11054,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),X59))))) )),
  inference(forward_demodulation,[],[f11015,f338])).
fof(f11015,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun('type.list.list'(X56),bool),X58),s('type.list.list'(X56),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),'const.list.CONS_2'(s(X56,sK3('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56)),s('type.list.list'(X56),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58)))),X56))))))) | p(s(bool,happ(s(fun('type.list.list'(X56),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X56),bool),X57),'const.sum.INL_1'(s(fun('type.list.list'(X56),bool),X58))))),s('type.list.list'(X56),X59))))) )),
  inference(superposition,[],[f309,f338])).
fof(f11053,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),X55)))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11052,f343])).
fof(f11052,plain,(
  ( ! [X54,X52,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),X55))))) )),
  inference(forward_demodulation,[],[f11014,f343])).
fof(f11014,plain,(
  ( ! [X54,X52,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.list.list'(X51),bool),X53),s('type.list.list'(X51),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),'const.list.CONS_2'(s(X51,sK3('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51)),s('type.list.list'(X51),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54)))),X51))))))) | p(s(bool,happ(s(fun('type.list.list'(X51),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X51),bool),X52),'const.pair.,_2'(s(fun('type.list.list'(X51),bool),X53),s(X52,X54))))),s('type.list.list'(X51),X55))))) )),
  inference(superposition,[],[f309,f343])).
fof(f11051,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),X50)))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11050,f339])).
fof(f11050,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),X50))))) )),
  inference(forward_demodulation,[],[f11013,f339])).
fof(f11013,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun('type.list.list'(X47),bool),X49),s('type.list.list'(X47),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),'const.list.CONS_2'(s(X47,sK3('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47)),s('type.list.list'(X47),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49)))),X47))))))) | p(s(bool,happ(s(fun('type.list.list'(X47),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X48,fun('type.list.list'(X47),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X47),bool),X49))))),s('type.list.list'(X47),X50))))) )),
  inference(superposition,[],[f309,f339])).
fof(f11047,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),X41)))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11046,f314])).
fof(f11046,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),X41))))) )),
  inference(forward_demodulation,[],[f11011,f314])).
fof(f11011,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.list.list'(X38),bool),X39),s('type.list.list'(X38),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),'const.list.CONS_2'(s(X38,sK3('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38)),s('type.list.list'(X38),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0')),X38))))))) | p(s(bool,happ(s(fun('type.list.list'(X38),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X38),bool),X39),s(fun('type.num.num',fun('type.list.list'(X38),bool)),X40),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X38),X41))))) )),
  inference(superposition,[],[f309,f314])).
fof(f11041,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),X26)))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11040,f359])).
fof(f11040,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),X26))))) )),
  inference(forward_demodulation,[],[f11008,f359])).
fof(f11008,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X22),bool),X25),s('type.list.list'(X22),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),'const.list.CONS_2'(s(X22,sK3(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22)),s('type.list.list'(X22),sK2(happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0')),X22))))))) | p(s(bool,happ(s(fun('type.list.list'(X22),bool),happ(s(fun('type.list.list'(X23),fun('type.list.list'(X22),bool)),sK18(X24,X25,X23,fun('type.list.list'(X22),bool))),s('type.list.list'(X23),'const.list.NIL_0'))),s('type.list.list'(X22),X26))))) )),
  inference(superposition,[],[f309,f359])).
fof(f11039,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),X7)))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f11038,f307])).
fof(f11038,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),X7))))) )),
  inference(forward_demodulation,[],[f11005,f307])).
fof(f11005,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),'const.list.CONS_2'(s(X5,sK3(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5)),s('type.list.list'(X5),sK2(happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6)),X5))))))) | p(s(bool,happ(s(fun('type.list.list'(X5),bool),happ(s(fun(fun('type.list.list'(X5),bool),fun('type.list.list'(X5),bool)),'_i_0_0'),s(fun('type.list.list'(X5),bool),X6))),s('type.list.list'(X5),X7))))) )),
  inference(superposition,[],[f309,f307])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f6047,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(forward_demodulation,[],[f5959,f323])).
fof(f5959,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,X46)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK6('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42))))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK5('const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X47,X42)))))) )),
  inference(superposition,[],[f319,f323])).
fof(f9354,plain,(
  ( ! [X573,X569,X571,X567,X572,X574,X568,X570] : (p(s(bool,happ(s(fun(X567,bool),happ(s(fun(X569,fun(X567,bool)),happ(s(fun(X568,fun(X569,fun(X567,bool))),X570),s(X568,X571))),s(X569,X572))),s(X567,sK5(happ(s(fun('type.pair.prod'(X568,X569),fun(X567,bool)),sK19(X570,X569,X568,fun(X567,bool))),s('type.pair.prod'(X568,X569),'const.pair.,_2'(s(X568,X571),s(X569,X572)))),X574,X567))))) | ~p(s(bool,happ(s(fun(X567,bool),happ(s(fun(X569,fun(X567,bool)),happ(s(fun(X568,fun(X569,fun(X567,bool))),X570),s(X568,X571))),s(X569,X572))),s(X567,X573)))) | p(s(bool,happ(s(fun(X567,bool),X574),s(X567,sK6(happ(s(fun('type.pair.prod'(X568,X569),fun(X567,bool)),sK19(X570,X569,X568,fun(X567,bool))),s('type.pair.prod'(X568,X569),'const.pair.,_2'(s(X568,X571),s(X569,X572)))),X574,X567)))))) )),
  inference(forward_demodulation,[],[f9265,f367])).
fof(f9265,plain,(
  ( ! [X573,X569,X571,X567,X572,X574,X568,X570] : (~p(s(bool,happ(s(fun(X567,bool),happ(s(fun(X569,fun(X567,bool)),happ(s(fun(X568,fun(X569,fun(X567,bool))),X570),s(X568,X571))),s(X569,X572))),s(X567,X573)))) | p(s(bool,happ(s(fun(X567,bool),X574),s(X567,sK6(happ(s(fun('type.pair.prod'(X568,X569),fun(X567,bool)),sK19(X570,X569,X568,fun(X567,bool))),s('type.pair.prod'(X568,X569),'const.pair.,_2'(s(X568,X571),s(X569,X572)))),X574,X567))))) | p(s(bool,happ(s(fun(X567,bool),happ(s(fun('type.pair.prod'(X568,X569),fun(X567,bool)),sK19(X570,X569,X568,fun(X567,bool))),s('type.pair.prod'(X568,X569),'const.pair.,_2'(s(X568,X571),s(X569,X572))))),s(X567,sK5(happ(s(fun('type.pair.prod'(X568,X569),fun(X567,bool)),sK19(X570,X569,X568,fun(X567,bool))),s('type.pair.prod'(X568,X569),'const.pair.,_2'(s(X568,X571),s(X569,X572)))),X574,X567)))))) )),
  inference(superposition,[],[f319,f367])).
fof(f6045,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(forward_demodulation,[],[f5957,f340])).
fof(f5957,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(superposition,[],[f319,f340])).
fof(f6042,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f5952,f368])).
fof(f5952,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(superposition,[],[f319,f368])).
fof(f5955,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,X24)))) | p(s(bool,happ(s(fun(X18,bool),X25),s(X18,sK6(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18))))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK5(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18)))))) )),
  inference(superposition,[],[f319,f369])).
fof(f5954,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,X16)))) | p(s(bool,happ(s(fun(X10,bool),X17),s(X10,sK6(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10))))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK5(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10)))))) )),
  inference(superposition,[],[f319,f369])).
fof(f6030,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70)))) | p(s(bool,happ(s(fun(X66,bool),X71),s(X66,sK6(X67,X71,X66))))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X67,X71,X66)))))) )),
  inference(superposition,[],[f319,f369])).
fof(f5400,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun(X97,X95),X98))),s(X97,X99)))) | p(s(bool,happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96))))))) )),
  inference(superposition,[],[f316,f369])).
fof(f9587,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK10(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f753])).
fof(f10646,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,X110)))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,sK13(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f350,f753])).
fof(f9608,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f9579,f340])).
fof(f9579,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48))))))) )),
  inference(superposition,[],[f347,f340])).
fof(f10667,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK13(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f10638,f340])).
fof(f10638,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK13(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(superposition,[],[f350,f340])).
fof(f9605,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f9574,f368])).
fof(f9574,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f10664,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK13(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f10633,f368])).
fof(f10633,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK13(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(superposition,[],[f350,f368])).
fof(f9609,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f9580,f367])).
fof(f9580,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK10(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,X58))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f10668,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,X58)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK13(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f10639,f367])).
fof(f10639,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,X58)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,sK13(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))) )),
  inference(superposition,[],[f350,f367])).
fof(f9577,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK10(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f10636,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33)))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK13(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))) )),
  inference(superposition,[],[f350,f369])).
fof(f9576,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK10(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f10635,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23)))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK13(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))) )),
  inference(superposition,[],[f350,f369])).
fof(f9611,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f9582,f323])).
fof(f9582,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK10(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73))))))) )),
  inference(superposition,[],[f347,f323])).
fof(f10670,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,X73)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK13(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(forward_demodulation,[],[f10641,f323])).
fof(f10641,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,X73)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,sK13(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))) )),
  inference(superposition,[],[f350,f323])).
fof(f7325,plain,(
  ( ! [X70,X68,X72,X71,X69,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,happ(s(fun('type.pair.prod'(X68,X69),X70),X72),s('type.pair.prod'(X68,X69),'const.pair.,_2'(s(X68,sK7('_dst_x0x1_2'(s(fun(X70,bool),X71),s(fun('type.pair.prod'(X68,X69),X70),X72)),X69,X68)),s(X69,sK8('_dst_x0x1_2'(s(fun(X70,bool),X71),s(fun('type.pair.prod'(X68,X69),X70),X72)),X69,X68))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X68,X69),bool),'_dst_x0x1_2'(s(fun(X70,bool),X71),s(fun('type.pair.prod'(X68,X69),X70),X72))),s('type.pair.prod'(X68,X69),X73))))) )),
  inference(superposition,[],[f341,f369])).
fof(f11035,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,happ(s(fun(X64,bool),X65),s(X64,happ(s(fun('type.list.list'(X63),X64),X66),s('type.list.list'(X63),'const.list.NIL_0')))))) | ~p(s(bool,happ(s(fun('type.list.list'(X63),bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66))),s('type.list.list'(X63),'const.list.CONS_2'(s(X63,sK3('_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66)),X63)),s('type.list.list'(X63),sK2('_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66)),X63))))))) | p(s(bool,happ(s(fun('type.list.list'(X63),bool),'_dst_x0x1_2'(s(fun(X64,bool),X65),s(fun('type.list.list'(X63),X64),X66))),s('type.list.list'(X63),X67))))) )),
  inference(superposition,[],[f309,f369])).
fof(f5479,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,happ(s(fun('type.list.list'(X96),X97),X99),s('type.list.list'(X96),sK2('_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99)),X96))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X96),bool),'_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99))),s('type.list.list'(X96),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f412,f369])).
fof(f11240,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X78,bool),X79),s(X78,happ(s(fun(X77,X78),X80),s(X77,X81))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X77,bool),'_dst_x0x1_2'(s(fun(X78,bool),X79),s(fun(X77,X78),X80))),s(X77,X81)))) ) | $spl10),
  inference(superposition,[],[f10870,f369])).
fof(f10677,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK13(X13,X11,X12,X10,X9))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f10649,f261])).
fof(f10649,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK13(X13,X11,X12,X10,X9))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f350,f437])).
fof(f10674,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK13(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f10645,f344])).
fof(f10645,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103)))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK13(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(superposition,[],[f350,f344])).
fof(f10673,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK13(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f10644,f338])).
fof(f10644,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95)))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK13(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(superposition,[],[f350,f338])).
fof(f10672,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK13(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f10643,f343])).
fof(f10643,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88)))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK13(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(superposition,[],[f350,f343])).
fof(f10671,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK13(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f10642,f339])).
fof(f10642,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80)))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK13(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(superposition,[],[f350,f339])).
fof(f10669,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK13(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f10640,f314])).
fof(f10640,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65)))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK13(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(superposition,[],[f350,f314])).
fof(f10666,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK13(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f10637,f359])).
fof(f10637,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41)))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK13(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(superposition,[],[f350,f359])).
fof(f10665,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK13(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f10634,f307])).
fof(f10634,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13)))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK13(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(superposition,[],[f350,f307])).
fof(f10632,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK13(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f350,f753])).
fof(f10631,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103)))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK13(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))) )),
  inference(superposition,[],[f350,f344])).
fof(f10630,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95)))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK13(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))) )),
  inference(superposition,[],[f350,f338])).
fof(f10629,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88)))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK13(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))) )),
  inference(superposition,[],[f350,f343])).
fof(f10628,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80)))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK13(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))) )),
  inference(superposition,[],[f350,f339])).
fof(f10627,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73)))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK13(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))) )),
  inference(superposition,[],[f350,f323])).
fof(f10626,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65)))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK13(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))) )),
  inference(superposition,[],[f350,f314])).
fof(f10625,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,X58)))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,sK13(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49)))))) )),
  inference(superposition,[],[f350,f367])).
fof(f10624,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48)))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK13(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))) )),
  inference(superposition,[],[f350,f340])).
fof(f10623,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41)))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK13(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))) )),
  inference(superposition,[],[f350,f359])).
fof(f10622,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33)))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK13(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))) )),
  inference(superposition,[],[f350,f369])).
fof(f10621,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23)))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK13(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))) )),
  inference(superposition,[],[f350,f369])).
fof(f10620,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13)))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK13(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))) )),
  inference(superposition,[],[f350,f307])).
fof(f10619,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))) )),
  inference(superposition,[],[f350,f368])).
fof(f10529,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X120),'const.sum.INL_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2130,f369])).
fof(f10368,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X115,bool),X116),s(fun(X117,X115),X118))),s(X117,X119))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X120,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X115,bool),X116),s(X115,happ(s(fun(X117,X115),X118),s(X117,X119)))))))))) ) | $spl10),
  inference(superposition,[],[f2128,f369])).
fof(f10140,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(X71,bool),'_dst_x0x1_2'(s(fun(X69,bool),X70),s(fun(X71,X69),X72))),s(X71,X73))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X69,bool),X70),s(X69,happ(s(fun(X71,X69),X72),s(X71,X73)))))))) ) | $spl10),
  inference(superposition,[],[f2124,f369])).
fof(f9812,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X124),'const.sum.INL_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6343,f369])).
fof(f9618,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9590,f261])).
fof(f9590,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK10(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f347,f437])).
fof(f9615,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,X103)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))) )),
  inference(forward_demodulation,[],[f9586,f344])).
fof(f9586,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,sK10(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,X103))))))) )),
  inference(superposition,[],[f347,f344])).
fof(f9614,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,X95)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))) )),
  inference(forward_demodulation,[],[f9585,f338])).
fof(f9585,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,sK10(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,X95))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f9613,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,X88)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))) )),
  inference(forward_demodulation,[],[f9584,f343])).
fof(f9584,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,sK10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,X88))))))) )),
  inference(superposition,[],[f347,f343])).
fof(f9612,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,X80)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))) )),
  inference(forward_demodulation,[],[f9583,f339])).
fof(f9583,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,sK10(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,X80))))))) )),
  inference(superposition,[],[f347,f339])).
fof(f9610,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,X65)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))) )),
  inference(forward_demodulation,[],[f9581,f314])).
fof(f9581,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK10(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,X65))))))) )),
  inference(superposition,[],[f347,f314])).
fof(f9607,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f9578,f359])).
fof(f9578,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41))))))) )),
  inference(superposition,[],[f347,f359])).
fof(f9606,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f9575,f307])).
fof(f9575,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13))))))) )),
  inference(superposition,[],[f347,f307])).
fof(f9573,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,sK10(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,X110)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f347,f753])).
fof(f9572,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,sK10(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,X103))))))) )),
  inference(superposition,[],[f347,f344])).
fof(f9571,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,sK10(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,X95))))))) )),
  inference(superposition,[],[f347,f338])).
fof(f9570,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,sK10(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,X88))))))) )),
  inference(superposition,[],[f347,f343])).
fof(f9569,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,sK10(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,X80))))))) )),
  inference(superposition,[],[f347,f339])).
fof(f9568,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,sK10(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,X73))))))) )),
  inference(superposition,[],[f347,f323])).
fof(f9567,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,sK10(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,X65))))))) )),
  inference(superposition,[],[f347,f314])).
fof(f9566,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,sK10(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,X58))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f9565,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK10(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48))))))) )),
  inference(superposition,[],[f347,f340])).
fof(f9564,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK10(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41))))))) )),
  inference(superposition,[],[f347,f359])).
fof(f9563,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK10(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9562,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK10(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))))))) )),
  inference(superposition,[],[f347,f369])).
fof(f9561,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK10(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13))))))) )),
  inference(superposition,[],[f347,f307])).
fof(f9560,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK10(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9509,plain,(
  ( ! [X123,X121,X124,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X124,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X119,X120),X122),s(X119,X123))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X119,X120),X122))),s(X119,X123)))) )),
  inference(superposition,[],[f6195,f369])).
fof(f6031,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | p(s(bool,happ(s(fun(X72,bool),X77),s(X72,sK6('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72))))) | p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,sK5('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72)))))) )),
  inference(superposition,[],[f319,f369])).
fof(f5513,plain,(
  ( ! [X292,X294,X288,X290,X293,X289,X291] : (~p(s(bool,happ(s(fun(X289,bool),X290),s(X289,happ(s(fun(X288,X289),X291),s(X288,X292)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X293),bool),sK18(X294,happ(s(fun(X288,bool),'_dst_x0x1_2'(s(fun(X289,bool),X290),s(fun(X288,X289),X291))),s(X288,X292)),X293,bool))))) )),
  inference(superposition,[],[f4695,f369])).
fof(f5495,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X184,X185),X187),s(X184,X188)))))) | s(bool,X189) = s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X184,X185),X187))),s(X184,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f487,f369])).
fof(f5493,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X173,X174),X176),s(X173,X177)))))) | s(bool,X178) != s(bool,happ(s(fun(X173,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X173,X174),X176))),s(X173,X177))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f439,f369])).
fof(f5488,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X144,bool),X145),s(X144,happ(s(fun(X143,X144),X146),s(X143,X147)))))) | s(bool,X148) != s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X144,bool),X145),s(fun(X143,X144),X146))),s(X143,X147))) | p(s(bool,X148))) )),
  inference(superposition,[],[f410,f369])).
fof(f5487,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | ~p(s(bool,X142)) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f337,f369])).
fof(f5485,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129)))))) | p(s(bool,X130)) | s(bool,X130) = s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f334,f369])).
fof(f5482,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114)))) )),
  inference(superposition,[],[f270,f369])).
fof(f5481,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f269,f369])).
fof(f5478,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun(X91,X92),X94),s(X91,X95)))))) | p(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))),s(X91,'const.min.@_1'(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))))))))) )),
  inference(superposition,[],[f316,f369])).
fof(f5477,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))) )),
  inference(superposition,[],[f4520,f369])).
fof(f5416,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (p(s(bool,happ(s(fun(X189,bool),'_dst_x0x1_2'(s(fun(X187,bool),X188),s(fun(X189,X187),X190))),s(X189,X191)))) | s(bool,X192) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f487,f369])).
fof(f5414,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X178,bool),'_dst_x0x1_2'(s(fun(X176,bool),X177),s(fun(X178,X176),X179))),s(X178,X180)))) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),X177),s(X176,happ(s(fun(X178,X176),X179),s(X178,X180))))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f439,f369])).
fof(f5409,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,X151) != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150))))) | p(s(bool,X151))) )),
  inference(superposition,[],[f410,f369])).
fof(f5408,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | ~p(s(bool,X145)) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) )),
  inference(superposition,[],[f337,f369])).
fof(f5406,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X130,X128),X131))),s(X130,X132)))) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f334,f369])).
fof(f5403,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117)))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f270,f369])).
fof(f5402,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X108,bool),X109),s(fun(X110,X108),X111))),s(X110,X112)))) | s(bool,'T_0') != s(bool,happ(s(fun(X108,bool),X109),s(X108,happ(s(fun(X110,X108),X111),s(X110,X112)))))) )),
  inference(superposition,[],[f269,f369])).
fof(f350,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK13(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f347,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK10(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f10676,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f10675,f307])).
fof(f10675,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(subsumption_resolution,[],[f10647,f261])).
fof(f10647,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(superposition,[],[f350,f432])).
fof(f10288,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f10260,f261])).
fof(f10260,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK12(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f349,f437])).
fof(f10287,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f10286,f307])).
fof(f10286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(subsumption_resolution,[],[f10258,f261])).
fof(f10258,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK12(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(superposition,[],[f349,f432])).
fof(f10285,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f10256,f344])).
fof(f10256,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK12(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97)))))))) )),
  inference(superposition,[],[f349,f344])).
fof(f10284,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f10255,f338])).
fof(f10255,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK12(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90)))))))) )),
  inference(superposition,[],[f349,f338])).
fof(f10283,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f10254,f343])).
fof(f10254,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK12(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82)))))))) )),
  inference(superposition,[],[f349,f343])).
fof(f10282,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f10253,f339])).
fof(f10253,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK12(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75)))))))) )),
  inference(superposition,[],[f349,f339])).
fof(f10281,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f10252,f323])).
fof(f10252,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73)))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK12(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67)))))))) )),
  inference(superposition,[],[f349,f323])).
fof(f10280,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f10251,f314])).
fof(f10251,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK12(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60)))))))) )),
  inference(superposition,[],[f349,f314])).
fof(f10279,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f10250,f367])).
fof(f10250,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58)))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK12(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50)))))))) )),
  inference(superposition,[],[f349,f367])).
fof(f10278,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f10249,f340])).
fof(f10249,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK12(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))))) )),
  inference(superposition,[],[f349,f340])).
fof(f10277,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f10248,f359])).
fof(f10248,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK12(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))))) )),
  inference(superposition,[],[f349,f359])).
fof(f10276,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f10245,f307])).
fof(f10245,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK12(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))))) )),
  inference(superposition,[],[f349,f307])).
fof(f10275,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f10244,f368])).
fof(f10244,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK12(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f10257,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK12(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f349,f753])).
fof(f10247,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33)))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK12(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))))) )),
  inference(superposition,[],[f349,f369])).
fof(f10246,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23)))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK12(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))))) )),
  inference(superposition,[],[f349,f369])).
fof(f10243,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK12(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f349,f753])).
fof(f10242,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103)))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK12(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96)))))))) )),
  inference(superposition,[],[f349,f344])).
fof(f10241,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95)))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK12(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89)))))))) )),
  inference(superposition,[],[f349,f338])).
fof(f10240,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88)))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK12(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81)))))))) )),
  inference(superposition,[],[f349,f343])).
fof(f10239,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80)))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK12(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74)))))))) )),
  inference(superposition,[],[f349,f339])).
fof(f10238,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73)))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK12(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66)))))))) )),
  inference(superposition,[],[f349,f323])).
fof(f10237,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65)))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK12(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59)))))))) )),
  inference(superposition,[],[f349,f314])).
fof(f10236,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58)))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK12(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49)))))))) )),
  inference(superposition,[],[f349,f367])).
fof(f10235,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK12(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))))) )),
  inference(superposition,[],[f349,f340])).
fof(f10234,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41)))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK12(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))))) )),
  inference(superposition,[],[f349,f359])).
fof(f10233,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33)))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK12(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))))) )),
  inference(superposition,[],[f349,f369])).
fof(f10232,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23)))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK12(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))))) )),
  inference(superposition,[],[f349,f369])).
fof(f10231,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK12(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))))) )),
  inference(superposition,[],[f349,f307])).
fof(f10230,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f9922,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9894,f261])).
fof(f9894,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK11(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f348,f437])).
fof(f9921,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f9920,f307])).
fof(f9920,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(subsumption_resolution,[],[f9892,f261])).
fof(f9892,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK11(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(superposition,[],[f348,f432])).
fof(f9919,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),X100),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(forward_demodulation,[],[f9890,f344])).
fof(f9890,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X97,X96),X101),s(fun(X96,X97),X100)))) | ~p(s(bool,happ(s(fun(X97,bool),X102),s(X97,happ(s(fun(X96,X97),'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100))))),s(X96,sK11(X102,X101,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X96,X97))))))) | p(s(bool,happ(s(fun(X97,bool),X102),s(X97,X103))))) )),
  inference(superposition,[],[f348,f344])).
fof(f9918,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),X92),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(forward_demodulation,[],[f9889,f338])).
fof(f9889,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X90,X89),X93),s(fun(X89,X90),X92)))) | ~p(s(bool,happ(s(fun(X90,bool),X94),s(X90,happ(s(fun(X89,X90),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92))))),s(X89,sK11(X94,X93,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X89,X90))))))) | p(s(bool,happ(s(fun(X90,bool),X94),s(X90,X95))))) )),
  inference(superposition,[],[f348,f338])).
fof(f9917,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(forward_demodulation,[],[f9888,f343])).
fof(f9888,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X82,X81),X86),s(fun(X81,X82),X84)))) | ~p(s(bool,happ(s(fun(X82,bool),X87),s(X82,happ(s(fun(X81,X82),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85))))),s(X81,sK11(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X81,X82))))))) | p(s(bool,happ(s(fun(X82,bool),X87),s(X82,X88))))) )),
  inference(superposition,[],[f348,f343])).
fof(f9916,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),X77),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(forward_demodulation,[],[f9887,f339])).
fof(f9887,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X75,X74),X78),s(fun(X74,X75),X77)))) | ~p(s(bool,happ(s(fun(X75,bool),X79),s(X75,happ(s(fun(X74,X75),'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77))))),s(X74,sK11(X79,X78,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X74,X75))))))) | p(s(bool,happ(s(fun(X75,bool),X79),s(X75,X80))))) )),
  inference(superposition,[],[f348,f339])).
fof(f9915,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(forward_demodulation,[],[f9886,f323])).
fof(f9886,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X67,X66),X71),s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70)))))) | ~p(s(bool,happ(s(fun(X67,bool),X72),s(X67,happ(s(fun(X66,X67),'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70))))),s(X66,sK11(X72,X71,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X66,X67))))))) | p(s(bool,happ(s(fun(X67,bool),X72),s(X67,X73))))) )),
  inference(superposition,[],[f348,f323])).
fof(f9914,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),X61),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(forward_demodulation,[],[f9885,f314])).
fof(f9885,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X60,X59),X63),s(fun(X59,X60),X61)))) | ~p(s(bool,happ(s(fun(X60,bool),X64),s(X60,happ(s(fun(X59,X60),'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0'))),s(X59,sK11(X64,X63,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X59,X60))))))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,X65))))) )),
  inference(superposition,[],[f348,f314])).
fof(f9913,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(forward_demodulation,[],[f9884,f367])).
fof(f9884,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X56),s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55)))))) | ~p(s(bool,happ(s(fun(X50,bool),X57),s(X50,happ(s(fun(X49,X50),happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))),s(X49,sK11(X57,X56,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X49,X50))))))) | p(s(bool,happ(s(fun(X50,bool),X57),s(X50,X58))))) )),
  inference(superposition,[],[f348,f367])).
fof(f9912,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f9883,f340])).
fof(f9883,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK11(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(superposition,[],[f348,f340])).
fof(f9911,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f9882,f359])).
fof(f9882,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK11(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(superposition,[],[f348,f359])).
fof(f9910,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f9879,f307])).
fof(f9879,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK11(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(superposition,[],[f348,f307])).
fof(f9909,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f9878,f368])).
fof(f9878,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK11(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9891,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X105,X104),X108),s(fun(X104,X105),X106)))) | ~p(s(bool,happ(s(fun(X105,bool),X109),s(X105,happ(s(fun(X104,X105),sK4(X106,X107,fun(X104,X105))),s(X104,sK11(X109,X108,sK4(X106,X107,fun(X104,X105)),X104,X105))))))) | p(s(bool,happ(s(fun(X105,bool),X109),s(X105,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f348,f753])).
fof(f9881,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK11(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9880,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK11(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9877,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X104,X105),X106),s(fun(X105,X104),X108)))) | ~p(s(bool,happ(s(fun(X104,bool),X109),s(X104,happ(s(fun(X105,X104),X108),s(X105,sK11(X109,sK4(X106,X107,fun(X104,X105)),X108,X105,X104))))))) | p(s(bool,happ(s(fun(X104,bool),X109),s(X104,X110)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X104,X105),bool),X107)))) )),
  inference(superposition,[],[f348,f753])).
fof(f9876,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X96,X97),X100),s(fun(X97,X96),X101)))) | ~p(s(bool,happ(s(fun(X96,bool),X102),s(X96,happ(s(fun(X97,X96),X101),s(X97,sK11(X102,'const.pair.SND_1'(s('type.pair.prod'(X98,fun(X96,X97)),'const.pair.,_2'(s(X98,X99),s(fun(X96,X97),X100)))),X101,X97,X96))))))) | p(s(bool,happ(s(fun(X96,bool),X102),s(X96,X103))))) )),
  inference(superposition,[],[f348,f344])).
fof(f9875,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X89,X90),X92),s(fun(X90,X89),X93)))) | ~p(s(bool,happ(s(fun(X89,bool),X94),s(X89,happ(s(fun(X90,X89),X93),s(X90,sK11(X94,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X89,X90),X91),'const.sum.INL_1'(s(fun(X89,X90),X92)))),X93,X90,X89))))))) | p(s(bool,happ(s(fun(X89,bool),X94),s(X89,X95))))) )),
  inference(superposition,[],[f348,f338])).
fof(f9874,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X82),X84),s(fun(X82,X81),X86)))) | ~p(s(bool,happ(s(fun(X81,bool),X87),s(X81,happ(s(fun(X82,X81),X86),s(X82,sK11(X87,'const.pair.FST_1'(s('type.pair.prod'(fun(X81,X82),X83),'const.pair.,_2'(s(fun(X81,X82),X84),s(X83,X85)))),X86,X82,X81))))))) | p(s(bool,happ(s(fun(X81,bool),X87),s(X81,X88))))) )),
  inference(superposition,[],[f348,f343])).
fof(f9873,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X74,X75),X77),s(fun(X75,X74),X78)))) | ~p(s(bool,happ(s(fun(X74,bool),X79),s(X74,happ(s(fun(X75,X74),X78),s(X75,sK11(X79,'const.sum.OUTR_1'(s('type.sum.sum'(X76,fun(X74,X75)),'const.sum.INR_1'(s(fun(X74,X75),X77)))),X78,X75,X74))))))) | p(s(bool,happ(s(fun(X74,bool),X79),s(X74,X80))))) )),
  inference(superposition,[],[f348,f339])).
fof(f9872,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X66,X67),happ(s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num',X70))),s(fun(X67,X66),X71)))) | ~p(s(bool,happ(s(fun(X66,bool),X72),s(X66,happ(s(fun(X67,X66),X71),s(X67,sK11(X72,'const.ind_type.FCONS_3'(s(fun(X66,X67),X68),s(fun('type.num.num',fun(X66,X67)),X69),s('type.num.num','const.num.SUC_1'(s('type.num.num',X70)))),X71,X67,X66))))))) | p(s(bool,happ(s(fun(X66,bool),X72),s(X66,X73))))) )),
  inference(superposition,[],[f348,f323])).
fof(f9871,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X59,X60),X61),s(fun(X60,X59),X63)))) | ~p(s(bool,happ(s(fun(X59,bool),X64),s(X59,happ(s(fun(X60,X59),X63),s(X60,sK11(X64,'const.ind_type.FCONS_3'(s(fun(X59,X60),X61),s(fun('type.num.num',fun(X59,X60)),X62),s('type.num.num','const.num.|48|_0')),X63,X60,X59))))))) | p(s(bool,happ(s(fun(X59,bool),X64),s(X59,X65))))) )),
  inference(superposition,[],[f348,f314])).
fof(f9870,plain,(
  ( ! [X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),happ(s(fun(X52,fun(X49,X50)),happ(s(fun(X51,fun(X52,fun(X49,X50))),X53),s(X51,X54))),s(X52,X55))),s(fun(X50,X49),X56)))) | ~p(s(bool,happ(s(fun(X49,bool),X57),s(X49,happ(s(fun(X50,X49),X56),s(X50,sK11(X57,happ(s(fun('type.pair.prod'(X51,X52),fun(X49,X50)),sK19(X53,X52,X51,fun(X49,X50))),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55)))),X56,X50,X49))))))) | p(s(bool,happ(s(fun(X49,bool),X57),s(X49,X58))))) )),
  inference(superposition,[],[f348,f367])).
fof(f9869,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK11(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48))))) )),
  inference(superposition,[],[f348,f340])).
fof(f9868,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK11(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK18(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41))))) )),
  inference(superposition,[],[f348,f359])).
fof(f9867,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK11(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9866,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK11(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23))))) )),
  inference(superposition,[],[f348,f369])).
fof(f9865,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK11(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13))))) )),
  inference(superposition,[],[f348,f307])).
fof(f9864,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK11(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9617,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f9616,f307])).
fof(f9616,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f9588,f261])).
fof(f9588,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f347,f432])).
fof(f9355,plain,(
  ( ! [X575,X580,X582,X576,X578,X581,X577,X579] : (p(s(bool,happ(s(fun(X575,bool),happ(s(fun(X577,fun(X575,bool)),happ(s(fun(X576,fun(X577,fun(X575,bool))),X578),s(X576,X579))),s(X577,X580))),s(X575,sK6(X581,happ(s(fun('type.pair.prod'(X576,X577),fun(X575,bool)),sK19(X578,X577,X576,fun(X575,bool))),s('type.pair.prod'(X576,X577),'const.pair.,_2'(s(X576,X579),s(X577,X580)))),X575))))) | ~p(s(bool,happ(s(fun(X575,bool),happ(s(fun(X577,fun(X575,bool)),happ(s(fun(X576,fun(X577,fun(X575,bool))),X578),s(X576,X579))),s(X577,X580))),s(X575,sK5(X581,happ(s(fun('type.pair.prod'(X576,X577),fun(X575,bool)),sK19(X578,X577,X576,fun(X575,bool))),s('type.pair.prod'(X576,X577),'const.pair.,_2'(s(X576,X579),s(X577,X580)))),X575))))) | ~p(s(bool,happ(s(fun(X575,bool),X581),s(X575,X582))))) )),
  inference(forward_demodulation,[],[f9266,f367])).
fof(f9266,plain,(
  ( ! [X575,X580,X582,X576,X578,X581,X577,X579] : (~p(s(bool,happ(s(fun(X575,bool),happ(s(fun(X577,fun(X575,bool)),happ(s(fun(X576,fun(X577,fun(X575,bool))),X578),s(X576,X579))),s(X577,X580))),s(X575,sK5(X581,happ(s(fun('type.pair.prod'(X576,X577),fun(X575,bool)),sK19(X578,X577,X576,fun(X575,bool))),s('type.pair.prod'(X576,X577),'const.pair.,_2'(s(X576,X579),s(X577,X580)))),X575))))) | ~p(s(bool,happ(s(fun(X575,bool),X581),s(X575,X582)))) | p(s(bool,happ(s(fun(X575,bool),happ(s(fun('type.pair.prod'(X576,X577),fun(X575,bool)),sK19(X578,X577,X576,fun(X575,bool))),s('type.pair.prod'(X576,X577),'const.pair.,_2'(s(X576,X579),s(X577,X580))))),s(X575,sK6(X581,happ(s(fun('type.pair.prod'(X576,X577),fun(X575,bool)),sK19(X578,X577,X576,fun(X575,bool))),s('type.pair.prod'(X576,X577),'const.pair.,_2'(s(X576,X579),s(X577,X580)))),X575)))))) )),
  inference(superposition,[],[f320,f367])).
fof(f6417,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6386,f261])).
fof(f6386,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f320,f437])).
fof(f6410,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69))))) )),
  inference(forward_demodulation,[],[f6369,f344])).
fof(f6369,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X68),s(X64,X69)))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK6(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X64)))))) )),
  inference(superposition,[],[f320,f344])).
fof(f6409,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63))))) )),
  inference(forward_demodulation,[],[f6368,f338])).
fof(f6368,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X62),s(X59,X63)))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK6(X62,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X59)))))) )),
  inference(superposition,[],[f320,f338])).
fof(f6408,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58))))) )),
  inference(forward_demodulation,[],[f6367,f343])).
fof(f6367,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X57),s(X53,X58)))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK6(X57,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X53)))))) )),
  inference(superposition,[],[f320,f343])).
fof(f6407,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52))))) )),
  inference(forward_demodulation,[],[f6366,f339])).
fof(f6366,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X51),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK6(X51,'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X48)))))) )),
  inference(superposition,[],[f320,f339])).
fof(f6406,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47))))) )),
  inference(forward_demodulation,[],[f6365,f323])).
fof(f6365,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (~p(s(bool,happ(s(fun(X42,bool),happ(s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num',X45))),s(X42,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42))))) | ~p(s(bool,happ(s(fun(X42,bool),X46),s(X42,X47)))) | p(s(bool,happ(s(fun(X42,bool),'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s(X42,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X42,bool),X43),s(fun('type.num.num',fun(X42,bool)),X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45)))),X42)))))) )),
  inference(superposition,[],[f320,f323])).
fof(f6405,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41))))) )),
  inference(forward_demodulation,[],[f6364,f314])).
fof(f6364,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK6(X40,'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X37)))))) )),
  inference(superposition,[],[f320,f314])).
fof(f6404,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36))))) )),
  inference(forward_demodulation,[],[f6363,f340])).
fof(f6363,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36)))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32)))))) )),
  inference(superposition,[],[f320,f340])).
fof(f6403,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31))))) )),
  inference(forward_demodulation,[],[f6362,f359])).
fof(f6362,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31)))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26)))))) )),
  inference(superposition,[],[f320,f359])).
fof(f6402,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f6359,f307])).
fof(f6359,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6)))))) )),
  inference(superposition,[],[f320,f307])).
fof(f6401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f6358,f368])).
fof(f6358,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0)))))) )),
  inference(superposition,[],[f320,f368])).
fof(f6399,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,happ(s(fun(X66,X67),X69),s(X66,sK5(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66))))))) | ~p(s(bool,happ(s(fun(X66,bool),X70),s(X66,X71)))) | p(s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69))),s(X66,sK6(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6370,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | ~p(s(bool,happ(s(fun(X70,bool),X73),s(X70,X74)))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK6(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f320,f753])).
fof(f6361,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,sK5(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18))))) | ~p(s(bool,happ(s(fun(X18,bool),X24),s(X18,X25)))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK6(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6360,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,sK5(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10))))) | ~p(s(bool,happ(s(fun(X10,bool),X16),s(X10,X17)))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK6(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10)))))) )),
  inference(superposition,[],[f320,f369])).
fof(f6056,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f6017,f261])).
fof(f6017,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f319,f437])).
fof(f6051,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (p(s(bool,happ(s(fun(X64,bool),X67),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | ~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(forward_demodulation,[],[f5963,f344])).
fof(f5963,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (~p(s(bool,happ(s(fun(X64,bool),X67),s(X64,X68)))) | p(s(bool,happ(s(fun(X64,bool),X69),s(X64,sK6('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64))))) | p(s(bool,happ(s(fun(X64,bool),'const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67))))),s(X64,sK5('const.pair.SND_1'(s('type.pair.prod'(X65,fun(X64,bool)),'const.pair.,_2'(s(X65,X66),s(fun(X64,bool),X67)))),X69,X64)))))) )),
  inference(superposition,[],[f319,f344])).
fof(f6050,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,happ(s(fun(X59,bool),X61),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | ~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(forward_demodulation,[],[f5962,f338])).
fof(f5962,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(X59,bool),X61),s(X59,X62)))) | p(s(bool,happ(s(fun(X59,bool),X63),s(X59,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59))))) | p(s(bool,happ(s(fun(X59,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61))))),s(X59,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,bool),X60),'const.sum.INL_1'(s(fun(X59,bool),X61)))),X63,X59)))))) )),
  inference(superposition,[],[f319,f338])).
fof(f6049,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),X55),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | ~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(forward_demodulation,[],[f5961,f343])).
fof(f5961,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun(X53,bool),X55),s(X53,X57)))) | p(s(bool,happ(s(fun(X53,bool),X58),s(X53,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53))))) | p(s(bool,happ(s(fun(X53,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56))))),s(X53,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X53,bool),X54),'const.pair.,_2'(s(fun(X53,bool),X55),s(X54,X56)))),X58,X53)))))) )),
  inference(superposition,[],[f319,f343])).
fof(f6048,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),X50),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(forward_demodulation,[],[f5960,f339])).
fof(f5960,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),X50),s(X48,X51)))) | p(s(bool,happ(s(fun(X48,bool),X52),s(X48,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48))))) | p(s(bool,happ(s(fun(X48,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50))))),s(X48,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X49,fun(X48,bool)),'const.sum.INR_1'(s(fun(X48,bool),X50)))),X52,X48)))))) )),
  inference(superposition,[],[f319,f339])).
fof(f6046,plain,(
  ( ! [X39,X37,X41,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),X38),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(forward_demodulation,[],[f5958,f314])).
fof(f5958,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X40)))) | p(s(bool,happ(s(fun(X37,bool),X41),s(X37,sK6('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37))))) | p(s(bool,happ(s(fun(X37,bool),'const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0'))),s(X37,sK5('const.ind_type.FCONS_3'(s(fun(X37,bool),X38),s(fun('type.num.num',fun(X37,bool)),X39),s('type.num.num','const.num.|48|_0')),X41,X37)))))) )),
  inference(superposition,[],[f319,f314])).
fof(f6044,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(forward_demodulation,[],[f5956,f359])).
fof(f5956,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK18(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(superposition,[],[f319,f359])).
fof(f6043,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(forward_demodulation,[],[f5953,f307])).
fof(f5953,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(superposition,[],[f319,f307])).
fof(f5964,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,X73)))) | p(s(bool,happ(s(fun(X70,bool),X74),s(X70,sK6(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK5(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f319,f753])).
fof(f4615,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4582,f261])).
fof(f4582,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4520,f437])).
fof(f4536,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4520,f753])).
fof(f4520,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4399,f269])).
fof(f4399,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f318,f753])).
fof(f4451,plain,(
  ( ! [X165,X163,X166,X164] : (~p(s(bool,happ(s(fun(X163,bool),X164),s(X163,X166)))) | p(s(bool,happ(s(fun(X163,bool),sK4(X164,X165,fun(X163,bool))),s(X163,'const.min.@_1'(s(fun(X163,bool),sK4(X164,X165,fun(X163,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X163,bool),bool),X165)))) )),
  inference(superposition,[],[f316,f753])).
fof(f1596,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1576,f261])).
fof(f1576,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f316,f437])).
fof(f349,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK12(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f348,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f320,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X55,X27] : (! [X33] : (p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33)))) => p(s(bool,happ(s(fun(X9,bool),X55),s(X9,X33))))) => (? [X33] : p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33)))) => ? [X33] : p(s(bool,happ(s(fun(X9,bool),X55),s(X9,X33))))))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.MONO_EXISTS')).
fof(f319,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X27,X33] : (p(s(bool,happ(s(fun(X9,bool),X27),s(X9,X33)))) => p(s(bool,happ(s(fun(X9,bool),X27),s(X9,'const.min.@_1'(s(fun(X9,bool),X27)))))))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.SELECT_AX')).
fof(f11245,plain,(
  ( ! [X94,X95,X96] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X95),s('type.num.num',X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X94),s(fun('type.num.num',bool),X95),s('type.num.num','const.num.SUC_1'(s('type.num.num',X96)))))) ) | $spl10),
  inference(superposition,[],[f10870,f323])).
fof(f9817,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X142),s('type.num.num',X143))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X141),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.SUC_1'(s('type.num.num',X143)))))) )),
  inference(superposition,[],[f6343,f323])).
fof(f9514,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X144,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X142),s('type.num.num',X143))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X141),s(fun('type.num.num',bool),X142),s('type.num.num','const.num.SUC_1'(s('type.num.num',X143)))))) )),
  inference(superposition,[],[f6195,f323])).
fof(f5247,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun('type.num.num',bool),X124),s('type.num.num',X125)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X126),bool),sK18(X127,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.SUC_1'(s('type.num.num',X125)))),X126,bool))))) )),
  inference(superposition,[],[f4695,f323])).
fof(f2455,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X111),s('type.num.num',X112))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.SUC_1'(s('type.num.num',X112)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2431,f1756])).
fof(f2431,plain,(
  ( ! [X111,X112,X110] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun('type.num.num',bool),X111),s('type.num.num',X112))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X110),s(fun('type.num.num',bool),X111),s('type.num.num','const.num.SUC_1'(s('type.num.num',X112)))))) ) | $spl10),
  inference(superposition,[],[f1849,f323])).
fof(f2402,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f323,f437])).
fof(f2400,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f323,f437])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f318,f437])).
fof(f6425,plain,(
  ( ! [X64,X65] : (p(s(bool,sK6(X64,'_i_0_0',bool))) | ~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))))) )),
  inference(forward_demodulation,[],[f6398,f307])).
fof(f6398,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X64,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f320,f307])).
fof(f6373,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X7,X8,bool))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool))) )),
  inference(superposition,[],[f320,f437])).
fof(f6372,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,sK6(X3,X4,bool))))) | s(bool,X6) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f320,f437])).
fof(f1288,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f318,f437])).
fof(f4055,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK14(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK14(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f353,f437])).
fof(f4056,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK14(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK14(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f353,f437])).
fof(f11287,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(subsumption_resolution,[],[f11208,f424])).
fof(f11208,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(superposition,[],[f10870,f2671])).
fof(f11061,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f11060,f261])).
fof(f11060,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f11024,f1756])).
fof(f11024,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.CONS_2'(s(X23,sK3(X24,X23)),s('type.list.list'(X23),sK2(X24,X23))))))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X23),bool),X24),s('type.list.list'(X23),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f309,f2498])).
fof(f10993,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f10902,f424])).
fof(f10902,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK18(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5268,f2671])).
fof(f10880,plain,(
  ( ! [X54,X53] : (s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(subsumption_resolution,[],[f10798,f264])).
fof(f10798,plain,(
  ( ! [X54,X53] : (p(s(bool,'F_0')) | s(bool,X54) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X53))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X53)) ) | $spl10),
  inference(superposition,[],[f2351,f2781])).
fof(f10820,plain,(
  ( ! [X98] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X98),s(bool,sK15('_i_0_0',X98,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X98) | s(bool,'T_0') = s(bool,sK15('_i_0_0',X98,bool,bool))) ) | $spl10),
  inference(superposition,[],[f355,f2781])).
fof(f10819,plain,(
  ( ! [X97] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X97),s(bool,sK15(X97,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X97) | s(bool,'T_0') = s(bool,sK15(X97,'_i_0_0',bool,bool))) ) | $spl10),
  inference(superposition,[],[f355,f2781])).
fof(f10749,plain,(
  ( ! [X103,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103))) | s(bool,'T_0') = s(bool,sK4(X103,X104,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X104)))) ) | $spl10),
  inference(superposition,[],[f2781,f753])).
fof(f10744,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) ) | $spl10),
  inference(superposition,[],[f2781,f323])).
fof(f10742,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X84,bool),happ(s(fun(X83,fun(X84,bool)),X85),s(X83,X86))),s(X84,X87))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X83,X84),bool),sK19(X85,X84,X83,bool)),s('type.pair.prod'(X83,X84),'const.pair.,_2'(s(X83,X86),s(X84,X87)))))) ) | $spl10),
  inference(superposition,[],[f2781,f367])).
fof(f10741,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),X82))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X81),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X81),X82)))) ) | $spl10),
  inference(superposition,[],[f2781,f340])).
fof(f10739,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X74,bool),X75),s(X74,happ(s(fun(X73,X74),X76),s(X73,X77))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X74,bool),X75),s(fun(X73,X74),X76))),s(X73,X77)))) ) | $spl10),
  inference(superposition,[],[f2781,f369])).
fof(f10738,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X70,bool),'_dst_x0x1_2'(s(fun(X68,bool),X69),s(fun(X70,X68),X71))),s(X70,X72))))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,happ(s(fun(X70,X68),X71),s(X70,X72)))))) ) | $spl10),
  inference(superposition,[],[f2781,f369])).
fof(f10680,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f10679,f261])).
fof(f10679,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f10654,f1756])).
fof(f10654,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK13(X48,X46,X47,X45,X44))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f350,f2498])).
fof(f10291,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f10290,f261])).
fof(f10290,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f10265,f1756])).
fof(f10265,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK12(X48,X46,X47,X45,X44))))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f349,f2498])).
fof(f9925,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9924,f261])).
fof(f9924,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9899,f1756])).
fof(f9899,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,sK11(X48,X46,X47,X45,X44))))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,X49)))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f348,f2498])).
fof(f9858,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9783,f424])).
fof(f9783,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X48),'const.sum.INL_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6343,f2671])).
fof(f9621,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f9620,f261])).
fof(f9620,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f9595,f1756])).
fof(f9595,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X44,bool),X48),s(X44,sK10(X48,X46,X47,X45,X44))))) | p(s(bool,happ(s(fun(X44,bool),X48),s(X44,happ(s(fun(X45,X44),X47),s(X45,X49)))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f347,f2498])).
fof(f9554,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(subsumption_resolution,[],[f9480,f424])).
fof(f9480,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X48,bool),'const.sum.INR_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | $spl10),
  inference(superposition,[],[f6195,f2671])).
fof(f9308,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(forward_demodulation,[],[f9171,f1756])).
fof(f9171,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X42,bool),happ(s(fun(X41,fun(X42,bool)),X43),s(X41,X44))),s(X42,X45))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X41,X42),bool),sK19(X43,X42,X41,bool)),s('type.pair.prod'(X41,X42),'const.pair.,_2'(s(X41,X44),s(X42,X45)))))) ) | $spl10),
  inference(superposition,[],[f367,f2498])).
fof(f9233,plain,(
  ( ! [X325,X321,X323,X324,X326,X322] : (s(bool,'T_0') != s(bool,happ(s(fun(X322,bool),happ(s(fun(X321,fun(X322,bool)),X323),s(X321,X324))),s(X322,X325))) | s(bool,X326) = s(bool,happ(s(fun('type.pair.prod'(X321,X322),bool),sK19(X323,X322,X321,bool)),s('type.pair.prod'(X321,X322),'const.pair.,_2'(s(X321,X324),s(X322,X325))))) | s(bool,'F_0') = s(bool,X326)) ) | $spl10),
  inference(superposition,[],[f3504,f367])).
fof(f9221,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X256),'const.sum.INL_1'(s(bool,happ(s(fun(X252,bool),happ(s(fun(X251,fun(X252,bool)),X253),s(X251,X254))),s(X252,X255))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X251,X252),bool),sK19(X253,X252,X251,bool)),s('type.pair.prod'(X251,X252),'const.pair.,_2'(s(X251,X254),s(X252,X255)))))) )),
  inference(superposition,[],[f2343,f367])).
fof(f9220,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X250,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X246,bool),happ(s(fun(X245,fun(X246,bool)),X247),s(X245,X248))),s(X246,X249))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X245,X246),bool),sK19(X247,X246,X245,bool)),s('type.pair.prod'(X245,X246),'const.pair.,_2'(s(X245,X248),s(X246,X249)))))) )),
  inference(superposition,[],[f2341,f367])).
fof(f9216,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X228,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X224,bool),happ(s(fun(X223,fun(X224,bool)),X225),s(X223,X226))),s(X224,X227))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X223,X224),bool),sK19(X225,X224,X223,bool)),s('type.pair.prod'(X223,X224),'const.pair.,_2'(s(X223,X226),s(X224,X227)))))) )),
  inference(superposition,[],[f535,f367])).
fof(f9215,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X222),'const.sum.INL_1'(s(bool,happ(s(fun(X218,bool),happ(s(fun(X217,fun(X218,bool)),X219),s(X217,X220))),s(X218,X221))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X217,X218),bool),sK19(X219,X218,X217,bool)),s('type.pair.prod'(X217,X218),'const.pair.,_2'(s(X217,X220),s(X218,X221)))))) )),
  inference(superposition,[],[f507,f367])).
fof(f9213,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X207,bool),happ(s(fun(X206,fun(X207,bool)),X208),s(X206,X209))),s(X207,X210))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X206,X207),bool),sK19(X208,X207,X206,bool)),s('type.pair.prod'(X206,X207),'const.pair.,_2'(s(X206,X209),s(X207,X210)))))) )),
  inference(superposition,[],[f465,f367])).
fof(f9174,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X58,bool),happ(s(fun(X57,fun(X58,bool)),X59),s(X57,X60))),s(X58,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X57,X58),bool),sK19(X59,X58,X57,bool)),s('type.pair.prod'(X57,X58),'const.pair.,_2'(s(X57,X60),s(X58,X61)))))) ) | $spl10),
  inference(superposition,[],[f367,f2671])).
fof(f9173,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,X54))),s(X52,X55))) | s(bool,X56) = s(bool,happ(s(fun('type.pair.prod'(X51,X52),bool),sK19(X53,X52,X51,bool)),s('type.pair.prod'(X51,X52),'const.pair.,_2'(s(X51,X54),s(X52,X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl10),
  inference(superposition,[],[f367,f2671])).
fof(f9165,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),sK19(X8,X7,X6,bool)),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10)))))) )),
  inference(superposition,[],[f367,f268])).
fof(f8104,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8103,f261])).
fof(f8103,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f8072,f1756])).
fof(f8072,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X45,X48) != s(X45,happ(s(fun(X44,X45),X46),s(X44,X49))) | s(X44,X49) = s(X44,happ(s(fun(X45,X44),X47),s(X45,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f352,f2498])).
fof(f7563,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(subsumption_resolution,[],[f7562,f261])).
fof(f7562,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(forward_demodulation,[],[f7531,f1756])).
fof(f7531,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X44,X48) != s(X44,happ(s(fun(X45,X44),X47),s(X45,X49))) | s(X45,X49) = s(X45,happ(s(fun(X44,X45),X46),s(X44,X48))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X44,X45),X46),s(fun(X45,X44),X47)))) ) | $spl10),
  inference(superposition,[],[f351,f2498])).
fof(f12912,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | ($spl10 | $spl26)),
  inference(global_subsumption,[],[f7371,f260,f259,f261,f263,f265,f266,f267,f268,f270,f269,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f487,f753,f783,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f2400,f2402,f2458,f2436,f2491,f2589,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f4043,f4044,f4055,f4056,f4367,f4368,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4419,f4421,f4422,f4427,f4429,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4594,f4595,f4631,f4641,f4675,f4681,f4695,f4725,f4724,f4718,f4780,f4790,f4824,f4906,f4931,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5264,f5269,f5273,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5287,f5288,f5298,f5299,f5399,f5400,f5402,f5403,f5406,f5408,f5409,f5414,f5416,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5485,f5487,f5488,f5493,f5495,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6062,f6030,f6031,f6032,f6063,f6195,f6343,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6372,f6373,f6417,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6572,f6573,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6601,f6602,f6613,f6614,f6615,f6618,f6619,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7283,f7284,f7297,f7298,f7346,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8144,f8230,f8231,f8241,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8441,f8444,f8446,f8448,f8449,f8450,f8451,f8452,f8803,f8804,f8861,f8864,f8866,f8868,f8869,f8870,f8871,f8872,f9165,f9166,f9167,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9220,f9221,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9550,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9712,f9853,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f10019,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11034,f11035,f11399,f11400,f11617,f11620,f11623,f11625,f11626,f11627,f11628,f11629,f7353,f7368,f8261,f6599,f6566,f4449,f264,f273,f274,f465,f507,f535,f615,f664,f717,f829,f1290,f1398,f2291,f2339,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2428,f2496,f2497,f4042,f4054,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4417,f4418,f4428,f4430,f4431,f4435,f4436,f5263,f5265,f5277,f5286,f5404,f5405,f5415,f5417,f5418,f5422,f5423,f5483,f5484,f5494,f5496,f5497,f5501,f5502,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6353,f6371,f6416,f6570,f6571,f6600,f6616,f6617,f7282,f7296])).
fof(f7371,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(subsumption_resolution,[],[f7370,f261])).
fof(f7370,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f7360,f1756])).
fof(f7360,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X17,bool))) ) | $spl10),
  inference(superposition,[],[f7353,f2498])).
fof(f7349,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f7348,f261])).
fof(f7348,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(forward_demodulation,[],[f7316,f1756])).
fof(f7316,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),X32)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X29,X30),bool),X31),s('type.pair.prod'(X29,X30),'const.pair.,_2'(s(X29,sK7(X31,X30,X29)),s(X30,sK8(X31,X30,X29))))))) ) | $spl10),
  inference(superposition,[],[f341,f2498])).
fof(f7342,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(forward_demodulation,[],[f7302,f1756])).
fof(f7302,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,sK7(X23,bool,X24)),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X24,bool),bool),X23),s('type.pair.prod'(X24,bool),X25)))) | s(bool,'F_0') = s(bool,sK8(X23,bool,X24))) ) | $spl10),
  inference(superposition,[],[f341,f2498])).
fof(f7337,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'T_0'),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f7288,f1756])).
fof(f7288,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))),s(X24,sK8(X23,X24,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X24),bool),X23),s('type.pair.prod'(bool,X24),X25)))) | s(bool,'F_0') = s(bool,sK7(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f341,f2498])).
fof(f7305,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,sK7(X33,bool,X34)),s(bool,X35)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X34,bool),bool),X33),s('type.pair.prod'(X34,bool),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK8(X33,bool,X34))) ) | $spl10),
  inference(superposition,[],[f341,f2671])).
fof(f7304,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,sK7(X29,bool,X30)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X30,bool),bool),X29),s('type.pair.prod'(X30,bool),X31)))) | s(bool,X32) = s(bool,sK8(X29,bool,X30)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f341,f2671])).
fof(f7296,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK7(X0,bool,X1)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,'T_0') = s(bool,sK8(X0,bool,X1))) )),
  inference(superposition,[],[f341,f268])).
fof(f7291,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,sK8(X33,X34,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X34),bool),X33),s('type.pair.prod'(bool,X34),X36)))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK7(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f341,f2671])).
fof(f7290,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,sK8(X29,X30,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X30),bool),X29),s('type.pair.prod'(bool,X30),X31)))) | s(bool,X32) = s(bool,sK7(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f341,f2671])).
fof(f7282,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,sK8(X0,X1,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool))) )),
  inference(superposition,[],[f341,f268])).
fof(f6663,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6623,f1756])).
fof(f6623,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X36,X37,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X36,X37,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f355,f2498])).
fof(f6659,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'T_0'))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(forward_demodulation,[],[f6606,f1756])).
fof(f6606,plain,(
  ( ! [X24,X23,X25] : (s(X25,happ(s(fun(bool,X25),X23),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,happ(s(fun(bool,X25),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK15(X23,X24,bool,X25))) ) | $spl10),
  inference(superposition,[],[f355,f2498])).
fof(f6645,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(forward_demodulation,[],[f6577,f1756])).
fof(f6577,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X35,bool),X37),s(X35,sK15(X37,X36,X35,bool)))) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,sK15(X37,X36,X35,bool))))) ) | $spl10),
  inference(superposition,[],[f355,f2498])).
fof(f6626,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X46,X47,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X46,X47,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6625,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X42,X43,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X42,X43,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6617,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X13,X14,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X13,X14,X12,bool))))) )),
  inference(superposition,[],[f355,f268])).
fof(f6616,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15('_i_0_0',X11,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15('_i_0_0',X11,bool,bool))) )),
  inference(superposition,[],[f355,f465])).
fof(f6609,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,happ(s(fun(bool,X35),X33),s(bool,X36))) != s(X35,happ(s(fun(bool,X35),X34),s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK15(X33,X34,bool,X35))) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6608,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,happ(s(fun(bool,X31),X29),s(bool,'F_0'))) != s(X31,happ(s(fun(bool,X31),X30),s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK15(X29,X30,bool,X31)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6600,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,X2))) )),
  inference(superposition,[],[f355,f268])).
fof(f6580,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X48) != s(bool,happ(s(fun(X45,bool),X47),s(X45,sK15(X47,X46,X45,bool)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,happ(s(fun(X45,bool),X46),s(X45,sK15(X47,X46,X45,bool))))) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6579,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,happ(s(fun(X41,bool),X43),s(X41,sK15(X43,X42,X41,bool)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,X44) = s(bool,happ(s(fun(X41,bool),X42),s(X41,sK15(X43,X42,X41,bool)))) | s(bool,'F_0') = s(bool,X44)) ) | $spl10),
  inference(superposition,[],[f355,f2671])).
fof(f6571,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK15(X14,X13,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK15(X14,X13,X12,bool))))) )),
  inference(superposition,[],[f355,f268])).
fof(f6570,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK15(X11,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK15(X11,'_i_0_0',bool,bool))) )),
  inference(superposition,[],[f355,f465])).
fof(f6420,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6419,f261])).
fof(f6419,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(forward_demodulation,[],[f6391,f1756])).
fof(f6391,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,happ(s(fun(X31,bool),X33),s(X31,X34)))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK6(X33,X32,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X32,X31))))) ) | $spl10),
  inference(superposition,[],[f320,f2498])).
fof(f6416,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6(X0,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(forward_demodulation,[],[f6415,f307])).
fof(f6415,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f6384,f261])).
fof(f6384,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f320,f465])).
fof(f6412,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f6377,f1756])).
fof(f6377,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X25)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,sK6(X23,X24,bool))))) | s(bool,'F_0') = s(bool,sK5(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f320,f2498])).
fof(f6380,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X34),s(bool,X35)))) | ~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X36)))) | p(s(bool,happ(s(fun(bool,bool),X34),s(bool,sK6(X33,X34,bool))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,sK5(X33,X34,bool))) ) | $spl10),
  inference(superposition,[],[f320,f2671])).
fof(f6379,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X30),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X30),s(bool,sK6(X29,X30,bool))))) | s(bool,X32) = s(bool,sK5(X29,X30,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f320,f2671])).
fof(f6371,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f320,f268])).
fof(f6353,plain,(
  ( ! [X64,X65,X63] : (s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6293,f264])).
fof(f6293,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'F_0')) | s(bool,X65) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X63),'const.sum.INL_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2351,f2343])).
fof(f6257,plain,(
  ( ! [X130,X131,X129] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f2343,f753])).
fof(f6256,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X125,bool),'const.pair.,_2'(s(X125,X126),s(bool,X127)))))) )),
  inference(superposition,[],[f2343,f344])).
fof(f6254,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X121),'const.sum.INL_1'(s(bool,X119))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,X119),s(X118,X120)))))) )),
  inference(superposition,[],[f2343,f343])).
fof(f6252,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X112),s('type.num.num',X113))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X111),s(fun('type.num.num',bool),X112),s('type.num.num','const.num.SUC_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f2343,f323])).
fof(f6251,plain,(
  ( ! [X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X110),'const.sum.INL_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2343,f314])).
fof(f6250,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f2343,f340])).
fof(f6249,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X104),'const.sum.INL_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK18(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2343,f359])).
fof(f6248,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f2343,f369])).
fof(f6247,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f2343,f369])).
fof(f6205,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f6145,f264])).
fof(f6145,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'F_0')) | s(bool,X62) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f2351,f2341])).
fof(f6110,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X126))))) | s(bool,'T_0') = s(bool,sK4(X126,X127,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X127)))) )),
  inference(superposition,[],[f2341,f753])).
fof(f6109,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X125,bool),'const.sum.INR_1'(s(bool,X124))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X122,bool),'const.pair.,_2'(s(X122,X123),s(bool,X124)))))) )),
  inference(superposition,[],[f2341,f344])).
fof(f6107,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,X116))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),'const.pair.,_2'(s(bool,X116),s(X115,X117)))))) )),
  inference(superposition,[],[f2341,f343])).
fof(f6105,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X109),s('type.num.num',X110))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X108),s(fun('type.num.num',bool),X109),s('type.num.num','const.num.SUC_1'(s('type.num.num',X110)))))) )),
  inference(superposition,[],[f2341,f323])).
fof(f6104,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,X105))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X105),s(fun('type.num.num',bool),X106),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2341,f314])).
fof(f6103,plain,(
  ( ! [X103,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X102),X103)))) )),
  inference(superposition,[],[f2341,f340])).
fof(f6102,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X100))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X98),bool),sK18(X99,X100,X98,bool)),s('type.list.list'(X98),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2341,f359])).
fof(f6101,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,happ(s(fun(X92,X93),X95),s(X92,X96))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X93,bool),X94),s(fun(X92,X93),X95))),s(X92,X96)))) )),
  inference(superposition,[],[f2341,f369])).
fof(f6100,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X88,bool),'_dst_x0x1_2'(s(fun(X86,bool),X87),s(fun(X88,X86),X89))),s(X88,X90))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90)))))) )),
  inference(superposition,[],[f2341,f369])).
fof(f6059,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(subsumption_resolution,[],[f6058,f261])).
fof(f6058,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(forward_demodulation,[],[f6022,f1756])).
fof(f6022,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X31,bool),X34),s(X31,sK6(X32,X34,X31))))) | p(s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X32,X34,X31))))) | s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33)))) ) | $spl10),
  inference(superposition,[],[f319,f2498])).
fof(f5551,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(forward_demodulation,[],[f5292,f1756])).
fof(f5292,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun(X63,bool),'_dst_x0x1_2'(s(fun(X61,bool),X62),s(fun(X63,X61),X64))),s(X63,X65))) | s(bool,'F_0') = s(bool,happ(s(fun(X61,bool),X62),s(X61,happ(s(fun(X63,X61),X64),s(X63,X65)))))) ) | $spl10),
  inference(superposition,[],[f369,f2498])).
fof(f5512,plain,(
  ( ! [X282,X285,X287,X283,X284,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X283,bool),X284),s(X283,happ(s(fun(X282,X283),X285),s(X282,X286))))) | s(bool,X287) = s(bool,happ(s(fun(X282,bool),'_dst_x0x1_2'(s(fun(X283,bool),X284),s(fun(X282,X283),X285))),s(X282,X286))) | s(bool,'F_0') = s(bool,X287)) ) | $spl10),
  inference(superposition,[],[f3504,f369])).
fof(f5502,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (p(s(bool,happ(s(fun(X225,bool),X226),s(X225,happ(s(fun(X224,X225),X227),s(X224,X228)))))) | s(bool,X229) != s(bool,happ(s(fun(X224,bool),'_dst_x0x1_2'(s(fun(X225,bool),X226),s(fun(X224,X225),X227))),s(X224,X228))) | s(bool,'F_0') = s(bool,X229)) )),
  inference(superposition,[],[f2351,f369])).
fof(f5501,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (~p(s(bool,happ(s(fun(X219,bool),X220),s(X219,happ(s(fun(X218,X219),X221),s(X218,X222)))))) | s(bool,X223) = s(bool,happ(s(fun(X218,bool),'_dst_x0x1_2'(s(fun(X219,bool),X220),s(fun(X218,X219),X221))),s(X218,X222))) | s(bool,'F_0') = s(bool,X223)) )),
  inference(superposition,[],[f2349,f369])).
fof(f5497,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X197,bool),X198),s(X197,happ(s(fun(X196,X197),X199),s(X196,X200))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X197,bool),X198),s(fun(X196,X197),X199))),s(X196,X200)))) )),
  inference(superposition,[],[f535,f369])).
fof(f5496,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X195),'const.sum.INL_1'(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f507,f369])).
fof(f5494,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183)))) )),
  inference(superposition,[],[f465,f369])).
fof(f5484,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f274,f369])).
fof(f5483,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) )),
  inference(superposition,[],[f273,f369])).
fof(f5433,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X287,bool),'_dst_x0x1_2'(s(fun(X285,bool),X286),s(fun(X287,X285),X288))),s(X287,X289))) | s(bool,X290) = s(bool,happ(s(fun(X285,bool),X286),s(X285,happ(s(fun(X287,X285),X288),s(X287,X289))))) | s(bool,'F_0') = s(bool,X290)) ) | $spl10),
  inference(superposition,[],[f3504,f369])).
fof(f5423,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (p(s(bool,happ(s(fun(X229,bool),'_dst_x0x1_2'(s(fun(X227,bool),X228),s(fun(X229,X227),X230))),s(X229,X231)))) | s(bool,X232) != s(bool,happ(s(fun(X227,bool),X228),s(X227,happ(s(fun(X229,X227),X230),s(X229,X231))))) | s(bool,'F_0') = s(bool,X232)) )),
  inference(superposition,[],[f2351,f369])).
fof(f5422,plain,(
  ( ! [X222,X225,X223,X221,X226,X224] : (~p(s(bool,happ(s(fun(X223,bool),'_dst_x0x1_2'(s(fun(X221,bool),X222),s(fun(X223,X221),X224))),s(X223,X225)))) | s(bool,X226) = s(bool,happ(s(fun(X221,bool),X222),s(X221,happ(s(fun(X223,X221),X224),s(X223,X225))))) | s(bool,'F_0') = s(bool,X226)) )),
  inference(superposition,[],[f2349,f369])).
fof(f5418,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),X200),s(fun(X201,X199),X202))),s(X201,X203))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X199,bool),X200),s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))))) )),
  inference(superposition,[],[f535,f369])).
fof(f5417,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X198),'const.sum.INL_1'(s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X193,bool),X194),s(fun(X195,X193),X196))),s(X195,X197))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,happ(s(fun(X195,X193),X196),s(X195,X197)))))) )),
  inference(superposition,[],[f507,f369])).
fof(f5415,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X184,X182),X185))),s(X184,X186))))) | s(bool,'F_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X184,X182),X185),s(X184,X186)))))) )),
  inference(superposition,[],[f465,f369])).
fof(f5405,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X123,bool),X124),s(fun(X125,X123),X126))),s(X125,X127)))) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,happ(s(fun(X125,X123),X126),s(X125,X127)))))) )),
  inference(superposition,[],[f274,f369])).
fof(f5404,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'F_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) )),
  inference(superposition,[],[f273,f369])).
fof(f5295,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,X82) = s(bool,happ(s(fun(X79,bool),'_dst_x0x1_2'(s(fun(X77,bool),X78),s(fun(X79,X77),X80))),s(X79,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,happ(s(fun(X77,bool),X78),s(X77,happ(s(fun(X79,X77),X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f369,f2671])).
fof(f5294,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0x1_2'(s(fun(X71,bool),X72),s(fun(X73,X71),X74))),s(X73,X75))) | s(bool,X76) = s(bool,happ(s(fun(X71,bool),X72),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f369,f2671])).
fof(f5286,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X28,X26),X29))),s(X28,X30))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30)))))) )),
  inference(superposition,[],[f369,f268])).
fof(f5277,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5240,f261])).
fof(f5240,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK18(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4695,f507])).
fof(f5265,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f5232,f261])).
fof(f5232,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK18(X73,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71)))),X72,bool)))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f4695,f535])).
fof(f5263,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5230,f261])).
fof(f5230,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK18(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4695,f465])).
fof(f5256,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f5255,f261])).
fof(f5255,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f5219,f1756])).
fof(f5219,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f4695,f2498])).
fof(f4827,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4826,f261])).
fof(f4826,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4798,f1756])).
fof(f4798,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f357,f2498])).
fof(f4770,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4735,f264])).
fof(f4735,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK18(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4681,f2671])).
fof(f4700,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4691,f264])).
fof(f4691,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK18(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4594,f2671])).
fof(f4678,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4677,f261])).
fof(f4677,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f4649,f1756])).
fof(f4649,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f356,f2498])).
fof(f4618,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4617,f261])).
fof(f4617,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f4587,f1756])).
fof(f4587,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f4520,f2498])).
fof(f4519,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4372,f1756])).
fof(f4372,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f753,f2498])).
fof(f4446,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3504,f753])).
fof(f4436,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | s(bool,X113) != s(bool,sK4(X111,X112,bool)) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X112)))) )),
  inference(superposition,[],[f2351,f753])).
fof(f4435,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | s(bool,X110) = s(bool,sK4(X108,X109,bool)) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X109)))) )),
  inference(superposition,[],[f2349,f753])).
fof(f4431,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f535,f753])).
fof(f4430,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f507,f753])).
fof(f4428,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f465,f753])).
fof(f4418,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f274,f753])).
fof(f4417,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f273,f753])).
fof(f4375,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f753,f2671])).
fof(f4374,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f753,f2671])).
fof(f4366,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f753,f268])).
fof(f4305,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4248,f261])).
fof(f4248,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f439,f535])).
fof(f4225,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94)))))) )),
  inference(superposition,[],[f535,f344])).
fof(f4224,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90)))))) )),
  inference(superposition,[],[f535,f338])).
fof(f4223,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X88,bool),'const.sum.INR_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f535,f343])).
fof(f4222,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83)))))) )),
  inference(superposition,[],[f535,f339])).
fof(f4221,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f535,f323])).
fof(f4220,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f535,f314])).
fof(f4219,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f535,f340])).
fof(f4218,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK18(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f535,f359])).
fof(f4186,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4130,f261])).
fof(f4130,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f439,f507])).
fof(f4107,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f507,f344])).
fof(f4106,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X86),'const.sum.INL_1'(s(bool,X87)))))) )),
  inference(superposition,[],[f507,f338])).
fof(f4105,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f507,f343])).
fof(f4104,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,X80)))))) )),
  inference(superposition,[],[f507,f339])).
fof(f4103,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f507,f323])).
fof(f4102,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f507,f314])).
fof(f4101,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f507,f340])).
fof(f4100,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK18(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f507,f359])).
fof(f4070,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f4060,f1756])).
fof(f4060,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK14(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK14(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f353,f2498])).
fof(f4067,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4048,f1756])).
fof(f4048,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK14(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2498])).
fof(f4063,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,f(s(X33,sK14(X34,X35,bool,X33)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(X33,sK14(X34,X35,bool,X33))))) ) | $spl10),
  inference(superposition,[],[f353,f2671])).
fof(f4062,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') != s(bool,f(s(X29,sK14(X30,X31,bool,X29)))) | s(fun(X29,bool),X30) = s(fun(X29,bool),X31) | s(bool,X32) = s(bool,g(s(X29,sK14(X30,X31,bool,X29)))) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f353,f2671])).
fof(f4054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK14(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK14(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f353,f268])).
fof(f4051,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,f(s(bool,X36))) != s(X35,g(s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK14(X33,X34,X35,bool))) ) | $spl10),
  inference(superposition,[],[f353,f2671])).
fof(f4050,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK14(X29,X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f353,f2671])).
fof(f4042,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK14(X0,X1,X2,bool))) )),
  inference(superposition,[],[f353,f268])).
fof(f4015,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3504,f344])).
fof(f4014,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,X87)) ) | $spl10),
  inference(superposition,[],[f3504,f338])).
fof(f4013,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3504,f343])).
fof(f4012,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X80) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3504,f339])).
fof(f4011,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X75),s('type.num.num',X76))) | s(bool,X77) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl10),
  inference(superposition,[],[f3504,f323])).
fof(f4010,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3504,f314])).
fof(f4009,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3504,f340])).
fof(f4008,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK18(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3504,f359])).
fof(f4007,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3504,f307])).
fof(f3947,plain,(
  ( ! [X94,X95,X93,X96] : (s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(subsumption_resolution,[],[f3572,f264])).
fof(f3572,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,'F_0')) | s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(superposition,[],[f317,f2671])).
fof(f3945,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(subsumption_resolution,[],[f3567,f264])).
fof(f3567,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(superposition,[],[f412,f2671])).
fof(f3940,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3539,f264])).
fof(f3539,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f334,f2671])).
fof(f3598,plain,(
  ( ! [X177,X175,X176] : (~p(s(bool,happ(s(fun(bool,bool),X176),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176)))) | s(bool,X177) = s(bool,sK4(X175,X176,bool)) | s(bool,'F_0') = s(bool,X177)) ) | $spl10),
  inference(superposition,[],[f318,f2671])).
fof(f3503,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f273,f2671])).
fof(f3432,plain,(
  ( ! [X189,X187,X188] : (~p(s(bool,happ(s(fun(bool,bool),X188),s(bool,X189)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X188)))) | s(bool,'F_0') = s(bool,X189) | s(bool,'F_0') = s(bool,sK4(X187,X188,bool))) ) | $spl10),
  inference(superposition,[],[f318,f2671])).
fof(f3326,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2671,f344])).
fof(f3325,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') = s(bool,X71) | s(bool,X72) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X70),'const.sum.INL_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f2671,f338])).
fof(f3324,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2671,f343])).
fof(f3323,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X63,bool),'const.sum.INR_1'(s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2671,f339])).
fof(f3322,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61))) | s(bool,X62) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl10),
  inference(superposition,[],[f2671,f323])).
fof(f3321,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2671,f314])).
fof(f3320,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2671,f340])).
fof(f3319,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK18(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2671,f359])).
fof(f3255,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2671,f344])).
fof(f3254,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X83),'const.sum.INL_1'(s(bool,X84)))))) ) | $spl10),
  inference(superposition,[],[f2671,f338])).
fof(f3253,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2671,f343])).
fof(f3252,plain,(
  ( ! [X78,X76,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X76,bool),'const.sum.INR_1'(s(bool,X77)))))) ) | $spl10),
  inference(superposition,[],[f2671,f339])).
fof(f3251,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl10),
  inference(superposition,[],[f2671,f323])).
fof(f3250,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2671,f314])).
fof(f3249,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2671,f340])).
fof(f3248,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK18(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2671,f359])).
fof(f3247,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2671,f307])).
fof(f2833,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2734,f1756])).
fof(f2734,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f318,f2498])).
fof(f2819,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2818,f261])).
fof(f2818,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2703,f1756])).
fof(f2703,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f316,f2498])).
fof(f2817,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2816,f261])).
fof(f2816,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2702,f1756])).
fof(f2702,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f318,f2498])).
fof(f2785,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2666,f1756])).
fof(f2666,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(superposition,[],[f2498,f323])).
fof(f2783,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2664,f1756])).
fof(f2664,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2498,f340])).
fof(f2781,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2662,f1756])).
fof(f2662,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2498,f307])).
fof(f2498,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2475,f261])).
fof(f2475,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f2349,f1756])).
fof(f2497,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2474,f261])).
fof(f2474,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2349,f432])).
fof(f2496,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2473,f261])).
fof(f2473,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2349,f465])).
fof(f2428,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f465,f323])).
fof(f2352,plain,(
  ( ! [X21,X22] : (s(bool,X22) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2315,f261])).
fof(f2315,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X21)) )),
  inference(superposition,[],[f439,f465])).
fof(f2351,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X13)) )),
  inference(forward_demodulation,[],[f2350,f307])).
fof(f2350,plain,(
  ( ! [X14,X13] : (s(bool,X14) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X13))) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f2312,f261])).
fof(f2312,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,X14) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X13))) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f410,f465])).
fof(f2349,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X11)) )),
  inference(forward_demodulation,[],[f2348,f307])).
fof(f2348,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2311,f261])).
fof(f2311,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X11))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f337,f465])).
fof(f2344,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(forward_demodulation,[],[f2296,f307])).
fof(f2296,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f465,f344])).
fof(f2343,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(forward_demodulation,[],[f2295,f307])).
fof(f2295,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X35))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X34),'const.sum.INL_1'(s(bool,X35)))))) )),
  inference(superposition,[],[f465,f338])).
fof(f2342,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f2294,f307])).
fof(f2294,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f465,f343])).
fof(f2341,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f2293,f307])).
fof(f2293,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30)))))) )),
  inference(superposition,[],[f465,f339])).
fof(f2340,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2292,f307])).
fof(f2292,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f314])).
fof(f2339,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2290,f307])).
fof(f2290,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK18(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f465,f359])).
fof(f2291,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f465,f340])).
fof(f1398,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f340,f268])).
fof(f1290,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f318,f268])).
fof(f829,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool)),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f359,f268])).
fof(f717,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) )),
  inference(superposition,[],[f344,f268])).
fof(f664,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6)))))) )),
  inference(superposition,[],[f343,f268])).
fof(f615,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f314,f268])).
fof(f535,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,X4))))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f339,f268])).
fof(f507,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,X4))))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f338,f268])).
fof(f465,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f307,f268])).
fof(f274,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X28] : (s(bool,'F_0') = s(bool,X28) <=> ~p(s(bool,X28)))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.EQ_CLAUSES|split|2')).
fof(f273,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f264,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.modSem.v_|49||49||split|5',aHL_FALSITY)).
fof(f12045,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22)))) ) | $spl10),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4449,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6599,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8261,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10151,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10379,f10475,f10550,f10529,f10530,f10540,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140])).
fof(f8140,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f8127,f264])).
fof(f8127,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X21,bool)) | s(bool,X23) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),X22))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f7371,f2671])).
fof(f8135,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8119,f264])).
fof(f8119,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X0,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl10),
  inference(superposition,[],[f7371,f268])).
fof(f7368,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7355,f261])).
fof(f7355,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) | s(bool,X4) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f7353,f437])).
fof(f7353,plain,(
  ( ! [X66,X67] : (~p(s(bool,sK7('const.pair.FST_0',X66,bool))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))))) )),
  inference(forward_demodulation,[],[f7352,f340])).
fof(f7352,plain,(
  ( ! [X66,X67] : (~p(s(bool,sK7('const.pair.FST_0',X66,bool))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67))))) )),
  inference(forward_demodulation,[],[f7324,f343])).
fof(f7324,plain,(
  ( ! [X66,X67] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,sK7('const.pair.FST_0',X66,bool)),s(X66,sK8('const.pair.FST_0',X66,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67))))) )),
  inference(superposition,[],[f341,f340])).
fof(f11630,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f9712,f2671])).
fof(f11629,plain,(
  ( ! [X175,X173,X174] : (s(bool,X173) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X175,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X173,X174,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X174)))) ) | $spl26),
  inference(superposition,[],[f9712,f753])).
fof(f11628,plain,(
  ( ! [X171,X169,X172,X170] : (s(bool,X171) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X172,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X169,bool),'const.pair.,_2'(s(X169,X170),s(bool,X171)))))) ) | $spl26),
  inference(superposition,[],[f9712,f344])).
fof(f11627,plain,(
  ( ! [X167,X166,X168] : (s(bool,X167) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X168,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X166),'const.sum.INL_1'(s(bool,X167)))))) ) | $spl26),
  inference(superposition,[],[f9712,f338])).
fof(f11626,plain,(
  ( ! [X165,X163,X164,X162] : (s(bool,X163) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X165,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X163),s(X162,X164)))))) ) | $spl26),
  inference(superposition,[],[f9712,f343])).
fof(f11625,plain,(
  ( ! [X161,X159,X160] : (s(bool,X160) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X161,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X159,bool),'const.sum.INR_1'(s(bool,X160)))))) ) | $spl26),
  inference(superposition,[],[f9712,f339])).
fof(f11623,plain,(
  ( ! [X154,X152,X153] : (s(bool,X152) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X154,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X152),s(fun('type.num.num',bool),X153),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f9712,f314])).
fof(f11620,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,X141) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X142,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X139),bool),sK18(X140,X141,X139,bool)),s('type.list.list'(X139),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f9712,f359])).
fof(f11617,plain,(
  ( ! [X125,X126] : (s(bool,X125) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X126,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X125)))) ) | $spl26),
  inference(superposition,[],[f9712,f307])).
fof(f11517,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f11500])).
fof(f11500,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f11281,f437])).
fof(f11400,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.ind_type.recspace'(X16),'const.ind_type.CONSTR_3'(s('type.num.num',X17),s(X16,X18),s(fun('type.num.num','type.ind_type.recspace'(X16)),sK15(happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)),X19,fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))))) != s('type.ind_type.recspace'(X16),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),X19),s(fun('type.num.num','type.ind_type.recspace'(X16)),sK15(happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)),X19,fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))))) | s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),X19) = s(fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16)),happ(s(fun(X16,fun(fun('type.num.num','type.ind_type.recspace'(X16)),'type.ind_type.recspace'(X16))),'const.ind_type.CONSTR_1'(s('type.num.num',X17))),s(X16,X18)))) )),
  inference(superposition,[],[f355,f324])).
fof(f11399,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.ind_type.recspace'(X12),'const.ind_type.CONSTR_3'(s('type.num.num',X13),s(X12,X14),s(fun('type.num.num','type.ind_type.recspace'(X12)),sK15(X15,happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)),fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))))) != s('type.ind_type.recspace'(X12),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),X15),s(fun('type.num.num','type.ind_type.recspace'(X12)),sK15(X15,happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)),fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))))) | s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),X15) = s(fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12)),happ(s(fun(X12,fun(fun('type.num.num','type.ind_type.recspace'(X12)),'type.ind_type.recspace'(X12))),'const.ind_type.CONSTR_1'(s('type.num.num',X13))),s(X12,X14)))) )),
  inference(superposition,[],[f355,f324])).
fof(f11250,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107))) | s(bool,'T_0') = s(bool,sK4(X107,X108,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X108)))) ) | $spl10),
  inference(superposition,[],[f10870,f753])).
fof(f11281,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f11188,f8144])).
fof(f11188,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('const.pair.FST_0',X0,bool))))) ) | ($spl10 | $spl26)),
  inference(unit_resulting_resolution,[],[f8231,f10870])).
fof(f11003,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f10949,f261])).
fof(f10949,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X112),bool),sK18(X113,'|dollar|exists_1'(s(fun('type.list.list'(X110),bool),sK18(X111,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X110,bool))),X112,bool))))) ) | $spl10),
  inference(superposition,[],[f4695,f5268])).
fof(f10998,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f10923,f261])).
fof(f10923,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f439,f5268])).
fof(f10870,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f10772,f264])).
fof(f10772,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f269,f2781])).
fof(f10475,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2130,f437])).
fof(f10314,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2128,f437])).
fof(f10086,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f2124,f437])).
fof(f10019,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f10002])).
fof(f10002,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X3),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f9853,f437])).
fof(f9822,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X161),'const.sum.INL_1'(s(bool,X159))))) | s(bool,'T_0') = s(bool,sK4(X159,X160,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X160)))) )),
  inference(superposition,[],[f6343,f753])).
fof(f9821,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X158),'const.sum.INL_1'(s(bool,X157))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))))) )),
  inference(superposition,[],[f6343,f344])).
fof(f9820,plain,(
  ( ! [X154,X152,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X154),'const.sum.INL_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X153)))))) )),
  inference(superposition,[],[f6343,f338])).
fof(f9819,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X151),'const.sum.INL_1'(s(bool,X149))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X149),s(X148,X150)))))) )),
  inference(superposition,[],[f6343,f343])).
fof(f9818,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X147),'const.sum.INL_1'(s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X146)))))) )),
  inference(superposition,[],[f6343,f339])).
fof(f9816,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X140),'const.sum.INL_1'(s(bool,X138))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6343,f314])).
fof(f9813,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X125),bool),sK18(X126,X127,X125,bool)),s('type.list.list'(X125),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6343,f359])).
fof(f9853,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9763,f8144])).
fof(f9763,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8231,f6343])).
fof(f9712,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f9695])).
fof(f9695,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f9550,f437])).
fof(f9519,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X161,bool),'const.sum.INR_1'(s(bool,X159))))) | s(bool,'T_0') = s(bool,sK4(X159,X160,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X160)))) )),
  inference(superposition,[],[f6195,f753])).
fof(f9518,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X158,bool),'const.sum.INR_1'(s(bool,X157))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),'const.pair.,_2'(s(X155,X156),s(bool,X157)))))) )),
  inference(superposition,[],[f6195,f344])).
fof(f9517,plain,(
  ( ! [X154,X152,X153] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X154,bool),'const.sum.INR_1'(s(bool,X153))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X152),'const.sum.INL_1'(s(bool,X153)))))) )),
  inference(superposition,[],[f6195,f338])).
fof(f9516,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X151,bool),'const.sum.INR_1'(s(bool,X149))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X148),'const.pair.,_2'(s(bool,X149),s(X148,X150)))))) )),
  inference(superposition,[],[f6195,f343])).
fof(f9515,plain,(
  ( ! [X146,X147,X145] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X147,bool),'const.sum.INR_1'(s(bool,X146))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X145,bool),'const.sum.INR_1'(s(bool,X146)))))) )),
  inference(superposition,[],[f6195,f339])).
fof(f9513,plain,(
  ( ! [X140,X138,X139] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X140,bool),'const.sum.INR_1'(s(bool,X138))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6195,f314])).
fof(f9510,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X125),bool),sK18(X126,X127,X125,bool)),s('type.list.list'(X125),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6195,f359])).
fof(f9550,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(forward_demodulation,[],[f9460,f8144])).
fof(f9460,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK7('const.pair.FST_0',X1,bool))))))) ) | $spl26),
  inference(unit_resulting_resolution,[],[f8231,f6195])).
fof(f9285,plain,(
  ( ! [X707,X712,X708,X710,X706,X713,X709,X711] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X712),s('type.list.list'('type.modSem.v'),X713))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X707,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),happ(s(fun(X706,fun(X707,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),X708),s(X706,X709))),s(X707,X710))),s('type.list.list'('type.modSem.v'),X711))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X712) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'(X706,X707),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),sK19(X708,X707,X706,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.pair.prod'(X706,X707),'const.pair.,_2'(s(X706,X709),s(X707,X710)))))) )),
  inference(superposition,[],[f394,f367])).
fof(f9279,plain,(
  ( ! [X666,X668,X662,X664,X667,X669,X663,X665] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X668),s('type.list.list'('type.modSem.v'),X669))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X667),s('type.list.list'('type.modSem.v'),happ(s(fun(X663,'type.list.list'('type.modSem.v')),happ(s(fun(X662,fun(X663,'type.list.list'('type.modSem.v'))),X664),s(X662,X665))),s(X663,X666))))) | s('type.list.list'('type.modSem.v'),X669) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'(X662,X663),'type.list.list'('type.modSem.v')),sK19(X664,X663,X662,'type.list.list'('type.modSem.v'))),s('type.pair.prod'(X662,X663),'const.pair.,_2'(s(X662,X665),s(X663,X666)))))) )),
  inference(superposition,[],[f395,f367])).
fof(f9347,plain,(
  ( ! [X466,X461,X463,X459,X465,X460,X462,X464] : (s(fun(X459,X460),X466) = s(fun(X459,X460),happ(s(fun(X462,fun(X459,X460)),happ(s(fun(X461,fun(X462,fun(X459,X460))),X463),s(X461,X464))),s(X462,X465))) | s(X460,happ(s(fun(X459,X460),X466),s(X459,sK15(happ(s(fun('type.pair.prod'(X461,X462),fun(X459,X460)),sK19(X463,X462,X461,fun(X459,X460))),s('type.pair.prod'(X461,X462),'const.pair.,_2'(s(X461,X464),s(X462,X465)))),X466,X459,X460)))) != s(X460,happ(s(fun(X459,X460),happ(s(fun(X462,fun(X459,X460)),happ(s(fun(X461,fun(X462,fun(X459,X460))),X463),s(X461,X464))),s(X462,X465))),s(X459,sK15(happ(s(fun('type.pair.prod'(X461,X462),fun(X459,X460)),sK19(X463,X462,X461,fun(X459,X460))),s('type.pair.prod'(X461,X462),'const.pair.,_2'(s(X461,X464),s(X462,X465)))),X466,X459,X460))))) )),
  inference(forward_demodulation,[],[f9252,f367])).
fof(f9252,plain,(
  ( ! [X466,X461,X463,X459,X465,X460,X462,X464] : (s(X460,happ(s(fun(X459,X460),X466),s(X459,sK15(happ(s(fun('type.pair.prod'(X461,X462),fun(X459,X460)),sK19(X463,X462,X461,fun(X459,X460))),s('type.pair.prod'(X461,X462),'const.pair.,_2'(s(X461,X464),s(X462,X465)))),X466,X459,X460)))) != s(X460,happ(s(fun(X459,X460),happ(s(fun(X462,fun(X459,X460)),happ(s(fun(X461,fun(X462,fun(X459,X460))),X463),s(X461,X464))),s(X462,X465))),s(X459,sK15(happ(s(fun('type.pair.prod'(X461,X462),fun(X459,X460)),sK19(X463,X462,X461,fun(X459,X460))),s('type.pair.prod'(X461,X462),'const.pair.,_2'(s(X461,X464),s(X462,X465)))),X466,X459,X460)))) | s(fun(X459,X460),X466) = s(fun(X459,X460),happ(s(fun('type.pair.prod'(X461,X462),fun(X459,X460)),sK19(X463,X462,X461,fun(X459,X460))),s('type.pair.prod'(X461,X462),'const.pair.,_2'(s(X461,X464),s(X462,X465)))))) )),
  inference(superposition,[],[f355,f367])).
fof(f9346,plain,(
  ( ! [X453,X455,X451,X457,X452,X454,X456,X458] : (s(fun(X451,X452),X458) = s(fun(X451,X452),happ(s(fun(X454,fun(X451,X452)),happ(s(fun(X453,fun(X454,fun(X451,X452))),X455),s(X453,X456))),s(X454,X457))) | s(X452,happ(s(fun(X451,X452),X458),s(X451,sK15(X458,happ(s(fun('type.pair.prod'(X453,X454),fun(X451,X452)),sK19(X455,X454,X453,fun(X451,X452))),s('type.pair.prod'(X453,X454),'const.pair.,_2'(s(X453,X456),s(X454,X457)))),X451,X452)))) != s(X452,happ(s(fun(X451,X452),happ(s(fun(X454,fun(X451,X452)),happ(s(fun(X453,fun(X454,fun(X451,X452))),X455),s(X453,X456))),s(X454,X457))),s(X451,sK15(X458,happ(s(fun('type.pair.prod'(X453,X454),fun(X451,X452)),sK19(X455,X454,X453,fun(X451,X452))),s('type.pair.prod'(X453,X454),'const.pair.,_2'(s(X453,X456),s(X454,X457)))),X451,X452))))) )),
  inference(forward_demodulation,[],[f9251,f367])).
fof(f9251,plain,(
  ( ! [X453,X455,X451,X457,X452,X454,X456,X458] : (s(X452,happ(s(fun(X451,X452),X458),s(X451,sK15(X458,happ(s(fun('type.pair.prod'(X453,X454),fun(X451,X452)),sK19(X455,X454,X453,fun(X451,X452))),s('type.pair.prod'(X453,X454),'const.pair.,_2'(s(X453,X456),s(X454,X457)))),X451,X452)))) != s(X452,happ(s(fun(X451,X452),happ(s(fun(X454,fun(X451,X452)),happ(s(fun(X453,fun(X454,fun(X451,X452))),X455),s(X453,X456))),s(X454,X457))),s(X451,sK15(X458,happ(s(fun('type.pair.prod'(X453,X454),fun(X451,X452)),sK19(X455,X454,X453,fun(X451,X452))),s('type.pair.prod'(X453,X454),'const.pair.,_2'(s(X453,X456),s(X454,X457)))),X451,X452)))) | s(fun(X451,X452),X458) = s(fun(X451,X452),happ(s(fun('type.pair.prod'(X453,X454),fun(X451,X452)),sK19(X455,X454,X453,fun(X451,X452))),s('type.pair.prod'(X453,X454),'const.pair.,_2'(s(X453,X456),s(X454,X457)))))) )),
  inference(superposition,[],[f355,f367])).
fof(f9181,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))) = s(X6,happ(s(fun('type.pair.prod'(X1,X2),X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)))),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(superposition,[],[f369,f367])).
fof(f8872,plain,(
  ( ! [X146,X147,X145] : (s(bool,X145) != s(bool,sK7('const.pair.FST_0',X147,bool)) | s(bool,'F_0') = s(bool,sK4(X145,X146,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X146)))) ) | $spl26),
  inference(superposition,[],[f8241,f753])).
fof(f8871,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X143) != s(bool,sK7('const.pair.FST_0',X144,bool)) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))))) ) | $spl26),
  inference(superposition,[],[f8241,f344])).
fof(f8870,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) != s(bool,sK7('const.pair.FST_0',X140,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X138),'const.sum.INL_1'(s(bool,X139)))))) ) | $spl26),
  inference(superposition,[],[f8241,f338])).
fof(f8869,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X135) != s(bool,sK7('const.pair.FST_0',X137,bool)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X134),'const.pair.,_2'(s(bool,X135),s(X134,X136)))))) ) | $spl26),
  inference(superposition,[],[f8241,f343])).
fof(f8868,plain,(
  ( ! [X132,X133,X131] : (s(bool,X132) != s(bool,sK7('const.pair.FST_0',X133,bool)) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,X132)))))) ) | $spl26),
  inference(superposition,[],[f8241,f339])).
fof(f8866,plain,(
  ( ! [X125,X126,X124] : (s(bool,X124) != s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X124),s(fun('type.num.num',bool),X125),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8241,f314])).
fof(f8864,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X119) != s(bool,sK7('const.pair.FST_0',X120,bool)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X117),bool),sK18(X118,X119,X117,bool)),s('type.list.list'(X117),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8241,f359])).
fof(f8861,plain,(
  ( ! [X103,X104] : (s(bool,X103) != s(bool,sK7('const.pair.FST_0',X104,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103)))) ) | $spl26),
  inference(superposition,[],[f8241,f307])).
fof(f8804,plain,(
  ( ! [X14,X13] : (s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X14),s('type.list.list'('type.modSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13)),X14,'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.modSem.v'),sK15(happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13)),X14,'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X14) = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13)))) )),
  inference(superposition,[],[f355,f332])).
fof(f8803,plain,(
  ( ! [X12,X11] : (s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X12),s('type.list.list'('type.modSem.v'),sK15(X12,happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X11)),'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X11),s('type.list.list'('type.modSem.v'),sK15(X12,happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X11)),'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X12) = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),'type.modSem.v')),'const.modSem.Conv_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X11)))) )),
  inference(superposition,[],[f355,f332])).
fof(f8452,plain,(
  ( ! [X146,X147,X145] : (s(bool,X145) = s(bool,sK7('const.pair.FST_0',X147,bool)) | s(bool,'T_0') = s(bool,sK4(X145,X146,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X146)))) ) | $spl26),
  inference(superposition,[],[f8230,f753])).
fof(f8451,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X143) = s(bool,sK7('const.pair.FST_0',X144,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))))) ) | $spl26),
  inference(superposition,[],[f8230,f344])).
fof(f8450,plain,(
  ( ! [X140,X138,X139] : (s(bool,X139) = s(bool,sK7('const.pair.FST_0',X140,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X138),'const.sum.INL_1'(s(bool,X139)))))) ) | $spl26),
  inference(superposition,[],[f8230,f338])).
fof(f8449,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X135) = s(bool,sK7('const.pair.FST_0',X137,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X134),'const.pair.,_2'(s(bool,X135),s(X134,X136)))))) ) | $spl26),
  inference(superposition,[],[f8230,f343])).
fof(f8448,plain,(
  ( ! [X132,X133,X131] : (s(bool,X132) = s(bool,sK7('const.pair.FST_0',X133,bool)) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,X132)))))) ) | $spl26),
  inference(superposition,[],[f8230,f339])).
fof(f8446,plain,(
  ( ! [X125,X126,X124] : (s(bool,X124) = s(bool,sK7('const.pair.FST_0',X126,bool)) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X124),s(fun('type.num.num',bool),X125),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f8230,f314])).
fof(f8444,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X119) = s(bool,sK7('const.pair.FST_0',X120,bool)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X117),bool),sK18(X118,X119,X117,bool)),s('type.list.list'(X117),'const.list.NIL_0')))) ) | $spl26),
  inference(superposition,[],[f8230,f359])).
fof(f8441,plain,(
  ( ! [X103,X104] : (s(bool,X103) = s(bool,sK7('const.pair.FST_0',X104,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X103)))) ) | $spl26),
  inference(superposition,[],[f8230,f307])).
fof(f8260,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X60),s('type.list.list'('type.modSem.v'),X61))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X62),s('type.list.list'('type.modSem.v'),X63))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X62) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.SND_1'(s('type.pair.prod'(X58,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.,_2'(s(X58,X59),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X60)))))) )),
  inference(superposition,[],[f394,f344])).
fof(f8259,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X54),s('type.list.list'('type.modSem.v'),X55))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X56),s('type.list.list'('type.modSem.v'),X57))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X56) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.sum.OUTL_1'(s('type.sum.sum'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X53),'const.sum.INL_1'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X54)))))) )),
  inference(superposition,[],[f394,f338])).
fof(f8258,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X48),s('type.list.list'('type.modSem.v'),X50))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51),s('type.list.list'('type.modSem.v'),X52))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X47),'const.pair.,_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X48),s(X47,X49)))))) )),
  inference(superposition,[],[f394,f343])).
fof(f8257,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X43),s('type.list.list'('type.modSem.v'),X44))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45),s('type.list.list'('type.modSem.v'),X46))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.sum.OUTR_1'(s('type.sum.sum'(X42,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.sum.INR_1'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X43)))))) )),
  inference(superposition,[],[f394,f339])).
fof(f8256,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.modSem.v'),X41))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X37),s('type.num.num',X38))),s('type.list.list'('type.modSem.v'),X39))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X36),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X37),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38)))))) )),
  inference(superposition,[],[f394,f323])).
fof(f8255,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X31),s('type.list.list'('type.modSem.v'),X33))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34),s('type.list.list'('type.modSem.v'),X35))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.ind_type.FCONS_3'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X31),s(fun('type.num.num','type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f394,f314])).
fof(f8254,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.modSem.v'),X30))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'const.pair.FST_1'(s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),X27))),s('type.list.list'('type.modSem.v'),X28))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'const.pair.FST_0'),s('type.pair.prod'('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X26),X27)))) )),
  inference(superposition,[],[f394,f340])).
fof(f8253,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X22),s('type.list.list'('type.modSem.v'),X23))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.modSem.v'),X25))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.list.list'(X20),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),sK18(X21,X22,X20,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')))),s('type.list.list'(X20),'const.list.NIL_0')))) )),
  inference(superposition,[],[f394,f359])).
fof(f8252,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.modSem.v'),X19))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(X13,happ(s(fun(X12,X13),X15),s(X12,X16))))),s('type.list.list'('type.modSem.v'),X17))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X12,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X13,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X14),s(fun(X12,X13),X15))),s(X12,X16)))) )),
  inference(superposition,[],[f394,f369])).
fof(f8251,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10),s('type.list.list'('type.modSem.v'),X11))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X6,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_dst_x0x1_2'(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(fun(X6,X4),X7))),s(X6,X8))),s('type.list.list'('type.modSem.v'),X9))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X10) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun(X4,'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),X5),s(X4,happ(s(fun(X6,X4),X7),s(X6,X8)))))) )),
  inference(superposition,[],[f394,f369])).
fof(f8250,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),'type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'))),'_i_0_0'),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0)))) )),
  inference(superposition,[],[f394,f307])).
fof(f8241,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(subsumption_resolution,[],[f8188,f264])).
fof(f8188,plain,(
  ( ! [X41,X40] : (p(s(bool,'F_0')) | s(bool,X41) != s(bool,sK7('const.pair.FST_0',X40,bool)) | s(bool,'F_0') = s(bool,X41)) ) | $spl26),
  inference(superposition,[],[f2351,f8144])).
fof(f8231,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(subsumption_resolution,[],[f8165,f264])).
fof(f8165,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7('const.pair.FST_0',X1,bool))) ) | $spl26),
  inference(superposition,[],[f269,f8144])).
fof(f8230,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(subsumption_resolution,[],[f8153,f424])).
fof(f8153,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7('const.pair.FST_0',X2,bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f8144,f437])).
fof(f8101,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f8067,f261])).
fof(f8067,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f352,f437])).
fof(f7560,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f7526,f261])).
fof(f7526,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f351,f437])).
fof(f7559,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(subsumption_resolution,[],[f7524,f261])).
fof(f7524,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f351,f432])).
fof(f6840,plain,(
  ( ! [X0] : (s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X0),s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v',X0))),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f331,f263])).
fof(f6632,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.list.list_size_2'(s(fun(X67,'type.num.num'),X68),s('type.list.list'(X67),sK15('const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)),X69,'type.list.list'(X67),'type.num.num')))) != s('type.num.num',happ(s(fun('type.list.list'(X67),'type.num.num'),X69),s('type.list.list'(X67),sK15('const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)),X69,'type.list.list'(X67),'type.num.num')))) | s(fun('type.list.list'(X67),'type.num.num'),X69) = s(fun('type.list.list'(X67),'type.num.num'),'const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)))) )),
  inference(superposition,[],[f355,f315])).
fof(f6631,plain,(
  ( ! [X66] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK15('const.modSem.Vectorv_0',X66,'type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X66),s('type.list.list'('type.modSem.v'),sK15('const.modSem.Vectorv_0',X66,'type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0') = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X66)) )),
  inference(superposition,[],[f355,f266])).
fof(f6630,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f355,f368])).
fof(f6619,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X20,X21,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X20,X21,X19,bool))))) )),
  inference(superposition,[],[f355,f437])).
fof(f6618,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f355,f437])).
fof(f6615,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f355,f340])).
fof(f6614,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f355,f369])).
fof(f6613,plain,(
  ( ! [X0,X1] : (s(X0,sK15('_i_0_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15('_i_0_0',X1,X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f355,f307])).
fof(f6602,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,X9))) )),
  inference(superposition,[],[f355,f437])).
fof(f6601,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK15(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f355,f437])).
fof(f6657,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65))))) )),
  inference(forward_demodulation,[],[f6598,f344])).
fof(f6598,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15('const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X69,X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f355,f344])).
fof(f6656,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60))))) )),
  inference(forward_demodulation,[],[f6597,f338])).
fof(f6597,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15('const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X63,X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f355,f338])).
fof(f6655,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54))))) )),
  inference(forward_demodulation,[],[f6596,f343])).
fof(f6596,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X58,X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f355,f343])).
fof(f6654,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49))))) )),
  inference(forward_demodulation,[],[f6595,f339])).
fof(f6595,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15('const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X52,X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f355,f339])).
fof(f6653,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43))))) )),
  inference(forward_demodulation,[],[f6594,f323])).
fof(f6594,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15('const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X47,X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f355,f323])).
fof(f6652,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38))))) )),
  inference(forward_demodulation,[],[f6593,f314])).
fof(f6593,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15('const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X41,X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f355,f314])).
fof(f6651,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33))))) )),
  inference(forward_demodulation,[],[f6592,f340])).
fof(f6592,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f355,f340])).
fof(f6650,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27))))) )),
  inference(forward_demodulation,[],[f6591,f359])).
fof(f6591,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f355,f359])).
fof(f6590,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f355,f369])).
fof(f6589,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f6649,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7))))) )),
  inference(forward_demodulation,[],[f6588,f307])).
fof(f6588,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f355,f307])).
fof(f6648,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1))))) )),
  inference(forward_demodulation,[],[f6587,f368])).
fof(f6587,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f355,f368])).
fof(f6586,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.list.list_size_2'(s(fun(X67,'type.num.num'),X68),s('type.list.list'(X67),sK15(X69,'const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)),'type.list.list'(X67),'type.num.num')))) != s('type.num.num',happ(s(fun('type.list.list'(X67),'type.num.num'),X69),s('type.list.list'(X67),sK15(X69,'const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)),'type.list.list'(X67),'type.num.num')))) | s(fun('type.list.list'(X67),'type.num.num'),X69) = s(fun('type.list.list'(X67),'type.num.num'),'const.list.list_size_1'(s(fun(X67,'type.num.num'),X68)))) )),
  inference(superposition,[],[f355,f315])).
fof(f6585,plain,(
  ( ! [X66] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK15(X66,'const.modSem.Vectorv_0','type.list.list'('type.modSem.v'),'type.modSem.v')))) != s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X66),s('type.list.list'('type.modSem.v'),sK15(X66,'const.modSem.Vectorv_0','type.list.list'('type.modSem.v'),'type.modSem.v')))) | s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0') = s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),X66)) )),
  inference(superposition,[],[f355,f266])).
fof(f6584,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK15(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f355,f368])).
fof(f6573,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK15(X21,X20,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK15(X21,X20,X19,bool))))) )),
  inference(superposition,[],[f355,f437])).
fof(f6572,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK15(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK15(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f355,f437])).
fof(f6569,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK15(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f355,f340])).
fof(f6568,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK15(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f355,f369])).
fof(f6567,plain,(
  ( ! [X0,X1] : (s(X0,sK15(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK15(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f355,f307])).
fof(f6643,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(fun(X64,X65),X68) = s(fun(X64,X65),X69) | s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65))))) )),
  inference(forward_demodulation,[],[f6565,f344])).
fof(f6565,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(X65,happ(s(fun(X64,X65),X68),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) != s(X65,happ(s(fun(X64,X65),X69),s(X64,sK15(X69,'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))),X64,X65)))) | s(fun(X64,X65),X69) = s(fun(X64,X65),'const.pair.SND_1'(s('type.pair.prod'(X66,fun(X64,X65)),'const.pair.,_2'(s(X66,X67),s(fun(X64,X65),X68)))))) )),
  inference(superposition,[],[f355,f344])).
fof(f6642,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(fun(X59,X60),X62) = s(fun(X59,X60),X63) | s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60))))) )),
  inference(forward_demodulation,[],[f6564,f338])).
fof(f6564,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(X60,happ(s(fun(X59,X60),X62),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) != s(X60,happ(s(fun(X59,X60),X63),s(X59,sK15(X63,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))),X59,X60)))) | s(fun(X59,X60),X63) = s(fun(X59,X60),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X59,X60),X61),'const.sum.INL_1'(s(fun(X59,X60),X62)))))) )),
  inference(superposition,[],[f355,f338])).
fof(f6641,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(fun(X53,X54),X56) = s(fun(X53,X54),X58) | s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54))))) )),
  inference(forward_demodulation,[],[f6563,f343])).
fof(f6563,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X54,happ(s(fun(X53,X54),X56),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) != s(X54,happ(s(fun(X53,X54),X58),s(X53,sK15(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))),X53,X54)))) | s(fun(X53,X54),X58) = s(fun(X53,X54),'const.pair.FST_1'(s('type.pair.prod'(fun(X53,X54),X55),'const.pair.,_2'(s(fun(X53,X54),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f355,f343])).
fof(f6640,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(fun(X48,X49),X51) = s(fun(X48,X49),X52) | s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49))))) )),
  inference(forward_demodulation,[],[f6562,f339])).
fof(f6562,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(X49,happ(s(fun(X48,X49),X51),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),X52),s(X48,sK15(X52,'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))),X48,X49)))) | s(fun(X48,X49),X52) = s(fun(X48,X49),'const.sum.OUTR_1'(s('type.sum.sum'(X50,fun(X48,X49)),'const.sum.INR_1'(s(fun(X48,X49),X51)))))) )),
  inference(superposition,[],[f355,f339])).
fof(f6639,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(fun(X42,X43),X47) = s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))) | s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43))))) )),
  inference(forward_demodulation,[],[f6561,f323])).
fof(f6561,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s(X43,happ(s(fun(X42,X43),X47),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) != s(X43,happ(s(fun(X42,X43),happ(s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num',X46))),s(X42,sK15(X47,'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))),X42,X43)))) | s(fun(X42,X43),X47) = s(fun(X42,X43),'const.ind_type.FCONS_3'(s(fun(X42,X43),X44),s(fun('type.num.num',fun(X42,X43)),X45),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f355,f323])).
fof(f6638,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(fun(X37,X38),X39) = s(fun(X37,X38),X41) | s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38))))) )),
  inference(forward_demodulation,[],[f6560,f314])).
fof(f6560,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,happ(s(fun(X37,X38),X39),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X41),s(X37,sK15(X41,'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')),X37,X38)))) | s(fun(X37,X38),X41) = s(fun(X37,X38),'const.ind_type.FCONS_3'(s(fun(X37,X38),X39),s(fun('type.num.num',fun(X37,X38)),X40),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f355,f314])).
fof(f6637,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33))))) )),
  inference(forward_demodulation,[],[f6559,f340])).
fof(f6559,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK15(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f355,f340])).
fof(f6636,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27))))) )),
  inference(forward_demodulation,[],[f6558,f359])).
fof(f6558,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK15(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK18(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f355,f359])).
fof(f6557,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK15(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f355,f369])).
fof(f6556,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK15(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f355,f369])).
fof(f6635,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7))))) )),
  inference(forward_demodulation,[],[f6555,f307])).
fof(f6555,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK15(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f355,f307])).
fof(f6634,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) )),
  inference(forward_demodulation,[],[f6554,f368])).
fof(f6554,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK15(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f355,f368])).
fof(f6343,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6271,f264])).
fof(f6271,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f269,f2343])).
fof(f6195,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6124,f264])).
fof(f6124,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f269,f2341])).
fof(f5531,plain,(
  ( ! [X430,X436,X432,X434,X431,X437,X433,X435] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X435),s('type.list.list'('type.modSem.v'),X436))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X437),s('type.list.list'('type.modSem.v'),happ(s(fun(X431,'type.list.list'('type.modSem.v')),X432),s(X431,happ(s(fun(X430,X431),X433),s(X430,X434))))))) | s('type.list.list'('type.modSem.v'),X436) = s('type.list.list'('type.modSem.v'),happ(s(fun(X430,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X431,'type.list.list'('type.modSem.v')),X432),s(fun(X430,X431),X433))),s(X430,X434)))) )),
  inference(superposition,[],[f395,f369])).
fof(f5525,plain,(
  ( ! [X389,X391,X393,X395,X390,X392,X394] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X389,bool),happ(s(fun(X391,fun(X389,bool)),X392),s(X391,happ(s(fun(X390,X391),X393),s(X390,X394))))))) | s(X389,X395) = s(X389,sK4(X395,happ(s(fun(X390,fun(X389,bool)),'_dst_x0x1_2'(s(fun(X391,fun(X389,bool)),X392),s(fun(X390,X391),X393))),s(X390,X394)),X389))) )),
  inference(superposition,[],[f753,f369])).
fof(f5516,plain,(
  ( ! [X313,X315,X310,X316,X318,X312,X314,X311,X317] : (s(fun(X310,X311),X317) != s(fun(X310,X311),happ(s(fun(X313,fun(X310,X311)),X314),s(X313,happ(s(fun(X312,X313),X315),s(X312,X316))))) | s(X311,happ(s(fun(X310,X311),X317),s(X310,X318))) = s(X311,happ(s(fun(X310,X311),happ(s(fun(X312,fun(X310,X311)),'_dst_x0x1_2'(s(fun(X313,fun(X310,X311)),X314),s(fun(X312,X313),X315))),s(X312,X316))),s(X310,X318)))) )),
  inference(superposition,[],[f354,f369])).
fof(f5563,plain,(
  ( ! [X212,X210,X208,X211,X209] : (s(bool,'T_0') != s(bool,happ(s(fun(X209,bool),X210),s(X209,happ(s(fun(X208,X209),X211),s(X208,X212))))) | s(bool,'T_0') = s(bool,happ(s(fun(X208,bool),'_dst_x0x1_2'(s(fun(X209,bool),X210),s(fun(X208,X209),X211))),s(X208,X212)))) ) | $spl10),
  inference(forward_demodulation,[],[f5499,f1756])).
fof(f5499,plain,(
  ( ! [X212,X210,X208,X211,X209] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X209,bool),X210),s(X209,happ(s(fun(X208,X209),X211),s(X208,X212))))) | s(bool,'T_0') = s(bool,happ(s(fun(X208,bool),'_dst_x0x1_2'(s(fun(X209,bool),X210),s(fun(X208,X209),X211))),s(X208,X212)))) ) | $spl10),
  inference(superposition,[],[f1849,f369])).
fof(f5476,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK4(X85,'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84)),X81))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f318,f369])).
fof(f5453,plain,(
  ( ! [X445,X447,X443,X449,X444,X446,X448,X450] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X448),s('type.list.list'('type.modSem.v'),X449))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X450),s('type.list.list'('type.modSem.v'),happ(s(fun(X445,'type.list.list'('type.modSem.v')),'_dst_x0x1_2'(s(fun(X443,'type.list.list'('type.modSem.v')),X444),s(fun(X445,X443),X446))),s(X445,X447))))) | s('type.list.list'('type.modSem.v'),X449) = s('type.list.list'('type.modSem.v'),happ(s(fun(X443,'type.list.list'('type.modSem.v')),X444),s(X443,happ(s(fun(X445,X443),X446),s(X445,X447)))))) )),
  inference(superposition,[],[f395,f369])).
fof(f5447,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X399,bool),happ(s(fun(X402,fun(X399,bool)),'_dst_x0x1_2'(s(fun(X400,fun(X399,bool)),X401),s(fun(X402,X400),X403))),s(X402,X404))))) | s(X399,X405) = s(X399,sK4(X405,happ(s(fun(X400,fun(X399,bool)),X401),s(X400,happ(s(fun(X402,X400),X403),s(X402,X404)))),X399))) )),
  inference(superposition,[],[f753,f369])).
fof(f5438,plain,(
  ( ! [X325,X327,X321,X323,X324,X326,X320,X322,X328] : (s(fun(X320,X321),X327) != s(fun(X320,X321),happ(s(fun(X324,fun(X320,X321)),'_dst_x0x1_2'(s(fun(X322,fun(X320,X321)),X323),s(fun(X324,X322),X325))),s(X324,X326))) | s(X321,happ(s(fun(X320,X321),X327),s(X320,X328))) = s(X321,happ(s(fun(X320,X321),happ(s(fun(X322,fun(X320,X321)),X323),s(X322,happ(s(fun(X324,X322),X325),s(X324,X326))))),s(X320,X328)))) )),
  inference(superposition,[],[f354,f369])).
fof(f5557,plain,(
  ( ! [X212,X215,X213,X211,X214] : (s(bool,'T_0') != s(bool,happ(s(fun(X213,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X213,X211),X214))),s(X213,X215))) | s(bool,'T_0') = s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X213,X211),X214),s(X213,X215)))))) ) | $spl10),
  inference(forward_demodulation,[],[f5420,f1756])).
fof(f5420,plain,(
  ( ! [X212,X215,X213,X211,X214] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,happ(s(fun(X213,bool),'_dst_x0x1_2'(s(fun(X211,bool),X212),s(fun(X213,X211),X214))),s(X213,X215))) | s(bool,'T_0') = s(bool,happ(s(fun(X211,bool),X212),s(X211,happ(s(fun(X213,X211),X214),s(X213,X215)))))) ) | $spl10),
  inference(superposition,[],[f1849,f369])).
fof(f5299,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s('type.num.num','const.list.list_size_2'(s(fun(X105,'type.num.num'),X106),s('type.list.list'(X105),happ(s(fun(X107,'type.list.list'(X105)),X108),s(X107,X109))))) = s('type.num.num',happ(s(fun(X107,'type.num.num'),'_dst_x0x1_2'(s(fun('type.list.list'(X105),'type.num.num'),'const.list.list_size_1'(s(fun(X105,'type.num.num'),X106))),s(fun(X107,'type.list.list'(X105)),X108))),s(X107,X109)))) )),
  inference(superposition,[],[f369,f315])).
fof(f5298,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s(fun(X95,X96),'_dst_x0x1_2'(s(fun(X97,X96),happ(s(fun(X99,fun(X97,X96)),X100),s(X99,X101))),s(fun(X95,X97),X98))) = s(fun(X95,X96),happ(s(fun(X99,fun(X95,X96)),'_dst_x0x1_2'(s(fun(fun(X97,X96),fun(X95,X96)),'_dstsel_0_x01_1'(s(fun(X95,X97),X98))),s(fun(X99,fun(X97,X96)),X100))),s(X99,X101)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5288,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X39,X37),X40))),s(X39,X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X39,X37),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f369,f437])).
fof(f5287,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X33,X31),X34))),s(X33,X35))) | s(bool,X36) = s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X33,X31),X34),s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f369,f437])).
fof(f5285,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f369,f340])).
fof(f5284,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5283,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(X104,happ(s(fun(fun(X99,X100),X104),X105),s(fun(X99,X100),'_dst_x0x1_2'(s(fun(X101,X100),X103),s(fun(X99,X101),X102))))) = s(X104,happ(s(fun(fun(X101,X100),X104),'_dst_x0x1_2'(s(fun(fun(X99,X100),X104),X105),s(fun(fun(X101,X100),fun(X99,X100)),'_dstsel_0_x01_1'(s(fun(X99,X101),X102))))),s(fun(X101,X100),X103)))) )),
  inference(superposition,[],[f369,f368])).
fof(f5282,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK18(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f369,f359])).
fof(f5281,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f369,f369])).
fof(f5280,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f369,f369])).
fof(f5279,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f369,f354])).
fof(f5273,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f5238,f261])).
fof(f5238,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK18(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK18(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f4695,f4724])).
fof(f5269,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f5235,f261])).
fof(f5235,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK18(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4695,f753])).
fof(f5268,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5234,f261])).
fof(f5234,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK18(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4695,f1756])).
fof(f5264,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(subsumption_resolution,[],[f5231,f261])).
fof(f5231,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK18(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(superposition,[],[f4695,f432])).
fof(f5253,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5214,f261])).
fof(f5214,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4695,f437])).
fof(f4959,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X50),s('type.list.list'('type.modSem.v'),X48))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X51),s('type.list.list'('type.modSem.v'),X52))) | s('type.list.list'('type.modSem.v'),X52) = s('type.list.list'('type.modSem.v'),sK4(X48,X49,'type.list.list'('type.modSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.modSem.v'),bool),X49)))) )),
  inference(superposition,[],[f395,f753])).
fof(f4958,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X45),s('type.list.list'('type.modSem.v'),X44))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X46),s('type.list.list'('type.modSem.v'),X47))) | s('type.list.list'('type.modSem.v'),X47) = s('type.list.list'('type.modSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X42,'type.list.list'('type.modSem.v')),'const.pair.,_2'(s(X42,X43),s('type.list.list'('type.modSem.v'),X44)))))) )),
  inference(superposition,[],[f395,f344])).
fof(f4957,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X39),s('type.list.list'('type.modSem.v'),X38))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X40),s('type.list.list'('type.modSem.v'),X41))) | s('type.list.list'('type.modSem.v'),X41) = s('type.list.list'('type.modSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.modSem.v'),X37),'const.sum.INL_1'(s('type.list.list'('type.modSem.v'),X38)))))) )),
  inference(superposition,[],[f395,f338])).
fof(f4956,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X34),s('type.list.list'('type.modSem.v'),X32))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X35),s('type.list.list'('type.modSem.v'),X36))) | s('type.list.list'('type.modSem.v'),X36) = s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.modSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f395,f343])).
fof(f4955,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X28),s('type.list.list'('type.modSem.v'),X27))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X29),s('type.list.list'('type.modSem.v'),X30))) | s('type.list.list'('type.modSem.v'),X30) = s('type.list.list'('type.modSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.list.list'('type.modSem.v')),'const.sum.INR_1'(s('type.list.list'('type.modSem.v'),X27)))))) )),
  inference(superposition,[],[f395,f339])).
fof(f4954,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X24),s('type.list.list'('type.modSem.v'),X25))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X23),s('type.list.list'('type.modSem.v'),happ(s(fun('type.num.num','type.list.list'('type.modSem.v')),X21),s('type.num.num',X22))))) | s('type.list.list'('type.modSem.v'),X25) = s('type.list.list'('type.modSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.modSem.v'),X20),s(fun('type.num.num','type.list.list'('type.modSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f395,f323])).
fof(f4953,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X17),s('type.list.list'('type.modSem.v'),X15))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X18),s('type.list.list'('type.modSem.v'),X19))) | s('type.list.list'('type.modSem.v'),X19) = s('type.list.list'('type.modSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.modSem.v'),X15),s(fun('type.num.num','type.list.list'('type.modSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f395,f314])).
fof(f4952,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X13),s('type.list.list'('type.modSem.v'),X14))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X12),s('type.list.list'('type.modSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.modSem.v'),X10),X11))))) | s('type.list.list'('type.modSem.v'),X14) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.modSem.v'),X10),'type.list.list'('type.modSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.modSem.v'),X10),X11)))) )),
  inference(superposition,[],[f395,f340])).
fof(f4951,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X7),s('type.list.list'('type.modSem.v'),X6))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X8),s('type.list.list'('type.modSem.v'),X9))) | s('type.list.list'('type.modSem.v'),X9) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.modSem.v')),sK18(X5,X6,X4,'type.list.list'('type.modSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f395,f359])).
fof(f4950,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),X0))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.list.list'('type.modSem.v'),X3) = s('type.list.list'('type.modSem.v'),happ(s(fun('type.list.list'('type.modSem.v'),'type.list.list'('type.modSem.v')),'_i_0_0'),s('type.list.list'('type.modSem.v'),X0)))) )),
  inference(superposition,[],[f395,f307])).
fof(f4934,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4933,f261])).
fof(f4933,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f4924,f1756])).
fof(f4924,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f4595,f2498])).
fof(f4931,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f4919,f261])).
fof(f4919,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4595,f437])).
fof(f4906,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4859,f261])).
fof(f4859,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK18(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f439,f4724])).
fof(f4718,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f4681,f439])).
fof(f4724,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f424,f4681,f2349])).
fof(f4725,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK18(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4681,f4594])).
fof(f4681,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK18(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f261,f4594])).
fof(f4595,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4520,f340])).
fof(f4525,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4475,f261])).
fof(f4475,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f439,f753])).
fof(f4452,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,happ(s(fun(X167,bool),X168),s(X167,sK4(X170,sK4(X168,X169,fun(X167,bool)),X167))))) | p(s(bool,'|dollar|exists_1'(s(fun(X167,bool),sK4(X168,X169,fun(X167,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X167,bool),bool),X169)))) )),
  inference(superposition,[],[f318,f753])).
fof(f4521,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'T_0') = s(bool,sK4(X104,X105,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | $spl10),
  inference(forward_demodulation,[],[f4433,f1756])).
fof(f4433,plain,(
  ( ! [X105,X104] : (s(bool,X104) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X104,X105,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105)))) ) | $spl10),
  inference(superposition,[],[f1849,f753])).
fof(f4368,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f753,f437])).
fof(f4367,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f753,f437])).
fof(f4044,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK14(X7,X8,X9,bool))) )),
  inference(superposition,[],[f353,f437])).
fof(f4043,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK14(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f353,f437])).
fof(f3974,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3504,f437])).
fof(f3504,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f424,f2671])).
fof(f3290,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2671,f2671])).
fof(f2671,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2498,f2498])).
fof(f2589,plain,(
  s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),'const.list.NIL_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f267,f263])).
fof(f2491,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2465,f261])).
fof(f2465,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2349,f437])).
fof(f2458,plain,(
  ( ! [X127,X130,X128,X126,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),happ(s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num',X129)))))) | ~p(s(bool,happ(s(fun(X126,bool),happ(s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num',X129))),s(X126,sK4(X130,'const.ind_type.FCONS_3'(s(fun(X126,bool),X127),s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129)))),X126)))))) )),
  inference(forward_demodulation,[],[f2435,f323])).
fof(f2435,plain,(
  ( ! [X127,X130,X128,X126,X129] : (~p(s(bool,happ(s(fun(X126,bool),happ(s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num',X129))),s(X126,sK4(X130,'const.ind_type.FCONS_3'(s(fun(X126,bool),X127),s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129)))),X126))))) | p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),'const.ind_type.FCONS_3'(s(fun(X126,bool),X127),s(fun('type.num.num',fun(X126,bool)),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129))))))))) )),
  inference(superposition,[],[f318,f323])).
fof(f2131,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2096,f1756])).
fof(f2096,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X33,bool),'const.pair.,_2'(s(X33,X34),s(bool,X35)))))) ) | $spl10),
  inference(superposition,[],[f1849,f344])).
fof(f2130,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2095,f1756])).
fof(f2095,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32)))))) ) | $spl10),
  inference(superposition,[],[f1849,f338])).
fof(f2129,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2094,f1756])).
fof(f2094,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30)))))) ) | $spl10),
  inference(superposition,[],[f1849,f343])).
fof(f2128,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2093,f1756])).
fof(f2093,plain,(
  ( ! [X26,X27] : (s(bool,X27) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X26,bool),'const.sum.INR_1'(s(bool,X27)))))) ) | $spl10),
  inference(superposition,[],[f1849,f339])).
fof(f2127,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X24) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2092,f1756])).
fof(f2092,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f1849,f314])).
fof(f2125,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(forward_demodulation,[],[f2090,f1756])).
fof(f2090,plain,(
  ( ! [X21,X19,X20] : (s(bool,X21) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X19),bool),sK18(X20,X21,X19,bool)),s('type.list.list'(X19),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f1849,f359])).
fof(f2124,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f2089,f1756])).
fof(f2089,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X18)))) ) | $spl10),
  inference(superposition,[],[f1849,f307])).
fof(f1849,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f1825,f261])).
fof(f1825,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f439,f1756])).
fof(f1737,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK18(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f359,f1600,f439])).
fof(f1756,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1600,f270])).
fof(f1304,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1291,f261])).
fof(f1291,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f318,f437])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f1297,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1279,f307])).
fof(f1279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f318,f307])).
fof(f1298,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1280,f359])).
fof(f1280,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK18(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f318,f359])).
fof(f1437,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1424,f340])).
fof(f1424,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f318,f340])).
fof(f1299,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1281,f314])).
fof(f1281,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f318,f314])).
fof(f1300,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12)))))) )),
  inference(forward_demodulation,[],[f1282,f339])).
fof(f1282,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14))))))))) )),
  inference(superposition,[],[f318,f339])).
fof(f1301,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1283,f343])).
fof(f1283,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f318,f343])).
fof(f1302,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | ~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21)))))) )),
  inference(forward_demodulation,[],[f1284,f338])).
fof(f1284,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun(X21,bool),X23),s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23))))))))) )),
  inference(superposition,[],[f318,f338])).
fof(f1303,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X28)))) | ~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25)))))) )),
  inference(forward_demodulation,[],[f1285,f344])).
fof(f1285,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X25,bool),X28),s(X25,sK4(X29,'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28)))),X25))))) | p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),'const.pair.SND_1'(s('type.pair.prod'(X26,fun(X25,bool)),'const.pair.,_2'(s(X26,X27),s(fun(X25,bool),X28))))))))) )),
  inference(superposition,[],[f318,f344])).
fof(f1407,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f318,f340])).
fof(f1272,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1029,f261])).
fof(f1029,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f439,f437])).
fof(f1271,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1028,f261])).
fof(f1028,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f410,f437])).
fof(f1270,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1027,f261])).
fof(f1027,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f337,f437])).
fof(f918,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f437,f344])).
fof(f917,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,X32))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f437,f338])).
fof(f916,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f437,f343])).
fof(f915,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X24,bool),'const.sum.INR_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f437,f339])).
fof(f914,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f437,f314])).
fof(f913,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK18(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f437,f359])).
fof(f912,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f437,f307])).
fof(f902,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f437,f437])).
fof(f885,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,X43)))))) )),
  inference(superposition,[],[f437,f344])).
fof(f884,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X38),'const.sum.INL_1'(s(bool,X39)))))) )),
  inference(superposition,[],[f437,f338])).
fof(f883,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f437,f343])).
fof(f882,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,X32)))))) )),
  inference(superposition,[],[f437,f339])).
fof(f881,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f437,f314])).
fof(f880,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK18(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f437,f359])).
fof(f879,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f437,f307])).
fof(f783,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f770,f261])).
fof(f770,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f439,f432])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f752,f264])).
fof(f752,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f317,f268])).
fof(f487,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f482,f264])).
fof(f482,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f334,f268])).
fof(f439,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f273,f268])).
fof(f432,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f265,f270])).
fof(f424,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f264,f269])).
fof(f397,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X21,X22,X50,X23] : ! [X24,X25,X26] : s(fun(X22,X21),'_dst_x0x1_2'(s(fun(X50,X21),X24),s(fun(X22,X50),happ(s(fun(X23,fun(X22,X50)),X25),s(X23,X26))))) = s(fun(X22,X21),happ(s(fun(X23,fun(X22,X21)),'_dst_x00x11_2'(s(fun(X50,X21),X24),s(fun(X23,fun(X22,X50)),X25))),s(X23,X26)))),
  file('thm.modSem.v_|49||49||split|5','_dst_x00x11_ax')).
fof(f394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) != s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3))) | s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) = s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2)) )),
  inference(cnf_transformation,[],[f258])).
fof(f396,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0) != s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2) | s('type.list.list'('type.modSem.v'),X1) != s('type.list.list'('type.modSem.v'),X3) | s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))) = s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X2),s('type.list.list'('type.modSem.v'),X3)))) )),
  inference(cnf_transformation,[],[f258])).
fof(f410,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X33,X53,X45,X54] : ((s(bool,X33) = s(bool,X53) & (p(s(bool,X53)) => s(bool,X45) = s(bool,X54))) => ((p(s(bool,X33)) => p(s(bool,X45))) <=> (p(s(bool,X53)) => p(s(bool,X54)))))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.IMP_CONG')).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X21,X22,X23] : ! [X24,X25,X26] : s(X21,happ(s(fun(X22,X21),X24),s(X22,happ(s(fun(X23,X22),X25),s(X23,X26))))) = s(X21,happ(s(fun(X23,X21),'_dst_x0x1_2'(s(fun(X22,X21),X24),s(fun(X23,X22),X25))),s(X23,X26)))),
  file('thm.modSem.v_|49||49||split|5','_dst_x0x1_ax')).
fof(f368,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X21,X22,X23] : ! [X25,X26] : s(fun(X22,X21),'_dst_x0x1_2'(s(fun(X23,X21),X26),s(fun(X22,X23),X25))) = s(fun(X22,X21),happ(s(fun(fun(X23,X21),fun(X22,X21)),'_dstsel_0_x01_1'(s(fun(X22,X23),X25))),s(fun(X23,X21),X26)))),
  file('thm.modSem.v_|49||49||split|5','_dstsel_0_x01_ax')).
fof(f367,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),sK19(X3,X2,X1,X0)),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ! [X3] : ? [X4] : ! [X5,X6] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X6))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),X4),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X5),s(X2,X6)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X47,X9,X31] : ! [X10] : ? [X41] : ! [X33,X45] : s(X47,happ(s(fun(X31,X47),happ(s(fun(X9,fun(X31,X47)),X10),s(X9,X33))),s(X31,X45))) = s(X47,happ(s(fun('type.pair.prod'(X9,X31),X47),X41),s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X33),s(X31,X45)))))),
  file('thm.modSem.v_|49||49||split|5','thm.pair.pair_Axiom')).
fof(f359,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X31,X9] : ! [X40,X11] : ? [X41] : (s(X31,X40) = s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X31,happ(s(fun(X31,X31),happ(s(fun('type.list.list'(X9),fun(X31,X31)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X31,X31))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X31,happ(s(fun('type.list.list'(X9),X31),X41),s('type.list.list'(X9),X17))))))),
  file('thm.modSem.v_|49||49||split|5','thm.list.list_Axiom')).
fof(f360,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK18(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f354,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X31,X9] : ! [X10,X32] : (s(fun(X9,X31),X10) = s(fun(X9,X31),X32) <=> ! [X33] : s(X31,happ(s(fun(X9,X31),X10),s(X9,X33))) = s(X31,happ(s(fun(X9,X31),X32),s(X9,X33))))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.FUN_EQ_THM')).
fof(f355,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK15(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f246])).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK14(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK14(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.modSem.v_|49||49||split|5',aHL_EXT)).
fof(f344,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X31] : ! [X33,X45] : s(X31,X45) = s(X31,'const.pair.SND_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X33),s(X31,X45)))))),
  file('thm.modSem.v_|49||49||split|5','thm.pair.SND')).
fof(f343,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X31,X9] : ! [X33,X45] : s(X9,X33) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),'const.pair.,_2'(s(X9,X33),s(X31,X45)))))),
  file('thm.modSem.v_|49||49||split|5','thm.pair.FST')).
fof(f342,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK9(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X31,X9] : ! [X36] : ? [X10] : ! [X37,X38,X39] : s(X31,happ(s(fun('type.ind_type.recspace'(X9),X31),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X37),s(X9,X38),s(fun('type.num.num','type.ind_type.recspace'(X9)),X39))))) = s(X31,happ(s(fun(fun('type.num.num',X31),X31),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X31),X31)))),X36),s('type.num.num',X37))),s(X9,X38))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X39))),s(fun('type.num.num',X31),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X31),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X39)))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.CONSTR_REC')).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9,X31] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X31),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X31),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X31),X8)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_0_1_const.pair.FST')).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X31] : ! [X33] : s(X31,X33) = s(X31,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X31),'const.sum.INR_1'(s(X31,X33)))))),
  file('thm.modSem.v_|49||49||split|5','thm.sum.OUTR')).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X31,X9] : ! [X33] : s(X9,X33) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X31),'const.sum.INL_1'(s(X9,X33)))))),
  file('thm.modSem.v_|49||49||split|5','thm.sum.OUTL')).
fof(f334,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X42,X43] : ((p(s(bool,X42)) => p(s(bool,X43))) => ((p(s(bool,X43)) => p(s(bool,X42))) => s(bool,X42) = s(bool,X43)))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.IMP_ANTISYM_AX')).
fof(f337,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f333,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X16,X17] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.modSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.basicSize.option_size_2'(s(fun('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'type.num.num'),'const.basicSize.pair_size_2'(s(fun('type.list.list'('type.string.char'),'type.num.num'),'const.list.list_size_1'(s(fun('type.string.char','type.num.num'),'const.string.char_size_0'))),s(fun('type.semanticPrimitives.tid_or_exn','type.num.num'),'const.semanticPrimitives.tid_or_exn_size_0'))),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X17)))))))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_size_def|split|1')).
fof(f331,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v',X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X0),s('type.list.list'('type.modSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v',X0))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X1)))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16,X17] : s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.CONS_2'(s('type.modSem.v',X16),s('type.list.list'('type.modSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v',X16))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X17)))))))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_size_def|split|11')).
fof(f330,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun('type.list.list'('type.string.char'),X0),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2))),s('type.list.list'('type.string.char'),X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2),s('type.list.list'('type.string.char'),X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.list.list'('type.string.char'),X0),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2))),s('type.list.list'('type.string.char'),X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2),s('type.list.list'('type.string.char'),X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.list.list'('type.string.char'),X0),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2))),s('type.list.list'('type.string.char'),X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X2),s('type.list.list'('type.string.char'),X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X9] : ! [X16,X17,X60,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.string.char'),X9),happ(s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16))),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X17))),s('type.list.list'('type.string.char'),X60))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Recclosure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16),s('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),X17),s('type.list.list'('type.string.char'),X60))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_case_def|split|3')).
fof(f329,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun('type.modLang.exp',X0),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.string.char'),X2))),s('type.modLang.exp',X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.string.char'),X2),s('type.modLang.exp',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.modLang.exp',X0),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.string.char'),X2))),s('type.modLang.exp',X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.string.char'),X2),s('type.modLang.exp',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.modLang.exp',X0),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1))),s('type.list.list'('type.string.char'),X2))),s('type.modLang.exp',X3))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X1),s('type.list.list'('type.string.char'),X2),s('type.modLang.exp',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X6),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.modSem.v'),X0),X9)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X9] : ! [X16,X17,X60,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.modLang.exp',X9),happ(s(fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9)),happ(s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16))),s('type.list.list'('type.string.char'),X17))),s('type.modLang.exp',X60))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Closure_3'(s('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),X16),s('type.list.list'('type.string.char'),X17),s('type.modLang.exp',X60))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_case_def|split|2')).
fof(f328,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1))),s('type.list.list'('type.modSem.v'),X2))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.modSem.v'),X0),X8)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1))),s('type.list.list'('type.modSem.v'),X2))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.modSem.v'),X0),X8)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.list.list'('type.modSem.v'),X0),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1))),s('type.list.list'('type.modSem.v'),X2))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X1),s('type.list.list'('type.modSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X5),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.modSem.v'),X0),X8)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ! [X16,X17,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.modSem.v'),X9),happ(s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16))),s('type.list.list'('type.modSem.v'),X17))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Conv_2'(s('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),X16),s('type.list.list'('type.modSem.v'),X17))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_case_def|split|1')).
fof(f326,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X0)),X3),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X0))),X4),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.modSem.v'),X0),X7)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.modSem.v_CASE_7'(s('type.modSem.v','const.modSem.Loc_1'(s('type.num.num',X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.option.option'('type.pair.prod'('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn')),fun('type.list.list'('type.modSem.v'),X9)),X11),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.string.char'),fun('type.modLang.exp',X9))),X12),s(fun('type.pair.prod'('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn'))))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.num.num','type.semanticPrimitives.tid_or_exn')))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.modSem.v'))),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.modLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.modSem.v'),X9),X15)))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_case_def|split|4')).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : ! [X8,X20,X35] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X20),s('type.num.num',X35))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X20))),s('type.num.num',X35)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_0_3_const.ind_type.FCONS')).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : ! [X8,X20,X35] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X20),s(fun('type.num.num','type.ind_type.recspace'(X9)),X35))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X20))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X35)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_1_3_const.ind_type.CONSTR')).
fof(f323,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X10,X44] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X44))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X44)))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.FCONS|split|1')).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),X2))) = s('type.num.num',happ(s(fun('type.list.list'(X0),'type.num.num'),'const.list.list_size_1'(s(fun(X0,'type.num.num'),X1))),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),X2))) = s('type.num.num',happ(s(fun('type.list.list'(X0),'type.num.num'),'const.list.list_size_1'(s(fun(X0,'type.num.num'),X1))),s('type.list.list'(X0),X2)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.list.list_size_2'(s(fun(X0,'type.num.num'),X1),s('type.list.list'(X0),X2))) = s('type.num.num',happ(s(fun('type.list.list'(X0),'type.num.num'),'const.list.list_size_1'(s(fun(X0,'type.num.num'),X1))),s('type.list.list'(X0),X2)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X9] : ! [X8,X20] : s('type.num.num','const.list.list_size_2'(s(fun(X9,'type.num.num'),X8),s('type.list.list'(X9),X20))) = s('type.num.num',happ(s(fun('type.list.list'(X9),'type.num.num'),'const.list.list_size_1'(s(fun(X9,'type.num.num'),X8))),s('type.list.list'(X9),X20)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_1_2_const.list.list_size')).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.FCONS|split|0')).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) <=> p(s(bool,X28)))),
  file('thm.modSem.v_|49||49||split|5','thm.bool.EQ_CLAUSES|split|0')).
fof(f270,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f267,plain,(
  ( ! [X0] : (s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X0)))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.modSem.v_size_1'(s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),X6)))))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_size_def|split|5')).
fof(f266,plain,(
  ( ! [X0] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X8))) = s('type.modSem.v',happ(s(fun('type.list.list'('type.modSem.v'),'type.modSem.v'),'const.modSem.Vectorv_0'),s('type.list.list'('type.modSem.v'),X8)))),
  file('thm.modSem.v_|49||49||split|5','ARITY_0_1_const.modSem.Vectorv')).
fof(f265,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.modSem.v_|49||49||split|5','thm.ind_type.ISO_REFL')).
fof(f263,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f66])).
fof(f66,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.modSem.v4_size_1'(s('type.list.list'('type.modSem.v'),'const.list.NIL_0')))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_size_def|split|10')).
fof(f261,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.modSem.v_|49||49||split|5',aHL_TRUTH)).
fof(f18875,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(forward_demodulation,[],[f18778,f338])).
fof(f18778,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(duplicate_literal_removal,[],[f18614])).
fof(f18614,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X130),'const.sum.INL_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) ) | ($spl10 | $spl34)),
  inference(superposition,[],[f10870,f13483])).
fof(f13483,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(trivial_inequality_removal,[],[f13472])).
fof(f13472,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(superposition,[],[f12040,f437])).
fof(f18934,plain,(
  ( ! [X74,X72] : (s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(subsumption_resolution,[],[f18933,f264])).
fof(f18933,plain,(
  ( ! [X74,X72] : (p(s(bool,'F_0')) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(forward_demodulation,[],[f18593,f338])).
fof(f18593,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X73),'const.sum.INL_1'(s(bool,'F_0')))))) | s(bool,X72) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X72)) ) | $spl34),
  inference(superposition,[],[f2351,f13483])).
fof(f1306,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1306_D])).
fof(f1306_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1307,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1307_D])).
fof(f1307_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1912,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f318,f307])).
fof(f19123,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl42),
  inference(cnf_transformation,[],[f19123_D])).
fof(f19123_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f19122,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2302,f261])).
fof(f2302,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f318,f465])).
fof(f6054,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f6054_D])).
fof(f6054_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4613,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4580,f261])).
fof(f4580,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4520,f465])).
fof(f1598,plain,(
  ( ! [X19] : (~p(s(bool,X19))) ) | $spl8),
  inference(cnf_transformation,[],[f1598_D])).
fof(f1598_D,plain,(
  ( ! [X19] : (~p(s(bool,X19))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f4593,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4520,f307])).
fof(f16420,plain,(
  ( ! [X55] : (~p(s(bool,X55)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4520,f307])).
fof(f17231,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f17126,f261])).
fof(f17126,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f4520,f465])).
fof(f18858,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl40),
  inference(cnf_transformation,[],[f18858_D])).
fof(f18858_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f18785,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(duplicate_literal_removal,[],[f18596])).
fof(f18596,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X80),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl34)),
  inference(superposition,[],[f2563,f13483])).
fof(f2563,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2562,f264])).
fof(f2562,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2528,f1321])).
fof(f1321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1307,f274])).
fof(f2528,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2351,f1345])).
fof(f1345,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1323,f261])).
fof(f1323,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1307,f437])).
fof(f14583,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) | $spl38),
  inference(cnf_transformation,[],[f14583_D])).
fof(f14583_D,plain,(
  ( ! [X57,X56] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f14533,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f14394])).
fof(f14394,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) != s(bool,sK7('const.pair.FST_0',X56,bool)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | ($spl6 | $spl10 | $spl26)),
  inference(superposition,[],[f2563,f13097])).
fof(f12911,plain,(
  $false | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12040,f12767])).
fof(f12767,plain,(
  ( ! [X89] : (s(bool,'T_0') = s(bool,X89)) ) | $spl28),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2458,f2436,f2491,f2496,f2497,f2589,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f4042,f4043,f4044,f4054,f4055,f4056,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4435,f4436,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4594,f4595,f4631,f4641,f4675,f4681,f4695,f4725,f4724,f4718,f4780,f4790,f4824,f4906,f4931,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5263,f5264,f5265,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5422,f5423,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5501,f5502,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6416,f6417,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7296,f7297,f7298,f7346,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9220,f9221,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11034,f11035,f11399,f11400,f7353,f7368,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f12766])).
fof(f12766,plain,(
  ( ! [X88,X87,X89] : (s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4427,f8145])).
fof(f12762,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK18(X5,X6,X4,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4695,f8145])).
fof(f12761,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK18(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5243,f8145])).
fof(f12760,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK18(X114,happ(s(fun('type.list.list'(X110),bool),sK18(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5244,f8145])).
fof(f12759,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK18(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5246,f8145])).
fof(f12758,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK18(X131,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X129)))),X130,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5248,f8145])).
fof(f12757,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK18(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5249,f8145])).
fof(f12756,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X139),bool),sK18(X140,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X137),'const.sum.INL_1'(s(bool,X138)))),X139,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5250,f8145])).
fof(f12755,plain,(
  ( ! [X144,X142,X145,X143,X141] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK18(X145,'const.pair.SND_1'(s('type.pair.prod'(X141,bool),'const.pair.,_2'(s(X141,X142),s(bool,X143)))),X144,bool))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5251,f8145])).
fof(f12910,plain,(
  $false | ($spl28 | $spl36)),
  inference(subsumption_resolution,[],[f12042,f12767])).
fof(f12909,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f264,f8145])).
fof(f12908,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2458,f2436,f2491,f2496,f2497,f2589,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f4042,f4043,f4044,f4054,f4055,f4056,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4435,f4436,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4594,f4595,f4631,f4641,f4675,f4681,f4695,f4725,f4724,f4718,f4780,f4790,f4824,f4906,f4931,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5263,f5264,f5265,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5422,f5423,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5501,f5502,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6416,f6417,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7296,f7297,f7298,f7346,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9220,f9221,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11034,f11035,f11399,f11400,f7353,f7368,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12904,f12907])).
fof(f12907,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl28),
  inference(subsumption_resolution,[],[f273,f8145])).
fof(f12904,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f12903])).
fof(f12903,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X11)) ) | $spl28),
  inference(subsumption_resolution,[],[f2349,f8145])).
fof(f12782,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f12781,f8145])).
fof(f12781,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5408,f8145])).
fof(f12780,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f12779,f8145])).
fof(f12779,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) ) | $spl28),
  inference(subsumption_resolution,[],[f5487,f8145])).
fof(f12778,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl28),
  inference(subsumption_resolution,[],[f352,f8145])).
fof(f12777,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f357,f8145])).
fof(f12776,plain,(
  ( ! [X344,X346,X341,X343,X339,X345,X340,X342,X338] : (s(X339,X346) = s(X339,happ(s(fun(X338,X339),happ(s(fun(X340,fun(X338,X339)),X341),s(X340,happ(s(fun(X342,X340),X343),s(X342,X344))))),s(X338,happ(s(fun(X339,X338),X345),s(X339,X346)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5440,f8145])).
fof(f12775,plain,(
  ( ! [X356,X358,X364,X360,X362,X357,X359,X361,X363] : (s(X356,X364) = s(X356,happ(s(fun(X357,X356),X363),s(X357,happ(s(fun(X356,X357),happ(s(fun(X358,fun(X356,X357)),X359),s(X358,happ(s(fun(X360,X358),X361),s(X360,X362))))),s(X356,X364)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5442,f8145])).
fof(f12774,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7500,f8145])).
fof(f12773,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8041,f8145])).
fof(f12772,plain,(
  ( ! [X333,X335,X329,X331,X332,X334,X328,X330,X336] : (s(X329,X336) = s(X329,happ(s(fun(X328,X329),happ(s(fun(X330,fun(X328,X329)),'_dst_x0x1_2'(s(fun(X331,fun(X328,X329)),X332),s(fun(X330,X331),X333))),s(X330,X334))),s(X328,happ(s(fun(X329,X328),X335),s(X329,X336)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5518,f8145])).
fof(f12771,plain,(
  ( ! [X350,X346,X352,X354,X349,X351,X347,X353,X348] : (s(X346,X354) = s(X346,happ(s(fun(X347,X346),X353),s(X347,happ(s(fun(X346,X347),happ(s(fun(X348,fun(X346,X347)),'_dst_x0x1_2'(s(fun(X349,fun(X346,X347)),X350),s(fun(X348,X349),X351))),s(X348,X352))),s(X346,X354)))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5520,f8145])).
fof(f12770,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f7501,f8145])).
fof(f12769,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl28),
  inference(subsumption_resolution,[],[f8042,f8145])).
fof(f12768,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) ) | $spl28),
  inference(subsumption_resolution,[],[f4523,f8145])).
fof(f12905,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2458,f2436,f2491,f2496,f2497,f2589,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f4042,f4043,f4044,f4054,f4055,f4056,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4435,f4436,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4594,f4595,f4631,f4641,f4675,f4681,f4695,f4725,f4724,f4718,f4780,f4790,f4824,f4906,f4931,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5263,f5264,f5265,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5422,f5423,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5501,f5502,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6416,f6417,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7296,f7297,f7298,f7346,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9220,f9221,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11034,f11035,f11399,f11400,f7353,f7368,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12904])).
fof(f12870,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379,f12764,f10151,f12765,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12788,f12793,f12796,f12869])).
fof(f12869,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f3503,f8145])).
fof(f12796,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f12795])).
fof(f12795,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f9554,f12767])).
fof(f12793,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f12792])).
fof(f12792,plain,(
  ( ! [X47,X49] : (s(bool,X47) = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f9858,f12767])).
fof(f12788,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl28)),
  inference(condensation,[],[f12787])).
fof(f12787,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30)) ) | ($spl10 | $spl28)),
  inference(subsumption_resolution,[],[f11287,f12767])).
fof(f12765,plain,(
  ( ! [X105,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X105))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK4(X104,X105,bool))))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379,f12764,f10151])).
fof(f12764,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X163,bool),'const.sum.INR_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10151,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379])).
fof(f12763,plain,(
  ( ! [X163,X161,X162] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X163),'const.sum.INL_1'(s(bool,sK4(X161,X162,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162)))) ) | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10151,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10379,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540])).
fof(f12833,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f12832,f8145])).
fof(f12832,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5404,f12780])).
fof(f12831,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f12830,f8145])).
fof(f12830,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl28),
  inference(subsumption_resolution,[],[f5483,f12782])).
fof(f12797,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379,f12764,f10151,f12765,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12788,f12793,f12796])).
fof(f12794,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379,f12764,f10151,f12765,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12788,f12793])).
fof(f12789,plain,(
  $false | ($spl10 | $spl28)),
  inference(global_subsumption,[],[f260,f259,f261,f263,f264,f265,f266,f267,f268,f270,f269,f274,f273,f307,f314,f315,f323,f324,f325,f326,f327,f328,f329,f330,f331,f332,f333,f337,f334,f338,f339,f340,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f369,f410,f396,f395,f394,f397,f424,f432,f437,f439,f465,f487,f507,f535,f615,f664,f717,f753,f783,f829,f879,f880,f881,f882,f883,f884,f885,f902,f912,f913,f914,f915,f916,f917,f918,f1270,f1271,f1272,f1296,f1597,f1394,f1396,f1398,f1601,f398,f358,f357,f356,f352,f351,f1582,f341,f350,f347,f1596,f349,f348,f320,f319,f316,f412,f309,f1407,f1286,f1288,f1290,f751,f1303,f750,f1302,f749,f1301,f748,f1300,f747,f1299,f1437,f1425,f850,f1298,f746,f1297,f318,f317,f1304,f1600,f1756,f1737,f1849,f2124,f2125,f2126,f2127,f2128,f2129,f2130,f2131,f2339,f2291,f2340,f2341,f2342,f2343,f2344,f2349,f2351,f2352,f2400,f2402,f2428,f2455,f2458,f2436,f2491,f2496,f2497,f2498,f2589,f2781,f2783,f2785,f2671,f2817,f2819,f2833,f2862,f2863,f2864,f2865,f2866,f2867,f2868,f2869,f2870,f2871,f3247,f3248,f3249,f3250,f3251,f3252,f3253,f3254,f3255,f3290,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3432,f3503,f3504,f3940,f3945,f3947,f3598,f3974,f4007,f4008,f4009,f4010,f4011,f4012,f4013,f4014,f4015,f4042,f4043,f4044,f4067,f4050,f4051,f4054,f4055,f4056,f4070,f4062,f4063,f4100,f4101,f4102,f4103,f4104,f4105,f4106,f4107,f4186,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4305,f4366,f4367,f4368,f4519,f4374,f4375,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4520,f4415,f4416,f4417,f4418,f4419,f4421,f4422,f4427,f4428,f4429,f4430,f4431,f4521,f4435,f4436,f4446,f4451,f4452,f4453,f4454,f4523,f4524,f4525,f4536,f4615,f4618,f4594,f4595,f4631,f4641,f4675,f4678,f4681,f4695,f4700,f4725,f4724,f4718,f4770,f4780,f4790,f4824,f4827,f4906,f4931,f4934,f4930,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f5165,f5168,f5191,f5253,f5256,f5263,f5264,f5265,f5266,f5268,f5269,f5273,f5277,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5279,f5280,f5281,f5282,f5283,f5284,f5285,f5286,f5287,f5288,f5551,f5294,f5295,f5298,f5299,f5399,f5400,f5402,f5403,f5404,f5405,f5406,f5408,f5409,f5414,f5415,f5416,f5417,f5418,f5557,f5422,f5423,f5433,f5434,f5438,f5440,f5441,f5442,f5443,f5446,f5447,f5448,f5450,f5451,f5453,f5476,f5477,f5478,f5479,f5481,f5482,f5483,f5484,f5485,f5487,f5488,f5493,f5494,f5495,f5496,f5497,f5563,f5501,f5502,f5512,f5513,f5516,f5518,f5519,f5520,f5521,f5524,f5525,f5526,f5528,f5529,f5531,f6042,f6043,f5954,f5955,f6044,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f5964,f6053,f6056,f6059,f6062,f6030,f6031,f6032,f6063,f6100,f6101,f6102,f6103,f6104,f6105,f6107,f6109,f6110,f6195,f6205,f6247,f6248,f6249,f6250,f6251,f6252,f6254,f6256,f6257,f6343,f6353,f6401,f6402,f6360,f6361,f6403,f6404,f6405,f6406,f6407,f6408,f6409,f6410,f6370,f6371,f6372,f6373,f6412,f6379,f6380,f6416,f6417,f6420,f6425,f6399,f6426,f6634,f6635,f6556,f6557,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6645,f6579,f6580,f6584,f6585,f6586,f6648,f6649,f6589,f6590,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6600,f6601,f6602,f6659,f6608,f6609,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6663,f6625,f6626,f6630,f6631,f6632,f6840,f7326,f7327,f7271,f7272,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7281,f7282,f7283,f7284,f7337,f7290,f7291,f7296,f7297,f7298,f7342,f7304,f7305,f7346,f7349,f7325,f7500,f7501,f7510,f7513,f7514,f7523,f7559,f7560,f7563,f8041,f8042,f8051,f8054,f8055,f8064,f8101,f8104,f8134,f8250,f8251,f8252,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8803,f8804,f9165,f9166,f9167,f9308,f9173,f9174,f9181,f9195,f9196,f9197,f9198,f9213,f9215,f9216,f9325,f9220,f9221,f9233,f9234,f9346,f9347,f9263,f9354,f9355,f9267,f9357,f9358,f9279,f9285,f9554,f9508,f9509,f9510,f9511,f9512,f9513,f9514,f9515,f9516,f9517,f9518,f9519,f9560,f9561,f9562,f9563,f9564,f9565,f9566,f9567,f9568,f9569,f9570,f9571,f9572,f9573,f9605,f9606,f9576,f9577,f9607,f9608,f9609,f9610,f9611,f9612,f9613,f9614,f9615,f9587,f9617,f9618,f9621,f9858,f9811,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9864,f9865,f9866,f9867,f9868,f9869,f9870,f9871,f9872,f9873,f9874,f9875,f9876,f9877,f9909,f9910,f9880,f9881,f9911,f9912,f9913,f9914,f9915,f9916,f9917,f9918,f9919,f9891,f9921,f9922,f9925,f10086,f10161,f10140,f10141,f10230,f10231,f10232,f10233,f10234,f10235,f10236,f10237,f10238,f10239,f10240,f10241,f10242,f10243,f10275,f10276,f10246,f10247,f10277,f10278,f10279,f10280,f10281,f10282,f10283,f10284,f10285,f10257,f10287,f10288,f10291,f10314,f10389,f10368,f10369,f10475,f10550,f10529,f10530,f10619,f10620,f10621,f10622,f10623,f10624,f10625,f10626,f10627,f10628,f10629,f10630,f10631,f10632,f10664,f10665,f10635,f10636,f10666,f10667,f10668,f10669,f10670,f10671,f10672,f10673,f10674,f10646,f10676,f10677,f10680,f10738,f10739,f10741,f10742,f10744,f10749,f10870,f10880,f10819,f10820,f10993,f10998,f11003,f11037,f11039,f11006,f11007,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f11055,f11057,f11017,f11058,f11061,f11034,f11035,f11287,f11239,f11240,f11242,f11243,f11245,f11250,f11399,f11400,f12039,f12041,f7353,f7368,f7371,f8135,f8140,f12045,f8145,f8261,f12755,f12756,f12757,f12758,f12759,f12760,f12761,f12762,f6599,f6566,f4449,f10540,f12763,f10379,f12764,f10151,f12765,f12766,f12767,f12768,f12769,f12770,f12771,f12772,f12773,f12774,f12775,f12776,f12777,f12778,f12780,f12782,f12788])).
fof(f12786,plain,(
  $false | ($spl6 | $spl10 | $spl28)),
  inference(subsumption_resolution,[],[f12038,f12767])).
fof(f12038,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f11231,f1322])).
fof(f1322,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1307,f269])).
fof(f11231,plain,(
  ( ! [X65] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK4(X65,'_i_0_0',bool))) ) | ($spl6 | $spl10)),
  inference(superposition,[],[f10870,f1321])).
fof(f12785,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f12043,f12767])).
fof(f12043,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9803,f1322])).
fof(f9803,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6343,f1321])).
fof(f12784,plain,(
  $false | ($spl6 | $spl28)),
  inference(subsumption_resolution,[],[f12044,f12767])).
fof(f12044,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(subsumption_resolution,[],[f9500,f1322])).
fof(f9500,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,sK4(X98,'_i_0_0',bool))) ) | $spl6),
  inference(superposition,[],[f6195,f1321])).
fof(f12750,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f12325,f11960])).
fof(f12325,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f11960,f12044])).
fof(f12749,plain,(
  $false | ($spl6 | $spl30)),
  inference(subsumption_resolution,[],[f12324,f11960])).
fof(f12324,plain,(
  ( ! [X99] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl6 | $spl30)),
  inference(backward_demodulation,[],[f11960,f12043])).
fof(f12748,plain,(
  $false | ($spl6 | $spl10 | $spl30)),
  inference(subsumption_resolution,[],[f12323,f11960])).
fof(f12323,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl6 | $spl10 | $spl30)),
  inference(backward_demodulation,[],[f11960,f12038])).
fof(f12740,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f12739,f8145])).
fof(f12739,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12207,f11960])).
fof(f12207,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl30),
  inference(backward_demodulation,[],[f11960,f5483])).
fof(f12738,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f12737,f8145])).
fof(f12737,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f12197,f11960])).
fof(f12197,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl30),
  inference(backward_demodulation,[],[f11960,f5404])).
fof(f12730,plain,(
  $false | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f12729,f8145])).
fof(f12729,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(subsumption_resolution,[],[f12047,f11960])).
fof(f12047,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl30),
  inference(backward_demodulation,[],[f11960,f273])).
fof(f12728,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f12046,f261])).
fof(f12046,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(backward_demodulation,[],[f11960,f264])).
fof(f12042,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl36),
  inference(cnf_transformation,[],[f12042_D])).
fof(f12042_D,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f12041,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(subsumption_resolution,[],[f10810,f10870])).
fof(f10810,plain,(
  ( ! [X78,X79] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X79,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X78))) | s(bool,'T_0') = s(bool,X78)) ) | $spl10),
  inference(superposition,[],[f6195,f2781])).
fof(f12040,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl34),
  inference(cnf_transformation,[],[f12040_D])).
fof(f12040_D,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f12039,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f10811,f10870])).
fof(f10811,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f6343,f2781])).
fof(f11961,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl32),
  inference(cnf_transformation,[],[f11961_D])).
fof(f11961_D,plain,(
  ( ! [X80,X81] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f11887,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(duplicate_literal_removal,[],[f11693])).
fof(f11693,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl26)),
  inference(superposition,[],[f2563,f9712])).
fof(f8144,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl26),
  inference(cnf_transformation,[],[f8144_D])).
fof(f8144_D,plain,(
  ( ! [X40] : (s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8145,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | $spl28),
  inference(cnf_transformation,[],[f8145_D])).
fof(f8145_D,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8134,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'F_0') = s(bool,sK7('const.pair.FST_0',X40,bool))) ) | $spl10),
  inference(superposition,[],[f7371,f343])).
fof(f6483,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(cnf_transformation,[],[f6483_D])).
fof(f6483_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6485,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl24),
  inference(cnf_transformation,[],[f6485_D])).
fof(f6485_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6467,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f6055,f307])).
fof(f6055,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl20),
  inference(cnf_transformation,[],[f6055_D])).
fof(f6055_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6062,plain,(
  ( ! [X64,X65] : (p(s(bool,sK5('_i_0_0',X65,bool))) | ~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool)))))) )),
  inference(forward_demodulation,[],[f6029,f307])).
fof(f6029,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X65,bool)))))) )),
  inference(superposition,[],[f319,f307])).
fof(f6053,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f6052,f307])).
fof(f6052,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6015,f261])).
fof(f6015,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f319,f465])).
fof(f5271,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5271_D])).
fof(f5271_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5270,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5236,f273])).
fof(f5236,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK18(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4695,f2499])).
fof(f2499,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2476,f261])).
fof(f2476,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f2349,f1512])).
fof(f1512,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1459,f270])).
fof(f1459,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1322,f317])).
fof(f5267,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl14),
  inference(cnf_transformation,[],[f5267_D])).
fof(f5267_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5266,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5233,f273])).
fof(f5233,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK18(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4695,f2498])).
fof(f4937,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl12),
  inference(cnf_transformation,[],[f4937_D])).
fof(f4937_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4930,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4595,f343])).
fof(f1733,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1646,f264])).
fof(f1646,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f424,f1598,f334])).
fof(f1732,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1648,f1598])).
fof(f1648,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1322,f1598,f334])).
fof(f1731,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1656,f264])).
fof(f1656,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f424,f1598,f334])).
fof(f1730,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f1658,f1598])).
fof(f1658,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1322,f1598,f334])).
fof(f1666,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1322,f1598,f317])).
fof(f1663,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f432,f1598,f269])).
fof(f1661,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f424,f1598,f274])).
fof(f1659,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1598,f1322,f1598,f334])).
fof(f1657,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f424,f1598,f334])).
fof(f1655,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1598,f1322,f1598,f334])).
fof(f1653,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1598,f424,f1598,f334])).
fof(f1649,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1598,f1322,f1598,f334])).
fof(f1647,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f264,f424,f1598,f334])).
fof(f1645,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1598,f1322,f1598,f334])).
fof(f1643,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1598,f424,f1598,f334])).
fof(f1628,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f344,f1598,f410])).
fof(f1626,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f338,f1598,f410])).
fof(f1624,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f343,f1598,f410])).
fof(f1622,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f339,f1598,f410])).
fof(f1620,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f314,f1598,f410])).
fof(f1618,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f359,f1598,f410])).
fof(f1616,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f307,f1598,f410])).
fof(f1613,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f432,f1598,f410])).
fof(f1610,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1322,f1322,f1598,f487])).
fof(f1608,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f424,f424,f1598,f487])).
fof(f1602,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f261,f1598])).
fof(f1600,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f1600_D])).
fof(f1600_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1597,plain,(
  ( ! [X19] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X19))) )),
  inference(forward_demodulation,[],[f1581,f307])).
fof(f1581,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f316,f307])).
fof(f1296,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f318,f307])).
fof(f423,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f422])).
fof(f422,plain,(
  s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f419,f420])).
fof(f418,plain,(
  s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) | $spl0),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f420,plain,(
  s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1) | $spl2),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f259,plain,(
  s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1) | s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1)))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) | s('type.list.list'('type.modSem.v'),sK0) = s('type.list.list'('type.modSem.v'),sK1)) & (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) | s('type.list.list'('type.modSem.v'),sK0) != s('type.list.list'('type.modSem.v'),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f202])).
fof(f202,plain,(
  ? [X0,X1] : ((s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))) | s('type.list.list'('type.modSem.v'),X0) = s('type.list.list'('type.modSem.v'),X1)) & (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))) | s('type.list.list'('type.modSem.v'),X0) != s('type.list.list'('type.modSem.v'),X1)))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ? [X0,X1] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))) <~> s('type.list.list'('type.modSem.v'),X0) = s('type.list.list'('type.modSem.v'),X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X0))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X1))) <=> s('type.list.list'('type.modSem.v'),X0) = s('type.list.list'('type.modSem.v'),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X6))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X7))) <=> s('type.list.list'('type.modSem.v'),X6) = s('type.list.list'('type.modSem.v'),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X6))) = s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),X7))) <=> s('type.list.list'('type.modSem.v'),X6) = s('type.list.list'('type.modSem.v'),X7))),
  file('thm.modSem.v_|49||49||split|5','thm.modSem.v_|49||49||split|5')).
fof(f419,plain,(
  s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) | $spl1),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f421,plain,(
  s('type.list.list'('type.modSem.v'),sK0) != s('type.list.list'('type.modSem.v'),sK1) | $spl3),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s('type.list.list'('type.modSem.v'),sK0) != s('type.list.list'('type.modSem.v'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f260,plain,(
  s('type.list.list'('type.modSem.v'),sK0) != s('type.list.list'('type.modSem.v'),sK1) | s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK0))) != s('type.modSem.v','const.modSem.Vectorv_1'(s('type.list.list'('type.modSem.v'),sK1)))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.modSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 207
% Passive clauses: 2444
% Generated clauses: 30772
% Final active clauses: 134
% Final passive clauses: 1384
% Input formulas: 69
% Initial clauses: 140
% 
% Duplicate literals: 974
% Trivial inequalities: 222
% Fw subsumption resolutions: 1492
% Fw demodulations: 5556
% Bw demodulations: 709
% Condensations: 1342
% Global subsumptions: 14
% 
% Simple tautologies: 349
% Equational tautologies: 182
% Forward subsumptions: 17040
% Fw demodulations to eq. taut.: 1082
% Bw demodulations to eq. taut.: 61
% 
% Unit resulting resolution: 2726
% Factoring: 1
% Forward superposition: 8336
% Backward superposition: 7280
% Self superposition: 175
% Equality factoring: 90
% Equality resolution: 31
% 
% Split clauses: 23
% Split components: 46
% SAT solver clauses: 2558
% SAT solver unit clauses: 154
% SAT solver binary clauses: 799
% 
% Sat splits: 23
% Sat splitting refutations: 52
% 
% Memory used [KB]: 33645
% Time elapsed: 1.374 s
% ------------------------------
% ------------------------------
% Success in time 2.72 s
