Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.patSem
% (137109)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 161
% Passive clauses: 3182
% Generated clauses: 17884
% Final active clauses: 153
% Final passive clauses: 2858
% Input formulas: 69
% Initial clauses: 145
% 
% Duplicate literals: 671
% Trivial inequalities: 475
% Fw subsumption resolutions: 1956
% Fw demodulations: 2637
% Bw demodulations: 6
% 
% Simple tautologies: 434
% Equational tautologies: 54
% Forward subsumptions: 8160
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 317
% Factoring: 1
% Forward superposition: 5990
% Backward superposition: 5342
% Self superposition: 133
% Equality factoring: 88
% Equality resolution: 41
% 
% Split clauses: 200
% Split components: 412
% SAT solver clauses: 478
% SAT solver unit clauses: 12
% SAT solver binary clauses: 106
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 200
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 477
% TWLsolver calls for satisfiability: 150
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.patSem
% (137131)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 1337
% Generated clauses: 15666
% Final active clauses: 138
% Final passive clauses: 1154
% Input formulas: 69
% Initial clauses: 145
% 
% Duplicate literals: 451
% Trivial inequalities: 442
% Fw subsumption resolutions: 1117
% Fw demodulations: 2225
% Bw demodulations: 38
% Condensations: 763
% Global subsumptions: 8
% 
% Simple tautologies: 292
% Equational tautologies: 45
% Forward subsumptions: 6836
% 
% Binary resolution: 172
% Factoring: 1
% Forward superposition: 5083
% Backward superposition: 3604
% Self superposition: 113
% Equality factoring: 60
% Equality resolution: 37
% 
% Split clauses: 47
% Split components: 94
% SAT solver clauses: 2222
% SAT solver unit clauses: 154
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 1
% 
% Sat splits: 47
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2189
% TWLsolver calls for satisfiability: 641
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.patSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.patSem
% SZS output start Proof for thm.patSem
fof(f93094,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f257,f423,f421,f256,f422,f420,f425,f1305,f1316,f2014,f2020,f2017,f4917,f4924,f4925,f4927,f4929,f4933,f4935,f4937,f4939,f4943,f4945,f4947,f4949,f4952,f4954,f4956,f4958,f4961,f4963,f4965,f4967,f4969,f4971,f4973,f4995,f4997,f5027,f5031,f5033,f5124,f5126,f5128,f5130,f5225,f5226,f5227,f5228,f5229,f5714,f5715,f5718,f5719,f9174,f9289,f9291,f9293,f9294,f9290,f9298,f9299,f9292,f5598,f5618,f4594,f4614,f9324,f1315,f9325,f9547,f9549,f9556,f9557,f9591,f9593,f9595,f9597,f13162,f2015,f13169,f13163,f2021,f13209,f13226,f13224,f16442,f16517,f16860,f9239,f17548,f17625,f9238,f92976,f92991,f92993,f93042,f93045,f93073])).
fof(f93073,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f93072,f2510])).
fof(f2510,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),sK1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),sK0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f423,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3)) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3)))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X16,X17,X18,X19] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.patSem.v'),X17))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) <=> (s('type.num.num',X16) = s('type.num.num',X18) & s('type.list.list'('type.patSem.v'),X17) = s('type.list.list'('type.patSem.v'),X19)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_|49||49||split|1')).
fof(f93072,plain,(
  ( ! [X428] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X428),s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X428),s('type.list.list'('type.patSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f82930,f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.patSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.patSem.v'),X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8,X22] : s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.patSem.v'),X22))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X8))),s('type.list.list'('type.patSem.v'),X22)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_2_const.patSem.Conv')).
fof(f82930,plain,(
  ( ! [X428] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X428),s('type.list.list'('type.patSem.v'),sK1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X428))),s('type.list.list'('type.patSem.v'),sK0)))) ) | $spl0),
  inference(superposition,[],[f327,f77545])).
fof(f77545,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X1),s('type.list.list'('type.patSem.v'),sK0))) = s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X1),s('type.list.list'('type.patSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f68610,f322])).
fof(f322,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X7),s('type.list.list'('type.patSem.v'),X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X7),s('type.list.list'('type.patSem.v'),X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X7),s('type.list.list'('type.patSem.v'),X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.patSem.v'),X9),X15),s('type.list.list'('type.patSem.v'),X6))) = s(X9,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.patSem.v'),X9),X15)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_case_def|split|5')).
fof(f68610,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),X1),s('type.list.list'('type.patSem.v'),sK1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X1)))) ) | $spl0),
  inference(superposition,[],[f322,f420])).
fof(f93045,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f93044,f2771])).
fof(f2771,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK0),s('type.patLang.exp',X0))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK1),s('type.patLang.exp',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f423,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3))) & (s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2) | s('type.patLang.exp',X1) != s('type.patLang.exp',X3) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3))) & ((s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2) | s('type.patLang.exp',X1) != s('type.patLang.exp',X3)) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3)))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) <=> (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X16),s('type.patLang.exp',X17))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X18),s('type.patLang.exp',X19))) <=> (s('type.list.list'('type.patSem.v'),X16) = s('type.list.list'('type.patSem.v'),X18) & s('type.patLang.exp',X17) = s('type.patLang.exp',X19)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_|49||49||split|2')).
fof(f93044,plain,(
  ( ! [X187] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK0),s('type.patLang.exp',X187))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK1),s('type.patLang.exp',X187)))) ) | $spl0),
  inference(forward_demodulation,[],[f82881,f326])).
fof(f326,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.patLang.exp',X1)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.patLang.exp',X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X22] : s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X22))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X8))),s('type.patLang.exp',X22)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_2_const.patSem.Closure')).
fof(f82881,plain,(
  ( ! [X187] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK1),s('type.patLang.exp',X187))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),sK0))),s('type.patLang.exp',X187)))) ) | $spl0),
  inference(superposition,[],[f326,f77545])).
fof(f93042,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f93041,f8583])).
fof(f8583,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK0),s('type.list.list'('type.patLang.exp'),X0),s('type.num.num',X1))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK1),s('type.list.list'('type.patLang.exp'),X2),s('type.num.num',X3)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f423,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5))) & (s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3) | s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5)))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5))) & ((s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3) | s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5)) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5)))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) <=> (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16,X17,X20,X18,X19,X21] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X16),s('type.list.list'('type.patLang.exp'),X17),s('type.num.num',X20))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patLang.exp'),X19),s('type.num.num',X21))) <=> (s('type.list.list'('type.patSem.v'),X16) = s('type.list.list'('type.patSem.v'),X18) & s('type.list.list'('type.patLang.exp'),X17) = s('type.list.list'('type.patLang.exp'),X19) & s('type.num.num',X20) = s('type.num.num',X21)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_|49||49||split|3')).
fof(f93041,plain,(
  ( ! [X184,X183] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK0),s('type.list.list'('type.patLang.exp'),X183),s('type.num.num',X184))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK1),s('type.list.list'('type.patLang.exp'),X183),s('type.num.num',X184)))) ) | $spl0),
  inference(forward_demodulation,[],[f82879,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8,X22,X23] : s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X22),s('type.num.num',X23))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X8))),s('type.list.list'('type.patLang.exp'),X22))),s('type.num.num',X23)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_3_const.patSem.Recclosure')).
fof(f82879,plain,(
  ( ! [X184,X183] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK1),s('type.list.list'('type.patLang.exp'),X183),s('type.num.num',X184))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),sK0))),s('type.list.list'('type.patLang.exp'),X183))),s('type.num.num',X184)))) ) | $spl0),
  inference(superposition,[],[f355,f77545])).
fof(f92993,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f92992,f423])).
fof(f92992,plain,(
  s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1) | $spl0),
  inference(forward_demodulation,[],[f82797,f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X26] : ! [X28] : s(X26,X28) = s(X26,happ(s(fun(X26,X26),'_i_0_0'),s(X26,X28)))),
  file('thm.patSem.v_|49||49||split|5','_i0_ax')).
fof(f82797,plain,(
  s('type.list.list'('type.patSem.v'),sK1) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),sK0))) | $spl0),
  inference(superposition,[],[f304,f77545])).
fof(f92991,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f92990,f2510])).
fof(f92990,plain,(
  ( ! [X44] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.patSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f82796,f327])).
fof(f82796,plain,(
  ( ! [X44] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.patSem.v'),sK1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X44))),s('type.list.list'('type.patSem.v'),sK0)))) ) | $spl0),
  inference(superposition,[],[f77545,f327])).
fof(f92976,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f92975,f423])).
fof(f92975,plain,(
  s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1) | $spl0),
  inference(forward_demodulation,[],[f82773,f304])).
fof(f82773,plain,(
  s('type.list.list'('type.patSem.v'),sK1) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),sK0))) | $spl0),
  inference(superposition,[],[f77545,f304])).
fof(f9238,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl20),
  inference(cnf_transformation,[],[f9238_D])).
fof(f9238_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f17625,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) | $spl38),
  inference(cnf_transformation,[],[f17625_D])).
fof(f17625_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f17548,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f17377])).
fof(f17377,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6098,f16063])).
fof(f16063,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f16048])).
fof(f16048,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f15965,f439])).
fof(f439,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f265,f265])).
fof(f265,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.patSem.v_|49||49||split|5',aHL_BOOL_CASES)).
fof(f15965,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl10),
  inference(forward_demodulation,[],[f15883,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X30,X9] : ! [X31] : s(X9,X31) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X30),'const.sum.INL_1'(s(X9,X31)))))),
  file('thm.patSem.v_|49||49||split|5','thm.sum.OUTL')).
fof(f15883,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f9297,f8569])).
fof(f8569,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f8489,f260])).
fof(f260,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.patSem.v_|49||49||split|5',aHL_FALSITY)).
fof(f8489,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(superposition,[],[f266,f2705])).
fof(f2705,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2586,f2045])).
fof(f2045,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2017,f267])).
fof(f267,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X33] : (s(bool,'T_0') = s(bool,X33) <=> p(s(bool,X33)))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.EQ_CLAUSES|split|1')).
fof(f2586,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2279,f304])).
fof(f2279,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2256,f258])).
fof(f258,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  p(s(bool,'T_0'))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.TRUTH')).
fof(f2256,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f1954,f2045])).
fof(f1954,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f1953,f304])).
fof(f1953,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1921,f258])).
fof(f1921,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f332,f466])).
fof(f466,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f265])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X37,X38] : ((p(s(bool,X37)) => p(s(bool,X38))) => ((p(s(bool,X38)) => p(s(bool,X37))) => s(bool,X37) = s(bool,X38)))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.IMP_ANTISYM_AX')).
fof(f266,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f9297,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2017,f2045,f2026,f2126,f2239,f2215,f2217,f2272,f2277,f2278,f2279,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2705,f2595,f2741,f2757,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3199,f3200,f3202,f3203,f3205,f3206,f3207,f3242,f3271,f3273,f3274,f3276,f3277,f3278,f3384,f3455,f3456,f3892,f3899,f3550,f3914,f3915,f3916,f3939,f3922,f3923,f3956,f3989,f3990,f3992,f3993,f3995,f3996,f3997,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4463,f4299,f4300,f4353,f4355,f4356,f4465,f4371,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4745,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4642,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4619,f4616,f4464,f2743,f2019,f343,f342,f316,f315,f312,f3897,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f3942,f3934,f3935,f4921,f4918,f5230,f4596,f4199,f4053,f3991,f3272,f3201,f2707,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f3994,f3275,f3204,f2710,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5704,f5711,f5712,f5713,f5714,f5716,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6052,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6350,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6572,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7606,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8466,f8469,f8473,f8569,f8577,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8774,f8779,f8781,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9289,f9291,f9293,f9294,f9296])).
fof(f9296,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9295])).
fof(f9295,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f7544,f3456])).
fof(f7544,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6482,f2595])).
fof(f8801,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X68),s('type.list.list'('type.patLang.exp'),X66),s('type.num.num',X69))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X70),s('type.list.list'('type.patLang.exp'),X71),s('type.num.num',X72))) | s('type.list.list'('type.patLang.exp'),X71) = s('type.list.list'('type.patLang.exp'),sK4(X66,X67,'type.list.list'('type.patLang.exp'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patLang.exp'),bool),X67)))) )),
  inference(superposition,[],[f398,f761])).
fof(f8800,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X61),s('type.list.list'('type.patLang.exp'),X60),s('type.num.num',X62))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X63),s('type.list.list'('type.patLang.exp'),X64),s('type.num.num',X65))) | s('type.list.list'('type.patLang.exp'),X64) = s('type.list.list'('type.patLang.exp'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patLang.exp'),X59),'const.sum.INL_1'(s('type.list.list'('type.patLang.exp'),X60)))))) )),
  inference(superposition,[],[f398,f333])).
fof(f8799,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X54),s('type.list.list'('type.patLang.exp'),X53),s('type.num.num',X55))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X56),s('type.list.list'('type.patLang.exp'),X57),s('type.num.num',X58))) | s('type.list.list'('type.patLang.exp'),X57) = s('type.list.list'('type.patLang.exp'),'const.pair.SND_1'(s('type.pair.prod'(X51,'type.list.list'('type.patLang.exp')),'const.pair.,_2'(s(X51,X52),s('type.list.list'('type.patLang.exp'),X53)))))) )),
  inference(superposition,[],[f398,f338])).
fof(f8798,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X49] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X46),s('type.list.list'('type.patLang.exp'),X44),s('type.num.num',X47))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X48),s('type.list.list'('type.patLang.exp'),X49),s('type.num.num',X50))) | s('type.list.list'('type.patLang.exp'),X49) = s('type.list.list'('type.patLang.exp'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patLang.exp'),X43),'const.pair.,_2'(s('type.list.list'('type.patLang.exp'),X44),s(X43,X45)))))) )),
  inference(superposition,[],[f398,f337])).
fof(f8797,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X40),s('type.list.list'('type.patLang.exp'),X41),s('type.num.num',X42))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X38),s('type.list.list'('type.patLang.exp'),happ(s(fun('type.num.num','type.list.list'('type.patLang.exp')),X36),s('type.num.num',X37))),s('type.num.num',X39))) | s('type.list.list'('type.patLang.exp'),X41) = s('type.list.list'('type.patLang.exp'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patLang.exp'),X35),s(fun('type.num.num','type.list.list'('type.patLang.exp')),X36),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f398,f319])).
fof(f8796,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X30),s('type.list.list'('type.patLang.exp'),X28),s('type.num.num',X31))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X32),s('type.list.list'('type.patLang.exp'),X33),s('type.num.num',X34))) | s('type.list.list'('type.patLang.exp'),X33) = s('type.list.list'('type.patLang.exp'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patLang.exp'),X28),s(fun('type.num.num','type.list.list'('type.patLang.exp')),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f398,f311])).
fof(f8795,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X23),s('type.list.list'('type.patLang.exp'),X22),s('type.num.num',X24))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X25),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X27))) | s('type.list.list'('type.patLang.exp'),X26) = s('type.list.list'('type.patLang.exp'),'const.sum.OUTR_1'(s('type.sum.sum'(X21,'type.list.list'('type.patLang.exp')),'const.sum.INR_1'(s('type.list.list'('type.patLang.exp'),X22)))))) )),
  inference(superposition,[],[f398,f334])).
fof(f8794,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patLang.exp'),X19),s('type.num.num',X20))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X16),s('type.list.list'('type.patLang.exp'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patLang.exp'),X14),X15))),s('type.num.num',X17))) | s('type.list.list'('type.patLang.exp'),X19) = s('type.list.list'('type.patLang.exp'),happ(s(fun('type.pair.prod'('type.list.list'('type.patLang.exp'),X14),'type.list.list'('type.patLang.exp')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patLang.exp'),X14),X15)))) )),
  inference(superposition,[],[f398,f335])).
fof(f8793,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X10))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X11),s('type.list.list'('type.patLang.exp'),X12),s('type.num.num',X13))) | s('type.list.list'('type.patLang.exp'),X12) = s('type.list.list'('type.patLang.exp'),happ(s(fun('type.list.list'(X6),'type.list.list'('type.patLang.exp')),sK16(X7,X8,X6,'type.list.list'('type.patLang.exp'))),s('type.list.list'(X6),'const.list.NIL_0')))) )),
  inference(superposition,[],[f398,f353])).
fof(f8792,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X1),s('type.list.list'('type.patLang.exp'),X0),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patLang.exp'),X4) = s('type.list.list'('type.patLang.exp'),happ(s(fun('type.list.list'('type.patLang.exp'),'type.list.list'('type.patLang.exp')),'_i_0_0'),s('type.list.list'('type.patLang.exp'),X0)))) )),
  inference(superposition,[],[f398,f304])).
fof(f8781,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f8746,f258])).
fof(f8746,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(superposition,[],[f4637,f5716])).
fof(f8779,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f8723,f258])).
fof(f8723,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f441,f5716])).
fof(f8774,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f8705,f426])).
fof(f8705,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5716,f2595])).
fof(f8593,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X66),s('type.list.list'('type.patLang.exp'),X68),s('type.num.num',X69))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X70),s('type.list.list'('type.patLang.exp'),X71),s('type.num.num',X72))) | s('type.list.list'('type.patSem.v'),X70) = s('type.list.list'('type.patSem.v'),sK4(X66,X67,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X67)))) )),
  inference(superposition,[],[f397,f761])).
fof(f8592,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X60),s('type.list.list'('type.patLang.exp'),X61),s('type.num.num',X62))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X63),s('type.list.list'('type.patLang.exp'),X64),s('type.num.num',X65))) | s('type.list.list'('type.patSem.v'),X63) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X59),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X60)))))) )),
  inference(superposition,[],[f397,f333])).
fof(f8591,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X53),s('type.list.list'('type.patLang.exp'),X54),s('type.num.num',X55))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X56),s('type.list.list'('type.patLang.exp'),X57),s('type.num.num',X58))) | s('type.list.list'('type.patSem.v'),X56) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X51,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X51,X52),s('type.list.list'('type.patSem.v'),X53)))))) )),
  inference(superposition,[],[f397,f338])).
fof(f8590,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X49] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X44),s('type.list.list'('type.patLang.exp'),X46),s('type.num.num',X47))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X48),s('type.list.list'('type.patLang.exp'),X49),s('type.num.num',X50))) | s('type.list.list'('type.patSem.v'),X48) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X43),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X44),s(X43,X45)))))) )),
  inference(superposition,[],[f397,f337])).
fof(f8589,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X42,X40] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X40),s('type.list.list'('type.patLang.exp'),X41),s('type.num.num',X42))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X36),s('type.num.num',X37))),s('type.list.list'('type.patLang.exp'),X38),s('type.num.num',X39))) | s('type.list.list'('type.patSem.v'),X40) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X35),s(fun('type.num.num','type.list.list'('type.patSem.v')),X36),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f397,f319])).
fof(f8588,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X28),s('type.list.list'('type.patLang.exp'),X30),s('type.num.num',X31))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X32),s('type.list.list'('type.patLang.exp'),X33),s('type.num.num',X34))) | s('type.list.list'('type.patSem.v'),X32) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X28),s(fun('type.num.num','type.list.list'('type.patSem.v')),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f397,f311])).
fof(f8587,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X22),s('type.list.list'('type.patLang.exp'),X23),s('type.num.num',X24))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X25),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X27))) | s('type.list.list'('type.patSem.v'),X25) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X21,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X22)))))) )),
  inference(superposition,[],[f397,f334])).
fof(f8586,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patLang.exp'),X19),s('type.num.num',X20))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X14),X15))),s('type.list.list'('type.patLang.exp'),X16),s('type.num.num',X17))) | s('type.list.list'('type.patSem.v'),X18) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X14),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X14),X15)))) )),
  inference(superposition,[],[f397,f335])).
fof(f8585,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X9),s('type.num.num',X10))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X11),s('type.list.list'('type.patLang.exp'),X12),s('type.num.num',X13))) | s('type.list.list'('type.patSem.v'),X11) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X6),'type.list.list'('type.patSem.v')),sK16(X7,X8,X6,'type.list.list'('type.patSem.v'))),s('type.list.list'(X6),'const.list.NIL_0')))) )),
  inference(superposition,[],[f397,f353])).
fof(f8584,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X3) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f397,f304])).
fof(f8577,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8511,f260])).
fof(f8511,plain,(
  ( ! [X43,X42] : (p(s(bool,'F_0')) | s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f1956,f2705])).
fof(f8473,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,sK4(X75,X76,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76)))) ) | $spl10),
  inference(superposition,[],[f2705,f761])).
fof(f8469,plain,(
  ( ! [X66,X64,X65] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X64),s(fun('type.num.num',bool),X65),s('type.num.num','const.num.SUC_1'(s('type.num.num',X66)))))) ) | $spl10),
  inference(superposition,[],[f2705,f319])).
fof(f8466,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) ) | $spl10),
  inference(superposition,[],[f2705,f335])).
fof(f7958,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X1))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2)) )),
  inference(equality_resolution,[],[f394])).
fof(f7947,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.patLang.exp',X48) != s('type.patLang.exp',X50) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X52),s('type.patLang.exp',X50))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X51),s('type.patLang.exp',sK4(X48,X49,'type.patLang.exp')))) | s('type.list.list'('type.patSem.v'),X51) != s('type.list.list'('type.patSem.v'),X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.patLang.exp',bool),X49)))) )),
  inference(superposition,[],[f394,f761])).
fof(f7576,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f6482,f761])).
fof(f7575,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f6482,f333])).
fof(f7574,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f6482,f338])).
fof(f7573,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f6482,f337])).
fof(f7572,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f6482,f319])).
fof(f7571,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6482,f311])).
fof(f7570,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f6482,f334])).
fof(f7569,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f6482,f335])).
fof(f7568,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6482,f353])).
fof(f7606,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f7548,f426])).
fof(f7548,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f6482,f2595])).
fof(f6996,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X1))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(equality_resolution,[],[f391])).
fof(f6985,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.list.list'('type.patSem.v'),X48) != s('type.list.list'('type.patSem.v'),X50) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X52),s('type.list.list'('type.patSem.v'),X50))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X51),s('type.list.list'('type.patSem.v'),sK4(X48,X49,'type.list.list'('type.patSem.v'))))) | s('type.num.num',X51) != s('type.num.num',X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X49)))) )),
  inference(superposition,[],[f391,f761])).
fof(f6543,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f6278,f761])).
fof(f6542,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f6278,f333])).
fof(f6541,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f6278,f338])).
fof(f6540,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f6278,f337])).
fof(f6539,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f6278,f319])).
fof(f6538,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6278,f311])).
fof(f6537,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X97))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X96,bool),'const.sum.INR_1'(s(bool,X97)))))) )),
  inference(superposition,[],[f6278,f334])).
fof(f6536,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f6278,f335])).
fof(f6535,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6278,f353])).
fof(f6572,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f6515,f426])).
fof(f6515,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f6278,f2595])).
fof(f6490,plain,(
  ( ! [X61,X62,X63] : (s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f6433,f260])).
fof(f6433,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'F_0')) | s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f1956,f1952])).
fof(f6482,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6412,f260])).
fof(f6412,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f266,f1952])).
fof(f6398,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f1952,f761])).
fof(f6396,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f1952,f338])).
fof(f6395,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,X107),s(X106,X108)))))) )),
  inference(superposition,[],[f1952,f337])).
fof(f6394,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X105),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X103),s('type.num.num',X104))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X102),s(fun('type.num.num',bool),X103),s('type.num.num','const.num.SUC_1'(s('type.num.num',X104)))))) )),
  inference(superposition,[],[f1952,f319])).
fof(f6393,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X101),'const.sum.INL_1'(s(bool,X99))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1952,f311])).
fof(f6391,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f1952,f335])).
fof(f6390,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1952,f353])).
fof(f6350,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f6349,f258])).
fof(f6349,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f6320,f2045])).
fof(f6320,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X37,X41) = s(X37,happ(s(fun(X38,X37),X40),s(X38,happ(s(fun(X37,X38),X39),s(X37,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f351,f2279])).
fof(f6347,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f6315,f258])).
fof(f6315,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f351,f439])).
fof(f6312,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f351,f761])).
fof(f6302,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f351,f761])).
fof(f6286,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f6229,f260])).
fof(f6229,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f1956,f1948])).
fof(f6278,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6209,f260])).
fof(f6209,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f266,f1948])).
fof(f6195,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f1948,f761])).
fof(f6193,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X108),s(bool,X109)))))) )),
  inference(superposition,[],[f1948,f338])).
fof(f6192,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X106,bool),'const.sum.INR_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f1948,f337])).
fof(f6191,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f1948,f319])).
fof(f6190,plain,(
  ( ! [X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1948,f311])).
fof(f6188,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f1948,f335])).
fof(f6187,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X86),bool),sK16(X87,X88,X86,bool)),s('type.list.list'(X86),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1948,f353])).
fof(f6052,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f6051,f258])).
fof(f6051,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'T_0')) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(forward_demodulation,[],[f6022,f2045])).
fof(f6022,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(X38,X41) = s(X38,happ(s(fun(X37,X38),X39),s(X37,happ(s(fun(X38,X37),X40),s(X38,X41))))) | s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X37,X38),X39),s(fun(X38,X37),X40)))) ) | $spl10),
  inference(superposition,[],[f350,f2279])).
fof(f6049,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f6017,f258])).
fof(f6017,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f350,f439])).
fof(f6014,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f350,f761])).
fof(f6004,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f350,f761])).
fof(f5700,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK16(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4637,f761])).
fof(f5699,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK16(X145,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X142),'const.sum.INL_1'(s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4637,f333])).
fof(f5698,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,X139)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK16(X141,'const.pair.SND_1'(s('type.pair.prod'(X137,bool),'const.pair.,_2'(s(X137,X138),s(bool,X139)))),X140,bool))))) )),
  inference(superposition,[],[f4637,f338])).
fof(f5697,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK16(X136,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X133),s(X132,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4637,f337])).
fof(f5696,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,happ(s(fun('type.num.num',bool),X128),s('type.num.num',X129)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X130),bool),sK16(X131,'const.ind_type.FCONS_3'(s(bool,X127),s(fun('type.num.num',bool),X128),s('type.num.num','const.num.SUC_1'(s('type.num.num',X129)))),X130,bool))))) )),
  inference(superposition,[],[f4637,f319])).
fof(f5695,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,X123)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X125),bool),sK16(X126,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.|48|_0')),X125,bool))))) )),
  inference(superposition,[],[f4637,f311])).
fof(f5694,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK16(X122,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X120)))),X121,bool))))) )),
  inference(superposition,[],[f4637,f334])).
fof(f5693,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK16(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4637,f335])).
fof(f5692,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK16(X114,happ(s(fun('type.list.list'(X110),bool),sK16(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4637,f353])).
fof(f5691,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK16(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4637,f304])).
fof(f5724,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5688,f258])).
fof(f5688,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4637,f507])).
fof(f5721,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(subsumption_resolution,[],[f5686,f258])).
fof(f5686,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X91),bool),sK16(X92,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'T_0',X89,bool))),X91,bool))))) )),
  inference(superposition,[],[f4637,f4699])).
fof(f5717,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(subsumption_resolution,[],[f5683,f258])).
fof(f5683,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(X79,bool),X80)),X81,bool)))) | s(X79,X83) = s(X79,sK4(X83,X80,X79))) )),
  inference(superposition,[],[f4637,f761])).
fof(f5716,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5682,f258])).
fof(f5682,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X77,bool))))) ) | $spl10),
  inference(superposition,[],[f4637,f2045])).
fof(f5713,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(subsumption_resolution,[],[f5680,f258])).
fof(f5680,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X72),bool),sK16(X73,'const.ind_type.ISO_2'(s(fun(X71,X71),'_i_0_0'),s(fun(X71,X71),'_i_0_0')),X72,bool))))) )),
  inference(superposition,[],[f4637,f434])).
fof(f5712,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f5679,f258])).
fof(f5679,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68)))),X69,bool)))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f4637,f534])).
fof(f5711,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5678,f258])).
fof(f5678,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4637,f466])).
fof(f5704,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f5703,f258])).
fof(f5703,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(forward_demodulation,[],[f5667,f2045])).
fof(f5667,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,X23,X24,bool)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f4637,f2279])).
fof(f5701,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5662,f258])).
fof(f5662,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4637,f439])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9] : ! [X32,X33] : (! [X31] : (s(X9,X31) = s(X9,X33) => p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X32)))))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.JRH_INDUCT_UTIL')).
fof(f4467,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4415,f258])).
fof(f4415,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f332,f761])).
fof(f4468,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4416,f258])).
fof(f4416,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f412,f761])).
fof(f754,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f313,f304])).
fof(f858,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK16(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f313,f353])).
fof(f1434,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f313,f335])).
fof(f755,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X6) = s(X3,sK4(X6,'const.sum.OUTR_1'(s('type.sum.sum'(X4,fun(X3,bool)),'const.sum.INR_1'(s(fun(X3,bool),X5)))),X3))) )),
  inference(superposition,[],[f313,f334])).
fof(f756,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8)))) | s(X7,X10) = s(X7,sK4(X10,'const.ind_type.FCONS_3'(s(fun(X7,bool),X8),s(fun('type.num.num',fun(X7,bool)),X9),s('type.num.num','const.num.|48|_0')),X7))) )),
  inference(superposition,[],[f313,f311])).
fof(f2209,plain,(
  ( ! [X132,X130,X128,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X128,bool),happ(s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num',X131)))))) | s(X128,X132) = s(X128,sK4(X132,'const.ind_type.FCONS_3'(s(fun(X128,bool),X129),s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num','const.num.SUC_1'(s('type.num.num',X131)))),X128))) )),
  inference(superposition,[],[f313,f319])).
fof(f757,plain,(
  ( ! [X14,X12,X15,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X13)))) | s(X11,X15) = s(X11,sK4(X15,'const.pair.FST_1'(s('type.pair.prod'(fun(X11,bool),X12),'const.pair.,_2'(s(fun(X11,bool),X13),s(X12,X14)))),X11))) )),
  inference(superposition,[],[f313,f337])).
fof(f758,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X19)))) | s(X16,X20) = s(X16,sK4(X20,'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))) )),
  inference(superposition,[],[f313,f338])).
fof(f759,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f313,f333])).
fof(f4376,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f313,f761])).
fof(f4340,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f266,f761])).
fof(f4341,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f267,f761])).
fof(f4342,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f270,f761])).
fof(f4343,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f271,f761])).
fof(f4344,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f329,f761])).
fof(f4346,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f332,f761])).
fof(f4347,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f412,f761])).
fof(f4352,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f441,f761])).
fof(f4354,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f488,f761])).
fof(f4359,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f1954,f761])).
fof(f4360,plain,(
  ( ! [X111,X109,X110] : (p(s(bool,X109)) | s(bool,X111) != s(bool,sK4(X109,X110,bool)) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f1956,f761])).
fof(f4537,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4464,f761])).
fof(f4374,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f312,f761])).
fof(f4377,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f414,f761])).
fof(f4498,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(fun(X57,X58),X59) != s(fun(X57,X58),X61) | s(X58,happ(s(fun(X57,X58),X61),s(X57,X62))) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f348,f761])).
fof(f2174,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f319,f439])).
fof(f2176,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f319,f439])).
fof(f2202,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f466,f319])).
fof(f2710,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2591,f2045])).
fof(f2591,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))) ) | $spl10),
  inference(superposition,[],[f2279,f319])).
fof(f3204,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X78) = s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) ) | $spl10),
  inference(superposition,[],[f2595,f319])).
fof(f3275,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X63),s('type.num.num',X64))) | s(bool,X65) = s(bool,'const.ind_type.FCONS_3'(s(bool,X62),s(fun('type.num.num',bool),X63),s('type.num.num','const.num.SUC_1'(s('type.num.num',X64))))) | s(bool,'F_0') = s(bool,X65)) ) | $spl10),
  inference(superposition,[],[f2595,f319])).
fof(f3994,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X78),s('type.num.num',X79))) | s(bool,X80) = s(bool,'const.ind_type.FCONS_3'(s(bool,X77),s(fun('type.num.num',bool),X78),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))) | s(bool,'F_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f3456,f319])).
fof(f4056,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f507,f319])).
fof(f4202,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X82),s('type.num.num',X83))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X81),s(fun('type.num.num',bool),X82),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83)))))) )),
  inference(superposition,[],[f534,f319])).
fof(f1403,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f335,f439])).
fof(f1405,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f335,f439])).
fof(f1407,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f335,f265])).
fof(f1902,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f466,f335])).
fof(f2022,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29))))) )),
  inference(forward_demodulation,[],[f2000,f335])).
fof(f2000,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f312,f335])).
fof(f2707,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2588,f2045])).
fof(f2588,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2279,f335])).
fof(f3201,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2595,f335])).
fof(f3272,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2595,f335])).
fof(f3991,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3456,f335])).
fof(f4053,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f507,f335])).
fof(f4199,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f534,f335])).
fof(f4596,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4464,f335])).
fof(f5230,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0'))))) )),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f313,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f754,f755,f756,f757,f758,f759,f761,f791,f837,f858,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1403,f1405,f1407,f1416,f1447,f1434,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1902,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2022,f2174,f2176,f2202,f2239,f2209,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4053,f4054,f4055,f4056,f4057,f4058,f4059,f4138,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4340,f4341,f4342,f4343,f4344,f4346,f4347,f4352,f4353,f4354,f4355,f4356,f4359,f4360,f4374,f4375,f4376,f4377,f4383,f4385,f4390,f4400,f4467,f4468,f4469,f4498,f4537,f4596,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918])).
fof(f4918,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f4906,f258])).
fof(f4906,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4596,f439])).
fof(f4921,plain,(
  ( ! [X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4920,f258])).
fof(f4920,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(forward_demodulation,[],[f4911,f2045])).
fof(f4911,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X17),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),X18)))) ) | $spl10),
  inference(superposition,[],[f4596,f2279])).
fof(f3935,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,f(s(X33,sK12(X34,X35,bool,X33)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(X33,sK12(X34,X35,bool,X33))))) ) | $spl10),
  inference(superposition,[],[f347,f2595])).
fof(f3934,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') != s(bool,f(s(X29,sK12(X30,X31,bool,X29)))) | s(fun(X29,bool),X30) = s(fun(X29,bool),X31) | s(bool,X32) = s(bool,g(s(X29,sK12(X30,X31,bool,X29)))) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f347,f2595])).
fof(f3942,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') != s(bool,f(s(X23,sK12(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK12(X24,X25,bool,X23))))) ) | $spl10),
  inference(forward_demodulation,[],[f3932,f2045])).
fof(f3932,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) != s(bool,f(s(X23,sK12(X24,X25,bool,X23)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,'F_0') = s(bool,g(s(X23,sK12(X24,X25,bool,X23))))) ) | $spl10),
  inference(superposition,[],[f347,f2279])).
fof(f3928,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK12(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK12(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f347,f439])).
fof(f3927,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK12(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK12(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f347,f439])).
fof(f3926,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f347,f265])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : ! [X32] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X32),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X33] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X32),s('type.list.list'(X9),X33)))) => ! [X34] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X32),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X34),s('type.list.list'(X9),X33)))))))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X32),s('type.list.list'(X9),X35)))))),
  file('thm.patSem.v_|49||49||split|5','thm.list.list_INDUCT')).
fof(f414,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f3043,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f3026,f304])).
fof(f3026,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f304])).
fof(f3044,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f3027,f353])).
fof(f3027,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f353])).
fof(f3046,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9)))))) )),
  inference(forward_demodulation,[],[f3029,f334])).
fof(f3029,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X11),s('type.list.list'(X9),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11)))),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X10,fun('type.list.list'(X9),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X9),bool),X11))))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f334])).
fof(f3047,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f3030,f311])).
fof(f3030,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),X13),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f311])).
fof(f3049,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19)))))) )),
  inference(forward_demodulation,[],[f3032,f337])).
fof(f3032,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X19),bool),X21),s('type.list.list'(X19),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22)))),X19))))) | ~p(s(bool,happ(s(fun('type.list.list'(X19),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X19),bool),X20),'const.pair.,_2'(s(fun('type.list.list'(X19),bool),X21),s(X20,X22))))),s('type.list.list'(X19),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f337])).
fof(f3050,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23)))))) )),
  inference(forward_demodulation,[],[f3033,f338])).
fof(f3033,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f338])).
fof(f3051,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27)))))) )),
  inference(forward_demodulation,[],[f3034,f333])).
fof(f3034,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29))))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f333])).
fof(f3052,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f3035,f260])).
fof(f3035,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f414,f265])).
fof(f3897,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(subsumption_resolution,[],[f3519,f260])).
fof(f3519,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),'const.list.NIL_0')))) | s(bool,X85) = s(bool,happ(s(fun('type.list.list'(X83),bool),X84),s('type.list.list'(X83),sK2(X84,X83)))) | s(bool,'F_0') = s(bool,X85)) ) | $spl10),
  inference(superposition,[],[f414,f2595])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : ! [X32,X31] : (p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31)))) => p(s(bool,happ(s(fun(X9,bool),X32),s(X9,'const.min.@_1'(s(fun(X9,bool),X32)))))))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.SELECT_AX')).
fof(f315,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X32] : (! [X31] : (p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X31))))) => (? [X31] : p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31)))) => ? [X31] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X31))))))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.MONO_EXISTS')).
fof(f316,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f342,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X30] : ! [X54,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X30),X10),s(fun(X30,X9),X54)))) => (! [X32] : (! [X31] : p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31)))) <=> ! [X31] : p(s(bool,happ(s(fun(X9,bool),X32),s(X9,happ(s(fun(X30,X9),X54),s(X30,X31))))))) & ! [X32] : (? [X31] : p(s(bool,happ(s(fun(X9,bool),X32),s(X9,X31)))) <=> ? [X31] : p(s(bool,happ(s(fun(X9,bool),X32),s(X9,happ(s(fun(X30,X9),X54),s(X30,X31))))))) & ! [X6,X56] : (s(X9,X6) = s(X9,happ(s(fun(X30,X9),X54),s(X30,X56))) <=> s(X30,X56) = s(X30,happ(s(fun(X9,X30),X10),s(X9,X6))))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.ISO_USAGE')).
fof(f343,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2019,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1993,f258])).
fof(f1993,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f312,f439])).
fof(f2743,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2742,f258])).
fof(f2742,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2627,f2045])).
fof(f2627,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f312,f2279])).
fof(f4464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4324,f266])).
fof(f4324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f314,f761])).
fof(f4616,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4583,f258])).
fof(f4583,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4464,f439])).
fof(f4619,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(subsumption_resolution,[],[f4618,f258])).
fof(f4618,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f4588,f2045])).
fof(f4588,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) ) | $spl10),
  inference(superposition,[],[f4464,f2279])).
fof(f341,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f344,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f3048,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15)))))) )),
  inference(forward_demodulation,[],[f3031,f319])).
fof(f3031,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X15),bool),happ(s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num',X18))),s('type.list.list'(X15),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))),X15))))) | ~p(s(bool,happ(s(fun('type.list.list'(X15),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X15),bool),X16),s(fun('type.num.num',fun('type.list.list'(X15),bool)),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.list.list'(X15),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f319])).
fof(f3045,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f3028,f335])).
fof(f3028,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f414,f335])).
fof(f1999,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.min.@_1'(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)))))))) )),
  inference(superposition,[],[f312,f353])).
fof(f345,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f236])).
fof(f346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f236])).
fof(f350,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0)))))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X30,X9] : ! [X10,X54] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X30),X10),s(fun(X30,X9),X54)))) <=> (! [X31] : s(X30,X31) = s(X30,happ(s(fun(X9,X30),X10),s(X9,happ(s(fun(X30,X9),X54),s(X30,X31))))) & ! [X41] : s(X9,X41) = s(X9,happ(s(fun(X30,X9),X54),s(X30,happ(s(fun(X9,X30),X10),s(X9,X41)))))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.ISO')).
fof(f351,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f352,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f396,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X30,X51,X9,X52] : ! [X53,X54,X55,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X52),X10),s(fun(X52,X9),X55)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X30,X51),X54),s(fun(X51,X30),X53))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X30),fun(X52,X51)),'_dst_x00x11_2'(s(fun(X30,X51),X54),s(fun(fun(X9,X30),fun(X52,X30)),'_dstsel_0_x01_1'(s(fun(X52,X9),X55))))),s(fun(fun(X52,X51),fun(X9,X30)),'_dst_x00x11_2'(s(fun(X51,X30),X53),s(fun(fun(X52,X51),fun(X9,X51)),'_dstsel_0_x01_1'(s(fun(X9,X52),X10)))))))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.ISO_FUN')).
fof(f4303,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f761,f304])).
fof(f4304,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f761,f353])).
fof(f4306,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))) | s(X12,X15) = s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))) )),
  inference(superposition,[],[f761,f334])).
fof(f4307,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),X17))) | s(X16,X19) = s(X16,sK4(X19,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.|48|_0')),X16))) )),
  inference(superposition,[],[f761,f311])).
fof(f4309,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X27))) | s(X25,X29) = s(X25,sK4(X29,'const.pair.FST_1'(s('type.pair.prod'(fun(X25,bool),X26),'const.pair.,_2'(s(fun(X25,bool),X27),s(X26,X28)))),X25))) )),
  inference(superposition,[],[f761,f337])).
fof(f4310,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X33))) | s(X30,X34) = s(X30,sK4(X34,'const.pair.SND_1'(s('type.pair.prod'(X31,fun(X30,bool)),'const.pair.,_2'(s(X31,X32),s(fun(X30,bool),X33)))),X30))) )),
  inference(superposition,[],[f761,f338])).
fof(f4311,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X37))) | s(X35,X38) = s(X35,sK4(X38,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X35,bool),X36),'const.sum.INL_1'(s(fun(X35,bool),X37)))),X35))) )),
  inference(superposition,[],[f761,f333])).
fof(f4312,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f761,f761])).
fof(f4308,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X20,bool),happ(s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num',X23))))) | s(X20,X24) = s(X20,sK4(X24,'const.ind_type.FCONS_3'(s(fun(X20,bool),X21),s(fun('type.num.num',fun(X20,bool)),X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23)))),X20))) )),
  inference(superposition,[],[f761,f319])).
fof(f4305,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f761,f335])).
fof(f4642,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4633,f260])).
fof(f4633,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4595,f2595])).
fof(f4637,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4625,f260])).
fof(f4625,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4595,f265])).
fof(f4595,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4464,f353])).
fof(f4700,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4623,f4595])).
fof(f4855,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X80),s('type.list.list'('type.patLang.exp'),X81),s('type.num.num',X78))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X82),s('type.list.list'('type.patLang.exp'),X83),s('type.num.num',X84))) | s('type.num.num',X84) = s('type.num.num',sK4(X78,X79,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X79)))) )),
  inference(superposition,[],[f399,f761])).
fof(f4854,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X73),s('type.list.list'('type.patLang.exp'),X74),s('type.num.num',X72))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X75),s('type.list.list'('type.patLang.exp'),X76),s('type.num.num',X77))) | s('type.num.num',X77) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X71),'const.sum.INL_1'(s('type.num.num',X72)))))) )),
  inference(superposition,[],[f399,f333])).
fof(f4853,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X66),s('type.list.list'('type.patLang.exp'),X67),s('type.num.num',X65))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X68),s('type.list.list'('type.patLang.exp'),X69),s('type.num.num',X70))) | s('type.num.num',X70) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X63,'type.num.num'),'const.pair.,_2'(s(X63,X64),s('type.num.num',X65)))))) )),
  inference(superposition,[],[f399,f338])).
fof(f4852,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X58),s('type.list.list'('type.patLang.exp'),X59),s('type.num.num',X56))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X60),s('type.list.list'('type.patLang.exp'),X61),s('type.num.num',X62))) | s('type.num.num',X62) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X55),'const.pair.,_2'(s('type.num.num',X56),s(X55,X57)))))) )),
  inference(superposition,[],[f399,f337])).
fof(f4851,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X52),s('type.list.list'('type.patLang.exp'),X53),s('type.num.num',X54))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X50),s('type.list.list'('type.patLang.exp'),X51),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X48),s('type.num.num',X49))))) | s('type.num.num',X54) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X47),s(fun('type.num.num','type.num.num'),X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f399,f319])).
fof(f4850,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X42),s('type.list.list'('type.patLang.exp'),X43),s('type.num.num',X40))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X44),s('type.list.list'('type.patLang.exp'),X45),s('type.num.num',X46))) | s('type.num.num',X46) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X40),s(fun('type.num.num','type.num.num'),X41),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f399,f311])).
fof(f4849,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X35),s('type.list.list'('type.patLang.exp'),X36),s('type.num.num',X34))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X37),s('type.list.list'('type.patLang.exp'),X38),s('type.num.num',X39))) | s('type.num.num',X39) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X33,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f399,f334])).
fof(f4848,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X30),s('type.list.list'('type.patLang.exp'),X31),s('type.num.num',X32))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X28),s('type.list.list'('type.patLang.exp'),X29),s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X26),X27))))) | s('type.num.num',X32) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X26),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X26),X27)))) )),
  inference(superposition,[],[f399,f335])).
fof(f4847,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X22),s('type.num.num',X20))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X23),s('type.list.list'('type.patLang.exp'),X24),s('type.num.num',X25))) | s('type.num.num',X25) = s('type.num.num',happ(s(fun('type.list.list'(X18),'type.num.num'),sK16(X19,X20,X18,'type.num.num')),s('type.list.list'(X18),'const.list.NIL_0')))) )),
  inference(superposition,[],[f399,f353])).
fof(f4822,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4775,f258])).
fof(f4775,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f441,f4699])).
fof(f4745,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4710,f260])).
fof(f4710,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4623,f2595])).
fof(f4693,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f353,f4623,f441])).
fof(f4699,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f426,f4623,f1954])).
fof(f4623,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f258,f4595])).
fof(f4469,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4419,f258])).
fof(f4419,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f441,f761])).
fof(f4400,plain,(
  ( ! [X261,X263,X259,X260,X262] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X261),s('type.patLang.exp',X259))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X262),s('type.patLang.exp',X263))) | s('type.patLang.exp',X263) = s('type.patLang.exp',sK4(X259,X260,'type.patLang.exp')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.patLang.exp',bool),X260)))) )),
  inference(superposition,[],[f393,f761])).
fof(f4390,plain,(
  ( ! [X218,X222,X220,X219,X221] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X218),s('type.list.list'('type.patSem.v'),X220))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X221),s('type.list.list'('type.patSem.v'),X222))) | s('type.num.num',X221) = s('type.num.num',sK4(X218,X219,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X219)))) )),
  inference(superposition,[],[f389,f761])).
fof(f4385,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X193),s('type.patLang.exp',X195))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X196),s('type.patLang.exp',X197))) | s('type.list.list'('type.patSem.v'),X196) = s('type.list.list'('type.patSem.v'),sK4(X193,X194,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X194)))) )),
  inference(superposition,[],[f392,f761])).
fof(f4383,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X185),s('type.list.list'('type.patSem.v'),X183))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X186),s('type.list.list'('type.patSem.v'),X187))) | s('type.list.list'('type.patSem.v'),X187) = s('type.list.list'('type.patSem.v'),sK4(X183,X184,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X184)))) )),
  inference(superposition,[],[f390,f761])).
fof(f4375,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f314,f761])).
fof(f4371,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3456,f761])).
fof(f4465,plain,(
  ( ! [X112,X113] : (s(bool,'T_0') != s(bool,X112) | s(bool,'T_0') = s(bool,sK4(X112,X113,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) ) | $spl10),
  inference(forward_demodulation,[],[f4361,f2045])).
fof(f4361,plain,(
  ( ! [X112,X113] : (s(bool,X112) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X112,X113,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X113)))) ) | $spl10),
  inference(superposition,[],[f2126,f761])).
fof(f4356,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f534,f761])).
fof(f4355,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f507,f761])).
fof(f4353,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f466,f761])).
fof(f4300,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f761,f2595])).
fof(f4299,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f761,f2595])).
fof(f4463,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4297,f2045])).
fof(f4297,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,'F_0') = s(bool,sK4(X23,X24,bool))) ) | $spl10),
  inference(superposition,[],[f761,f2279])).
fof(f4293,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f761,f439])).
fof(f4292,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f761,f439])).
fof(f4291,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f761,f265])).
fof(f4285,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4228,f258])).
fof(f4228,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f441,f534])).
fof(f4205,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,X94)))))) )),
  inference(superposition,[],[f534,f333])).
fof(f4204,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f534,f338])).
fof(f4203,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X88,bool),'const.sum.INR_1'(s(bool,X86))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87)))))) )),
  inference(superposition,[],[f534,f337])).
fof(f4201,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,X78))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f534,f311])).
fof(f4200,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X77,bool),'const.sum.INR_1'(s(bool,X76))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X75,bool),'const.sum.INR_1'(s(bool,X76)))))) )),
  inference(superposition,[],[f534,f334])).
fof(f4198,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK16(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f534,f353])).
fof(f4138,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4082,f258])).
fof(f4082,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f441,f507])).
fof(f4059,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f507,f333])).
fof(f4058,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) )),
  inference(superposition,[],[f507,f338])).
fof(f4057,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f507,f337])).
fof(f4055,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X77),'const.sum.INL_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f507,f311])).
fof(f4054,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X73))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,X73)))))) )),
  inference(superposition,[],[f507,f334])).
fof(f4052,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK16(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f507,f353])).
fof(f3997,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3456,f333])).
fof(f3996,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') != s(bool,X87) | s(bool,X88) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87))))) | s(bool,'F_0') = s(bool,X88)) ) | $spl10),
  inference(superposition,[],[f3456,f338])).
fof(f3995,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X84) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3456,f337])).
fof(f3993,plain,(
  ( ! [X76,X74,X75] : (s(bool,'T_0') != s(bool,X74) | s(bool,X76) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f3456,f311])).
fof(f3992,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | s(bool,X73) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3456,f334])).
fof(f3990,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK16(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3456,f353])).
fof(f3989,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3456,f304])).
fof(f3956,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3456,f439])).
fof(f3923,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,f(s(bool,X36))) != s(X35,g(s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK12(X33,X34,X35,bool))) ) | $spl10),
  inference(superposition,[],[f347,f2595])).
fof(f3922,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK12(X29,X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f347,f2595])).
fof(f3939,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'T_0'))) != s(X25,g(s(bool,'T_0'))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK12(X23,X24,X25,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f3920,f2045])).
fof(f3920,plain,(
  ( ! [X24,X23,X25] : (s(X25,f(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) != s(X25,g(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))) | s(fun(bool,X25),X23) = s(fun(bool,X25),X24) | s(bool,'F_0') = s(bool,sK12(X23,X24,X25,bool))) ) | $spl10),
  inference(superposition,[],[f347,f2279])).
fof(f3916,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK12(X7,X8,X9,bool))) )),
  inference(superposition,[],[f347,f439])).
fof(f3915,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK12(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f347,f439])).
fof(f3914,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,X2,bool))) )),
  inference(superposition,[],[f347,f265])).
fof(f3550,plain,(
  ( ! [X177,X175,X176] : (~p(s(bool,happ(s(fun(bool,bool),X176),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176)))) | s(bool,X177) = s(bool,sK4(X175,X176,bool)) | s(bool,'F_0') = s(bool,X177)) ) | $spl10),
  inference(superposition,[],[f314,f2595])).
fof(f3899,plain,(
  ( ! [X94,X95,X93,X96] : (s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(subsumption_resolution,[],[f3524,f260])).
fof(f3524,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,'F_0')) | s(X93,X95) = s(X93,sK4(X95,X94,X93)) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94))) | s(bool,'F_0') = s(bool,X96)) ) | $spl10),
  inference(superposition,[],[f313,f2595])).
fof(f3892,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3491,f260])).
fof(f3491,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f329,f2595])).
fof(f3456,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f426,f2595])).
fof(f3455,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f270,f2595])).
fof(f3384,plain,(
  ( ! [X189,X187,X188] : (~p(s(bool,happ(s(fun(bool,bool),X188),s(bool,X189)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X188)))) | s(bool,'F_0') = s(bool,X189) | s(bool,'F_0') = s(bool,sK4(X187,X188,bool))) ) | $spl10),
  inference(superposition,[],[f314,f2595])).
fof(f3278,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2595,f333])).
fof(f3277,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f2595,f338])).
fof(f3276,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2595,f337])).
fof(f3274,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X61)) ) | $spl10),
  inference(superposition,[],[f2595,f311])).
fof(f3273,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X57) | s(bool,X58) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X56,bool),'const.sum.INR_1'(s(bool,X57))))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2595,f334])).
fof(f3271,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK16(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2595,f353])).
fof(f3242,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2595,f2595])).
fof(f3207,plain,(
  ( ! [X88,X87,X89] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X87),'const.sum.INL_1'(s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2595,f333])).
fof(f3206,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,X85) = s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X85)))))) ) | $spl10),
  inference(superposition,[],[f2595,f338])).
fof(f3205,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X80) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) ) | $spl10),
  inference(superposition,[],[f2595,f337])).
fof(f3203,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2595,f311])).
fof(f3202,plain,(
  ( ! [X70,X71,X69] : (s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,X70)))))) ) | $spl10),
  inference(superposition,[],[f2595,f334])).
fof(f3200,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK16(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2595,f353])).
fof(f3199,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2595,f304])).
fof(f2780,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X44),s('type.patLang.exp',X45))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X46),s('type.patLang.exp',X47))) | s('type.list.list'('type.patSem.v'),X46) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X44)))))) )),
  inference(superposition,[],[f392,f333])).
fof(f2779,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X39),s('type.patLang.exp',X40))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X41),s('type.patLang.exp',X42))) | s('type.list.list'('type.patSem.v'),X41) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.patSem.v'),X39)))))) )),
  inference(superposition,[],[f392,f338])).
fof(f2778,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X32),s('type.patLang.exp',X34))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X35),s('type.patLang.exp',X36))) | s('type.list.list'('type.patSem.v'),X35) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f392,f337])).
fof(f2777,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X29),s('type.patLang.exp',X30))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X26),s('type.num.num',X27))),s('type.patLang.exp',X28))) | s('type.list.list'('type.patSem.v'),X29) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X25),s(fun('type.num.num','type.list.list'('type.patSem.v')),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) )),
  inference(superposition,[],[f392,f319])).
fof(f2776,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X20),s('type.patLang.exp',X22))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X23),s('type.patLang.exp',X24))) | s('type.list.list'('type.patSem.v'),X23) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X20),s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f392,f311])).
fof(f2775,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X16),s('type.patLang.exp',X17))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X18),s('type.patLang.exp',X19))) | s('type.list.list'('type.patSem.v'),X18) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X15,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X16)))))) )),
  inference(superposition,[],[f392,f334])).
fof(f2774,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X14))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11))),s('type.patLang.exp',X12))) | s('type.list.list'('type.patSem.v'),X13) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X10),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11)))) )),
  inference(superposition,[],[f392,f335])).
fof(f2773,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X6),s('type.patLang.exp',X7))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X9))) | s('type.list.list'('type.patSem.v'),X8) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.patSem.v')),sK16(X5,X6,X4,'type.list.list'('type.patSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f392,f353])).
fof(f2772,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X2) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f392,f304])).
fof(f2757,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2658,f2045])).
fof(f2658,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f314,f2279])).
fof(f2741,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2740,f258])).
fof(f2740,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2626,f2045])).
fof(f2626,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f314,f2279])).
fof(f2595,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2279,f2279])).
fof(f2530,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X45),s('type.list.list'('type.patSem.v'),X44))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X46),s('type.list.list'('type.patSem.v'),X47))) | s('type.list.list'('type.patSem.v'),X47) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X44)))))) )),
  inference(superposition,[],[f390,f333])).
fof(f2529,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X40),s('type.list.list'('type.patSem.v'),X39))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X41),s('type.list.list'('type.patSem.v'),X42))) | s('type.list.list'('type.patSem.v'),X42) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.patSem.v'),X39)))))) )),
  inference(superposition,[],[f390,f338])).
fof(f2528,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X34),s('type.list.list'('type.patSem.v'),X32))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X35),s('type.list.list'('type.patSem.v'),X36))) | s('type.list.list'('type.patSem.v'),X36) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X31),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f390,f337])).
fof(f2527,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X29),s('type.list.list'('type.patSem.v'),X30))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X28),s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X26),s('type.num.num',X27))))) | s('type.list.list'('type.patSem.v'),X30) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X25),s(fun('type.num.num','type.list.list'('type.patSem.v')),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))) )),
  inference(superposition,[],[f390,f319])).
fof(f2526,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X22),s('type.list.list'('type.patSem.v'),X20))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X23),s('type.list.list'('type.patSem.v'),X24))) | s('type.list.list'('type.patSem.v'),X24) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X20),s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f390,f311])).
fof(f2525,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X16))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) | s('type.list.list'('type.patSem.v'),X19) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X15,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X16)))))) )),
  inference(superposition,[],[f390,f334])).
fof(f2524,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X13),s('type.list.list'('type.patSem.v'),X14))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11))))) | s('type.list.list'('type.patSem.v'),X14) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X10),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11)))) )),
  inference(superposition,[],[f390,f335])).
fof(f2523,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X7),s('type.list.list'('type.patSem.v'),X6))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.patSem.v'),X9))) | s('type.list.list'('type.patSem.v'),X9) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.patSem.v')),sK16(X5,X6,X4,'type.list.list'('type.patSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f390,f353])).
fof(f2522,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.list.list'('type.patSem.v'),X3) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f390,f304])).
fof(f2371,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X6))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X4))))),s('type.list.list'('type.patSem.v'),X7))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X4)))))) )),
  inference(superposition,[],[f389,f264])).
fof(f2278,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2255,f258])).
fof(f2255,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f1954,f434])).
fof(f2277,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2254,f258])).
fof(f2254,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f1954,f466])).
fof(f2272,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2246,f258])).
fof(f2246,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1954,f439])).
fof(f2217,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X175),s('type.patLang.exp',X176))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X177),s('type.patLang.exp',happ(s(fun('type.num.num','type.patLang.exp'),X173),s('type.num.num',X174))))) | s('type.patLang.exp',X176) = s('type.patLang.exp','const.ind_type.FCONS_3'(s('type.patLang.exp',X172),s(fun('type.num.num','type.patLang.exp'),X173),s('type.num.num','const.num.SUC_1'(s('type.num.num',X174)))))) )),
  inference(superposition,[],[f393,f319])).
fof(f2215,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X163),s('type.list.list'('type.patSem.v'),X164))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X161),s('type.num.num',X162))),s('type.list.list'('type.patSem.v'),X165))) | s('type.num.num',X163) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X160),s(fun('type.num.num','type.num.num'),X161),s('type.num.num','const.num.SUC_1'(s('type.num.num',X162)))))) )),
  inference(superposition,[],[f389,f319])).
fof(f2239,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126)))))) | ~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123)))))) )),
  inference(forward_demodulation,[],[f2208,f319])).
fof(f2208,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123))))) | p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))))))))) )),
  inference(superposition,[],[f314,f319])).
fof(f2126,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2099,f258])).
fof(f2099,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f441,f2045])).
fof(f2026,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f353,f2017,f441])).
fof(f1957,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1925,f258])).
fof(f1925,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f441,f466])).
fof(f1956,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f1955,f304])).
fof(f1955,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1922,f258])).
fof(f1922,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f412,f466])).
fof(f1952,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f1907,f304])).
fof(f1907,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38)))))) )),
  inference(superposition,[],[f466,f333])).
fof(f1951,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(forward_demodulation,[],[f1906,f304])).
fof(f1906,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(superposition,[],[f466,f338])).
fof(f1950,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(forward_demodulation,[],[f1905,f304])).
fof(f1905,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f466,f337])).
fof(f1949,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1904,f304])).
fof(f1904,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f466,f311])).
fof(f1948,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28))))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(forward_demodulation,[],[f1903,f304])).
fof(f1903,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X27,bool),'const.sum.INR_1'(s(bool,X28)))))) )),
  inference(superposition,[],[f466,f334])).
fof(f1947,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f1901,f304])).
fof(f1901,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f466,f353])).
fof(f1794,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X39),s('type.patLang.exp',X38))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X40),s('type.patLang.exp',X41))) | s('type.patLang.exp',X41) = s('type.patLang.exp','const.sum.OUTL_1'(s('type.sum.sum'('type.patLang.exp',X37),'const.sum.INL_1'(s('type.patLang.exp',X38)))))) )),
  inference(superposition,[],[f393,f333])).
fof(f1793,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X34),s('type.patLang.exp',X33))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X35),s('type.patLang.exp',X36))) | s('type.patLang.exp',X36) = s('type.patLang.exp','const.pair.SND_1'(s('type.pair.prod'(X31,'type.patLang.exp'),'const.pair.,_2'(s(X31,X32),s('type.patLang.exp',X33)))))) )),
  inference(superposition,[],[f393,f338])).
fof(f1792,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X28),s('type.patLang.exp',X26))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X29),s('type.patLang.exp',X30))) | s('type.patLang.exp',X30) = s('type.patLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.patLang.exp',X25),'const.pair.,_2'(s('type.patLang.exp',X26),s(X25,X27)))))) )),
  inference(superposition,[],[f393,f337])).
fof(f1791,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X22),s('type.patLang.exp',X20))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X23),s('type.patLang.exp',X24))) | s('type.patLang.exp',X24) = s('type.patLang.exp','const.ind_type.FCONS_3'(s('type.patLang.exp',X20),s(fun('type.num.num','type.patLang.exp'),X21),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f393,f311])).
fof(f1790,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X17),s('type.patLang.exp',X16))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X18),s('type.patLang.exp',X19))) | s('type.patLang.exp',X19) = s('type.patLang.exp','const.sum.OUTR_1'(s('type.sum.sum'(X15,'type.patLang.exp'),'const.sum.INR_1'(s('type.patLang.exp',X16)))))) )),
  inference(superposition,[],[f393,f334])).
fof(f1789,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X14))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X12),s('type.patLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.patLang.exp',X10),X11))))) | s('type.patLang.exp',X14) = s('type.patLang.exp',happ(s(fun('type.pair.prod'('type.patLang.exp',X10),'type.patLang.exp'),'const.pair.FST_0'),s('type.pair.prod'('type.patLang.exp',X10),X11)))) )),
  inference(superposition,[],[f393,f335])).
fof(f1788,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X7),s('type.patLang.exp',X6))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X9))) | s('type.patLang.exp',X9) = s('type.patLang.exp',happ(s(fun('type.list.list'(X4),'type.patLang.exp'),sK16(X5,X6,X4,'type.patLang.exp')),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f393,f353])).
fof(f1787,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.patLang.exp',X3) = s('type.patLang.exp',happ(s(fun('type.patLang.exp','type.patLang.exp'),'_i_0_0'),s('type.patLang.exp',X0)))) )),
  inference(superposition,[],[f393,f304])).
fof(f1565,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X42),s('type.list.list'('type.patSem.v'),X43))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.patSem.v'),X45))) | s('type.num.num',X44) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X41),'const.sum.INL_1'(s('type.num.num',X42)))))) )),
  inference(superposition,[],[f389,f333])).
fof(f1564,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X37),s('type.list.list'('type.patSem.v'),X38))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X39),s('type.list.list'('type.patSem.v'),X40))) | s('type.num.num',X39) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X35,'type.num.num'),'const.pair.,_2'(s(X35,X36),s('type.num.num',X37)))))) )),
  inference(superposition,[],[f389,f338])).
fof(f1563,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X30),s('type.list.list'('type.patSem.v'),X32))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X33),s('type.list.list'('type.patSem.v'),X34))) | s('type.num.num',X33) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X29),'const.pair.,_2'(s('type.num.num',X30),s(X29,X31)))))) )),
  inference(superposition,[],[f389,f337])).
fof(f1562,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X24),s('type.list.list'('type.patSem.v'),X26))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X27),s('type.list.list'('type.patSem.v'),X28))) | s('type.num.num',X27) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X24),s(fun('type.num.num','type.num.num'),X25),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f389,f311])).
fof(f1561,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X20),s('type.list.list'('type.patSem.v'),X21))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X22),s('type.list.list'('type.patSem.v'),X23))) | s('type.num.num',X22) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X19,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X20)))))) )),
  inference(superposition,[],[f389,f334])).
fof(f1560,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X14),X15))),s('type.list.list'('type.patSem.v'),X16))) | s('type.num.num',X17) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X14),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X14),X15)))) )),
  inference(superposition,[],[f389,f335])).
fof(f1559,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X10),s('type.list.list'('type.patSem.v'),X11))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.patSem.v'),X13))) | s('type.num.num',X12) = s('type.num.num',happ(s(fun('type.list.list'(X8),'type.num.num'),sK16(X9,X10,X8,'type.num.num')),s('type.list.list'(X8),'const.list.NIL_0')))) )),
  inference(superposition,[],[f389,f353])).
fof(f1558,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X4),s('type.list.list'('type.patSem.v'),X5))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X6),s('type.list.list'('type.patSem.v'),X7))) | s('type.num.num',X6) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f389,f304])).
fof(f1447,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1433,f335])).
fof(f1433,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f314,f335])).
fof(f1416,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f314,f335])).
fof(f1313,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1300,f258])).
fof(f1300,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f314,f439])).
fof(f1299,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f314,f265])).
fof(f1297,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f314,f439])).
fof(f1295,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f314,f439])).
fof(f1312,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26)))))) )),
  inference(forward_demodulation,[],[f1294,f333])).
fof(f1294,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28))))))))) )),
  inference(superposition,[],[f314,f333])).
fof(f1311,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24)))) | ~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21)))))) )),
  inference(forward_demodulation,[],[f1293,f338])).
fof(f1293,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24))))))))) )),
  inference(superposition,[],[f314,f338])).
fof(f1310,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f1292,f337])).
fof(f1292,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f314,f337])).
fof(f1309,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13)))) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12)))))) )),
  inference(forward_demodulation,[],[f1291,f311])).
fof(f1291,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f314,f311])).
fof(f1308,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8)))))) )),
  inference(forward_demodulation,[],[f1290,f334])).
fof(f1290,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK4(X11,'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10)))),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X9,fun(X8,bool)),'const.sum.INR_1'(s(fun(X8,bool),X10))))))))) )),
  inference(superposition,[],[f314,f334])).
fof(f1307,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1289,f353])).
fof(f1289,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f314,f353])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1288,f304])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f314,f304])).
fof(f1281,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1038,f258])).
fof(f1038,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f441,f439])).
fof(f1280,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1037,f258])).
fof(f1037,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f412,f439])).
fof(f1279,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1036,f258])).
fof(f1036,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f332,f439])).
fof(f927,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f439,f333])).
fof(f926,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f439,f338])).
fof(f925,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f439,f337])).
fof(f924,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f439,f311])).
fof(f923,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X21,bool),'const.sum.INR_1'(s(bool,X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f439,f334])).
fof(f922,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK16(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f439,f353])).
fof(f921,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f439,f304])).
fof(f911,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f439,f439])).
fof(f894,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43)))))) )),
  inference(superposition,[],[f439,f333])).
fof(f893,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))) )),
  inference(superposition,[],[f439,f338])).
fof(f892,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),'const.pair.,_2'(s(bool,X35),s(X34,X36)))))) )),
  inference(superposition,[],[f439,f337])).
fof(f891,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f439,f311])).
fof(f890,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X28,bool),'const.sum.INR_1'(s(bool,X29)))))) )),
  inference(superposition,[],[f439,f334])).
fof(f889,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK16(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f439,f353])).
fof(f888,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f439,f304])).
fof(f837,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool)),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f265])).
fof(f791,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f778,f258])).
fof(f778,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f441,f434])).
fof(f761,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f760,f260])).
fof(f760,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f313,f265])).
fof(f724,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) )),
  inference(superposition,[],[f338,f265])).
fof(f670,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6)))))) )),
  inference(superposition,[],[f337,f265])).
fof(f620,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f311,f265])).
fof(f534,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f334,f265])).
fof(f507,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f333,f265])).
fof(f488,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f483,f260])).
fof(f483,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f329,f265])).
fof(f441,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f265])).
fof(f434,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f261,f267])).
fof(f426,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f260,f266])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4)) )),
  inference(cnf_transformation,[],[f255])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.num.num',X2) = s('type.num.num',X5)) )),
  inference(cnf_transformation,[],[f255])).
fof(f400,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f395,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X24,X25,X46,X26] : ! [X27,X28,X29] : s(fun(X25,X24),'_dst_x0x1_2'(s(fun(X46,X24),X27),s(fun(X25,X46),happ(s(fun(X26,fun(X25,X46)),X28),s(X26,X29))))) = s(fun(X25,X24),happ(s(fun(X26,fun(X25,X24)),'_dst_x00x11_2'(s(fun(X46,X24),X27),s(fun(X26,fun(X25,X46)),X28))),s(X26,X29)))),
  file('thm.patSem.v_|49||49||split|5','_dst_x00x11_ax')).
fof(f393,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.patLang.exp',X1) = s('type.patLang.exp',X3)) )),
  inference(cnf_transformation,[],[f253])).
fof(f394,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patLang.exp',X1) != s('type.patLang.exp',X3) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2)) )),
  inference(cnf_transformation,[],[f253])).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f391,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f372])).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X31,X49,X41,X50] : ((s(bool,X31) = s(bool,X49) & (p(s(bool,X49)) => s(bool,X41) = s(bool,X50))) => ((p(s(bool,X31)) => p(s(bool,X41))) <=> (p(s(bool,X49)) => p(s(bool,X50)))))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.IMP_CONG')).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X24,X25,X26] : ! [X27,X28,X29] : s(X24,happ(s(fun(X25,X24),X27),s(X25,happ(s(fun(X26,X25),X28),s(X26,X29))))) = s(X24,happ(s(fun(X26,X24),'_dst_x0x1_2'(s(fun(X25,X24),X27),s(fun(X26,X25),X28))),s(X26,X29)))),
  file('thm.patSem.v_|49||49||split|5','_dst_x0x1_ax')).
fof(f363,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X24,X25,X26] : ! [X28,X29] : s(fun(X25,X24),'_dst_x0x1_2'(s(fun(X26,X24),X29),s(fun(X25,X26),X28))) = s(fun(X25,X24),happ(s(fun(fun(X26,X24),fun(X25,X24)),'_dstsel_0_x01_1'(s(fun(X25,X26),X28))),s(fun(X26,X24),X29)))),
  file('thm.patSem.v_|49||49||split|5','_dstsel_0_x01_ax')).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X16,X17,X20] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X16),s('type.list.list'('type.patLang.exp'),X17),s('type.num.num',X20))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.list_size_2'(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.list.list'('type.patLang.exp'),X17))),s('type.num.num',X20)))))))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_size_def|split|3')).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X30,X9] : ! [X39,X11] : ? [X40] : (s(X30,X39) = s(X30,happ(s(fun('type.list.list'(X9),X30),X40),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X30,happ(s(fun('type.list.list'(X9),X30),X40),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X30,happ(s(fun(X30,X30),happ(s(fun('type.list.list'(X9),fun(X30,X30)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X30,X30))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X30,happ(s(fun('type.list.list'(X9),X30),X40),s('type.list.list'(X9),X17))))))),
  file('thm.patSem.v_|49||49||split|5','thm.list.list_Axiom')).
fof(f354,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f348,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X30,X9] : ! [X10,X54] : (s(fun(X9,X30),X10) = s(fun(X9,X30),X54) <=> ! [X31] : s(X30,happ(s(fun(X9,X30),X10),s(X9,X31))) = s(X30,happ(s(fun(X9,X30),X54),s(X9,X31))))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.FUN_EQ_THM')).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f240])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.patSem.v_|49||49||split|5',aHL_EXT)).
fof(f338,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X30] : ! [X31,X41] : s(X30,X41) = s(X30,'const.pair.SND_1'(s('type.pair.prod'(X9,X30),'const.pair.,_2'(s(X9,X31),s(X30,X41)))))),
  file('thm.patSem.v_|49||49||split|5','thm.pair.SND')).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X30,X9] : ! [X31,X41] : s(X9,X31) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X30),'const.pair.,_2'(s(X9,X31),s(X30,X41)))))),
  file('thm.patSem.v_|49||49||split|5','thm.pair.FST')).
fof(f336,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X30,X9] : ! [X42] : ? [X10] : ! [X43,X44,X45] : s(X30,happ(s(fun('type.ind_type.recspace'(X9),X30),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X43),s(X9,X44),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45))))) = s(X30,happ(s(fun(fun('type.num.num',X30),X30),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X30),X30)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X30),X30))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X30),X30)))),X42),s('type.num.num',X43))),s(X9,X44))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45))),s(fun('type.num.num',X30),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X30),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X45)))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.CONSTR_REC')).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9,X30] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X30),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X30),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X30),X8)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_1_const.pair.FST')).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9,X30] : ! [X31] : s(X30,X31) = s(X30,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X30),'const.sum.INR_1'(s(X30,X31)))))),
  file('thm.patSem.v_|49||49||split|5','thm.sum.OUTR')).
fof(f329,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X1)))))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X16,X17] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.patSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X17)))))))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_size_def|split|1')).
fof(f325,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0)),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s('type.list.list'('type.patSem.v'),X1))),s('type.list.list'('type.patLang.exp'),X2))),s('type.num.num',X3))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X1),s('type.list.list'('type.patLang.exp'),X2),s('type.num.num',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X6),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.patSem.v'),X0),X9)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0)),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s('type.list.list'('type.patSem.v'),X1))),s('type.list.list'('type.patLang.exp'),X2))),s('type.num.num',X3))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X1),s('type.list.list'('type.patLang.exp'),X2),s('type.num.num',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X6),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.patSem.v'),X0),X9)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun('type.num.num',X0),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0)),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s('type.list.list'('type.patSem.v'),X1))),s('type.list.list'('type.patLang.exp'),X2))),s('type.num.num',X3))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X1),s('type.list.list'('type.patLang.exp'),X2),s('type.num.num',X3))),s(fun('type.ast.lit',X0),X4),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X6),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X7),s(fun('type.num.num',X0),X8),s(fun('type.list.list'('type.patSem.v'),X0),X9)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X9] : ! [X16,X17,X20,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9)),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s('type.list.list'('type.patSem.v'),X16))),s('type.list.list'('type.patLang.exp'),X17))),s('type.num.num',X20))) = s(X9,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X16),s('type.list.list'('type.patLang.exp'),X17),s('type.num.num',X20))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.patSem.v'),X9),X15)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_case_def|split|3')).
fof(f324,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s('type.num.num',X1))),s('type.list.list'('type.patSem.v'),X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s('type.num.num',X1))),s('type.list.list'('type.patSem.v'),X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.list.list'('type.patSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s('type.num.num',X1))),s('type.list.list'('type.patSem.v'),X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X9] : ! [X16,X17,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.patSem.v'),X9),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s('type.num.num',X16))),s('type.list.list'('type.patSem.v'),X17))) = s(X9,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.patSem.v'),X17))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.patSem.v'),X9),X15)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_case_def|split|1')).
fof(f323,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.patLang.exp',X0),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s('type.list.list'('type.patSem.v'),X1))),s('type.patLang.exp',X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.patLang.exp',X0),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s('type.list.list'('type.patSem.v'),X1))),s('type.patLang.exp',X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.patLang.exp',X0),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s('type.list.list'('type.patSem.v'),X1))),s('type.patLang.exp',X2))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X2))),s(fun('type.ast.lit',X0),X3),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X5),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X6),s(fun('type.num.num',X0),X7),s(fun('type.list.list'('type.patSem.v'),X0),X8)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X9] : ! [X16,X17,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.patLang.exp',X9),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s('type.list.list'('type.patSem.v'),X16))),s('type.patLang.exp',X17))) = s(X9,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X16),s('type.patLang.exp',X17))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.patSem.v'),X9),X15)))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_case_def|split|2')).
fof(f321,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : ! [X8,X22,X23] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X22),s('type.num.num',X23))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X22))),s('type.num.num',X23)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_3_const.ind_type.FCONS')).
fof(f320,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X22,X23] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X22),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X22))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_1_3_const.ind_type.CONSTR')).
fof(f319,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : ! [X6,X10,X36] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X36))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.FCONS|split|1')).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.FCONS|split|0')).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X33] : (s(bool,'F_0') = s(bool,X33) <=> ~p(s(bool,X33)))),
  file('thm.patSem.v_|49||49||split|5','thm.bool.EQ_CLAUSES|split|2')).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f264,plain,(
  ( ! [X0] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X6)))))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_size_def|split|5')).
fof(f263,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X8))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X8)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_1_const.patSem.Vectorv')).
fof(f262,plain,(
  ( ! [X0] : (s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X0))) = s('type.num.num',happ(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.patLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X0))) = s('type.num.num',happ(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.patLang.exp',X0)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X8] : s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X8))) = s('type.num.num',happ(s(fun('type.patLang.exp','type.num.num'),'const.patLang.exp_size_0'),s('type.patLang.exp',X8)))),
  file('thm.patSem.v_|49||49||split|5','ARITY_0_1_const.patLang.exp_size')).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.patSem.v_|49||49||split|5','thm.ind_type.ISO_REFL')).
fof(f6098,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f5409,f5335])).
fof(f5335,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f5297,f260])).
fof(f5297,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f4924,f2595])).
fof(f5409,plain,(
  ( ! [X23,X22] : (s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(subsumption_resolution,[],[f5364,f258])).
fof(f5364,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(superposition,[],[f441,f5287])).
fof(f5287,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f426,f4924,f1954])).
fof(f9239,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl22),
  inference(cnf_transformation,[],[f9239_D])).
fof(f9239_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f16860,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f16699])).
fof(f16699,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6098,f10031])).
fof(f10031,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10021])).
fof(f10021,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9302,f439])).
fof(f9302,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2017,f2045,f2026,f2126,f2239,f2215,f2217,f2272,f2277,f2278,f2279,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2705,f2595,f2741,f2757,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3199,f3200,f3202,f3203,f3205,f3206,f3207,f3242,f3271,f3273,f3274,f3276,f3277,f3278,f3384,f3455,f3456,f3892,f3899,f3550,f3914,f3915,f3916,f3939,f3922,f3923,f3956,f3989,f3990,f3992,f3993,f3995,f3996,f3997,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4463,f4299,f4300,f4353,f4355,f4356,f4465,f4371,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4745,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4642,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4619,f4616,f4464,f2743,f2019,f343,f342,f316,f315,f312,f3897,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f3942,f3934,f3935,f4921,f4918,f5230,f4596,f4199,f4053,f3991,f3272,f3201,f2707,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f3994,f3275,f3204,f2710,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5704,f5711,f5712,f5713,f5714,f5716,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6052,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6350,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6572,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7606,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8466,f8469,f8473,f8569,f8577,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8774,f8779,f8781,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9289,f9291,f9293,f9294,f9296,f9297,f9298,f9299,f9301])).
fof(f9301,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9300])).
fof(f9300,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f6511,f3456])).
fof(f6511,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6278,f2595])).
fof(f16517,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) | $spl36),
  inference(cnf_transformation,[],[f16517_D])).
fof(f16517_D,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f16442,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f16285])).
fof(f16285,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6098,f9920])).
fof(f9920,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f9910])).
fof(f9910,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9297,f439])).
fof(f13224,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl32),
  inference(cnf_transformation,[],[f13224_D])).
fof(f13224_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13226,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl34),
  inference(cnf_transformation,[],[f13226_D])).
fof(f13226_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13209,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl30),
  inference(superposition,[],[f13163,f304])).
fof(f2021,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl12),
  inference(cnf_transformation,[],[f2021_D])).
fof(f2021_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f13163,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl30),
  inference(cnf_transformation,[],[f13163_D])).
fof(f13163_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f13169,plain,(
  ( ! [X59,X60] : (p(s(bool,sK5('_i_0_0',X60,bool))) | ~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool)))))) )),
  inference(forward_demodulation,[],[f13138,f304])).
fof(f13138,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X60,bool)))))) )),
  inference(superposition,[],[f315,f304])).
fof(f2015,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2015_D])).
fof(f2015_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13162,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f13161,f304])).
fof(f13161,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f13125,f258])).
fof(f13125,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f315,f466])).
fof(f9597,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2239,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4353,f4355,f4356,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918,f5230,f4596,f4199,f4053,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5711,f5712,f5713,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9293,f9294,f9298,f9299,f5598,f5618,f4594,f4614,f2188,f9324,f9238,f9548,f9590,f9592,f9594,f9596])).
fof(f9596,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9504,f258])).
fof(f9504,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9238,f351])).
fof(f9594,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9503,f258])).
fof(f9503,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9238,f350])).
fof(f9592,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(subsumption_resolution,[],[f9502,f258])).
fof(f9502,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(backward_demodulation,[],[f9238,f346])).
fof(f9590,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(subsumption_resolution,[],[f9501,f258])).
fof(f9501,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(backward_demodulation,[],[f9238,f345])).
fof(f9548,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f9327,f9238])).
fof(f9327,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f9238,f270])).
fof(f2188,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f265,f319])).
fof(f9595,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2239,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4353,f4355,f4356,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918,f5230,f4596,f4199,f4053,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5711,f5712,f5713,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9293,f9294,f9298,f9299,f5598,f5618,f4594,f4614,f2188,f9324,f9238,f9548,f9590,f9592,f9594])).
fof(f9593,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2239,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4353,f4355,f4356,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918,f5230,f4596,f4199,f4053,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5711,f5712,f5713,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9293,f9294,f9298,f9299,f5598,f5618,f4594,f4614,f2188,f9324,f9238,f9548,f9590,f9592])).
fof(f9591,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2239,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4353,f4355,f4356,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918,f5230,f4596,f4199,f4053,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5711,f5712,f5713,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9293,f9294,f9298,f9299,f5598,f5618,f4594,f4614,f2188,f9324,f9238,f9548,f9590])).
fof(f9557,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9466,f9238])).
fof(f9466,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9238,f9302])).
fof(f9556,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9465,f9238])).
fof(f9465,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9238,f9297])).
fof(f9549,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f257,f256,f258,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f356,f363,f364,f412,f391,f390,f389,f394,f393,f392,f395,f400,f399,f398,f397,f426,f434,f439,f441,f466,f488,f507,f534,f620,f670,f724,f761,f791,f837,f888,f889,f890,f891,f892,f893,f894,f911,f921,f922,f923,f924,f925,f926,f927,f1279,f1280,f1281,f1306,f1307,f1308,f1309,f1310,f1311,f1312,f1295,f1297,f1299,f1313,f1305,f1416,f1447,f1558,f1559,f1560,f1561,f1562,f1563,f1564,f1565,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1947,f1948,f1949,f1950,f1951,f1952,f1954,f1956,f1957,f2014,f2020,f2239,f2215,f2217,f2272,f2277,f2278,f2371,f2522,f2523,f2524,f2525,f2526,f2527,f2528,f2529,f2530,f2772,f2773,f2774,f2775,f2776,f2777,f2778,f2779,f2780,f3914,f3915,f3916,f4052,f4054,f4055,f4057,f4058,f4059,f4138,f4198,f4200,f4201,f4203,f4204,f4205,f4285,f4291,f4292,f4293,f4353,f4355,f4356,f4375,f4383,f4385,f4390,f4400,f4469,f4623,f4699,f4693,f4822,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4917,f4700,f4595,f4637,f4305,f4308,f4312,f4311,f4310,f4309,f4307,f4306,f4304,f4303,f396,f352,f351,f350,f346,f345,f1999,f3045,f3048,f344,f341,f4616,f4464,f2019,f343,f342,f316,f315,f312,f3052,f3051,f3050,f3049,f3047,f3046,f3044,f3043,f414,f306,f3926,f3927,f3928,f4918,f5230,f4596,f4199,f4053,f2022,f1902,f1407,f1405,f1403,f4202,f4056,f2202,f2176,f2174,f4498,f4377,f4374,f4537,f4360,f4359,f4354,f4352,f4347,f4346,f4344,f4343,f4342,f4341,f4340,f4376,f759,f758,f757,f2209,f756,f755,f1434,f858,f754,f4468,f4467,f313,f5701,f5711,f5712,f5713,f5717,f5721,f5724,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f5700,f6004,f6014,f6049,f6187,f6188,f6190,f6191,f6192,f6193,f6195,f6278,f6286,f6302,f6312,f6347,f6390,f6391,f6393,f6394,f6395,f6396,f6398,f6482,f6490,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6985,f6996,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7947,f7958,f8584,f8585,f8586,f8587,f8588,f8589,f8590,f8591,f8592,f8593,f8792,f8793,f8794,f8795,f8796,f8797,f8798,f8799,f8800,f8801,f9293,f9294,f9298,f9299,f5598,f5618,f4594,f4614,f2188,f9324,f9238,f9548])).
fof(f9547,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f9326,f258])).
fof(f9326,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f9238,f260])).
fof(f9325,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl28),
  inference(cnf_transformation,[],[f9325_D])).
fof(f9325_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1315,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9324,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1913,f258])).
fof(f1913,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f314,f466])).
fof(f4614,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4581,f258])).
fof(f4581,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4464,f466])).
fof(f4594,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4464,f304])).
fof(f5618,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f5585,f258])).
fof(f5585,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4464,f466])).
fof(f5598,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4464,f304])).
fof(f9292,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f9292_D])).
fof(f9292_D,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9299,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f6519,f6278])).
fof(f6519,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X52,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f6278,f1948])).
fof(f9298,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f6530,f6482])).
fof(f6530,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6278,f1952])).
fof(f9290,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl24),
  inference(cnf_transformation,[],[f9290_D])).
fof(f9290_D,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9294,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f7552,f6278])).
fof(f7552,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X52),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X50,bool),'const.sum.INR_1'(s(bool,X51))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f6482,f1948])).
fof(f9293,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f7563,f6482])).
fof(f7563,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6482,f1952])).
fof(f9291,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(subsumption_resolution,[],[f8524,f8569])).
fof(f8524,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f6278,f2705])).
fof(f9289,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(subsumption_resolution,[],[f8525,f8569])).
fof(f8525,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f6482,f2705])).
fof(f9174,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9041])).
fof(f9041,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2344,f7488])).
fof(f7488,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f7474])).
fof(f7474,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f6570,f439])).
fof(f6570,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f6498,f1330])).
fof(f1330,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1316,f271])).
fof(f6498,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1331,f6278])).
fof(f1331,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1316,f266])).
fof(f2344,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2343,f260])).
fof(f2343,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2309,f1330])).
fof(f2309,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f1956,f1354])).
fof(f1354,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1332,f258])).
fof(f1332,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1316,f439])).
fof(f5719,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl18),
  inference(cnf_transformation,[],[f5719_D])).
fof(f5719_D,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5718,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5684,f270])).
fof(f5684,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X85,bool)))) | s(bool,'F_0') = s(bool,X84)) ) | $spl6),
  inference(superposition,[],[f4637,f2280])).
fof(f2280,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f2257,f258])).
fof(f2257,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl6),
  inference(superposition,[],[f1954,f1522])).
fof(f1522,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1469,f267])).
fof(f1469,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1331,f313])).
fof(f5715,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5715_D])).
fof(f5715_D,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5714,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5681,f270])).
fof(f5681,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X75),bool),sK16(X76,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X75,bool)))) | s(bool,'F_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f4637,f2279])).
fof(f5229,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4936,f260])).
fof(f4936,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f426,f2021,f329])).
fof(f5228,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4938,f2021])).
fof(f4938,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1331,f2021,f329])).
fof(f5227,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4946,f260])).
fof(f4946,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f426,f2021,f329])).
fof(f5226,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4948,f2021])).
fof(f4948,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1331,f2021,f329])).
fof(f5225,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5217,f260])).
fof(f5217,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(backward_demodulation,[],[f5009,f4700])).
fof(f5009,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f333,f2021,f1956])).
fof(f5130,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5129,f426])).
fof(f5129,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5012,f434])).
fof(f5012,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f434,f2021,f1956])).
fof(f5128,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5127,f426])).
fof(f5127,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f5013,f1522])).
fof(f5013,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1522,f2021,f1956])).
fof(f5126,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5125,f426])).
fof(f5125,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5014,f4699])).
fof(f5014,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f4699,f2021,f1956])).
fof(f5124,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f5123,f426])).
fof(f5123,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f5015,f2045])).
fof(f5015,plain,(
  s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f2045,f2021,f1956])).
fof(f5033,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f2021,f4595])).
fof(f5031,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1331,f2021,f313])).
fof(f5027,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f434,f426,f2021,f1956])).
fof(f4997,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1331,f1331,f2021,f488])).
fof(f4995,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f426,f426,f2021,f488])).
fof(f4973,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f333,f2021,f412])).
fof(f4971,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f338,f2021,f412])).
fof(f4969,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f337,f2021,f412])).
fof(f4967,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f311,f2021,f412])).
fof(f4965,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f334,f2021,f412])).
fof(f4963,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f353,f2021,f412])).
fof(f4961,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f304,f2021,f412])).
fof(f4958,plain,(
  $false | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f258,f2045,f2021,f412])).
fof(f4956,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f4699,f2021,f412])).
fof(f4954,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f258,f1522,f2021,f412])).
fof(f4952,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f434,f2021,f412])).
fof(f4949,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2021,f1331,f2021,f329])).
fof(f4947,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f260,f426,f2021,f329])).
fof(f4945,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2021,f1331,f2021,f329])).
fof(f4943,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2021,f426,f2021,f329])).
fof(f4939,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2021,f1331,f2021,f329])).
fof(f4937,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f260,f426,f2021,f329])).
fof(f4935,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2021,f1331,f2021,f329])).
fof(f4933,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2021,f426,f2021,f329])).
fof(f4929,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f426,f2021,f271])).
fof(f4927,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f434,f2021,f266])).
fof(f4925,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f2021])).
fof(f4924,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl14),
  inference(cnf_transformation,[],[f4924_D])).
fof(f4924_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4917,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4596,f337])).
fof(f2017,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2017_D])).
fof(f2017_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2020,plain,(
  ( ! [X24] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f1998,f304])).
fof(f1998,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f312,f304])).
fof(f2014,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2013,f304])).
fof(f2013,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1991,f258])).
fof(f1991,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f312,f466])).
fof(f1316,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1316_D])).
fof(f1316_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1305,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f314,f304])).
fof(f425,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f424])).
fof(f424,plain,(
  s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f421,f422])).
fof(f420,plain,(
  s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) | $spl0),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f422,plain,(
  s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1) | $spl2),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f256,plain,(
  s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1) | s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1)))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) | s('type.list.list'('type.patSem.v'),sK0) = s('type.list.list'('type.patSem.v'),sK1)) & (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) | s('type.list.list'('type.patSem.v'),sK0) != s('type.list.list'('type.patSem.v'),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f198])).
fof(f198,plain,(
  ? [X0,X1] : ((s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))) | s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X1)) & (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X1)))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ? [X0,X1] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))) <~> s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))) <=> s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X7))) <=> s('type.list.list'('type.patSem.v'),X6) = s('type.list.list'('type.patSem.v'),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X7))) <=> s('type.list.list'('type.patSem.v'),X6) = s('type.list.list'('type.patSem.v'),X7))),
  file('thm.patSem.v_|49||49||split|5','thm.patSem.v_|49||49||split|5')).
fof(f421,plain,(
  s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) | $spl1),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f423,plain,(
  s('type.list.list'('type.patSem.v'),sK0) != s('type.list.list'('type.patSem.v'),sK1) | $spl3),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  s('type.list.list'('type.patSem.v'),sK0) != s('type.list.list'('type.patSem.v'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f257,plain,(
  s('type.list.list'('type.patSem.v'),sK0) != s('type.list.list'('type.patSem.v'),sK1) | s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK0))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK1)))),
  inference(cnf_transformation,[],[f199])).
% SZS output end Proof for thm.patSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 180
% Passive clauses: 1811
% Generated clauses: 93129
% Final active clauses: 130
% Final passive clauses: 1187
% Input formulas: 69
% Initial clauses: 145
% 
% Duplicate literals: 764
% Trivial inequalities: 181
% Fw subsumption resolutions: 940
% Fw demodulations: 51250
% Bw demodulations: 339
% Condensations: 964
% Global subsumptions: 8
% 
% Simple tautologies: 285
% Equational tautologies: 146
% Forward subsumptions: 16277
% Fw demodulations to eq. taut.: 20426
% Bw demodulations to eq. taut.: 51
% 
% Unit resulting resolution: 26472
% Factoring: 1
% Forward superposition: 6086
% Backward superposition: 5255
% Self superposition: 147
% Equality factoring: 80
% Equality resolution: 33
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 1924
% SAT solver unit clauses: 183
% SAT solver binary clauses: 704
% 
% Sat splits: 26
% Sat splitting refutations: 51
% 
% Memory used [KB]: 44007
% Time elapsed: 2.737 s
% ------------------------------
% ------------------------------
% Success in time 4.061 s
