Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.mod_to_conProof
% (121147)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1007
% Passive clauses: 4888
% Generated clauses: 26999
% Final active clauses: 852
% Final passive clauses: 3408
% Input formulas: 516
% Initial clauses: 989
% 
% Duplicate literals: 161
% Trivial inequalities: 133
% Fw subsumption resolutions: 1307
% Fw demodulations: 2613
% Bw demodulations: 154
% 
% Simple tautologies: 357
% Equational tautologies: 16
% Forward subsumptions: 17345
% 
% Binary resolution: 4270
% Factoring: 420
% Forward superposition: 8925
% Backward superposition: 7819
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 14
% 
% Split clauses: 319
% Split components: 640
% SAT solver clauses: 462
% SAT solver unit clauses: 22
% SAT solver binary clauses: 93
% 
% Sat splits: 319
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 462
% TWLsolver calls for satisfiability: 326
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.mod_to_conProof
% (121156)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Passive clauses: 648
% Generated clauses: 1908
% Final passive clauses: 639
% Input formulas: 516
% Initial clauses: 989
% 
% Duplicate literals: 11
% Trivial inequalities: 1
% Fw subsumption resolutions: 6
% Fw demodulations: 70
% Bw demodulations: 9
% Condensations: 19
% Global subsumptions: 39
% 
% Simple tautologies: 123
% Equational tautologies: 3
% Forward subsumptions: 11
% 
% SAT solver clauses: 1950
% SAT solver unit clauses: 1323
% SAT solver binary clauses: 287
% SAT solver learnt clauses: 78
% SAT solver learnt literals: 24
% 
% TWLsolver clauses: 1794
% TWLsolver calls for satisfiability: 6972
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.mod_to_conProof
% (121179)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Passive clauses: 269
% Generated clauses: 1363
% Final passive clauses: 250
% Input formulas: 516
% Initial clauses: 989
% 
% Duplicate literals: 11
% Trivial inequalities: 1
% Fw subsumption resolutions: 3
% Fw demodulations: 321
% Bw demodulations: 19
% Condensations: 19
% 
% Simple tautologies: 123
% Equational tautologies: 3
% Forward subsumptions: 45
% 
% SAT solver clauses: 271
% SAT solver unit clauses: 194
% SAT solver binary clauses: 49
% 
% Memory used [KB]: 33389
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 31 next slice time: 31
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.mod_to_conProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.mod_to_conProof
% SZS output start Proof for thm.mod_to_conProof
fof(f1745,plain,(
  $false),
  inference(equality_resolution,[],[f1242])).
fof(f1242,plain,(
  ( ! [X0] : (s(fun(sK15,bool),sK16) != s(fun(sK15,bool),X0)) )),
  inference(resolution,[],[f1229,f924])).
fof(f924,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f731])).
fof(f731,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))))),
  inference(nnf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f546])).
fof(f546,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f300])).
fof(f300,axiom,(
  ! [X6] : ! [X28,X55] : (p(s(bool,'const.bool.IN_2'(s(X6,X28),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X55),s(fun(X6,bool),'const.pred_set.EMPTY_0')))))) <=> s(X6,X28) = s(X6,X55))),
  file('thm.mod_to_conProof.mod_to_conProof_238_prove_ltb','thm.pred_set.IN_SING')).
fof(f1229,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK15,bool),X4),s(fun(fun(sK15,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK15,bool),sK16),s(fun(fun(sK15,bool),bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(trivial_inequality_removal,[],[f1228])).
fof(f1228,plain,(
  ( ! [X4] : (s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),X4))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),X4))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK15,bool),X4),s(fun(fun(sK15,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK15,bool),sK16),s(fun(fun(sK15,bool),bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(forward_demodulation,[],[f1227,f926])).
fof(f926,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6] : ! [X18] : s(fun(X6,bool),X18) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X18)))),
  file('thm.mod_to_conProof.mod_to_conProof_238_prove_ltb','thm.pred_set.UNION_EMPTY|split|0')).
fof(f1227,plain,(
  ( ! [X4] : (s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),X4))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),'const.pred_set.EMPTY_0'),s(fun(sK15,bool),X4))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK15,bool),X4),s(fun(fun(sK15,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK15,bool),sK16),s(fun(fun(sK15,bool),bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(forward_demodulation,[],[f1226,f932])).
fof(f932,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) )),
  inference(cnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(rectify,[],[f219])).
fof(f219,axiom,(
  ! [X6] : ! [X28,X18,X19] : s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X28),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X19))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X28),s(fun(X6,bool),X18))),s(fun(X6,bool),X19)))),
  file('thm.mod_to_conProof.mod_to_conProof_238_prove_ltb','thm.pred_set.INSERT_UNION_EQ')).
fof(f1226,plain,(
  ( ! [X4] : (s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),X4))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),'const.pred_set.EMPTY_0'))),s(fun(sK15,bool),X4))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK15,bool),X4),s(fun(fun(sK15,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK15,bool),sK16),s(fun(fun(sK15,bool),bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(forward_demodulation,[],[f1189,f931])).
fof(f931,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) )),
  inference(cnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(flattening,[],[f556])).
fof(f556,plain,(
  ! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X18,X19,X20] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X18),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X19),s(fun(X6,bool),X20))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X19))),s(fun(X6,bool),X20)))),
  file('thm.mod_to_conProof.mod_to_conProof_238_prove_ltb','thm.pred_set.UNION_ASSOC')).
fof(f1189,plain,(
  ( ! [X4] : (s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),X4))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),'const.pred_set.EMPTY_0'))))),s(fun(sK15,bool),X4))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK15,bool),X4),s(fun(fun(sK15,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK15,bool),sK16),s(fun(fun(sK15,bool),bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(superposition,[],[f855,f923])).
fof(f923,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(cnf_transformation,[],[f731])).
fof(f855,plain,(
  s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),sK16))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),'const.pred_set.EMPTY_0'))))),s(fun(sK15,bool),sK16)))),
  inference(cnf_transformation,[],[f698])).
fof(f698,plain,(
  s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),sK16))))) != s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),'const.pred_set.UNION_2'(s(fun(sK15,bool),sK18),s(fun(sK15,bool),'const.pred_set.INSERT_2'(s(sK15,sK17),s(fun(sK15,bool),'const.pred_set.EMPTY_0'))))),s(fun(sK15,bool),sK16)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17,sK18])],[f643])).
fof(f643,plain,(
  ? [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),X7))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),X7))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),s(fun(X6,bool),X7)))),
  file('thm.mod_to_conProof.mod_to_conProof_238_prove_ltb','thm.mod_to_conProof.mod_to_conProof_238_prove_ltb')).
% SZS output end Proof for thm.mod_to_conProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 207
% Passive clauses: 200
% Generated clauses: 563
% Final active clauses: 206
% Final passive clauses: 196
% Input formulas: 516
% Initial clauses: 290
% 
% Selected by SInE selection: 72
% SInE iterations: 1
% 
% Duplicate literals: 22
% Trivial inequalities: 4
% Fw subsumption resolutions: 46
% Fw demodulations: 38
% 
% Simple tautologies: 96
% Equational tautologies: 2
% Forward subsumptions: 178
% 
% Binary resolution: 319
% Forward superposition: 166
% Equality resolution: 1
% 
% Split clauses: 6
% Split components: 13
% SAT solver clauses: 9
% SAT solver unit clauses: 2
% SAT solver binary clauses: 5
% 
% Sat splits: 6
% Sat splitting refutations: 2
% 
% Memory used [KB]: 2942
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
% Success in time 11.88 s
