Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.closProps
% (86465)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 3143
% Generated clauses: 15136
% Final active clauses: 121
% Final passive clauses: 2862
% Input formulas: 16
% Initial clauses: 20
% 
% Duplicate literals: 890
% Trivial inequalities: 86
% Fw subsumption resolutions: 1672
% Fw demodulations: 2050
% Bw demodulations: 5
% 
% Simple tautologies: 243
% Equational tautologies: 20
% Forward subsumptions: 6346
% 
% Binary resolution: 202
% Forward superposition: 6323
% Backward superposition: 5138
% Self superposition: 236
% Equality factoring: 32
% Equality resolution: 19
% 
% Split clauses: 425
% Split components: 856
% SAT solver clauses: 834
% SAT solver unit clauses: 11
% SAT solver binary clauses: 105
% 
% Sat splits: 425
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 833
% TWLsolver calls for satisfiability: 170
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f1637,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1636,f50,f71,f417,f440,f64,f81,f66,f83,f1461,f1568,f69,f86,f1614,f1621,f67,f84,f1629])).
fof(f1629,plain,(
  $false | ($spl1 | $spl26 | $spl438)),
  inference(subsumption_resolution,[],[f1628,f86])).
fof(f1628,plain,(
  s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,happ(s(fun('type.semanticPrimitives.abort',sK2),sK4),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))))) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4))),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) | ($spl1 | $spl438)),
  inference(forward_demodulation,[],[f1627,f1621])).
fof(f1627,plain,(
  s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4))),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) != s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) | ($spl1 | $spl438)),
  inference(backward_demodulation,[],[f1621,f71])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))) ) | $spl22),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f67,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X8] : ! [X14,X12,X10] : s(X8,happ(s(fun('type.semanticPrimitives.abort',X8),X10),s('type.semanticPrimitives.abort',X14))) = s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X14))),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.abort',X8),X10)))),
  file('thm.closProps.closProps_78_prove_ltb','thm.semanticPrimitives.error_result_case_def|split|1')).
fof(f1621,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.abort',X2),X3),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.abort',X2),X3)))) ) | $spl438),
  inference(cnf_transformation,[],[f1621_D])).
fof(f1621_D,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.abort',X2),X3),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.abort',X2),X3)))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f1614,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.abort',X2),X3),s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.abort',X2),X3)))) ) | ($spl22 | $spl432)),
  inference(superposition,[],[f84,f1610])).
fof(f1610,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK3) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) | $spl432),
  inference(cnf_transformation,[],[f1610_D])).
fof(f1610_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK3) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f86,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) ) | $spl26),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f69,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))) = s(X16,happ(s(fun(X18,X16),'_dst_x0x1_2'(s(fun(X17,X16),X19),s(fun(X18,X17),X20))),s(X18,X21)))),
  file('thm.closProps.closProps_78_prove_ltb','_dst_x0x1_ax')).
fof(f1568,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X5,happ(s(fun(X1,X5),'_dst_x0x1_2'(s(fun(X0,X5),X6),s(fun(X1,X0),X2))),s(X1,sK7(X3,X1)))) = s(X5,happ(s(fun(X0,X5),X6),s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),X3),s(fun(X1,X0),X2),s(fun('type.semanticPrimitives.abort',X0),X4))))) | s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X3,X1))))) ) | $spl360),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X5,happ(s(fun(X1,X5),'_dst_x0x1_2'(s(fun(X0,X5),X6),s(fun(X1,X0),X2))),s(X1,sK7(X3,X1)))) = s(X5,happ(s(fun(X0,X5),X6),s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),X3),s(fun(X1,X0),X2),s(fun('type.semanticPrimitives.abort',X0),X4))))) | s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X3,X1))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f1461,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X5,happ(s(fun(X1,X5),'_dst_x0x1_2'(s(fun(X0,X5),X6),s(fun(X1,X0),X2))),s(X1,sK7(X3,X1)))) = s(X5,happ(s(fun(X0,X5),X6),s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),X3),s(fun(X1,X0),X2),s(fun('type.semanticPrimitives.abort',X0),X4))))) | s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X3,X1))))) ) | ($spl26 | $spl72)),
  inference(superposition,[],[f86,f440])).
fof(f83,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))) ) | $spl20),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f66,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X6] : ! [X14,X12,X10] : s(X8,happ(s(fun(X6,X8),X12),s(X6,X14))) = s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X14))),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.abort',X8),X10)))),
  file('thm.closProps.closProps_78_prove_ltb','thm.semanticPrimitives.error_result_case_def|split|0')).
fof(f81,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK7(X1,X0))))) ) | $spl16),
  inference(cnf_transformation,[],[f81_D])).
fof(f81_D,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK7(X1,X0))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f64,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK7(X1,X0))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK7(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X15] : (? [X14] : s('type.semanticPrimitives.error_result'(X6),X15) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X14))) | ? [X14] : s('type.semanticPrimitives.error_result'(X6),X15) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X14))))),
  file('thm.closProps.closProps_78_prove_ltb','thm.semanticPrimitives.error_result_nchotomy')).
fof(f440,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK7(X1,X0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X1),s(fun(X0,X2),X3),s(fun('type.semanticPrimitives.abort',X2),X4))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0))))) ) | $spl72),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK7(X1,X0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X1),s(fun(X0,X2),X3),s(fun('type.semanticPrimitives.abort',X2),X4))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f417,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK7(X1,X0)))) = s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X1),s(fun(X0,X2),X3),s(fun('type.semanticPrimitives.abort',X2),X4))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(X1,X0))))) ) | ($spl16 | $spl20)),
  inference(superposition,[],[f83,f81])).
fof(f71,plain,(
  s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun(sK0,sK2),sK6))),s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) | $spl1),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun(sK0,sK2),sK6))),s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f50,plain,(
  s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun(sK0,sK2),sK6))),s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4)))))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun(sK0,sK2),sK6))),s(fun('type.semanticPrimitives.abort',sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK2),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f41])).
fof(f41,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(X1,happ(s(fun(X2,X1),X5),s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.abort',X2),X4))))) != s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun(X0,X2),X6))),s(fun('type.semanticPrimitives.abort',X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun('type.semanticPrimitives.abort',X2),X4)))))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(X1,happ(s(fun(X2,X1),X5),s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.abort',X2),X4))))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun(X0,X2),X6))),s(fun('type.semanticPrimitives.abort',X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun('type.semanticPrimitives.abort',X2),X4)))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6] : s(X1,happ(s(fun(X2,X1),X5),s(X2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X2),X6),s(fun('type.semanticPrimitives.abort',X2),X4))))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X3),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun(X0,X2),X6))),s(fun('type.semanticPrimitives.abort',X1),'_dst_x0x1_2'(s(fun(X2,X1),X5),s(fun('type.semanticPrimitives.abort',X2),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12] : s(X7,happ(s(fun(X8,X7),X11),s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X9),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.abort',X8),X10))))) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X9),s(fun(X6,X7),'_dst_x0x1_2'(s(fun(X8,X7),X11),s(fun(X6,X8),X12))),s(fun('type.semanticPrimitives.abort',X7),'_dst_x0x1_2'(s(fun(X8,X7),X11),s(fun('type.semanticPrimitives.abort',X8),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12] : s(X7,happ(s(fun(X8,X7),X11),s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X9),s(fun(X6,X8),X12),s(fun('type.semanticPrimitives.abort',X8),X10))))) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X9),s(fun(X6,X7),'_dst_x0x1_2'(s(fun(X8,X7),X11),s(fun(X6,X8),X12))),s(fun('type.semanticPrimitives.abort',X7),'_dst_x0x1_2'(s(fun(X8,X7),X11),s(fun('type.semanticPrimitives.abort',X8),X10)))))),
  file('thm.closProps.closProps_78_prove_ltb','thm.closProps.closProps_78_prove_ltb')).
fof(f1636,plain,(
  $false | ($spl1 | $spl72 | $spl360 | $spl433)),
  inference(subsumption_resolution,[],[f1635,f1609])).
fof(f1609,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK3) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) | $spl433),
  inference(cnf_transformation,[],[f1609_D])).
fof(f1609_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK3) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f1635,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK3) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) | ($spl1 | $spl72 | $spl360)),
  inference(subsumption_resolution,[],[f1634,f1568])).
fof(f1634,plain,(
  s(sK1,happ(s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK2,sK1),sK5),s(fun(sK0,sK2),sK6))),s(sK0,sK7(sK3,sK0)))) != s(sK1,happ(s(fun(sK2,sK1),sK5),s(sK2,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK3),s(fun(sK0,sK2),sK6),s(fun('type.semanticPrimitives.abort',sK2),sK4))))) | s('type.semanticPrimitives.error_result'(sK0),sK3) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK8(sK3,sK0)))) | ($spl1 | $spl72)),
  inference(superposition,[],[f71,f440])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 44
% Passive clauses: 223
% Generated clauses: 1550
% Final active clauses: 35
% Final passive clauses: 163
% Input formulas: 16
% Initial clauses: 20
% 
% Duplicate literals: 94
% Trivial inequalities: 25
% Fw subsumption resolutions: 64
% Fw demodulations: 68
% Bw demodulations: 19
% Condensations: 125
% 
% Simple tautologies: 44
% Equational tautologies: 18
% Forward subsumptions: 546
% 
% Binary resolution: 15
% Forward superposition: 390
% Backward superposition: 416
% Self superposition: 51
% Equality factoring: 28
% Equality resolution: 8
% 
% Split clauses: 25
% Split components: 55
% SAT solver clauses: 361
% SAT solver unit clauses: 27
% SAT solver binary clauses: 26
% SAT solver learnt clauses: 1
% 
% Sat splits: 25
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 360
% TWLsolver calls for satisfiability: 98
% 
% Memory used [KB]: 3198
% Time elapsed: 0.182 s
% ------------------------------
% ------------------------------
% Success in time 0.8 s
