Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.infer_eComplete
% (144639)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 731
% Passive clauses: 4025
% Generated clauses: 14541
% Final active clauses: 721
% Final passive clauses: 2935
% Input formulas: 40
% Initial clauses: 182
% 
% Duplicate literals: 577
% Trivial inequalities: 85
% Fw subsumption resolutions: 1788
% Fw demodulations: 529
% Bw demodulations: 10
% 
% Simple tautologies: 226
% Equational tautologies: 47
% Forward subsumptions: 6652
% 
% Binary resolution: 5707
% Factoring: 1297
% Forward superposition: 2115
% Backward superposition: 1970
% Self superposition: 29
% Equality factoring: 30
% Equality resolution: 20
% 
% Split clauses: 548
% Split components: 1105
% SAT solver clauses: 895
% SAT solver unit clauses: 16
% SAT solver binary clauses: 88
% 
% Sat splits: 548
% 
% TWLsolver clauses: 894
% TWLsolver calls for satisfiability: 402
% 
% Memory used [KB]: 6012
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.infer_eComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.infer_eComplete
% SZS output start Proof for thm.infer_eComplete
fof(f9509,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9508,f237,f440,f438,f590,f595,f257,f458,f2233,f2237,f2235,f449,f451,f281,f465,f2240,f2262,f288,f468,f2797,f568,f572,f234,f446,f253,f456,f9038,f9040,f2060,f2061,f415,f558,f570,f576,f236,f442,f2137,f2138,f289,f467,f2864,f2871,f637,f645,f620,f627,f473,f474,f2816,f293,f471,f9309,f282,f464,f569,f574,f235,f444,f4770,f4772,f1890,f1902,f287,f469,f238,f448,f9375,f9429,f9037,f9356])).
fof(f9356,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | $spl3568),
  inference(cnf_transformation,[],[f9356_D])).
fof(f9356_D,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) <=> ~$spl3568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3568])])).
fof(f9037,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | ($spl10 | $spl760 | $spl818 | $spl1252)),
  inference(forward_demodulation,[],[f9036,f2138])).
fof(f9036,plain,(
  s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | ($spl10 | $spl760 | $spl1252)),
  inference(subsumption_resolution,[],[f9001,f448])).
fof(f9001,plain,(
  ~p(s(bool,'T_0')) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | ($spl760 | $spl1252)),
  inference(superposition,[],[f2061,f2871])).
fof(f9429,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | $spl3570),
  inference(cnf_transformation,[],[f9429_D])).
fof(f9429_D,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) <=> ~$spl3570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3570])])).
fof(f9375,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | ($spl212 | $spl3568)),
  inference(backward_demodulation,[],[f9356,f576])).
fof(f448,plain,(
  p(s(bool,'T_0')) | $spl10),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f238,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb',aHL_TRUTH)).
fof(f469,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl34),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f287,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X3] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))))))) & (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',sK17(X1,X0)))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',sK17(X1,X0)))))) | p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X3] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))))))) & (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ? [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))) | p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))))) & (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ? [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))) | p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2))))))) & ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ? [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2)))))) | p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) <=> (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ! [X2] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X2))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X2)))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X22,X26] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26)))) <=> (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22)))) & p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26)))) & ! [X10] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26),s('type.infer_t.infer_t',X10))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22),s('type.infer_t.infer_t',X10)))))))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.unify.t_compat_def')).
fof(f1902,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) ) | $spl632),
  inference(cnf_transformation,[],[f1902_D])).
fof(f1902_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f1890,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) ) | ($spl16 | $spl34)),
  inference(resolution,[],[f469,f456])).
fof(f4772,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))) | $spl1684),
  inference(cnf_transformation,[],[f4772_D])).
fof(f4772_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))) <=> ~$spl1684),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1684])])).
fof(f4770,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))) | ($spl10 | $spl208 | $spl632)),
  inference(subsumption_resolution,[],[f4741,f448])).
fof(f4741,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))) | ($spl208 | $spl632)),
  inference(superposition,[],[f1902,f572])).
fof(f444,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))))) | $spl6),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f235,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))))))),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) & (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) & (~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3)))))))) & (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3)))))))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & ? [X3] : ((p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3))))))))) & (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3)))))))))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) => ! [X3] : ((p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3))))))))) => (p(s(bool,'const.bool.IN_2'(s('type.num.num',X3),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X3)))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) => ! [X9] : ((p(s(bool,'const.bool.IN_2'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9))))))))) => (p(s(bool,'const.bool.IN_2'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9)))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) => ! [X9] : ((p(s(bool,'const.bool.IN_2'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X8),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9))))))))) => (p(s(bool,'const.bool.IN_2'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9)))))))))))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.infer_eComplete.infer_eComplete_29_prove_ltb')).
fof(f574,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))))) | $spl210),
  inference(cnf_transformation,[],[f574_D])).
fof(f574_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f569,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15))))) | ($spl6 | $spl16)),
  inference(resolution,[],[f456,f444])).
fof(f464,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl24),
  inference(cnf_transformation,[],[f464_D])).
fof(f464_D,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f282,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ((p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0)))))) & (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) <=> s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) => (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) <=> s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) => (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) <=> s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X9,X22] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22)))) => (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X9),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22)))))) <=> s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X9)))))))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.unify.t_walkstar_tuvar_props')).
fof(f9309,plain,(
  $false | ($spl10 | $spl24 | $spl210 | $spl1684 | $spl3390)),
  inference(subsumption_resolution,[],[f9308,f448])).
fof(f9308,plain,(
  ~p(s(bool,'T_0')) | ($spl10 | $spl24 | $spl210 | $spl1684 | $spl3390)),
  inference(forward_demodulation,[],[f9307,f4772])).
fof(f9307,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))) | ($spl10 | $spl24 | $spl210 | $spl3390)),
  inference(subsumption_resolution,[],[f9306,f448])).
fof(f9306,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))) | ($spl24 | $spl210 | $spl3390)),
  inference(forward_demodulation,[],[f9305,f574])).
fof(f9305,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))) | ($spl24 | $spl3390)),
  inference(trivial_inequality_removal,[],[f9297])).
fof(f9297,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) != s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15)))) | ($spl24 | $spl3390)),
  inference(superposition,[],[f464,f9040])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X0))) ) | $spl36),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f293,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2816,plain,(
  $false | ($spl36 | $spl882 | $spl899)),
  inference(subsumption_resolution,[],[f2289,f2234])).
fof(f2234,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | $spl882),
  inference(cnf_transformation,[],[f2234_D])).
fof(f2234_D,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) <=> ~$spl882),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl882])])).
fof(f2289,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | ($spl36 | $spl899)),
  inference(resolution,[],[f2269,f471])).
fof(f2269,plain,(
  ~sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)),'F_0') | $spl899),
  inference(cnf_transformation,[],[f2269_D])).
fof(f2269_D,plain,(
  ~sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)),'F_0') <=> ~$spl899),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl899])])).
fof(f474,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP0(X0,X1) | p(s(bool,X0))) ) | $spl40),
  inference(cnf_transformation,[],[f474_D])).
fof(f474_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP0(X0,X1) | p(s(bool,X0))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f473,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~sP0(X0,X1) | p(s(bool,X1))) )),
  inference(global_subsumption,[],[f291,f241,f238])).
fof(f241,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f291,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f179])).
fof(f627,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP0(X0,'F_0')) ) | $spl218),
  inference(cnf_transformation,[],[f627_D])).
fof(f627_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP0(X0,'F_0')) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f620,plain,(
  ( ! [X0] : (~sP0(X0,'F_0') | p(s(bool,X0))) ) | ($spl13 | $spl40)),
  inference(resolution,[],[f474,f451])).
fof(f645,plain,(
  ( ! [X1] : (~sP0(X1,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl230),
  inference(cnf_transformation,[],[f645_D])).
fof(f645_D,plain,(
  ( ! [X1] : (~sP0(X1,'F_0') | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f637,plain,(
  ( ! [X1] : (~sP0(X1,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | ($spl16 | $spl218)),
  inference(resolution,[],[f627,f456])).
fof(f2871,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) | $spl1252),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) <=> ~$spl1252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1252])])).
fof(f2864,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) | ($spl230 | $spl898)),
  inference(resolution,[],[f2268,f645])).
fof(f2268,plain,(
  sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)),'F_0') | $spl898),
  inference(cnf_transformation,[],[f2268_D])).
fof(f2268_D,plain,(
  sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)),'F_0') <=> ~$spl898),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl898])])).
fof(f467,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3)))))) ) | $spl30),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f289,plain,(
  ( ! [X0,X3,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t',X3))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X3))))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f2138,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t',X0)))))) ) | $spl818),
  inference(cnf_transformation,[],[f2138_D])).
fof(f2138_D,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t',X0)))))) ) <=> ~$spl818),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl818])])).
fof(f2137,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t',X0)))))) ) | ($spl10 | $spl30 | $spl208)),
  inference(subsumption_resolution,[],[f2118,f448])).
fof(f2118,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t',X0)))))) ) | ($spl30 | $spl208)),
  inference(superposition,[],[f467,f572])).
fof(f442,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) | $spl4),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f236,plain,(
  p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))))),
  inference(cnf_transformation,[],[f153])).
fof(f576,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | $spl212),
  inference(cnf_transformation,[],[f576_D])).
fof(f576_D,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f570,plain,(
  s(bool,'T_0') = s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))))) | ($spl4 | $spl16)),
  inference(resolution,[],[f456,f442])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))) ) | $spl200),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X1)))) | s('type.infer_t.infer_t',X1) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X1))))),
  inference(rectify,[],[f132])).
fof(f132,plain,(
  ! [X0,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2)))) | s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X2,X3] : ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) | ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2))))) | s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X2))))) => s('type.infer_t.infer_t',X2) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X3),s('type.infer_t.infer_t',X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X21,X10,X22] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22)))) & p(s(bool,'const.infer.check_t_3'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t',X10))))) => s('type.infer_t.infer_t',X10) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X22),s('type.infer_t.infer_t',X10))))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.inferProps.t_walkstar_no_vars')).
fof(f2061,plain,(
  ( ! [X0] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) ) | $spl760),
  inference(cnf_transformation,[],[f2061_D])).
fof(f2061_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) ) <=> ~$spl760),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl760])])).
fof(f2060,plain,(
  ( ! [X0] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) ) | ($spl10 | $spl200 | $spl212)),
  inference(subsumption_resolution,[],[f2023,f448])).
fof(f2023,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) ) | ($spl200 | $spl212)),
  inference(superposition,[],[f558,f576])).
fof(f9040,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | $spl3390),
  inference(cnf_transformation,[],[f9040_D])).
fof(f9040_D,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) <=> ~$spl3390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3390])])).
fof(f9038,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | ($spl10 | $spl760 | $spl818 | $spl884 | $spl1252)),
  inference(forward_demodulation,[],[f9037,f2237])).
fof(f456,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f446,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | $spl8),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f234,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))),
  inference(cnf_transformation,[],[f153])).
fof(f572,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) | $spl208),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f568,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK15),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))) | ($spl8 | $spl16)),
  inference(resolution,[],[f456,f446])).
fof(f2797,plain,(
  $false | ($spl10 | $spl208 | $spl888)),
  inference(subsumption_resolution,[],[f2774,f448])).
fof(f2774,plain,(
  ~p(s(bool,'T_0')) | ($spl208 | $spl888)),
  inference(superposition,[],[f2262,f572])).
fof(f468,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl32),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f288,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f2262,plain,(
  ( ! [X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) ) | $spl888),
  inference(cnf_transformation,[],[f2262_D])).
fof(f2262_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) ) <=> ~$spl888),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl888])])).
fof(f2240,plain,(
  ( ! [X1] : (~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) ) | ($spl32 | $spl883)),
  inference(resolution,[],[f2235,f468])).
fof(f465,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl26),
  inference(cnf_transformation,[],[f465_D])).
fof(f465_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f281,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f451,plain,(
  ~p(s(bool,'F_0')) | $spl13),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f449,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f241,f238])).
fof(f2235,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | $spl883),
  inference(cnf_transformation,[],[f2235_D])).
fof(f2235_D,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) <=> ~$spl883),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl883])])).
fof(f2237,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | $spl884),
  inference(cnf_transformation,[],[f2237_D])).
fof(f2237_D,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) <=> ~$spl884),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl884])])).
fof(f2233,plain,(
  s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | ($spl13 | $spl26 | $spl214)),
  inference(subsumption_resolution,[],[f2220,f451])).
fof(f2220,plain,(
  p(s(bool,'F_0')) | s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))) | ($spl26 | $spl214)),
  inference(superposition,[],[f465,f595])).
fof(f458,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f257,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.infer_eComplete.infer_eComplete_29_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f595,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) | $spl214),
  inference(cnf_transformation,[],[f595_D])).
fof(f595_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f590,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))) | ($spl1 | $spl20)),
  inference(resolution,[],[f458,f438])).
fof(f438,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))))) | $spl1),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f440,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) | $spl3),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f237,plain,(
  ~p(s(bool,'const.infer.check_t_3'(s('type.num.num',sK13),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',sK16)))))))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK16),s(fun('type.num.num',bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK14))))))),
  inference(cnf_transformation,[],[f153])).
fof(f9508,plain,(
  $false | ($spl3 | $spl10 | $spl3570)),
  inference(subsumption_resolution,[],[f9506,f448])).
fof(f9506,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl3570)),
  inference(backward_demodulation,[],[f9429,f440])).
% SZS output end Proof for thm.infer_eComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 439
% Passive clauses: 1709
% Generated clauses: 9317
% Final active clauses: 356
% Final passive clauses: 1186
% Input formulas: 40
% Initial clauses: 182
% 
% Duplicate literals: 212
% Trivial inequalities: 61
% Fw subsumption resolutions: 984
% Fw demodulations: 459
% Bw demodulations: 98
% Condensations: 232
% Global subsumptions: 14
% 
% Simple tautologies: 254
% Equational tautologies: 29
% Forward subsumptions: 2993
% 
% Binary resolution: 2835
% Factoring: 91
% Forward superposition: 958
% Backward superposition: 1149
% Self superposition: 29
% Equality factoring: 28
% Equality resolution: 14
% 
% Split clauses: 234
% Split components: 485
% SAT solver clauses: 2459
% SAT solver unit clauses: 193
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 12
% 
% Sat splits: 234
% Sat splitting refutations: 27
% 
% TWLsolver clauses: 2402
% TWLsolver calls for satisfiability: 1860
% 
% Memory used [KB]: 8315
% Time elapsed: 0.673 s
% ------------------------------
% ------------------------------
% Success in time 1.34 s
