Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ImplicitQueue
% (57613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 3832
% Generated clauses: 19721
% Final active clauses: 147
% Final passive clauses: 3478
% Input formulas: 68
% Initial clauses: 102
% 
% Duplicate literals: 559
% Trivial inequalities: 338
% Fw subsumption resolutions: 1546
% Fw demodulations: 2806
% Bw demodulations: 19
% 
% Simple tautologies: 275
% Equational tautologies: 44
% Forward subsumptions: 9809
% 
% Binary resolution: 543
% Factoring: 1
% Forward superposition: 7044
% Backward superposition: 6702
% Self superposition: 171
% Equality factoring: 70
% Equality resolution: 31
% 
% Split clauses: 261
% Split components: 524
% SAT solver clauses: 700
% SAT solver unit clauses: 9
% SAT solver binary clauses: 151
% 
% Sat splits: 261
% 
% TWLsolver clauses: 700
% TWLsolver calls for satisfiability: 190
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ImplicitQueue
% (57651)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 1385
% Generated clauses: 14515
% Final active clauses: 112
% Final passive clauses: 1032
% Input formulas: 68
% Initial clauses: 102
% 
% Duplicate literals: 463
% Trivial inequalities: 158
% Fw subsumption resolutions: 822
% Fw demodulations: 1503
% Bw demodulations: 215
% Condensations: 1123
% Global subsumptions: 8
% 
% Simple tautologies: 178
% Equational tautologies: 44
% Forward subsumptions: 6744
% Bw demodulations to eq. taut.: 23
% 
% Binary resolution: 282
% Factoring: 1
% Forward superposition: 4533
% Backward superposition: 3654
% Self superposition: 147
% Equality factoring: 78
% Equality resolution: 20
% 
% Split clauses: 86
% Split components: 172
% SAT solver clauses: 2178
% SAT solver unit clauses: 148
% SAT solver binary clauses: 53
% SAT solver learnt clauses: 16
% 
% Sat splits: 86
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2146
% TWLsolver calls for satisfiability: 395
% 
% Memory used [KB]: 16630
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ImplicitQueue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ImplicitQueue
% SZS output start Proof for thm.ImplicitQueue
fof(f44941,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1265,f1266,f1258,f1270,f1804,f1807,f1805,f3364,f3370,f3858,f3841,f3861,f3996,f3997,f4000,f4002,f4003,f4005,f4008,f4010,f4012,f4013,f4014,f4016,f4018,f4020,f4022,f4024,f4026,f4028,f4030,f4033,f4035,f4037,f4039,f4040,f4041,f4053,f4055,f4058,f4068,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4083,f4084,f4086,f4089,f4091,f4093,f4095,f4098,f4101,f4102,f4103,f4105,f4107,f4109,f4111,f4113,f4114,f7707,f7709,f7692,f7718,f7711,f8526,f8598,f8672,f8674,f8675,f8676,f8730,f8733,f8731,f9136,f9137,f9146,f9151,f9156,f9157,f9170,f9172,f9175,f9178,f9180,f9200,f9206,f9208,f9210,f9212,f9214,f9216,f9218,f9220,f9221,f9222,f9223,f9224,f9225,f9227,f9229,f11259,f11261,f11263,f11267,f11269,f11273,f11275,f11316,f11320,f11733,f11734,f11737,f11741,f11745,f11746,f11747,f11748,f11750,f11752,f11754,f11755,f11756,f11757,f11758,f11759,f11761,f11763,f11764,f11765,f11766,f11767,f11768,f11769,f11770,f11771,f22120,f22204,f15025,f14122,f4535,f4552,f3266,f3371,f3283,f1268,f3859,f23875,f25355,f23961,f35485,f35486,f35487,f35488,f35556,f35489,f35553,f35557,f35491,f35492,f35560,f35493,f35494,f35561,f35562,f35563,f35564,f35555,f35559,f35565,f35552,f35550,f35577,f40754,f40756,f40755,f40890,f40892,f40894,f40898,f40900,f40902,f40904,f40906,f40908,f40910,f40912,f40914,f40916,f40918,f40920,f40922,f40924,f40928,f40930,f40932,f40934,f40936,f40938,f40940,f40942,f40944,f40946,f40948,f40950,f40952,f40954,f40959,f41048,f41050,f41052,f41054,f41056,f41058,f41060,f41076,f41083,f41089,f41093,f41095,f41194,f41195,f41196,f41197,f41198,f41199,f41200,f41201,f41202,f41203,f41204,f41205,f41206,f41207,f41208,f41209,f41210,f41211,f41212,f41213,f41214,f41215,f41216,f41217,f41218,f41219,f41220,f41221,f41222,f41223,f41224,f41225,f41226,f41227,f41228,f41229,f41672,f42000,f42002,f42001,f1806,f43374,f41770,f8597,f44473,f44475,f44476,f44478,f44479,f44481,f44482,f44484,f44486,f44485,f44487,f44488,f44490,f44491,f44494,f44495,f44497,f44498,f44500,f44501,f44503,f44504,f44506,f44509,f44510,f44511,f44512,f44513,f44514,f44474,f44516,f44517,f44519,f44520,f44521,f44483,f44522,f44527,f44528,f44531,f44532,f44477,f44536,f44537,f44540,f44541,f44543,f44545,f44544,f44547,f44548,f44480,f44552,f44492,f44507,f44564,f44566,f44567,f44568,f44569,f44570,f44571,f44572,f44573,f44574,f44565,f44575,f44577,f44578,f44579,f44580,f44581,f44582,f44583,f44584,f44585,f44576,f44592,f44593,f44594,f44597,f44599,f44601,f44605,f44607,f44609,f44611,f44613,f44615,f44617,f44619,f44621,f44623,f44625,f44627,f44631,f44633,f44635,f44637,f44639,f44641,f44643,f44645,f44647,f44649,f44651,f44653,f44658,f44749,f44751,f44753,f44755,f44757,f44759,f44773,f44780,f44786,f44790,f44792,f44794,f44894,f44895,f44896,f44897,f44898,f44899,f44900,f44901,f44902,f44903,f44904,f44905,f44906,f44907,f44908,f44909,f44910,f44911,f44912,f44913,f44914,f44915,f44916,f44917,f44918,f44919,f44920,f44921,f44922,f44923,f44924,f44939,f44940])).
fof(f44940,plain,(
  $false | ($spl62 | $spl84 | $spl104)),
  inference(subsumption_resolution,[],[f44537,f44938])).
fof(f44938,plain,(
  ( ! [X386] : (p(s(bool,X386))) ) | ($spl62 | $spl84)),
  inference(subsumption_resolution,[],[f44937,f230])).
fof(f230,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ImplicitQueue.digit_distinct|split|2',aHL_TRUTH)).
fof(f44937,plain,(
  ( ! [X386] : (~p(s(bool,'T_0')) | p(s(bool,X386))) ) | ($spl62 | $spl84)),
  inference(forward_demodulation,[],[f44504,f44478])).
fof(f44939,plain,(
  $false | ($spl62 | $spl84 | $spl106)),
  inference(subsumption_resolution,[],[f44541,f44938])).
fof(f44924,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44600,f44779])).
fof(f44779,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X44,X46,X48] : (p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X44,fun(X44,bool)),'=_0'),s(X44,X46))),s(X44,X48)))) <=> s(X44,X46) = s(X44,X48))),
  file('thm.ImplicitQueue.digit_distinct|split|2',arity_0_HL_EQ)).
fof(f44595,plain,(
  ( ! [X386] : (p(s(bool,X386))) ) | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44504,f44492])).
fof(f44600,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f244])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f44923,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44604,f230])).
fof(f44604,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f44595,f295])).
fof(f295,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f341,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f334,f240])).
fof(f240,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f334,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f233,f230])).
fof(f233,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f47])).
fof(f47,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f44922,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44606,f230])).
fof(f44606,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12123,f44595,f295])).
fof(f12123,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl21),
  inference(unit_resulting_resolution,[],[f7710,f240])).
fof(f7710,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl21),
  inference(cnf_transformation,[],[f7710_D])).
fof(f7710_D,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f44921,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44608,f230])).
fof(f44608,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f30168,f44595,f295])).
fof(f30168,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f30047,f334])).
fof(f30047,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(superposition,[],[f240,f29032])).
fof(f29032,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f312,f28871,f1697])).
fof(f1697,plain,(
  ( ! [X43,X46] : (s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1696,f230])).
fof(f1696,plain,(
  ( ! [X43,X46] : (~p(s(bool,'T_0')) | s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(forward_demodulation,[],[f1610,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X20] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X6,X20)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.DATATYPE_TAG_THM')).
fof(f1610,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X44,X45)))) | s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f359,f1236])).
fof(f1236,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1223,f230])).
fof(f1223,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f1209,f265])).
fof(f1209,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1208,f266])).
fof(f266,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X14] : ! [X16] : s(X14,X16) = s(X14,happ(s(fun(X14,X14),'_i_0_0'),s(X14,X16)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_i0_ax')).
fof(f1208,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1179,f230])).
fof(f1179,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f295,f389])).
fof(f389,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f266,f237])).
fof(f237,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ImplicitQueue.digit_distinct|split|2',aHL_BOOL_CASES)).
fof(f359,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f244,f237])).
fof(f28871,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f28768,f7709])).
fof(f28768,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK9('_i_0_0',bool)))) | ($spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f11918,f28739])).
fof(f28739,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f28607,f334])).
fof(f28607,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f240,f22267])).
fof(f22267,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f18853,f7709])).
fof(f18853,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,sK9('_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f12801,f11917])).
fof(f11917,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f11871,f341])).
fof(f11871,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f7709,f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f237,f237])).
fof(f12801,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f12764,f7709])).
fof(f12764,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK9('_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl26),
  inference(superposition,[],[f8598,f266])).
fof(f11918,plain,(
  s(bool,'T_0') != s(bool,sK9('_i_0_0',bool)) | $spl18),
  inference(subsumption_resolution,[],[f11877,f334])).
fof(f11877,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK9('_i_0_0',bool)) | $spl18),
  inference(superposition,[],[f240,f7709])).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X12,X14] : ! [X16,X17] : s(X12,happ(s(fun(X14,X12),X17),s(X14,X16))) = s(X12,happ(s(fun(fun(X14,X12),X12),'_dstsel_0_x0_1'(s(X14,X16))),s(fun(X14,X12),X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dstsel_0_x0_ax')).
fof(f44920,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44610,f230])).
fof(f44610,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28871,f44595,f295])).
fof(f44919,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44612,f230])).
fof(f44612,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24579,f44595,f295])).
fof(f24579,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK7(X0,'_i_0_0',bool))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1270,f240])).
fof(f44918,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44614,f230])).
fof(f44614,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11918,f44595,f295])).
fof(f44917,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44618,f44595])).
fof(f44618,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12123,f44595,f295])).
fof(f44916,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44620,f44595])).
fof(f44620,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f30168,f44595,f295])).
fof(f44915,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44622,f44595])).
fof(f44622,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28871,f44595,f295])).
fof(f44914,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44624,f44595])).
fof(f44624,plain,(
  ( ! [X0] : (~p(s(bool,sK7(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24579,f44595,f295])).
fof(f44913,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44626,f44595])).
fof(f44626,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11918,f44595,f295])).
fof(f44912,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44630,f230])).
fof(f44630,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f44595,f295])).
fof(f44911,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44632,f230])).
fof(f44632,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12123,f44595,f295])).
fof(f44910,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44634,f230])).
fof(f44634,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f30168,f44595,f295])).
fof(f44909,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44636,f230])).
fof(f44636,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28871,f44595,f295])).
fof(f44908,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44638,f230])).
fof(f44638,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24579,f44595,f295])).
fof(f44907,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44640,f230])).
fof(f44640,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11918,f44595,f295])).
fof(f44906,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44644,f44595])).
fof(f44644,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12123,f44595,f295])).
fof(f44905,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44646,f44595])).
fof(f44646,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f30168,f44595,f295])).
fof(f44904,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44648,f44595])).
fof(f44648,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28871,f44595,f295])).
fof(f44903,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44650,f44595])).
fof(f44650,plain,(
  ( ! [X0] : (~p(s(bool,sK7(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24579,f44595,f295])).
fof(f44902,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44652,f44595])).
fof(f44652,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11918,f44595,f295])).
fof(f44901,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44656,f44595])).
fof(f44656,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0'))))) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X46,X48] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X46))),s(bool,X48)))) <=> (p(s(bool,X46)) => p(s(bool,X48))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','==>_HL_0')).
fof(f44900,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44748,f44779])).
fof(f44748,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f44595,f359])).
fof(f44899,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44750,f44779])).
fof(f44750,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) ) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12123,f44595,f359])).
fof(f44898,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44752,f44779])).
fof(f44752,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) ) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f30168,f44595,f359])).
fof(f44897,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44754,f44779])).
fof(f44754,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28871,f44595,f359])).
fof(f44896,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44756,f44779])).
fof(f44756,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK7(X1,'_i_0_0',bool))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24579,f44595,f359])).
fof(f44895,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44758,f44779])).
fof(f44758,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK9('_i_0_0',bool))) ) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11918,f44595,f359])).
fof(f44894,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f44789,f44595])).
fof(f44789,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595,f303])).
fof(f44794,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595,f8241])).
fof(f8241,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'F_0')))) | p(s(bool,X0))) )),
  inference(condensation,[],[f8240])).
fof(f8240,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f8119,f240])).
fof(f8119,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f300,f237])).
fof(f300,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X46,X48] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X46))),s(bool,X48)))) <=> (p(s(bool,X46)) | p(s(bool,X48))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','|bslash||slash|_HL_0')).
fof(f44792,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595,f8175])).
fof(f8175,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X0)))) | p(s(bool,X0))) )),
  inference(condensation,[],[f8174])).
fof(f8174,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f8077,f240])).
fof(f8077,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f300,f237])).
fof(f44790,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f334,f44595,f303])).
fof(f44786,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f334,f44595,f300])).
fof(f44780,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f44595,f317])).
fof(f44773,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f341,f44595,f1209])).
fof(f44759,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f7709,f11918,f44595,f359])).
fof(f44757,plain,(
  $false | ($spl4 | $spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f24610,f24579,f44595,f359])).
fof(f24610,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK7(X0,'_i_0_0',bool))) ) | ($spl4 | $spl18)),
  inference(forward_demodulation,[],[f24593,f7709])).
fof(f24593,plain,(
  ( ! [X0] : (s(bool,sK9('_i_0_0',bool)) = s(bool,sK7(X0,'_i_0_0',bool))) ) | ($spl4 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1270,f415])).
fof(f415,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f409,f334])).
fof(f409,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f292,f237])).
fof(f292,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f44755,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f266,f28871,f44595,f359])).
fof(f44753,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f29032,f30168,f44595,f359])).
fof(f44751,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12140,f12123,f44595,f359])).
fof(f12140,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl21)),
  inference(forward_demodulation,[],[f12130,f7709])).
fof(f12130,plain,(
  s(bool,sK9('_i_0_0',bool)) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f11918,f7710,f415])).
fof(f44749,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12140,f341,f44595,f359])).
fof(f44658,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595,f44595,f303])).
fof(f44653,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f11918,f44595,f295])).
fof(f44651,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f24579,f44595,f295])).
fof(f44649,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f28871,f44595,f295])).
fof(f44647,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f30168,f44595,f295])).
fof(f44645,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f12123,f44595,f295])).
fof(f44643,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f341,f44595,f295])).
fof(f44641,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f11918,f44595,f295])).
fof(f44639,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f24579,f44595,f295])).
fof(f44637,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f28871,f44595,f295])).
fof(f44635,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f30168,f44595,f295])).
fof(f44633,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f12123,f44595,f295])).
fof(f44631,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f341,f44595,f295])).
fof(f44627,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f11918,f44595,f295])).
fof(f44625,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f24579,f44595,f295])).
fof(f44623,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f28871,f44595,f295])).
fof(f44621,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f30168,f44595,f295])).
fof(f44619,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f12123,f44595,f295])).
fof(f44617,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f44595,f341,f44595,f295])).
fof(f44615,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f11918,f44595,f295])).
fof(f44613,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f24579,f44595,f295])).
fof(f44611,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f28871,f44595,f295])).
fof(f44609,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f30168,f44595,f295])).
fof(f44607,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f12123,f44595,f295])).
fof(f44605,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f230,f341,f44595,f295])).
fof(f44601,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12140,f44595,f244])).
fof(f44599,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f341,f44595,f241])).
fof(f241,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f44597,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f334,f44595])).
fof(f44594,plain,(
  $false | ($spl56 | $spl60)),
  inference(subsumption_resolution,[],[f44474,f44477])).
fof(f44593,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f44591,f334])).
fof(f44591,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(backward_demodulation,[],[f44483,f44315])).
fof(f44315,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.ImplicitQueue.times'(sK0),X0),X2),s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s('type.ImplicitQueue.times'(sK0),sK3))))) )),
  inference(forward_demodulation,[],[f44314,f27966])).
fof(f27966,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X16,happ(s(fun(X20,X16),'_dst_x0x1x0_3'(s(fun(X18,fun(X17,X16)),X19),s(fun(X20,X18),X21),s(X17,happ(s(fun(X23,X17),X24),s(X23,X25))))),s(X20,X22))) = s(X16,happ(s(fun(X23,X16),'_dst_x0x0x1_3'(s(fun(X18,fun(X17,X16)),X19),s(X18,happ(s(fun(X20,X18),X21),s(X20,X22))),s(fun(X23,X17),X24))),s(X23,X25)))) )),
  inference(superposition,[],[f324,f323])).
fof(f323,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X12,X13,X41,X14] : ! [X15,X16,X42,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x0x1_3'(s(fun(X13,fun(X41,X12)),X15),s(X13,X16),s(fun(X14,X41),X42))),s(X14,X17))) = s(X12,happ(s(fun(X41,X12),happ(s(fun(X13,fun(X41,X12)),X15),s(X13,X16))),s(X41,happ(s(fun(X14,X41),X42),s(X14,X17)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x0x0x1_ax')).
fof(f324,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X12,X13,X14,X41] : ! [X15,X16,X42,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x1x0_3'(s(fun(X13,fun(X41,X12)),X15),s(fun(X14,X13),X16),s(X41,X42))),s(X14,X17))) = s(X12,happ(s(fun(X41,X12),happ(s(fun(X13,fun(X41,X12)),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))),s(X41,X42)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x0x1x0_ax')).
fof(f44314,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),X2),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44313,f315])).
fof(f315,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),X4),s('type.ImplicitQueue.times'(X1),X2))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X1),X2))),s(X0,X3),s(fun('type.ImplicitQueue.times'(X1),X0),X4),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X5)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),X4),s('type.ImplicitQueue.times'(X1),X2))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X1),X2))),s(X0,X3),s(fun('type.ImplicitQueue.times'(X1),X0),X4),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X5)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),X4),s('type.ImplicitQueue.times'(X1),X2))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X1),X2))),s(X0,X3),s(fun('type.ImplicitQueue.times'(X1),X0),X4),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X5)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X24,X6] : ! [X9,X35,X26,X36] : s(X24,happ(s(fun('type.ImplicitQueue.times'(X6),X24),X26),s('type.ImplicitQueue.times'(X6),X9))) = s(X24,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9))),s(X24,X35),s(fun('type.ImplicitQueue.times'(X6),X24),X26),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X24)),X36)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_case_def|split|1')).
fof(f44313,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(X0,X1),s(fun('type.ImplicitQueue.times'(sK0),X0),X2),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3))),s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44312,f26377])).
fof(f26377,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X13,X8),'_dst_x0x0x1_3'(s(fun(X10,fun(X9,X8)),X11),s(X10,X12),s(fun(X13,X9),X14))),s(X13,X15))) = s(X8,happ(s(fun(X10,X8),'_dst_x1x0_2'(s(fun(X10,fun(X9,X8)),X11),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15))))),s(X10,X12)))) )),
  inference(superposition,[],[f323,f321])).
fof(f321,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X12,X14,X13] : ! [X15,X16,X17] : s(X12,happ(s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),X15),s(X14,X17))),s(X13,X16))) = s(X12,happ(s(fun(X14,X12),'_dst_x1x0_2'(s(fun(X14,fun(X13,X12)),X15),s(X13,X16))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x1x0_ax')).
fof(f44312,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(X0,X1),s(fun('type.ImplicitQueue.times'(sK0),X0),X2),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3))))))) )),
  inference(forward_demodulation,[],[f43661,f321])).
fof(f43661,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(X0,X1),s(fun('type.ImplicitQueue.times'(sK0),X0),X2),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3))))),s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X3),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))))) )),
  inference(unit_resulting_resolution,[],[f29483,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f29483,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.times'(sK0),X0),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X1),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(X0,X2),s(fun('type.ImplicitQueue.times'(sK0),X0),X3),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),X0)),X1)))) )),
  inference(superposition,[],[f316,f229])).
fof(f229,plain,(
  s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) = s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),sK2),s('type.ImplicitQueue.times'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) = s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),sK2),s('type.ImplicitQueue.times'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f182])).
fof(f182,plain,(
  ? [X0,X1,X2,X3] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X3))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X3))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X3))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9))) = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9))) = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_distinct|split|2')).
fof(f316,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.ImplicitQueue.times'(X1),X0),happ(s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6),s('type.ImplicitQueue.times'(X1),X2))),s('type.ImplicitQueue.times'(X1),X3))) = s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X1),X2),s('type.ImplicitQueue.times'(X1),X3))),s(X0,X4),s(fun('type.ImplicitQueue.times'(X1),X0),X5),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X24,X6] : ! [X8,X7,X35,X26,X36] : s(X24,happ(s(fun('type.ImplicitQueue.times'(X6),X24),happ(s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X24)),X36),s('type.ImplicitQueue.times'(X6),X8))),s('type.ImplicitQueue.times'(X6),X7))) = s(X24,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7))),s(X24,X35),s(fun('type.ImplicitQueue.times'(X6),X24),X26),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X24)),X36)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_case_def|split|2')).
fof(f44592,plain,(
  $false | ($spl68 | $spl76)),
  inference(subsumption_resolution,[],[f44590,f334])).
fof(f44590,plain,(
  p(s(bool,'F_0')) | ($spl68 | $spl76)),
  inference(backward_demodulation,[],[f44483,f44492])).
fof(f44576,plain,(
  ( ! [X1084,X1083] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1084),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1083),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) ) | $spl126),
  inference(cnf_transformation,[],[f44576_D])).
fof(f44576_D,plain,(
  ( ! [X1084,X1083] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1084),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1083),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f44585,plain,(
  ( ! [X1128,X1124,X1127] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1127,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1127,X1128),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1124),s('type.ImplicitQueue.times'(sK0),sK3)))))))) ) | $spl136),
  inference(cnf_transformation,[],[f44585_D])).
fof(f44585_D,plain,(
  ( ! [X1128,X1124,X1127] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1127,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1127,X1128),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1124),s('type.ImplicitQueue.times'(sK0),sK3)))))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f44584,plain,(
  ( ! [X1128,X1126,X1124,X1127,X1125] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1127,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1127,X1128),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1124),s('type.ImplicitQueue.times'(sK0),sK3))))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1126),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1125),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f44134,f315])).
fof(f44134,plain,(
  ( ! [X1128,X1126,X1124,X1123,X1127,X1125] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1126),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1125),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1127,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1127,X1128),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1123),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1124),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1125)))))))) )),
  inference(superposition,[],[f19313,f29483])).
fof(f19313,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X27),s('type.ImplicitQueue.times'(sK0),X26))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X24,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X24,X25),s('type.ImplicitQueue.times'(sK0),X26)))))) )),
  inference(superposition,[],[f19272,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X24] : ! [X20,X25] : s(X24,X25) = s(X24,'const.pair.SND_1'(s('type.pair.prod'(X6,X24),'const.pair.,_2'(s(X6,X20),s(X24,X25)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.pair.SND')).
fof(f19272,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X0),s('type.ImplicitQueue.times'(sK0),X1))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),X1)) )),
  inference(superposition,[],[f289,f229])).
fof(f289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) | s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),X4)) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) | (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X3) & s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),X4))) & (s('type.ImplicitQueue.times'(X0),X1) != s('type.ImplicitQueue.times'(X0),X3) | s('type.ImplicitQueue.times'(X0),X2) != s('type.ImplicitQueue.times'(X0),X4) | s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4)))))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) | (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X3) & s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),X4))) & ((s('type.ImplicitQueue.times'(X0),X1) != s('type.ImplicitQueue.times'(X0),X3) | s('type.ImplicitQueue.times'(X0),X2) != s('type.ImplicitQueue.times'(X0),X4)) | s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) <=> (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X3) & s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),X4)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) <=> (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X3) & s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),X4)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X8,X7,X32,X33] : (s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7))) = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X32),s('type.ImplicitQueue.times'(X6),X33))) <=> (s('type.ImplicitQueue.times'(X6),X8) = s('type.ImplicitQueue.times'(X6),X32) & s('type.ImplicitQueue.times'(X6),X7) = s('type.ImplicitQueue.times'(X6),X33)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_|49||49||split|1')).
fof(f44583,plain,(
  ( ! [X1117,X1114] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1114),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1117),s('type.num.num','const.num.|48|_0')))) ) | $spl134),
  inference(cnf_transformation,[],[f44583_D])).
fof(f44583_D,plain,(
  ( ! [X1117,X1114] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1114),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1117),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f44582,plain,(
  ( ! [X1116,X1115,X1117,X1114] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1114),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1117),s('type.num.num','const.num.|48|_0'))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1116),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1115),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f44132,f315])).
fof(f44132,plain,(
  ( ! [X1116,X1115,X1113,X1117,X1114] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1116),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1115),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1113),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1114),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1115))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1117),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f19311,f29483])).
fof(f19311,plain,(
  ( ! [X19,X17,X18] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X19),s('type.ImplicitQueue.times'(sK0),X17))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),X17),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X18),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f19272,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X9,X26] : s(X6,X9) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X26),s('type.num.num','const.num.|48|_0')))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ind_type.FCONS|split|0')).
fof(f44581,plain,(
  ( ! [X1106,X1105,X1102] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1105),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1102),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1105,X1106)))))) ) | $spl132),
  inference(cnf_transformation,[],[f44581_D])).
fof(f44581_D,plain,(
  ( ! [X1106,X1105,X1102] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1105),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1102),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1105,X1106)))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f44580,plain,(
  ( ! [X1106,X1104,X1103,X1105,X1102] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1105),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1102),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1105,X1106))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1104),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1103),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f44130,f315])).
fof(f44130,plain,(
  ( ! [X1106,X1104,X1103,X1101,X1105,X1102] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1104),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1103),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1105),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1101),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1102),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1103))),s(X1105,X1106)))))) )),
  inference(superposition,[],[f19310,f29483])).
fof(f19310,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X16),s('type.ImplicitQueue.times'(sK0),X14))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X13),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f19272,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X24,X6] : ! [X20,X25] : s(X6,X20) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X24),'const.pair.,_2'(s(X6,X20),s(X24,X25)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.pair.FST')).
fof(f44579,plain,(
  ( ! [X1093,X1090,X1094] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1093,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1090),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1093,X1094)))) ) | $spl130),
  inference(cnf_transformation,[],[f44579_D])).
fof(f44579_D,plain,(
  ( ! [X1093,X1090,X1094] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1093,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1090),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1093,X1094)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f44578,plain,(
  ( ! [X1091,X1093,X1090,X1094,X1092] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1093,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1090),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1093,X1094))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1092),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1091),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f44128,f315])).
fof(f44128,plain,(
  ( ! [X1091,X1089,X1093,X1090,X1094,X1092] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1092),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1091),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1093,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1089),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1090),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1091))))),s(X1093,X1094)))) )),
  inference(superposition,[],[f19307,f29483])).
fof(f19307,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X5),s('type.ImplicitQueue.times'(sK0),X3))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X2,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f19272,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X14,X12] : ! [X15,X17] : s(X12,X15) = s(X12,happ(s(fun(X14,X12),'_dst_x0_1'(s(X12,X15))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x0_ax')).
fof(f44577,plain,(
  ( ! [X1082] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1082),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) | $spl128),
  inference(cnf_transformation,[],[f44577_D])).
fof(f44577_D,plain,(
  ( ! [X1082] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1082),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f44575,plain,(
  ( ! [X1082,X1084,X1083] : (s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1082),s('type.ImplicitQueue.times'(sK0),sK3))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1084),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1083),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f44126,f315])).
fof(f44126,plain,(
  ( ! [X1081,X1082,X1084,X1083] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1084),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1083),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1081),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1082),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1083)))))) )),
  inference(superposition,[],[f19306,f29483])).
fof(f19306,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X1),s('type.ImplicitQueue.times'(sK0),X0))) | s('type.ImplicitQueue.times'(sK0),sK1) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),X0)))) )),
  inference(superposition,[],[f19272,f266])).
fof(f44565,plain,(
  ( ! [X1021,X1022] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1021),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1022)))) ) | $spl114),
  inference(cnf_transformation,[],[f44565_D])).
fof(f44565_D,plain,(
  ( ! [X1021,X1022] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1021),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1022)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f44574,plain,(
  ( ! [X1066,X1062,X1065] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1065,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1065,X1066),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1062),s('type.ImplicitQueue.times'(sK0),sK3)))))))) ) | $spl124),
  inference(cnf_transformation,[],[f44574_D])).
fof(f44574_D,plain,(
  ( ! [X1066,X1062,X1065] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1065,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1065,X1066),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1062),s('type.ImplicitQueue.times'(sK0),sK3)))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f44573,plain,(
  ( ! [X1066,X1064,X1062,X1065,X1063] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1065,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1065,X1066),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1062),s('type.ImplicitQueue.times'(sK0),sK3))))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1063),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1064)))) )),
  inference(forward_demodulation,[],[f44122,f315])).
fof(f44122,plain,(
  ( ! [X1066,X1064,X1062,X1065,X1063,X1061] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1063),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1064))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1065,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X1065,X1066),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1061),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1062),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1063)))))))) )),
  inference(superposition,[],[f19103,f29483])).
fof(f19103,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X26),s('type.ImplicitQueue.times'(sK0),X27))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X24,'type.ImplicitQueue.times'(sK0)),'const.pair.,_2'(s(X24,X25),s('type.ImplicitQueue.times'(sK0),X26)))))) )),
  inference(superposition,[],[f19072,f308])).
fof(f19072,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X0),s('type.ImplicitQueue.times'(sK0),X1))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),X0)) )),
  inference(superposition,[],[f288,f229])).
fof(f288,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4))) | s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X3)) )),
  inference(cnf_transformation,[],[f214])).
fof(f44572,plain,(
  ( ! [X1052,X1055] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1052),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1055),s('type.num.num','const.num.|48|_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f44572_D])).
fof(f44572_D,plain,(
  ( ! [X1052,X1055] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1052),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1055),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f44571,plain,(
  ( ! [X1054,X1052,X1055,X1053] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1052),s('type.ImplicitQueue.times'(sK0),sK3))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1055),s('type.num.num','const.num.|48|_0'))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1053),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1054)))) )),
  inference(forward_demodulation,[],[f44120,f315])).
fof(f44120,plain,(
  ( ! [X1054,X1052,X1051,X1055,X1053] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1053),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1054))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1051),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1052),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1053))),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X1055),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f19101,f29483])).
fof(f19101,plain,(
  ( ! [X19,X17,X18] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X17),s('type.ImplicitQueue.times'(sK0),X19))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(sK0),X17),s(fun('type.num.num','type.ImplicitQueue.times'(sK0)),X18),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f19072,f276])).
fof(f44570,plain,(
  ( ! [X1040,X1044,X1043] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1043),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1040),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1043,X1044)))))) ) | $spl120),
  inference(cnf_transformation,[],[f44570_D])).
fof(f44570_D,plain,(
  ( ! [X1040,X1044,X1043] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1043),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1040),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1043,X1044)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f44569,plain,(
  ( ! [X1042,X1040,X1044,X1043,X1041] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1043),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1040),s('type.ImplicitQueue.times'(sK0),sK3))),s(X1043,X1044))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1041),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1042)))) )),
  inference(forward_demodulation,[],[f44118,f315])).
fof(f44118,plain,(
  ( ! [X1042,X1040,X1044,X1039,X1043,X1041] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1041),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1042))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X1043),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1039),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1040),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1041))),s(X1043,X1044)))))) )),
  inference(superposition,[],[f19100,f29483])).
fof(f19100,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X14),s('type.ImplicitQueue.times'(sK0),X16))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(sK0),X13),'const.pair.,_2'(s('type.ImplicitQueue.times'(sK0),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f19072,f310])).
fof(f44568,plain,(
  ( ! [X1031,X1032,X1028] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1031,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1028),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1031,X1032)))) ) | $spl118),
  inference(cnf_transformation,[],[f44568_D])).
fof(f44568_D,plain,(
  ( ! [X1031,X1032,X1028] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1031,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1028),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1031,X1032)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f44567,plain,(
  ( ! [X1031,X1029,X1032,X1030,X1028] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1031,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1028),s('type.ImplicitQueue.times'(sK0),sK3))))),s(X1031,X1032))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1029),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1030)))) )),
  inference(forward_demodulation,[],[f44116,f315])).
fof(f44116,plain,(
  ( ! [X1027,X1031,X1029,X1032,X1030,X1028] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1029),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1030))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X1031,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1027),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1028),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1029))))),s(X1031,X1032)))) )),
  inference(superposition,[],[f19097,f29483])).
fof(f19097,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X3),s('type.ImplicitQueue.times'(sK0),X5))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun(X2,'type.ImplicitQueue.times'(sK0)),'_dst_x0_1'(s('type.ImplicitQueue.times'(sK0),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f19072,f309])).
fof(f44566,plain,(
  ( ! [X1020] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1020),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) | $spl116),
  inference(cnf_transformation,[],[f44566_D])).
fof(f44566_D,plain,(
  ( ! [X1020] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1020),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f44564,plain,(
  ( ! [X1021,X1022,X1020] : (s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1020),s('type.ImplicitQueue.times'(sK0),sK3))))) | s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1021),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1022)))) )),
  inference(forward_demodulation,[],[f44114,f315])).
fof(f44114,plain,(
  ( ! [X1019,X1021,X1022,X1020] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1021),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))),s('type.ImplicitQueue.times'(sK0),X1022))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s('type.ImplicitQueue.times'(sK0),X1019),s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),X1020),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0))),X1021)))))) )),
  inference(superposition,[],[f19096,f29483])).
fof(f19096,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))) != s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(sK0),X0),s('type.ImplicitQueue.times'(sK0),X1))) | s('type.ImplicitQueue.times'(sK0),sK2) = s('type.ImplicitQueue.times'(sK0),happ(s(fun('type.ImplicitQueue.times'(sK0),'type.ImplicitQueue.times'(sK0)),'_i_0_0'),s('type.ImplicitQueue.times'(sK0),X0)))) )),
  inference(superposition,[],[f19072,f266])).
fof(f44507,plain,(
  ( ! [X389,X390] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X390))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) | $spl86),
  inference(cnf_transformation,[],[f44507_D])).
fof(f44507_D,plain,(
  ( ! [X389,X390] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X390))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f44492,plain,(
  ( ! [X368] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X368),s('type.ImplicitQueue.times'(sK0),sK3))))) ) | $spl76),
  inference(cnf_transformation,[],[f44492_D])).
fof(f44492_D,plain,(
  ( ! [X368] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X368),s('type.ImplicitQueue.times'(sK0),sK3))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f44552,plain,(
  ( ! [X569,X571,X570] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X569),s('type.ImplicitQueue.times'(sK0),sK3)))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X570),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X571))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44551,f315])).
fof(f44551,plain,(
  ( ! [X569,X571,X568,X570] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X570),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X571))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X568),s(fun('type.ImplicitQueue.times'(sK0),bool),X569),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X570))))) )),
  inference(forward_demodulation,[],[f44053,f27965])).
fof(f27965,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X12,X8),'_dst_x0x1x0_3'(s(fun(X10,fun(X9,X8)),X11),s(fun(X12,X10),X13),s(X9,X15))),s(X12,X14))) = s(X8,happ(s(fun(X10,X8),'_dst_x1x0_2'(s(fun(X10,fun(X9,X8)),X11),s(X9,X15))),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))))) )),
  inference(superposition,[],[f324,f321])).
fof(f44053,plain,(
  ( ! [X569,X571,X568,X570] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X571))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X570),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X568),s(fun('type.ImplicitQueue.times'(sK0),bool),X569),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X570))))) )),
  inference(superposition,[],[f23188,f29483])).
fof(f23188,plain,(
  ( ! [X220,X219] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X220))),s(bool,X219)))) | p(s(bool,X219))) )),
  inference(superposition,[],[f304,f321])).
fof(f304,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f44480,plain,(
  ( ! [X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X351),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) | $spl64),
  inference(cnf_transformation,[],[f44480_D])).
fof(f44480_D,plain,(
  ( ! [X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X351),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f44548,plain,(
  ( ! [X557] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X557),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) | $spl112),
  inference(cnf_transformation,[],[f44548_D])).
fof(f44548_D,plain,(
  ( ! [X557] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X557),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f44547,plain,(
  ( ! [X557,X558] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X557),s('type.ImplicitQueue.times'(sK0),sK3))))) | s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X558),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f44049,f315])).
fof(f44049,plain,(
  ( ! [X557,X556,X558] : (s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X558),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X556),s(fun('type.ImplicitQueue.times'(sK0),bool),X557),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X558)))))) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f12801,f29483])).
fof(f44544,plain,(
  ( ! [X552] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X552),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) | $spl108),
  inference(cnf_transformation,[],[f44544_D])).
fof(f44544_D,plain,(
  ( ! [X552] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X552),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f44545,plain,(
  ( ! [X551] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X551),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) | $spl110),
  inference(cnf_transformation,[],[f44545_D])).
fof(f44545_D,plain,(
  ( ! [X551] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X551),s('type.ImplicitQueue.times'(sK0),sK3)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f44543,plain,(
  ( ! [X551,X552] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X551),s('type.ImplicitQueue.times'(sK0),sK3))))) | s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X552),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) ) | $spl16),
  inference(forward_demodulation,[],[f44047,f315])).
fof(f44047,plain,(
  ( ! [X551,X552,X550] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X552),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X550),s(fun('type.ImplicitQueue.times'(sK0),bool),X551),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X552)))))) ) | $spl16),
  inference(superposition,[],[f10425,f29483])).
fof(f10425,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl16),
  inference(forward_demodulation,[],[f10382,f7946])).
fof(f7946,plain,(
  s(bool,'T_0') = s(bool,sK10('_i_0_0',bool)) | $spl16),
  inference(unit_resulting_resolution,[],[f341,f7923,f1209])).
fof(f7923,plain,(
  p(s(bool,sK10('_i_0_0',bool))) | $spl16),
  inference(forward_demodulation,[],[f7922,f266])).
fof(f7922,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK10('_i_0_0',bool))))) | $spl16),
  inference(subsumption_resolution,[],[f7900,f230])).
fof(f7900,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK10('_i_0_0',bool))))) | $spl16),
  inference(superposition,[],[f298,f4197])).
fof(f4197,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f341,f3861,f1209])).
fof(f298,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,X0)))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(fun(X44,bool),bool),'?_1'),s(fun(X44,bool),X45)))) <=> ? [X46] : p(s(bool,happ(s(fun(X44,bool),X45),s(X44,X46)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','?_HL_QUANT')).
fof(f10382,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK10('_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl16),
  inference(superposition,[],[f8066,f266])).
fof(f8066,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,sK10('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f8008,f230])).
fof(f8008,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,sK10('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(superposition,[],[f359,f7946])).
fof(f44541,plain,(
  ( ! [X546] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,'F_0'))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) | $spl106),
  inference(cnf_transformation,[],[f44541_D])).
fof(f44541_D,plain,(
  ( ! [X546] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,'F_0'))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f44540,plain,(
  ( ! [X546,X545] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X545),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,'F_0'))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44539,f315])).
fof(f44539,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,'F_0'))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X544),s(fun('type.ImplicitQueue.times'(sK0),bool),X545),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546))))) )),
  inference(forward_demodulation,[],[f44538,f27965])).
fof(f44538,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X544),s(fun('type.ImplicitQueue.times'(sK0),bool),X545),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546))))) )),
  inference(forward_demodulation,[],[f44045,f321])).
fof(f44045,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(bool,'F_0')))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X544),s(fun('type.ImplicitQueue.times'(sK0),bool),X545),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X546))))) )),
  inference(superposition,[],[f8241,f29483])).
fof(f44537,plain,(
  ( ! [X543] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) | $spl104),
  inference(cnf_transformation,[],[f44537_D])).
fof(f44537_D,plain,(
  ( ! [X543] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f44536,plain,(
  ( ! [X542,X543] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X542),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44535,f315])).
fof(f44535,plain,(
  ( ! [X542,X541,X543] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X541),s(fun('type.ImplicitQueue.times'(sK0),bool),X542),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543))))) )),
  inference(forward_demodulation,[],[f44044,f323])).
fof(f44044,plain,(
  ( ! [X542,X541,X543] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X541),s(fun('type.ImplicitQueue.times'(sK0),bool),X542),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X543))))) )),
  inference(superposition,[],[f8175,f29483])).
fof(f44477,plain,(
  ( ! [X348] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X348),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) | $spl60),
  inference(cnf_transformation,[],[f44477_D])).
fof(f44477_D,plain,(
  ( ! [X348] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X348),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f44532,plain,(
  ( ! [X532,X534] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X532),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) | $spl102),
  inference(cnf_transformation,[],[f44532_D])).
fof(f44532_D,plain,(
  ( ! [X532,X534] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X532),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f44531,plain,(
  ( ! [X532,X534,X533] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X532),s('type.ImplicitQueue.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X533),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44041,f315])).
fof(f44041,plain,(
  ( ! [X532,X534,X533,X531] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X533),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X531),s(fun('type.ImplicitQueue.times'(sK0),bool),X532),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X533))))))))) )),
  inference(superposition,[],[f7693,f29483])).
fof(f7693,plain,(
  ( ! [X37,X38] : (~p(s(bool,X38)) | p(s(bool,happ(s(fun(fun(X37,bool),bool),'!_1'),s(fun(X37,bool),'_dst_x0_1'(s(bool,X38))))))) )),
  inference(superposition,[],[f297,f309])).
fof(f297,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X44,X45] : (p(s(bool,happ(s(fun(fun(X44,bool),bool),'!_1'),s(fun(X44,bool),X45)))) <=> ! [X46] : p(s(bool,happ(s(fun(X44,bool),X45),s(X44,X46)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','!_HL_QUANT')).
fof(f44528,plain,(
  ( ! [X516,X518] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X516),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) | $spl100),
  inference(cnf_transformation,[],[f44528_D])).
fof(f44528_D,plain,(
  ( ! [X516,X518] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X516),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f44527,plain,(
  ( ! [X516,X518,X517] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X516),s('type.ImplicitQueue.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X517),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44036,f315])).
fof(f44036,plain,(
  ( ! [X516,X518,X517,X515] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X517),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X515),s(fun('type.ImplicitQueue.times'(sK0),bool),X516),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X517))))))))) )),
  inference(superposition,[],[f3842,f29483])).
fof(f3842,plain,(
  ( ! [X50,X49] : (~p(s(bool,X50)) | p(s(bool,happ(s(fun(fun(X49,bool),bool),'?_1'),s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))))))) )),
  inference(superposition,[],[f299,f309])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f44522,plain,(
  ( ! [X466,X465] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X465),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))))))) ) | $spl98),
  inference(cnf_transformation,[],[f44522_D])).
fof(f44522_D,plain,(
  ( ! [X466,X465] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X465),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f44483,plain,(
  ( ! [X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X353),s('type.ImplicitQueue.times'(sK0),sK3)))) ) | $spl68),
  inference(cnf_transformation,[],[f44483_D])).
fof(f44483_D,plain,(
  ( ! [X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X353),s('type.ImplicitQueue.times'(sK0),sK3)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f44521,plain,(
  ( ! [X466,X465,X464] : (s(bool,'F_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X464),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X465),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))))))) )),
  inference(forward_demodulation,[],[f44023,f315])).
fof(f44023,plain,(
  ( ! [X466,X463,X465,X464] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X465),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))))) | s(bool,'F_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X463),s(fun('type.ImplicitQueue.times'(sK0),bool),X464),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X465)))) )),
  inference(superposition,[],[f1424,f29483])).
fof(f1424,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,X22)))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f1423,f266])).
fof(f1423,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1413,f230])).
fof(f1413,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1259,f389])).
fof(f1259,plain,(
  ( ! [X21,X22] : (~p(s(bool,X22)) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,X22))))))) )),
  inference(superposition,[],[f280,f309])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK7(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X18,X19] : (! [X20] : (s(X6,X19) = s(X6,X20) => p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X20))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X18)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.bool.JRH_INDUCT_UTIL')).
fof(f44520,plain,(
  ( ! [X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X460),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) | $spl96),
  inference(cnf_transformation,[],[f44520_D])).
fof(f44520_D,plain,(
  ( ! [X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X460),s('type.ImplicitQueue.times'(sK0),sK3))))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f44519,plain,(
  ( ! [X461,X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X460),s('type.ImplicitQueue.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X461),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44022,f315])).
fof(f44022,plain,(
  ( ! [X461,X459,X460,X462] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X461),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X459),s(fun('type.ImplicitQueue.times'(sK0),bool),X460),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X461))))))))) )),
  inference(superposition,[],[f1259,f29483])).
fof(f44517,plain,(
  ( ! [X453,X451] : (s(bool,X453) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X451),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'F_0') = s(bool,X453)) ) | $spl94),
  inference(cnf_transformation,[],[f44517_D])).
fof(f44517_D,plain,(
  ( ! [X453,X451] : (s(bool,X453) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X451),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'F_0') = s(bool,X453)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f44516,plain,(
  ( ! [X453,X451,X452] : (s(bool,X453) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X451),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X452),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'F_0') = s(bool,X453)) )),
  inference(forward_demodulation,[],[f44020,f315])).
fof(f44020,plain,(
  ( ! [X453,X451,X452,X450] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X452),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,X453) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X450),s(fun('type.ImplicitQueue.times'(sK0),bool),X451),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X452))) | s(bool,'F_0') = s(bool,X453)) )),
  inference(superposition,[],[f1209,f29483])).
fof(f44474,plain,(
  ( ! [X345] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X345),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) | $spl56),
  inference(cnf_transformation,[],[f44474_D])).
fof(f44474_D,plain,(
  ( ! [X345] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X345),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f44514,plain,(
  ( ! [X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X423),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X425)) ) | $spl92),
  inference(cnf_transformation,[],[f44514_D])).
fof(f44514_D,plain,(
  ( ! [X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X423),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X425)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f44513,plain,(
  ( ! [X424,X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X423),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X424),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,X425)) )),
  inference(forward_demodulation,[],[f44014,f315])).
fof(f44014,plain,(
  ( ! [X422,X424,X423,X425] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X424),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,X425) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X422),s(fun('type.ImplicitQueue.times'(sK0),bool),X423),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X424))) | s(bool,'T_0') = s(bool,X425)) )),
  inference(superposition,[],[f415,f29483])).
fof(f44512,plain,(
  ( ! [X416,X418] : (s(bool,X418) != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X416),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X418)) ) | $spl90),
  inference(cnf_transformation,[],[f44512_D])).
fof(f44512_D,plain,(
  ( ! [X416,X418] : (s(bool,X418) != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X416),s('type.ImplicitQueue.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X418)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f44511,plain,(
  ( ! [X416,X418,X417] : (s(bool,X418) != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X416),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X417),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,X418)) )),
  inference(forward_demodulation,[],[f44012,f315])).
fof(f44012,plain,(
  ( ! [X416,X418,X415,X417] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X417),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,X418) != s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X415),s(fun('type.ImplicitQueue.times'(sK0),bool),X416),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X417))) | s(bool,'T_0') = s(bool,X418)) )),
  inference(superposition,[],[f359,f29483])).
fof(f44510,plain,(
  ( ! [X396,X398] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.ImplicitQueue.times'(sK0),bool),X396))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) | $spl88),
  inference(cnf_transformation,[],[f44510_D])).
fof(f44510_D,plain,(
  ( ! [X396,X398] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.ImplicitQueue.times'(sK0),bool),X396))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f44509,plain,(
  ( ! [X397,X396,X398] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.ImplicitQueue.times'(sK0),bool),X396))),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X397),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44508,f323])).
fof(f44508,plain,(
  ( ! [X397,X396,X398] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X396),s('type.ImplicitQueue.times'(sK0),sK3)))))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X397),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44007,f315])).
fof(f44007,plain,(
  ( ! [X397,X395,X396,X398] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X397),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398))),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X395),s(fun('type.ImplicitQueue.times'(sK0),bool),X396),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X397))))))) )),
  inference(superposition,[],[f305,f29483])).
fof(f305,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f44506,plain,(
  ( ! [X389,X388,X390] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X388),s('type.ImplicitQueue.times'(sK0),sK3)))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X390))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44505,f315])).
fof(f44505,plain,(
  ( ! [X389,X387,X388,X390] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X390))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X387),s(fun('type.ImplicitQueue.times'(sK0),bool),X388),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389))))) )),
  inference(forward_demodulation,[],[f44005,f324])).
fof(f44005,plain,(
  ( ! [X389,X387,X388,X390] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(bool,X390)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X387),s(fun('type.ImplicitQueue.times'(sK0),bool),X388),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X389))))) )),
  inference(superposition,[],[f304,f29483])).
fof(f44504,plain,(
  ( ! [X384,X386] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X384),s(bool,X386))),s('type.ImplicitQueue.times'(sK0),sK3)))) | p(s(bool,X386))) ) | $spl84),
  inference(cnf_transformation,[],[f44504_D])).
fof(f44504_D,plain,(
  ( ! [X384,X386] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X384),s(bool,X386))),s('type.ImplicitQueue.times'(sK0),sK3)))) | p(s(bool,X386))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f44503,plain,(
  ( ! [X385,X384,X386] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X384),s(bool,X386))),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X385),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X386))) )),
  inference(forward_demodulation,[],[f44502,f324])).
fof(f44502,plain,(
  ( ! [X385,X384,X386] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X384),s('type.ImplicitQueue.times'(sK0),sK3))))),s(bool,X386)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X385),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X386))) )),
  inference(forward_demodulation,[],[f44004,f315])).
fof(f44004,plain,(
  ( ! [X383,X385,X384,X386] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X385),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X386)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X383),s(fun('type.ImplicitQueue.times'(sK0),bool),X384),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X385))))),s(bool,X386))))) )),
  inference(superposition,[],[f303,f29483])).
fof(f44501,plain,(
  ( ! [X380,X382] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.ImplicitQueue.times'(sK0),bool),X380))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) | $spl82),
  inference(cnf_transformation,[],[f44501_D])).
fof(f44501_D,plain,(
  ( ! [X380,X382] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.ImplicitQueue.times'(sK0),bool),X380))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f44500,plain,(
  ( ! [X380,X382,X381] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.ImplicitQueue.times'(sK0),bool),X380))),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X381),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44499,f323])).
fof(f44499,plain,(
  ( ! [X380,X382,X381] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X380),s('type.ImplicitQueue.times'(sK0),sK3)))))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X381),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44003,f315])).
fof(f44003,plain,(
  ( ! [X379,X380,X382,X381] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X381),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382))),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X379),s(fun('type.ImplicitQueue.times'(sK0),bool),X380),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X381))))))) )),
  inference(superposition,[],[f302,f29483])).
fof(f302,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f44498,plain,(
  ( ! [X376,X378] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X376),s(bool,X378))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) | $spl80),
  inference(cnf_transformation,[],[f44498_D])).
fof(f44498_D,plain,(
  ( ! [X376,X378] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X376),s(bool,X378))),s('type.ImplicitQueue.times'(sK0),sK3))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f44497,plain,(
  ( ! [X377,X376,X378] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),X376),s(bool,X378))),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X377),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44496,f324])).
fof(f44496,plain,(
  ( ! [X377,X376,X378] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X376),s('type.ImplicitQueue.times'(sK0),sK3))))),s(bool,X378)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X377),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f44002,f315])).
fof(f44002,plain,(
  ( ! [X377,X376,X378,X375] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X377),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X375),s(fun('type.ImplicitQueue.times'(sK0),bool),X376),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X377))))),s(bool,X378))))) )),
  inference(superposition,[],[f301,f29483])).
fof(f301,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f44495,plain,(
  ( ! [X374,X373] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X374))) ) | $spl78),
  inference(cnf_transformation,[],[f44495_D])).
fof(f44495_D,plain,(
  ( ! [X374,X373] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X374))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f44494,plain,(
  ( ! [X372,X374,X373] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X372),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X374))) )),
  inference(forward_demodulation,[],[f44493,f315])).
fof(f44493,plain,(
  ( ! [X372,X374,X373,X371] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373),s('type.ImplicitQueue.times'(sK0),sK2))))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X374)) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X371),s(fun('type.ImplicitQueue.times'(sK0),bool),X372),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373))))) )),
  inference(forward_demodulation,[],[f44001,f323])).
fof(f44001,plain,(
  ( ! [X372,X374,X373,X371] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374))),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))))) | p(s(bool,X374)) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X371),s(fun('type.ImplicitQueue.times'(sK0),bool),X372),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X373))))) )),
  inference(superposition,[],[f300,f29483])).
fof(f44491,plain,(
  ( ! [X370,X369] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X370))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X370))) ) | $spl74),
  inference(cnf_transformation,[],[f44491_D])).
fof(f44491_D,plain,(
  ( ! [X370,X369] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X370))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X370))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f44490,plain,(
  ( ! [X368,X370,X369] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X368),s('type.ImplicitQueue.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X370))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X370))) )),
  inference(forward_demodulation,[],[f44489,f315])).
fof(f44489,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369),s('type.ImplicitQueue.times'(sK0),sK2))),s(bool,X370))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X367),s(fun('type.ImplicitQueue.times'(sK0),bool),X368),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369)))) | p(s(bool,X370))) )),
  inference(forward_demodulation,[],[f44000,f324])).
fof(f44000,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))),s(bool,X370)))) | p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X367),s(fun('type.ImplicitQueue.times'(sK0),bool),X368),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X369)))) | p(s(bool,X370))) )),
  inference(superposition,[],[f300,f29483])).
fof(f44488,plain,(
  ( ! [X364,X366] : (s(bool,X366) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X364),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,X366))) ) | $spl72),
  inference(cnf_transformation,[],[f44488_D])).
fof(f44488_D,plain,(
  ( ! [X364,X366] : (s(bool,X366) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X364),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,X366))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f44487,plain,(
  ( ! [X364,X366,X365] : (s(bool,X366) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X364),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X365),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | ~p(s(bool,X366))) )),
  inference(forward_demodulation,[],[f43999,f315])).
fof(f43999,plain,(
  ( ! [X364,X366,X365,X363] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X365),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | ~p(s(bool,X366)) | s(bool,X366) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X363),s(fun('type.ImplicitQueue.times'(sK0),bool),X364),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X365)))) )),
  inference(superposition,[],[f295,f29483])).
fof(f44485,plain,(
  ( ! [X356,X358] : (s(bool,X358) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X356),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,X358))) ) | $spl70),
  inference(cnf_transformation,[],[f44485_D])).
fof(f44485_D,plain,(
  ( ! [X356,X358] : (s(bool,X358) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X356),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,X358))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f44486,plain,(
  ( ! [X360,X362,X361] : (s(bool,X362) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X360),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X361),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X362))) )),
  inference(forward_demodulation,[],[f43998,f315])).
fof(f43998,plain,(
  ( ! [X360,X362,X359,X361] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X361),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X362)) | s(bool,X362) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X359),s(fun('type.ImplicitQueue.times'(sK0),bool),X360),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X361)))) )),
  inference(superposition,[],[f292,f29483])).
fof(f44484,plain,(
  ( ! [X356,X358,X357] : (s(bool,X358) = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X356),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X357),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X358))) )),
  inference(forward_demodulation,[],[f43997,f315])).
fof(f43997,plain,(
  ( ! [X356,X358,X357,X355] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X357),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | p(s(bool,X358)) | s(bool,X358) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X355),s(fun('type.ImplicitQueue.times'(sK0),bool),X356),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X357)))) )),
  inference(superposition,[],[f292,f29483])).
fof(f44482,plain,(
  ( ! [X354,X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X353),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X354),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f43996,f315])).
fof(f43996,plain,(
  ( ! [X352,X354,X353] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X354),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'F_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X352),s(fun('type.ImplicitQueue.times'(sK0),bool),X353),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X354)))) )),
  inference(superposition,[],[f245,f29483])).
fof(f245,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f44481,plain,(
  ( ! [X350] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X350),s('type.ImplicitQueue.times'(sK0),sK3))))) ) | $spl66),
  inference(cnf_transformation,[],[f44481_D])).
fof(f44481_D,plain,(
  ( ! [X350] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X350),s('type.ImplicitQueue.times'(sK0),sK3))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f44479,plain,(
  ( ! [X350,X351] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X350),s('type.ImplicitQueue.times'(sK0),sK3)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X351),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f43995,f315])).
fof(f43995,plain,(
  ( ! [X350,X349,X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X351),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))) | ~p(s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X349),s(fun('type.ImplicitQueue.times'(sK0),bool),X350),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X351))))) )),
  inference(superposition,[],[f244,f29483])).
fof(f44478,plain,(
  ( ! [X347] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X347),s('type.ImplicitQueue.times'(sK0),sK3)))) ) | $spl62),
  inference(cnf_transformation,[],[f44478_D])).
fof(f44478_D,plain,(
  ( ! [X347] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X347),s('type.ImplicitQueue.times'(sK0),sK3)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f44476,plain,(
  ( ! [X347,X348] : (s(bool,'T_0') = s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X347),s('type.ImplicitQueue.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X348),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f43994,f315])).
fof(f43994,plain,(
  ( ! [X346,X347,X348] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X348),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X346),s(fun('type.ImplicitQueue.times'(sK0),bool),X347),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X348)))) )),
  inference(superposition,[],[f241,f29483])).
fof(f44475,plain,(
  ( ! [X344] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X344),s('type.ImplicitQueue.times'(sK0),sK3)))) ) | $spl58),
  inference(cnf_transformation,[],[f44475_D])).
fof(f44475_D,plain,(
  ( ! [X344] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X344),s('type.ImplicitQueue.times'(sK0),sK3)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f44473,plain,(
  ( ! [X344,X345] : (s(bool,'T_0') != s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),X344),s('type.ImplicitQueue.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X345),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f43993,f315])).
fof(f43993,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(sK0),bool),happ(s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X345),s('type.ImplicitQueue.times'(sK0),sK2))),s('type.ImplicitQueue.times'(sK0),sK1)))) | s(bool,'T_0') != s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(sK0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(sK0),sK3))),s(bool,X343),s(fun('type.ImplicitQueue.times'(sK0),bool),X344),s(fun('type.ImplicitQueue.times'(sK0),fun('type.ImplicitQueue.times'(sK0),bool)),X345)))) )),
  inference(superposition,[],[f240,f29483])).
fof(f8597,plain,(
  ( ! [X60] : (s(bool,'T_0') = s(bool,X60)) ) | $spl24),
  inference(cnf_transformation,[],[f8597_D])).
fof(f8597_D,plain,(
  ( ! [X60] : (s(bool,'T_0') = s(bool,X60)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f41770,plain,(
  ( ! [X76] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X76)) ) | $spl52),
  inference(cnf_transformation,[],[f41770_D])).
fof(f41770_D,plain,(
  ( ! [X76] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X76)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f43374,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76)) ) | $spl44),
  inference(duplicate_literal_removal,[],[f43066])).
fof(f43066,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) ) | $spl44),
  inference(superposition,[],[f1697,f36723])).
fof(f36723,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(trivial_inequality_removal,[],[f36713])).
fof(f36713,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(superposition,[],[f36630,f356])).
fof(f36630,plain,(
  s(bool,'T_0') != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(subsumption_resolution,[],[f36520,f334])).
fof(f36520,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(superposition,[],[f240,f35559])).
fof(f1806,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl8),
  inference(cnf_transformation,[],[f1806_D])).
fof(f1806_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f42001,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) ) | $spl54),
  inference(cnf_transformation,[],[f42001_D])).
fof(f42001_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f42002,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,'F_0')))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f41875,f240])).
fof(f41875,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,'F_0')))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f8241,f356])).
fof(f42000,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f41874,f240])).
fof(f41874,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f8241,f356])).
fof(f41672,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76)) ) | $spl38),
  inference(duplicate_literal_removal,[],[f41373])).
fof(f41373,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) ) | $spl38),
  inference(superposition,[],[f1697,f35553])).
fof(f41229,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f40893,f41082])).
fof(f41082,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1806,f317])).
fof(f40893,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1806,f244])).
fof(f41228,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f40897,f230])).
fof(f40897,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f341,f1806,f295])).
fof(f41227,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f40899,f230])).
fof(f40899,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f1806,f295])).
fof(f41226,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40901,f230])).
fof(f40901,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f30168,f1806,f295])).
fof(f41225,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40903,f230])).
fof(f40903,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28871,f1806,f295])).
fof(f41224,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f40905,f230])).
fof(f40905,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f24579,f1806,f295])).
fof(f41223,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f40907,f230])).
fof(f40907,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1806,f295])).
fof(f41222,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f40909,f230])).
fof(f40909,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f36630,f1806,f295])).
fof(f41221,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f40913,f1806])).
fof(f40913,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f1806,f295])).
fof(f41220,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40915,f1806])).
fof(f40915,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f30168,f1806,f295])).
fof(f41219,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40917,f1806])).
fof(f40917,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28871,f1806,f295])).
fof(f41218,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f40919,f1806])).
fof(f40919,plain,(
  ( ! [X0] : (~p(s(bool,sK7(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f24579,f1806,f295])).
fof(f41217,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f40921,f1806])).
fof(f40921,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1806,f295])).
fof(f41216,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f40923,f1806])).
fof(f40923,plain,(
  ~p(s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f36630,f1806,f295])).
fof(f41215,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f40927,f230])).
fof(f40927,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f341,f1806,f295])).
fof(f41214,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f40929,f230])).
fof(f40929,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f1806,f295])).
fof(f41213,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40931,f230])).
fof(f40931,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f30168,f1806,f295])).
fof(f41212,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40933,f230])).
fof(f40933,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28871,f1806,f295])).
fof(f41211,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f40935,f230])).
fof(f40935,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f24579,f1806,f295])).
fof(f41210,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f40937,f230])).
fof(f40937,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1806,f295])).
fof(f41209,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f40939,f230])).
fof(f40939,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f36630,f1806,f295])).
fof(f41208,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f40943,f1806])).
fof(f40943,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f1806,f295])).
fof(f41207,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40945,f1806])).
fof(f40945,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f30168,f1806,f295])).
fof(f41206,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f40947,f1806])).
fof(f40947,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28871,f1806,f295])).
fof(f41205,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f40949,f1806])).
fof(f40949,plain,(
  ( ! [X0] : (~p(s(bool,sK7(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f24579,f1806,f295])).
fof(f41204,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f40951,f1806])).
fof(f40951,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1806,f295])).
fof(f41203,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f40953,f1806])).
fof(f40953,plain,(
  ~p(s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f36630,f1806,f295])).
fof(f41202,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f40957,f1806])).
fof(f40957,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f334,f1806,f303])).
fof(f41201,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f41047,f41082])).
fof(f41047,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f341,f1806,f359])).
fof(f41200,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f41049,f41082])).
fof(f41049,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f1806,f359])).
fof(f41199,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f41051,f41082])).
fof(f41051,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) ) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f30168,f1806,f359])).
fof(f41198,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f41053,f41082])).
fof(f41053,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28871,f1806,f359])).
fof(f41197,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f41055,f41082])).
fof(f41055,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK7(X1,'_i_0_0',bool))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f24579,f1806,f359])).
fof(f41196,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f41057,f41082])).
fof(f41057,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK9('_i_0_0',bool))) ) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11918,f1806,f359])).
fof(f41195,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f41059,f41082])).
fof(f41059,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) ) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f36630,f1806,f359])).
fof(f41194,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f41092,f1806])).
fof(f41092,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f334,f1806,f303])).
fof(f41095,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f334,f1806,f8175])).
fof(f41093,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f230,f334,f1806,f303])).
fof(f41089,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f334,f334,f1806,f300])).
fof(f41083,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f341,f1806,f317])).
fof(f41076,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f341,f341,f1806,f1209])).
fof(f41060,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f35559,f36630,f1806,f359])).
fof(f41058,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f7709,f11918,f1806,f359])).
fof(f41056,plain,(
  $false | ($spl4 | $spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f24610,f24579,f1806,f359])).
fof(f41054,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f266,f28871,f1806,f359])).
fof(f41052,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f29032,f30168,f1806,f359])).
fof(f41050,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12140,f12123,f1806,f359])).
fof(f41048,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12140,f341,f1806,f359])).
fof(f40959,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f334,f1806,f1806,f303])).
fof(f40954,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f1806,f36630,f1806,f295])).
fof(f40952,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f1806,f11918,f1806,f295])).
fof(f40950,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1806,f24579,f1806,f295])).
fof(f40948,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1806,f28871,f1806,f295])).
fof(f40946,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1806,f30168,f1806,f295])).
fof(f40944,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1806,f12123,f1806,f295])).
fof(f40942,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1806,f341,f1806,f295])).
fof(f40940,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f230,f36630,f1806,f295])).
fof(f40938,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f230,f11918,f1806,f295])).
fof(f40936,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f230,f24579,f1806,f295])).
fof(f40934,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f230,f28871,f1806,f295])).
fof(f40932,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f230,f30168,f1806,f295])).
fof(f40930,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f230,f12123,f1806,f295])).
fof(f40928,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f230,f341,f1806,f295])).
fof(f40924,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f1806,f36630,f1806,f295])).
fof(f40922,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f1806,f11918,f1806,f295])).
fof(f40920,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1806,f24579,f1806,f295])).
fof(f40918,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1806,f28871,f1806,f295])).
fof(f40916,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1806,f30168,f1806,f295])).
fof(f40914,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1806,f12123,f1806,f295])).
fof(f40912,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1806,f341,f1806,f295])).
fof(f40910,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f230,f36630,f1806,f295])).
fof(f40908,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f230,f11918,f1806,f295])).
fof(f40906,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f230,f24579,f1806,f295])).
fof(f40904,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f230,f28871,f1806,f295])).
fof(f40902,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f230,f30168,f1806,f295])).
fof(f40900,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f230,f12123,f1806,f295])).
fof(f40898,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f230,f341,f1806,f295])).
fof(f40894,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12140,f1806,f244])).
fof(f40892,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f341,f1806,f241])).
fof(f40890,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f334,f1806])).
fof(f40755,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl50),
  inference(cnf_transformation,[],[f40755_D])).
fof(f40755_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f40756,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X6)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f40642,f240])).
fof(f40642,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X6)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f8175,f356])).
fof(f40754,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f40641,f240])).
fof(f40641,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f8175,f356])).
fof(f35577,plain,(
  $false | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f35576,f341])).
fof(f35576,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl34 | $spl44)),
  inference(backward_demodulation,[],[f35550,f35559])).
fof(f35550,plain,(
  s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl34),
  inference(cnf_transformation,[],[f35550_D])).
fof(f35550_D,plain,(
  s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f35552,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | $spl36),
  inference(cnf_transformation,[],[f35552_D])).
fof(f35552_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f35565,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | ($spl4 | $spl18)),
  inference(forward_demodulation,[],[f35498,f24610])).
fof(f35498,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,sK7(X12,'_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) ) | $spl4),
  inference(superposition,[],[f33626,f24635])).
fof(f24635,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,sK7(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(subsumption_resolution,[],[f24595,f230])).
fof(f24595,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,sK7(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(superposition,[],[f1270,f356])).
fof(f33626,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(X0,bool),sK10('?_1',fun(X0,bool))),s(X0,sK10(sK10('?_1',fun(X0,bool)),X0)))))) )),
  inference(unit_resulting_resolution,[],[f32604,f298])).
fof(f32604,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK10('?_1',fun(X0,bool))))))) )),
  inference(unit_resulting_resolution,[],[f31803,f298])).
fof(f31803,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))) )),
  inference(unit_resulting_resolution,[],[f31246,f299])).
fof(f31246,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f230,f3842])).
fof(f35559,plain,(
  s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(cnf_transformation,[],[f35559_D])).
fof(f35559_D,plain,(
  s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f35555,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | $spl40),
  inference(cnf_transformation,[],[f35555_D])).
fof(f35555_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f35564,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl2),
  inference(forward_demodulation,[],[f35497,f22288])).
fof(f22288,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl2),
  inference(unit_resulting_resolution,[],[f341,f1268,f1209])).
fof(f35497,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl2),
  inference(superposition,[],[f33626,f22332])).
fof(f22332,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl2),
  inference(subsumption_resolution,[],[f22297,f334])).
fof(f22297,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f1268,f1590])).
fof(f1590,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f1236,f1236])).
fof(f35563,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl18),
  inference(forward_demodulation,[],[f35496,f7709])).
fof(f35496,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,sK9('_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl18),
  inference(superposition,[],[f33626,f11917])).
fof(f35562,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl16),
  inference(forward_demodulation,[],[f35495,f7946])).
fof(f35495,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,sK10('_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | $spl16),
  inference(superposition,[],[f33626,f7979])).
fof(f7979,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,sK10('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f7953,f334])).
fof(f7953,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,sK10('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl16),
  inference(superposition,[],[f7923,f1590])).
fof(f35561,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl48),
  inference(cnf_transformation,[],[f35561_D])).
fof(f35561_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f35494,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X10)))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f33626,f1590])).
fof(f35493,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f33626,f1590])).
fof(f35560,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) ) | $spl46),
  inference(cnf_transformation,[],[f35560_D])).
fof(f35560_D,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f35492,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X8) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f33626,f1590])).
fof(f35491,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X7) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f33626,f1590])).
fof(f35557,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))),
  inference(forward_demodulation,[],[f35490,f265])).
fof(f35490,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'const.bool.DATATYPE_1'(s(X4,X5)))))) | s(bool,'F_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f33626,f1236])).
fof(f35553,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f35553_D])).
fof(f35553_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f35489,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X3) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f33626,f356])).
fof(f35556,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f35556_D])).
fof(f35556_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f35488,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f33626,f356])).
fof(f35487,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f33626,f356])).
fof(f35486,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33626,f356])).
fof(f35485,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK10('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10(sK10('?_1',fun(bool,bool)),bool))),
  inference(superposition,[],[f33626,f237])).
fof(f23961,plain,(
  ( ! [X99,X100] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X100)) ) | $spl32),
  inference(cnf_transformation,[],[f23961_D])).
fof(f23961_D,plain,(
  ( ! [X99,X100] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f25355,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f25158])).
fof(f25158,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X98)) ) | $spl4),
  inference(superposition,[],[f1697,f24635])).
fof(f23875,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f23683])).
fof(f23683,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK7(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X98)) ) | $spl0),
  inference(superposition,[],[f1697,f22588])).
fof(f22588,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK7(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f22485,f341])).
fof(f22485,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK7(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1266,f356])).
fof(f3859,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f3859_D])).
fof(f3859_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1268,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl2),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3283,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3254,f230])).
fof(f3254,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f3206,f389])).
fof(f3206,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f3092,f240])).
fof(f3092,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f280,f690])).
fof(f690,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK7(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f689,f334])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK7(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f279,f237])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK7(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f3371,plain,(
  ( ! [X33] : (~p(s(bool,X33))) ) | $spl12),
  inference(cnf_transformation,[],[f3371_D])).
fof(f3371_D,plain,(
  ( ! [X33] : (~p(s(bool,X33))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3266,plain,(
  ( ! [X44] : (~p(s(bool,X44)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3206,f266])).
fof(f4552,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4521,f230])).
fof(f4521,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f3206,f389])).
fof(f4535,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3206,f266])).
fof(f14122,plain,(
  ( ! [X49] : (~p(s(bool,X49)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3206,f266])).
fof(f15025,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f14934,f230])).
fof(f14934,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f3206,f389])).
fof(f22204,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl30),
  inference(cnf_transformation,[],[f22204_D])).
fof(f22204_D,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X67)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f22120,plain,(
  ( ! [X66,X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67)) ) | ($spl2 | $spl18 | $spl21)),
  inference(duplicate_literal_removal,[],[f21928])).
fof(f21928,plain,(
  ( ! [X66,X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) ) | ($spl2 | $spl18 | $spl21)),
  inference(superposition,[],[f1697,f19203])).
fof(f19203,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl18 | $spl21)),
  inference(trivial_inequality_removal,[],[f19187])).
fof(f19187,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl18 | $spl21)),
  inference(superposition,[],[f19039,f356])).
fof(f19039,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl2 | $spl18 | $spl21)),
  inference(forward_demodulation,[],[f18951,f12140])).
fof(f18951,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))) | ($spl2 | $spl21)),
  inference(unit_resulting_resolution,[],[f12123,f18747])).
fof(f18747,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f18638,f334])).
fof(f18638,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(superposition,[],[f240,f2037])).
fof(f2037,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) | $spl2),
  inference(forward_demodulation,[],[f1919,f1283])).
fof(f1283,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl2),
  inference(unit_resulting_resolution,[],[f1268,f241])).
fof(f1919,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) ) | $spl2),
  inference(superposition,[],[f1344,f266])).
fof(f1344,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f1321,f230])).
fof(f1321,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(superposition,[],[f1209,f1283])).
fof(f11771,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760,f11762,f11315,f11271,f11265,f11262])).
fof(f11262,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11256,f244])).
fof(f11265,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11256,f11256,f295])).
fof(f11271,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11256,f11256,f295])).
fof(f11315,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f359])).
fof(f11762,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl20),
  inference(subsumption_resolution,[],[f11677,f230])).
fof(f11677,plain,(
  ( ! [X24,X23,X21,X25] : (~p(s(bool,'T_0')) | s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f7030])).
fof(f11760,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl20),
  inference(subsumption_resolution,[],[f11671,f230])).
fof(f11671,plain,(
  ( ! [X24,X23,X21,X25] : (~p(s(bool,'T_0')) | s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f7002])).
fof(f11753,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl20),
  inference(subsumption_resolution,[],[f11622,f230])).
fof(f11622,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f6968])).
fof(f11751,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl20),
  inference(subsumption_resolution,[],[f11590,f230])).
fof(f11590,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f6925])).
fof(f11749,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl20),
  inference(subsumption_resolution,[],[f11573,f230])).
fof(f11573,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f317])).
fof(f11744,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(condensation,[],[f11743])).
fof(f11743,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl20),
  inference(condensation,[],[f11742])).
fof(f11742,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(subsumption_resolution,[],[f11501,f11256])).
fof(f11501,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'T_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f6954])).
fof(f11740,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(condensation,[],[f11739])).
fof(f11739,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl20),
  inference(condensation,[],[f11738])).
fof(f11738,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(subsumption_resolution,[],[f11496,f11256])).
fof(f11496,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f6911])).
fof(f11319,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f1209])).
fof(f11256,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl20),
  inference(forward_demodulation,[],[f11255,f266])).
fof(f11255,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11234,f230])).
fof(f11234,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl20),
  inference(superposition,[],[f296,f9616])).
fof(f11219,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl20),
  inference(subsumption_resolution,[],[f11166,f230])).
fof(f11166,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl20),
  inference(superposition,[],[f359,f9616])).
fof(f11218,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7),s(fun(fun(bool,bool),bool),'!_1'))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11217,f323])).
fof(f11217,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11163,f230])).
fof(f11163,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(superposition,[],[f305,f9616])).
fof(f11215,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4),s(fun(fun(bool,bool),bool),'!_1'))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11214,f323])).
fof(f11214,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11160,f230])).
fof(f11160,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(superposition,[],[f302,f9616])).
fof(f11213,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun(bool,bool),bool),'!_1'),s(bool,X3))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11212,f324])).
fof(f11212,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X3))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11159,f230])).
fof(f11159,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X3))))) ) | $spl20),
  inference(superposition,[],[f301,f9616])).
fof(f9650,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl20),
  inference(subsumption_resolution,[],[f9631,f334])).
fof(f9631,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl20),
  inference(superposition,[],[f7711,f1590])).
fof(f9616,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f230,f7711,f295])).
fof(f9618,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f7711,f359])).
fof(f9619,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))) | $spl20),
  inference(unit_resulting_resolution,[],[f266,f7711,f359])).
fof(f9620,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(X0,X1)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f309,f7711,f359])).
fof(f9621,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(X0,X1)))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f310,f7711,f359])).
fof(f9622,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f276,f7711,f359])).
fof(f9623,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f308,f7711,f359])).
fof(f9624,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(fun('type.ImplicitQueue.times'(X0),bool),X1),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),bool)),X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f314,f7711,f359])).
fof(f9625,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f7711,f1259])).
fof(f447,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f276,f237])).
fof(f489,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f308,f237])).
fof(f535,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f309,f237])).
fof(f585,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f310,f237])).
fof(f1204,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20)))) )),
  inference(forward_demodulation,[],[f1164,f266])).
fof(f1164,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20)))) )),
  inference(superposition,[],[f389,f309])).
fof(f1205,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X22),s(X21,X23)))))) )),
  inference(forward_demodulation,[],[f1165,f266])).
fof(f1165,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X22),s(X21,X23)))))) )),
  inference(superposition,[],[f389,f310])).
fof(f1206,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1166,f266])).
fof(f1166,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X24),s(fun('type.num.num',bool),X25),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f389,f276])).
fof(f1207,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(forward_demodulation,[],[f1167,f266])).
fof(f1167,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(superposition,[],[f389,f308])).
fof(f1210,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1182,f230])).
fof(f1182,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(superposition,[],[f359,f389])).
fof(f1235,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1222,f230])).
fof(f1222,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1209,f389])).
fof(f1355,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f306,f237])).
fof(f1380,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X80),X81))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X80),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X80),X81)))) )),
  inference(superposition,[],[f389,f306])).
fof(f1545,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f1517,f230])).
fof(f1517,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f1209,f1447])).
fof(f1561,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,'const.bool.DATATYPE_1'(s(X51,X52))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X50)))) )),
  inference(superposition,[],[f1236,f266])).
fof(f1562,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X54) = s(bool,'const.bool.DATATYPE_1'(s(X56,X57))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),'_dst_x0_1'(s(bool,X54))),s(X53,X55)))) )),
  inference(superposition,[],[f1236,f309])).
fof(f1564,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,X63) = s(bool,'const.bool.DATATYPE_1'(s(X65,X66))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X62),'const.pair.,_2'(s(bool,X63),s(X62,X64)))))) )),
  inference(superposition,[],[f1236,f310])).
fof(f1565,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) = s(bool,'const.bool.DATATYPE_1'(s(X69,X70))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X67),s(fun('type.num.num',bool),X68),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1236,f276])).
fof(f1566,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X73) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X71,bool),'const.pair.,_2'(s(X71,X72),s(bool,X73)))))) )),
  inference(superposition,[],[f1236,f308])).
fof(f1684,plain,(
  ( ! [X59,X58] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) )),
  inference(forward_demodulation,[],[f1563,f265])).
fof(f1563,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'const.bool.DATATYPE_1'(s(X60,X61))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) )),
  inference(superposition,[],[f1236,f306])).
fof(f1808,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1784,f334])).
fof(f1784,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18)))))) )),
  inference(superposition,[],[f1424,f237])).
fof(f1890,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1889,f230])).
fof(f1889,plain,(
  ( ! [X0,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1869,f265])).
fof(f1869,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X1,X2)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f301,f1236])).
fof(f1893,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X27))),s(bool,X26)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(forward_demodulation,[],[f1892,f324])).
fof(f1892,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26))))),s(bool,X27)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1878,f230])).
fof(f1878,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26))))),s(bool,X27)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f301,f389])).
fof(f2117,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)))) )),
  inference(superposition,[],[f1590,f266])).
fof(f2118,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63) | s(bool,'F_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62)))) )),
  inference(superposition,[],[f1590,f309])).
fof(f2119,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65)))) )),
  inference(superposition,[],[f1590,f306])).
fof(f2120,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X68) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69)))))) )),
  inference(superposition,[],[f1590,f310])).
fof(f2121,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1590,f276])).
fof(f2122,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76)))))) )),
  inference(superposition,[],[f1590,f308])).
fof(f2178,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49))),s(X48,X50))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f1590,f309])).
fof(f2179,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53))) | s(bool,X54) = s(bool,happ(s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X52),X53))) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f1590,f306])).
fof(f2180,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,X56),s(X55,X57))))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f1590,f310])).
fof(f2181,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1590,f276])).
fof(f2182,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X62,bool),'const.pair.,_2'(s(X62,X63),s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f1590,f308])).
fof(f2336,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f244,f1590])).
fof(f2719,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2372,f334])).
fof(f2372,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f292,f1590])).
fof(f2730,plain,(
  ( ! [X76,X74,X77,X75] : (s(X74,X76) = s(X74,sK7(X76,X75,X74)) | s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(X74,bool),X75))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f2396,f334])).
fof(f2396,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,'F_0')) | s(X74,X76) = s(X74,sK7(X76,X75,X74)) | s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(X74,bool),X75))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f279,f1590])).
fof(f2731,plain,(
  ( ! [X80,X81,X82] : (s(bool,'F_0') = s(bool,X81) | s(bool,X82) = s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2398,f334])).
fof(f2398,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X81) | s(bool,X82) = s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f1424,f1590])).
fof(f2777,plain,(
  ( ! [X14,X17] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f2776,f230])).
fof(f2776,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f2752,f265])).
fof(f2752,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f302,f1236])).
fof(f2781,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47),s(fun(bool,bool),'_i_0_0'))),s(bool,X46)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2780,f323])).
fof(f2780,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2763,f230])).
fof(f2763,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f302,f389])).
fof(f2834,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)))) )),
  inference(superposition,[],[f1697,f266])).
fof(f2835,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) != s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62)))) )),
  inference(superposition,[],[f1697,f309])).
fof(f2836,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65)))) )),
  inference(superposition,[],[f1697,f306])).
fof(f2837,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X68) != s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69)))))) )),
  inference(superposition,[],[f1697,f310])).
fof(f2838,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1697,f276])).
fof(f2839,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76)))))) )),
  inference(superposition,[],[f1697,f308])).
fof(f2862,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f1697,f266])).
fof(f2863,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) != s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f1697,f309])).
fof(f2864,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,X66)) )),
  inference(superposition,[],[f1697,f306])).
fof(f2865,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X68) != s(bool,X70) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69))))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f1697,f310])).
fof(f2866,plain,(
  ( ! [X72,X71,X73] : (s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f1697,f276])).
fof(f2867,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76))))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f1697,f308])).
fof(f2933,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,X61) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f2337,f266])).
fof(f2934,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | s(bool,X65) = s(bool,happ(s(fun(X62,bool),'_dst_x0_1'(s(bool,X63))),s(X62,X64))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f2337,f309])).
fof(f2935,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,X68) = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f2337,f306])).
fof(f2936,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'T_0') != s(bool,X70) | s(bool,X72) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),'const.pair.,_2'(s(bool,X70),s(X69,X71))))) | s(bool,'F_0') = s(bool,X72)) )),
  inference(superposition,[],[f2337,f310])).
fof(f2937,plain,(
  ( ! [X74,X75,X73] : (s(bool,'T_0') != s(bool,X73) | s(bool,X75) = s(bool,'const.ind_type.FCONS_3'(s(bool,X73),s(fun('type.num.num',bool),X74),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X75)) )),
  inference(superposition,[],[f2337,f276])).
fof(f2938,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X79) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X76,bool),'const.pair.,_2'(s(X76,X77),s(bool,X78))))) | s(bool,'F_0') = s(bool,X79)) )),
  inference(superposition,[],[f2337,f308])).
fof(f3047,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X38))),s(bool,X39)))) | s(bool,X38) = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f2979,f334])).
fof(f2979,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X38))),s(bool,X39)))) | p(s(bool,'F_0')) | s(bool,X38) = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f304,f1590])).
fof(f3063,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X20))),s(bool,X21))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f3029,f334])).
fof(f3029,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X20)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X20))),s(bool,X21))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f304,f1590])).
fof(f3064,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,sK7(X3,X4,bool))) )),
  inference(superposition,[],[f690,f237])).
fof(f3069,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X17) = s(bool,'const.bool.DATATYPE_1'(s(X19,X20))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X18))) | s(bool,'F_0') = s(bool,sK7(X17,X18,bool))) )),
  inference(superposition,[],[f690,f1236])).
fof(f3071,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,X25) = s(bool,sK7(X23,X24,bool)) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f690,f1590])).
fof(f3073,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,sK7(X29,X30,bool))) )),
  inference(superposition,[],[f690,f1590])).
fof(f3107,plain,(
  ( ! [X57,X58] : (s(bool,'F_0') != s(bool,X57) | ~p(s(bool,sK7(X57,X58,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f244,f690])).
fof(f3108,plain,(
  ( ! [X59,X60] : (p(s(bool,X59)) | s(bool,'F_0') = s(bool,sK7(X59,X60,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f245,f690])).
fof(f3121,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X97))) | s(bool,'F_0') = s(bool,sK7(X97,X98,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X98)))) )),
  inference(superposition,[],[f389,f690])).
fof(f3124,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK7(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f1209,f690])).
fof(f3129,plain,(
  ( ! [X121,X122,X120] : (p(s(bool,'|dollar|exists_1'(s(fun(X122,bool),'_dst_x0_1'(s(bool,X120)))))) | s(bool,'F_0') = s(bool,sK7(X120,X121,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X121)))) )),
  inference(superposition,[],[f1424,f690])).
fof(f3134,plain,(
  ( ! [X136,X135,X137] : (s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,sK7(X135,X136,bool)) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X136)))) )),
  inference(superposition,[],[f1697,f690])).
fof(f3135,plain,(
  ( ! [X140,X138,X139] : (s(bool,X138) != s(bool,X140) | s(bool,'T_0') = s(bool,X140) | s(bool,'F_0') = s(bool,sK7(X138,X139,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X139)))) )),
  inference(superposition,[],[f1697,f690])).
fof(f3136,plain,(
  ( ! [X142,X143,X141] : (s(bool,'T_0') != s(bool,X141) | s(bool,X143) = s(bool,sK7(X141,X142,bool)) | s(bool,'F_0') = s(bool,X143) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X142)))) )),
  inference(superposition,[],[f2337,f690])).
fof(f3287,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(subsumption_resolution,[],[f3286,f230])).
fof(f3286,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(forward_demodulation,[],[f3260,f265])).
fof(f3260,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X23,X24)))) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(superposition,[],[f3206,f1236])).
fof(f3324,plain,(
  ( ! [X14,X17] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f3323,f230])).
fof(f3323,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f3297,f265])).
fof(f3297,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f305,f1236])).
fof(f3328,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47),s(fun(bool,bool),'_i_0_0'))),s(bool,X46)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f3327,f323])).
fof(f3327,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f3308,f230])).
fof(f3308,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f305,f389])).
fof(f3451,plain,(
  ( ! [X23,X19,X20] : (p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(subsumption_resolution,[],[f3450,f230])).
fof(f3450,plain,(
  ( ! [X23,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(forward_demodulation,[],[f3437,f265])).
fof(f3437,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X21,X22)))) | p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(superposition,[],[f296,f1236])).
fof(f3864,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(subsumption_resolution,[],[f3863,f230])).
fof(f3863,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(forward_demodulation,[],[f3834,f265])).
fof(f3834,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X23,X24)))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(superposition,[],[f299,f1236])).
fof(f5478,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f5420,f334])).
fof(f5420,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f1868,f1590])).
fof(f5713,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f5645,f334])).
fof(f5645,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f2746,f1590])).
fof(f5856,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f5778,f334])).
fof(f5778,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f5575,f1590])).
fof(f6115,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f284,f237])).
fof(f6122,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,X34) = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f284,f1590])).
fof(f6124,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) )),
  inference(superposition,[],[f284,f1590])).
fof(f6159,plain,(
  ( ! [X138,X136,X137] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X137),s('type.num.num',X138))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X136),s(fun('type.num.num',bool),X137),s('type.num.num','const.num.SUC_1'(s('type.num.num',X138)))))) )),
  inference(superposition,[],[f389,f284])).
fof(f6172,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,X190) != s(bool,happ(s(fun('type.num.num',bool),X188),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X187),s(fun('type.num.num',bool),X188),s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))))) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f1697,f284])).
fof(f6173,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,X194) != s(bool,happ(s(fun('type.num.num',bool),X192),s('type.num.num',X193))) | s(bool,'T_0') = s(bool,X194) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X191),s(fun('type.num.num',bool),X192),s('type.num.num','const.num.SUC_1'(s('type.num.num',X193)))))) )),
  inference(superposition,[],[f1697,f284])).
fof(f6174,plain,(
  ( ! [X198,X196,X197,X195] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X196),s('type.num.num',X197))) | s(bool,X198) = s(bool,'const.ind_type.FCONS_3'(s(bool,X195),s(fun('type.num.num',bool),X196),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197))))) | s(bool,'F_0') = s(bool,X198)) )),
  inference(superposition,[],[f2337,f284])).
fof(f6203,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(forward_demodulation,[],[f6120,f265])).
fof(f6120,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'const.bool.DATATYPE_1'(s(X26,X27))) = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f284,f1236])).
fof(f6347,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f341,f2965,f1209])).
fof(f6432,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f6340,f323])).
fof(f6340,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f2965,f305])).
fof(f6433,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f6339,f323])).
fof(f6339,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f2965,f302])).
fof(f6434,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))),s(bool,X0))))) )),
  inference(forward_demodulation,[],[f6338,f324])).
fof(f6338,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f2965,f301])).
fof(f6537,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X5),'const.ImplicitQueue.Zero_0'),s(bool,X6),s(fun('type.ImplicitQueue.times'(X5),bool),X7),s(fun('type.ImplicitQueue.times'(X5),fun('type.ImplicitQueue.times'(X5),bool)),X8)))) )),
  inference(superposition,[],[f314,f237])).
fof(f6542,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,X30) = s(bool,'const.bool.DATATYPE_1'(s(X33,X34))) | s(bool,'F_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X29),'const.ImplicitQueue.Zero_0'),s(bool,X30),s(fun('type.ImplicitQueue.times'(X29),bool),X31),s(fun('type.ImplicitQueue.times'(X29),fun('type.ImplicitQueue.times'(X29),bool)),X32)))) )),
  inference(superposition,[],[f314,f1236])).
fof(f6544,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,X40) | s(bool,X43) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X39),'const.ImplicitQueue.Zero_0'),s(bool,X40),s(fun('type.ImplicitQueue.times'(X39),bool),X41),s(fun('type.ImplicitQueue.times'(X39),fun('type.ImplicitQueue.times'(X39),bool)),X42))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f314,f1590])).
fof(f6546,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X49),'const.ImplicitQueue.Zero_0'),s(bool,X50),s(fun('type.ImplicitQueue.times'(X49),bool),X51),s(fun('type.ImplicitQueue.times'(X49),fun('type.ImplicitQueue.times'(X49),bool)),X52)))) )),
  inference(superposition,[],[f314,f1590])).
fof(f6595,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,X239) != s(bool,X242) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X238),'const.ImplicitQueue.Zero_0'),s(bool,X239),s(fun('type.ImplicitQueue.times'(X238),bool),X240),s(fun('type.ImplicitQueue.times'(X238),fun('type.ImplicitQueue.times'(X238),bool)),X241))) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f1697,f314])).
fof(f6596,plain,(
  ( ! [X243,X246,X244,X247,X245] : (s(bool,X244) != s(bool,X247) | s(bool,'T_0') = s(bool,X247) | s(bool,'F_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X243),'const.ImplicitQueue.Zero_0'),s(bool,X244),s(fun('type.ImplicitQueue.times'(X243),bool),X245),s(fun('type.ImplicitQueue.times'(X243),fun('type.ImplicitQueue.times'(X243),bool)),X246)))) )),
  inference(superposition,[],[f1697,f314])).
fof(f6597,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'T_0') != s(bool,X249) | s(bool,X252) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X248),'const.ImplicitQueue.Zero_0'),s(bool,X249),s(fun('type.ImplicitQueue.times'(X248),bool),X250),s(fun('type.ImplicitQueue.times'(X248),fun('type.ImplicitQueue.times'(X248),bool)),X251))) | s(bool,'F_0') = s(bool,X252)) )),
  inference(superposition,[],[f2337,f314])).
fof(f6886,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f6815,f334])).
fof(f6815,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f3291,f1590])).
fof(f6892,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f341,f317])).
fof(f6904,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X42)))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f317,f237])).
fof(f6911,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f317,f1590])).
fof(f6913,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,X69)))) | s(bool,X67) = s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f317,f1590])).
fof(f6919,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X84))),s(bool,X85)))) | s(bool,'F_0') = s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f317,f1590])).
fof(f6921,plain,(
  ( ! [X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X90))),s(bool,X91)))) | s(bool,'F_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f317,f237])).
fof(f6947,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'F_0')))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f317,f237])).
fof(f6954,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'F_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f317,f1590])).
fof(f7005,plain,(
  ( ! [X58,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X58)))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f6909,f265])).
fof(f6909,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.DATATYPE_1'(s(X56,X57))))),s(bool,X58)))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f317,f1236])).
fof(f7033,plain,(
  ( ! [X58,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'T_0')))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f6952,f265])).
fof(f6952,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'const.bool.DATATYPE_1'(s(X56,X57)))))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f317,f1236])).
fof(f7054,plain,(
  ( ! [X21,X19,X20] : (s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(subsumption_resolution,[],[f7053,f230])).
fof(f7053,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(forward_demodulation,[],[f6986,f265])).
fof(f6986,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X22,X23)))) | s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(superposition,[],[f317,f1236])).
fof(f7207,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f7123,f334])).
fof(f7123,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f6741,f1590])).
fof(f7217,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f6889,f245])).
fof(f7218,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f6889,f240])).
fof(f7257,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f7211,f324])).
fof(f7211,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f6889,f304])).
fof(f7262,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7240,f230])).
fof(f7240,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6889,f356])).
fof(f7308,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X42)))) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f318,f237])).
fof(f7315,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X62)))) | s(bool,X61) != s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7317,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,X69)))) | s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7323,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X84))),s(bool,X85)))) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7325,plain,(
  ( ! [X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X90))),s(bool,X91)))) | s(bool,'F_0') != s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f318,f237])).
fof(f7351,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'F_0')))) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f318,f237])).
fof(f7358,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'F_0')))) | s(bool,X61) != s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7360,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X68)))) | s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7366,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X85))),s(bool,X84)))) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7368,plain,(
  ( ! [X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X90)))) | s(bool,'F_0') != s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f318,f237])).
fof(f7481,plain,(
  ( ! [X58,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X58)))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f7313,f265])).
fof(f7313,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.DATATYPE_1'(s(X56,X57))))),s(bool,X58)))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f318,f1236])).
fof(f7509,plain,(
  ( ! [X58,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'T_0')))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f7356,f265])).
fof(f7356,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'const.bool.DATATYPE_1'(s(X56,X57)))))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f318,f1236])).
fof(f7535,plain,(
  ( ! [X30,X28,X29,X27] : (s(X27,X28) != s(X27,X29) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),happ(s(fun(X27,fun(X27,bool)),'=_0'),s(X27,X28))),s(X27,X29))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f7392,f334])).
fof(f7392,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'F_0')) | s(X27,X28) != s(X27,X29) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),happ(s(fun(X27,fun(X27,bool)),'=_0'),s(X27,X28))),s(X27,X29))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f318,f1590])).
fof(f7595,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7569,f334])).
fof(f7569,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f7400,f1590])).
fof(f7653,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f7624,f334])).
fof(f7624,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f7542,f1590])).
fof(f7665,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK9(X0,bool))) )),
  inference(superposition,[],[f297,f237])).
fof(f7672,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X13)))) | s(bool,X14) = s(bool,sK9(X13,bool)) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f297,f1590])).
fof(f7674,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X17)))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,sK9(X17,bool))) )),
  inference(superposition,[],[f297,f1590])).
fof(f7703,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X9)))) | s(bool,'F_0') = s(bool,sK9(X9,bool))) )),
  inference(forward_demodulation,[],[f7670,f265])).
fof(f7670,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'const.bool.DATATYPE_1'(s(X10,X11)))))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X9)))) | s(bool,'F_0') = s(bool,sK9(X9,bool))) )),
  inference(superposition,[],[f297,f1236])).
fof(f7714,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK9(X15,X14))))) )),
  inference(subsumption_resolution,[],[f7713,f230])).
fof(f7713,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK9(X15,X14))))) )),
  inference(forward_demodulation,[],[f7684,f265])).
fof(f7684,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK9(X15,X14))))) )),
  inference(superposition,[],[f297,f1236])).
fof(f7926,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK10(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(subsumption_resolution,[],[f7925,f230])).
fof(f7925,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK10(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(forward_demodulation,[],[f7906,f265])).
fof(f7906,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK10(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(superposition,[],[f298,f1236])).
fof(f8076,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f334,f334,f300])).
fof(f8184,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X43))),s(bool,X44)))) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8092,f334])).
fof(f8092,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X43))),s(bool,X44)))) | p(s(bool,'F_0')) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f300,f1590])).
fof(f8250,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X44))),s(bool,X43)))) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8134,f334])).
fof(f8134,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X44))),s(bool,X43)))) | p(s(bool,X44)) | p(s(bool,'F_0')) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f300,f1590])).
fof(f8308,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(subsumption_resolution,[],[f8307,f230])).
fof(f8307,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(forward_demodulation,[],[f8166,f265])).
fof(f8166,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(superposition,[],[f300,f1236])).
fof(f8680,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f230,f334,f303])).
fof(f8724,plain,(
  ( ! [X14,X17] : (p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f8723,f230])).
fof(f8723,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f8686,f265])).
fof(f8686,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f303,f1236])).
fof(f8735,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X52))),s(bool,X51)))) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X51)) )),
  inference(forward_demodulation,[],[f8734,f324])).
fof(f8734,plain,(
  ( ! [X52,X51] : (p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X51))))),s(bool,X52)))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f8699,f230])).
fof(f8699,plain,(
  ( ! [X52,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X51))))),s(bool,X52)))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f303,f389])).
fof(f296,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f3220,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),sK7(X27,X28,fun(X26,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f3206,f690])).
fof(f3285,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3256,f230])).
fof(f3256,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f3206,f356])).
fof(f3449,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f3433,f230])).
fof(f3433,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f296,f356])).
fof(f3791,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | p(s(bool,happ(s(fun(fun(X26,bool),bool),'?_1'),s(fun(X26,bool),sK7(X27,X28,fun(X26,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f299,f690])).
fof(f3862,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3830,f230])).
fof(f3830,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f299,f356])).
fof(f7712,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK9(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7680,f230])).
fof(f7680,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK9(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f297,f356])).
fof(f7695,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f7656,f266])).
fof(f7656,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f297,f266])).
fof(f7696,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK9(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(forward_demodulation,[],[f7657,f309])).
fof(f7657,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK9(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))))))) )),
  inference(superposition,[],[f297,f309])).
fof(f7698,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),X11)))) | ~p(s(bool,happ(s(fun(X9,bool),X11),s(X9,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9)))))) )),
  inference(forward_demodulation,[],[f7659,f310])).
fof(f7659,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X11),s(X9,sK9('const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9))))) | p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12))))))))) )),
  inference(superposition,[],[f297,f310])).
fof(f7699,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun(fun(X13,bool),bool),'!_1'),s(fun(X13,bool),X14)))) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK9('const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13)))))) )),
  inference(forward_demodulation,[],[f7660,f276])).
fof(f7660,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK9('const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))))) | p(s(bool,happ(s(fun(fun(X13,bool),bool),'!_1'),s(fun(X13,bool),'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f297,f276])).
fof(f7701,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),X23)))) | ~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK9('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(forward_demodulation,[],[f7662,f308])).
fof(f7662,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK9('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))))))) )),
  inference(superposition,[],[f297,f308])).
fof(f7702,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),X26)))) | ~p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK9('const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28)),X24)))))) )),
  inference(forward_demodulation,[],[f7663,f314])).
fof(f7663,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK9('const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28)),X24))))) | p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28))))))) )),
  inference(superposition,[],[f297,f314])).
fof(f7664,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,sK9(sK7(X30,X31,fun(X29,bool)),X29))))) | p(s(bool,happ(s(fun(fun(X29,bool),bool),'!_1'),s(fun(X29,bool),sK7(X30,X31,fun(X29,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X29,bool),bool),X31)))) )),
  inference(superposition,[],[f297,f690])).
fof(f7924,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK10(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7902,f230])).
fof(f7902,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK10(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f298,f356])).
fof(f7914,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f7891,f266])).
fof(f7891,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))),s(X0,sK10(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(superposition,[],[f298,f266])).
fof(f7915,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK10(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | ~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4))))) )),
  inference(forward_demodulation,[],[f7892,f309])).
fof(f7892,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4)))) | p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))),s(X2,sK10(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(superposition,[],[f298,f309])).
fof(f7917,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X11),s(X9,sK10('const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9))))) | ~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X11))))) )),
  inference(forward_demodulation,[],[f7894,f310])).
fof(f7894,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X11)))) | p(s(bool,happ(s(fun(X9,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12))))),s(X9,sK10('const.pair.FST_1'(s('type.pair.prod'(fun(X9,bool),X10),'const.pair.,_2'(s(fun(X9,bool),X11),s(X10,X12)))),X9)))))) )),
  inference(superposition,[],[f298,f310])).
fof(f7918,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK10('const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))))) | ~p(s(bool,happ(s(fun(fun(X13,bool),bool),'?_1'),s(fun(X13,bool),X14))))) )),
  inference(forward_demodulation,[],[f7895,f276])).
fof(f7895,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(fun(X13,bool),bool),'?_1'),s(fun(X13,bool),X14)))) | p(s(bool,happ(s(fun(X13,bool),'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0'))),s(X13,sK10('const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13)))))) )),
  inference(superposition,[],[f298,f276])).
fof(f7920,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK10('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | ~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23))))) )),
  inference(forward_demodulation,[],[f7897,f308])).
fof(f7897,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23)))) | p(s(bool,happ(s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))),s(X20,sK10('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(superposition,[],[f298,f308])).
fof(f7921,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK10('const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28)),X24))))) | ~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X26))))) )),
  inference(forward_demodulation,[],[f7898,f314])).
fof(f7898,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X26)))) | p(s(bool,happ(s(fun(X24,bool),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28))),s(X24,sK10('const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X25),'const.ImplicitQueue.Zero_0'),s(fun(X24,bool),X26),s(fun('type.ImplicitQueue.times'(X25),fun(X24,bool)),X27),s(fun('type.ImplicitQueue.times'(X25),fun('type.ImplicitQueue.times'(X25),fun(X24,bool))),X28)),X24)))))) )),
  inference(superposition,[],[f298,f314])).
fof(f6903,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X38))),s(X37,X40)))) | s(X37,X40) = s(X37,sK7(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f317,f690])).
fof(f6946,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X40))),s(X37,X38)))) | s(X37,X40) = s(X37,sK7(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f317,f690])).
fof(f7307,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X38))),s(X37,X40)))) | s(X37,X40) != s(X37,sK7(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f318,f690])).
fof(f7350,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X40))),s(X37,X38)))) | s(X37,X40) != s(X37,sK7(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f318,f690])).
fof(f7437,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7436,f321])).
fof(f7436,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7435,f314])).
fof(f7435,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0'),s(X0,X2),s(fun('type.ImplicitQueue.times'(X1),X0),X3),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X4))))))) )),
  inference(forward_demodulation,[],[f7286,f321])).
fof(f7286,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X1),'const.ImplicitQueue.Zero_0'),s(X0,X2),s(fun('type.ImplicitQueue.times'(X1),X0),X3),s(fun('type.ImplicitQueue.times'(X1),fun('type.ImplicitQueue.times'(X1),X0)),X4))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f314,f318])).
fof(f7700,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19)))))) | ~p(s(bool,happ(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))),s(X16,sK9('const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f7661,f284])).
fof(f7661,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))),s(X16,sK9('const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))))) | p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))))))) )),
  inference(superposition,[],[f297,f284])).
fof(f7919,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))),s(X16,sK10('const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))))) | ~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))))) )),
  inference(forward_demodulation,[],[f7896,f284])).
fof(f7896,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19)))))) | p(s(bool,happ(s(fun(X16,bool),'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))),s(X16,sK10('const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16)))))) )),
  inference(superposition,[],[f298,f284])).
fof(f7697,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | ~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK9(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f7658,f306])).
fof(f7658,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK9(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(superposition,[],[f297,f306])).
fof(f7916,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK10(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(forward_demodulation,[],[f7893,f306])).
fof(f7893,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK10(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(superposition,[],[f298,f306])).
fof(f7471,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f7263,f323])).
fof(f7263,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f266,f318])).
fof(f3431,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(fun(X26,bool),bool),'!_1'),s(fun(X26,bool),X27)))) | p(s(bool,happ(s(fun(X26,bool),sK7(X27,X28,fun(X26,bool))),s(X26,X29)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f296,f690])).
fof(f7899,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(fun(X29,bool),bool),'?_1'),s(fun(X29,bool),X30)))) | p(s(bool,happ(s(fun(X29,bool),sK7(X30,X31,fun(X29,bool))),s(X29,sK10(sK7(X30,X31,fun(X29,bool)),X29))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X29,bool),bool),X31)))) )),
  inference(superposition,[],[f298,f690])).
fof(f7666,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK9(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f297,f356])).
fof(f7667,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK9(X3,bool))) )),
  inference(superposition,[],[f297,f356])).
fof(f8739,plain,(
  ( ! [X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X55)))) | p(s(bool,X55))) )),
  inference(forward_demodulation,[],[f8738,f1447])).
fof(f8738,plain,(
  ( ! [X54,X55] : (p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X54,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X55))))) )),
  inference(subsumption_resolution,[],[f8701,f230])).
fof(f8701,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X54,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X55))))) )),
  inference(superposition,[],[f303,f1447])).
fof(f7329,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X99)))) | s(bool,X99) != s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) )),
  inference(superposition,[],[f318,f690])).
fof(f7309,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X44)))) | s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f318,f356])).
fof(f6925,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X99)))) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) )),
  inference(superposition,[],[f317,f690])).
fof(f6905,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X44)))) | s(bool,X43) = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f317,f356])).
fof(f1891,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,X5)))) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1870,f230])).
fof(f1870,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,X5)))) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f301,f356])).
fof(f2775,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2748,f230])).
fof(f2748,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f302,f356])).
fof(f8118,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X110))),s(bool,X112)))) | p(s(bool,sK7(X110,X111,bool))) | p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X111)))) )),
  inference(superposition,[],[f300,f690])).
fof(f8160,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X112))),s(bool,X110)))) | p(s(bool,X112)) | p(s(bool,sK7(X110,X111,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X111)))) )),
  inference(superposition,[],[f300,f690])).
fof(f8177,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8176])).
fof(f8176,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f8079,f240])).
fof(f8079,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f300,f356])).
fof(f8243,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8242])).
fof(f8242,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f8121,f240])).
fof(f8121,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f300,f356])).
fof(f8306,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8162,f230])).
fof(f8162,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f300,f356])).
fof(f3044,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f3043])).
fof(f3043,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2968,f240])).
fof(f2968,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f304,f356])).
fof(f3060,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(subsumption_resolution,[],[f3022,f334])).
fof(f3022,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(superposition,[],[f304,f237])).
fof(f3114,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X76))),s(bool,X78)))) | p(s(bool,sK7(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) )),
  inference(superposition,[],[f304,f690])).
fof(f3322,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3293,f230])).
fof(f3293,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f305,f356])).
fof(f8722,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8682,f230])).
fof(f8682,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f303,f356])).
fof(f6906,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f317,f356])).
fof(f6917,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,'T_0') = s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f317,f356])).
fof(f7310,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,X46) != s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f318,f356])).
fof(f7321,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,'T_0') != s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f318,f356])).
fof(f7353,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X47)))) | s(bool,X46) != s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f318,f356])).
fof(f7364,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X79))),s(bool,X78)))) | s(bool,'T_0') != s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f318,f356])).
fof(f3112,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,sK7(X70,X71,bool)))),s(bool,X72)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f301,f690])).
fof(f3208,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))))),s(bool,X41)))) | s(X39,X42) = s(X39,sK7(X42,X40,X39))) )),
  inference(subsumption_resolution,[],[f3162,f230])).
fof(f3162,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))))),s(bool,X41)))) | s(X39,X42) = s(X39,sK7(X42,X40,X39))) )),
  inference(superposition,[],[f301,f690])).
fof(f8740,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X58)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))))),s(bool,X58)))) | s(X56,X59) = s(X56,sK7(X59,X57,X56))) )),
  inference(subsumption_resolution,[],[f8702,f230])).
fof(f8702,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))))),s(bool,X58)))) | s(X56,X59) = s(X56,sK7(X59,X57,X56))) )),
  inference(superposition,[],[f303,f690])).
fof(f8721,plain,(
  ( ! [X111,X109,X110] : (~p(s(bool,X109)) | p(s(bool,X111)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,sK7(X109,X110,bool)))),s(bool,X111)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f303,f690])).
fof(f7372,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X99))),s(bool,'T_0')))) | s(bool,X99) != s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) )),
  inference(superposition,[],[f318,f690])).
fof(f7352,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X44))),s(bool,'T_0')))) | s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f318,f356])).
fof(f6968,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X99))),s(bool,'T_0')))) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) )),
  inference(superposition,[],[f317,f690])).
fof(f6948,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X44))),s(bool,'T_0')))) | s(bool,X43) = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f317,f356])).
fof(f3209,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X45))),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))))) | s(X43,X46) = s(X43,sK7(X46,X44,X43))) )),
  inference(subsumption_resolution,[],[f3163,f230])).
fof(f3163,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X45))),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))))) | s(X43,X46) = s(X43,sK7(X46,X44,X43))) )),
  inference(superposition,[],[f302,f690])).
fof(f3335,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X55))),s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))))) | s(X53,X56) = s(X53,sK7(X56,X54,X53))) )),
  inference(subsumption_resolution,[],[f3312,f230])).
fof(f3312,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X55))),s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))))) | s(X53,X56) = s(X53,sK7(X56,X54,X53))) )),
  inference(superposition,[],[f305,f690])).
fof(f3113,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X75))),s(bool,sK7(X73,X74,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f302,f690])).
fof(f3321,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X86))),s(bool,sK7(X84,X85,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X85)))) )),
  inference(superposition,[],[f305,f690])).
fof(f1902,plain,(
  ( ! [X39,X38] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X39))),s(bool,X38)))) | ~p(s(bool,X38))) )),
  inference(forward_demodulation,[],[f1883,f324])).
fof(f1883,plain,(
  ( ! [X39,X38] : (~p(s(bool,X38)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))))),s(bool,X39))))) )),
  inference(superposition,[],[f301,f266])).
fof(f5447,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))),s(bool,X0))))) )),
  inference(forward_demodulation,[],[f5357,f324])).
fof(f5357,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f1868,f301])).
fof(f5672,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5582,f324])).
fof(f5582,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f2746,f301])).
fof(f6843,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6748,f324])).
fof(f6748,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f3291,f301])).
fof(f8765,plain,(
  ( ! [X80,X81] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X81))),s(bool,X80)))) | ~p(s(bool,X80)) | p(s(bool,X81))) )),
  inference(forward_demodulation,[],[f8713,f324])).
fof(f8713,plain,(
  ( ! [X80,X81] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))))),s(bool,X81))))) )),
  inference(superposition,[],[f303,f266])).
fof(f2792,plain,(
  ( ! [X61,X60] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X61),s(fun(bool,bool),'_i_0_0'))),s(bool,X60)))) | ~p(s(bool,X60))) )),
  inference(forward_demodulation,[],[f2769,f323])).
fof(f2769,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X61))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60))))))) )),
  inference(superposition,[],[f302,f266])).
fof(f3340,plain,(
  ( ! [X64,X65] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X65),s(fun(bool,bool),'_i_0_0'))),s(bool,X64)))) | ~p(s(bool,X64))) )),
  inference(forward_demodulation,[],[f3315,f323])).
fof(f3315,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X65))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64))))))) )),
  inference(superposition,[],[f305,f266])).
fof(f5445,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f5359,f323])).
fof(f5359,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f1868,f305])).
fof(f5446,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f5358,f323])).
fof(f5358,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f1868,f302])).
fof(f5670,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5584,f323])).
fof(f5584,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f2746,f305])).
fof(f5671,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5583,f323])).
fof(f5583,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f2746,f302])).
fof(f6841,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6750,f323])).
fof(f6750,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f3291,f305])).
fof(f6842,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6749,f323])).
fof(f6749,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f3291,f302])).
fof(f6740,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f3291,f299])).
fof(f7541,plain,(
  p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'?_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f7400,f299])).
fof(f3843,plain,(
  ( ! [X52,X53] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X52),bool),bool),'?_1'),s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f299,f306])).
fof(f7210,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f6889,f296])).
fof(f7694,plain,(
  ( ! [X39] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X39),sK9('const.pair.FST_0','type.pair.prod'(bool,X39)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X39),bool),bool),'!_1'),s(fun('type.pair.prod'(bool,X39),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f297,f306])).
fof(f6994,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f6894,f321])).
fof(f6894,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))) )),
  inference(unit_resulting_resolution,[],[f283,f317])).
fof(f6995,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f6893,f321])).
fof(f6893,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))) )),
  inference(unit_resulting_resolution,[],[f275,f317])).
fof(f6891,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f283,f317])).
fof(f6890,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f275,f317])).
fof(f7403,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7402,f321])).
fof(f7402,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7401,f268])).
fof(f7401,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7297,f321])).
fof(f7297,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f268,f318])).
fof(f7454,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.times'(X0),bool),'_dst_x0x0x1_3'(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))),s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'))),s('type.ImplicitQueue.times'(X0),X1))))) )),
  inference(forward_demodulation,[],[f7279,f323])).
fof(f7279,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.ImplicitQueue.digit'(X0),bool),happ(s(fun('type.ImplicitQueue.digit'(X0),fun('type.ImplicitQueue.digit'(X0),bool)),'=_0'),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))))),s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f268,f318])).
fof(f6116,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,X10) = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f284,f356])).
fof(f6117,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f284,f356])).
fof(f6217,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,happ(s(fun('type.num.num',bool),X105),s('type.num.num',X106)))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X105),s(bool,X107))),s('type.num.num',X106))))) )),
  inference(forward_demodulation,[],[f6216,f284])).
fof(f6216,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X105),s(bool,X107))),s('type.num.num',X106)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X104),s(fun('type.num.num',bool),X105),s('type.num.num','const.num.SUC_1'(s('type.num.num',X106))))))) )),
  inference(forward_demodulation,[],[f6151,f324])).
fof(f6151,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),X105),s('type.num.num',X106))))),s(bool,X107)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X104),s(fun('type.num.num',bool),X105),s('type.num.num','const.num.SUC_1'(s('type.num.num',X106))))))) )),
  inference(superposition,[],[f304,f284])).
fof(f6221,plain,(
  ( ! [X144,X145,X143] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',X145))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X143),s(fun('type.num.num',bool),X144),s('type.num.num','const.num.SUC_1'(s('type.num.num',X145)))))) )),
  inference(forward_demodulation,[],[f6161,f265])).
fof(f6161,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.bool.DATATYPE_1'(s(X146,X147))) != s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',X145))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X143),s(fun('type.num.num',bool),X144),s('type.num.num','const.num.SUC_1'(s('type.num.num',X145)))))) )),
  inference(superposition,[],[f640,f284])).
fof(f8237,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,happ(s(fun('type.num.num',bool),X98),s('type.num.num',X99)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X98),s(bool,X100))),s('type.num.num',X99)))) | p(s(bool,X100))) )),
  inference(forward_demodulation,[],[f8236,f284])).
fof(f8236,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X98),s(bool,X100))),s('type.num.num',X99)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X97),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.SUC_1'(s('type.num.num',X99)))))) | p(s(bool,X100))) )),
  inference(forward_demodulation,[],[f8115,f324])).
fof(f8115,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),X98),s('type.num.num',X99))))),s(bool,X100)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X97),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.SUC_1'(s('type.num.num',X99)))))) | p(s(bool,X100))) )),
  inference(superposition,[],[f300,f284])).
fof(f8303,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,happ(s(fun('type.num.num',bool),X98),s('type.num.num',X99)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X100),s(fun('type.num.num',bool),X98))),s('type.num.num',X99)))) | p(s(bool,X100))) )),
  inference(forward_demodulation,[],[f8302,f284])).
fof(f8302,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X100),s(fun('type.num.num',bool),X98))),s('type.num.num',X99)))) | p(s(bool,X100)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X97),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.SUC_1'(s('type.num.num',X99))))))) )),
  inference(forward_demodulation,[],[f8157,f323])).
fof(f8157,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X100))),s(bool,happ(s(fun('type.num.num',bool),X98),s('type.num.num',X99)))))) | p(s(bool,X100)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X97),s(fun('type.num.num',bool),X98),s('type.num.num','const.num.SUC_1'(s('type.num.num',X99))))))) )),
  inference(superposition,[],[f300,f284])).
fof(f7464,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f7463,f284])).
fof(f7463,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f7270,f324])).
fof(f7270,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) )),
  inference(unit_resulting_resolution,[],[f284,f318])).
fof(f7478,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) != s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(fun('type.num.num',X21),X23),s(X21,X25))),s('type.num.num',X24))))) )),
  inference(forward_demodulation,[],[f7477,f284])).
fof(f7477,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(fun('type.num.num',X21),X23),s(X21,X25))),s('type.num.num',X24)))) | s(X21,X25) != s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(forward_demodulation,[],[f7304,f324])).
fof(f7304,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))))),s(X21,X25)))) | s(X21,X25) != s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f318,f284])).
fof(f7002,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(fun('type.num.num',X21),X23),s(X21,X25))),s('type.num.num',X24))))) )),
  inference(forward_demodulation,[],[f7001,f284])).
fof(f7001,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(fun('type.num.num',X21),X23),s(X21,X25))),s('type.num.num',X24)))) | s(X21,X25) = s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(forward_demodulation,[],[f6900,f324])).
fof(f6900,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))))),s(X21,X25)))) | s(X21,X25) = s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f317,f284])).
fof(f7587,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(bool,X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7549,f324])).
fof(f7549,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f7400,f301])).
fof(f7506,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) != s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25),s(fun('type.num.num',X21),X23))),s('type.num.num',X24))))) )),
  inference(forward_demodulation,[],[f7505,f284])).
fof(f7505,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25),s(fun('type.num.num',X21),X23))),s('type.num.num',X24)))) | s(X21,X25) != s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(forward_demodulation,[],[f7347,f323])).
fof(f7347,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25))),s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))))) | s(X21,X25) != s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f318,f284])).
fof(f7030,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25),s(fun('type.num.num',X21),X23))),s('type.num.num',X24))))) )),
  inference(forward_demodulation,[],[f7029,f284])).
fof(f7029,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25),s(fun('type.num.num',X21),X23))),s('type.num.num',X24)))) | s(X21,X25) = s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(forward_demodulation,[],[f6943,f323])).
fof(f6943,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X25))),s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))))) | s(X21,X25) = s(X21,'const.ind_type.FCONS_3'(s(X21,X22),s(fun('type.num.num',X21),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f317,f284])).
fof(f7431,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f7430,f323])).
fof(f7430,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f7429,f321])).
fof(f7429,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f7428,f284])).
fof(f7428,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) )),
  inference(forward_demodulation,[],[f7288,f321])).
fof(f7288,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(unit_resulting_resolution,[],[f284,f318])).
fof(f7585,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7551,f323])).
fof(f7551,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f7400,f305])).
fof(f7586,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7550,f323])).
fof(f7550,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f7400,f302])).
fof(f684,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f279,f266])).
fof(f685,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X7) = s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f279,f309])).
fof(f686,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | s(X8,X12) = s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))) )),
  inference(superposition,[],[f279,f310])).
fof(f687,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),X14)))) | s(X13,X16) = s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))) )),
  inference(superposition,[],[f279,f276])).
fof(f688,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | s(X17,X21) = s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) )),
  inference(superposition,[],[f279,f308])).
fof(f3075,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f690,f266])).
fof(f3076,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5))) | s(X3,X7) = s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f690,f309])).
fof(f3078,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14))) | s(X12,X16) = s(X12,sK7(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) )),
  inference(superposition,[],[f690,f310])).
fof(f3079,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),X18))) | s(X17,X20) = s(X17,sK7(X20,'const.ind_type.FCONS_3'(s(fun(X17,bool),X18),s(fun('type.num.num',fun(X17,bool)),X19),s('type.num.num','const.num.|48|_0')),X17))) )),
  inference(superposition,[],[f690,f276])).
fof(f3080,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24))) | s(X21,X25) = s(X21,sK7(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))) )),
  inference(superposition,[],[f690,f308])).
fof(f3081,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))) | s(X26,X29) = s(X26,sK7(X29,sK7(X27,X28,fun(X26,bool)),X26)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f690,f690])).
fof(f3139,plain,(
  ( ! [X154,X152,X151,X153] : (p(s(bool,'|dollar|exists_1'(s(fun(X151,bool),X152)))) | s(X151,X154) = s(X151,sK7(X154,sK7(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f279,f690])).
fof(f3207,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | s(bool,X37) = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | s(X35,X38) = s(X35,sK7(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f3161,f230])).
fof(f3161,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | s(bool,X37) = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | s(X35,X38) = s(X35,sK7(X38,X36,X35))) )),
  inference(superposition,[],[f295,f690])).
fof(f3211,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))))))) | s(X92,X95) = s(X92,sK7(X95,X93,X92))) )),
  inference(subsumption_resolution,[],[f3174,f230])).
fof(f3174,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))))))) | s(X92,X95) = s(X92,sK7(X95,X93,X92))) )),
  inference(superposition,[],[f1259,f690])).
fof(f6606,plain,(
  ( ! [X294,X296,X298,X295,X297,X299] : (p(s(bool,'|dollar|exists_1'(s(fun(X294,bool),X296)))) | s(X294,X299) = s(X294,sK7(X299,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X295),'const.ImplicitQueue.Zero_0'),s(fun(X294,bool),X296),s(fun('type.ImplicitQueue.times'(X295),fun(X294,bool)),X297),s(fun('type.ImplicitQueue.times'(X295),fun('type.ImplicitQueue.times'(X295),fun(X294,bool))),X298)),X294))) )),
  inference(superposition,[],[f279,f314])).
fof(f6607,plain,(
  ( ! [X300,X302,X304,X301,X303,X305] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X300,bool),X302))) | s(X300,X305) = s(X300,sK7(X305,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X301),'const.ImplicitQueue.Zero_0'),s(fun(X300,bool),X302),s(fun('type.ImplicitQueue.times'(X301),fun(X300,bool)),X303),s(fun('type.ImplicitQueue.times'(X301),fun('type.ImplicitQueue.times'(X301),fun(X300,bool))),X304)),X300))) )),
  inference(superposition,[],[f690,f314])).
fof(f6184,plain,(
  ( ! [X241,X239,X237,X240,X238] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X237,bool),happ(s(fun('type.num.num',fun(X237,bool)),X239),s('type.num.num',X240))))) | s(X237,X241) = s(X237,sK7(X241,'const.ind_type.FCONS_3'(s(fun(X237,bool),X238),s(fun('type.num.num',fun(X237,bool)),X239),s('type.num.num','const.num.SUC_1'(s('type.num.num',X240)))),X237))) )),
  inference(superposition,[],[f690,f284])).
fof(f6183,plain,(
  ( ! [X235,X233,X236,X234,X232] : (p(s(bool,'|dollar|exists_1'(s(fun(X232,bool),happ(s(fun('type.num.num',fun(X232,bool)),X234),s('type.num.num',X235)))))) | s(X232,X236) = s(X232,sK7(X236,'const.ind_type.FCONS_3'(s(fun(X232,bool),X233),s(fun('type.num.num',fun(X232,bool)),X234),s('type.num.num','const.num.SUC_1'(s('type.num.num',X235)))),X232))) )),
  inference(superposition,[],[f279,f284])).
fof(f1386,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'|dollar|exists_1'(s(fun(X99,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X99,bool),X100),X101)))))) | s(X99,X102) = s(X99,sK7(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) )),
  inference(superposition,[],[f279,f306])).
fof(f3077,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK7(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f690,f306])).
fof(f1422,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1405,f230])).
fof(f1405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1259,f356])).
fof(f3126,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | p(s(bool,'|dollar|exists_1'(s(fun(X115,bool),'_dst_x0_1'(s(bool,sK7(X113,X114,bool))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X114)))) )),
  inference(superposition,[],[f1259,f690])).
fof(f3105,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') != s(bool,sK7(X53,X54,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) )),
  inference(superposition,[],[f240,f690])).
fof(f3106,plain,(
  ( ! [X56,X55] : (~p(s(bool,X55)) | s(bool,'T_0') = s(bool,sK7(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f241,f690])).
fof(f3109,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X61)) | p(s(bool,X63)) | s(bool,X63) = s(bool,sK7(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f292,f690])).
fof(f3111,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X69)) | s(bool,X69) = s(bool,sK7(X67,X68,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68)))) )),
  inference(superposition,[],[f295,f690])).
fof(f3120,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | s(bool,X96) != s(bool,sK7(X94,X95,bool)) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X95)))) )),
  inference(superposition,[],[f359,f690])).
fof(f3122,plain,(
  ( ! [X101,X99,X100] : (p(s(bool,X99)) | s(bool,X101) = s(bool,sK7(X99,X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) )),
  inference(superposition,[],[f415,f690])).
fof(f3145,plain,(
  ( ! [X175,X173,X176,X174] : (s('type.ImplicitQueue.times'(X173),X174) != s('type.ImplicitQueue.times'(X173),X176) | s('type.ImplicitQueue.digit'(X173),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X173),X176))) = s('type.ImplicitQueue.digit'(X173),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X173),sK7(X174,X175,'type.ImplicitQueue.times'(X173))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ImplicitQueue.times'(X173),bool),X175)))) )),
  inference(superposition,[],[f282,f690])).
fof(f1351,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f306,f356])).
fof(f1353,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f306,f356])).
fof(f1401,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) )),
  inference(forward_demodulation,[],[f1382,f265])).
fof(f1382,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'const.bool.DATATYPE_1'(s(X87,X88))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) )),
  inference(superposition,[],[f640,f306])).
fof(f336,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.|48|_0'),s('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'))))))),s(fun('type.ImplicitQueue.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),sK8(X0)))))) )),
  inference(forward_demodulation,[],[f291,f285])).
fof(f291,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),happ(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'))))))),s(fun('type.ImplicitQueue.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),sK8(X0)))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),happ(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'))))))),s(fun('type.ImplicitQueue.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),sK8(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f129])).
fof(f129,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),happ(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.ImplicitQueue.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),'_i_0_0'))))))),s(fun('type.ImplicitQueue.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)))),X1))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ? [X38] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),happ(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)),happ(s(fun('type.ImplicitQueue.times'(X6),'type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),happ(s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),'type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X6),'const.bool.ARB_0'))),s('type.ImplicitQueue.times'(X6),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ImplicitQueue.times'(X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X6),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.ImplicitQueue.times'(X6),'type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'_dst_x1x0_2'(s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),'type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'const.pair.,_0'),s('type.ImplicitQueue.times'(X6),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun('type.ImplicitQueue.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ImplicitQueue.times'(X6),bool)),'_dst_x00x11_2'(s(fun(fun('type.ImplicitQueue.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_i_0_0'),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),'type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),'_i_0_0'))))))),s(fun('type.ImplicitQueue.digit'(X6),'type.ind_type.recspace'('type.pair.prod'('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.times'(X6)))),X38))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_TY_DEF')).
fof(f7614,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f341,f7542,f1209])).
fof(f7615,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f7542,f1259])).
fof(f7542,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f7400,f3206])).
fof(f7553,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f284,f7400,f359])).
fof(f7560,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f341,f7400,f1209])).
fof(f7561,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f7400,f1259])).
fof(f7528,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(subsumption_resolution,[],[f7385,f334])).
fof(f7385,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(superposition,[],[f318,f237])).
fof(f7434,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7433,f321])).
fof(f7433,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7432,f306])).
fof(f7432,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7287,f321])).
fof(f7287,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(unit_resulting_resolution,[],[f306,f318])).
fof(f7400,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f7399,f321])).
fof(f7399,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f7398,f267])).
fof(f7398,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))))) )),
  inference(forward_demodulation,[],[f7298,f321])).
fof(f7298,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f267,f318])).
fof(f7073,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f341,f6741,f1209])).
fof(f7074,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f6741,f1259])).
fof(f7052,plain,(
  ( ! [X6,X4,X5,X3] : (s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6982,f230])).
fof(f6982,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f317,f356])).
fof(f6889,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f341,f317])).
fof(f6741,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f3291,f3206])).
fof(f6758,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f341,f3291,f1209])).
fof(f6759,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f3291,f1259])).
fof(f6737,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f6705,f334])).
fof(f6705,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f6331,f1590])).
fof(f6687,plain,(
  ( ! [X4] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X4)))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6331,f237])).
fof(f6685,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f6331,f1590])).
fof(f6683,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f341,f6331,f1209])).
fof(f6684,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f6331,f1259])).
fof(f6612,plain,(
  ( ! [X333,X329,X331,X332,X334,X330] : (s('type.ImplicitQueue.digit'(X329),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X329),X331))) != s('type.ImplicitQueue.digit'(X329),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X329),X334))) | s('type.ImplicitQueue.times'(X329),X334) = s('type.ImplicitQueue.times'(X329),'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X330),'const.ImplicitQueue.Zero_0'),s('type.ImplicitQueue.times'(X329),X331),s(fun('type.ImplicitQueue.times'(X330),'type.ImplicitQueue.times'(X329)),X332),s(fun('type.ImplicitQueue.times'(X330),fun('type.ImplicitQueue.times'(X330),'type.ImplicitQueue.times'(X329))),X333)))) )),
  inference(superposition,[],[f281,f314])).
fof(f6584,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,X184) != s(bool,'const.bool.DATATYPE_1'(s(X187,X188))) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X183),'const.ImplicitQueue.Zero_0'),s(bool,X184),s(fun('type.ImplicitQueue.times'(X183),bool),X185),s(fun('type.ImplicitQueue.times'(X183),fun('type.ImplicitQueue.times'(X183),bool)),X186)))) )),
  inference(superposition,[],[f640,f314])).
fof(f6539,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X14),'const.ImplicitQueue.Zero_0'),s(bool,X15),s(fun('type.ImplicitQueue.times'(X14),bool),X16),s(fun('type.ImplicitQueue.times'(X14),fun('type.ImplicitQueue.times'(X14),bool)),X17)))) )),
  inference(superposition,[],[f314,f356])).
fof(f6538,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X13) = s(bool,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X9),'const.ImplicitQueue.Zero_0'),s(bool,X10),s(fun('type.ImplicitQueue.times'(X9),bool),X11),s(fun('type.ImplicitQueue.times'(X9),fun('type.ImplicitQueue.times'(X9),bool)),X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f314,f356])).
fof(f6330,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f2965,f299])).
fof(f6331,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f2965,f3206])).
fof(f6348,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f2965,f1259])).
fof(f6188,plain,(
  ( ! [X255,X257,X259,X256,X258] : (s('type.ImplicitQueue.digit'(X255),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X255),X259))) != s('type.ImplicitQueue.digit'(X255),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X255),happ(s(fun('type.num.num','type.ImplicitQueue.times'(X255)),X257),s('type.num.num',X258))))) | s('type.ImplicitQueue.times'(X255),X259) = s('type.ImplicitQueue.times'(X255),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(X255),X256),s(fun('type.num.num','type.ImplicitQueue.times'(X255)),X257),s('type.num.num','const.num.SUC_1'(s('type.num.num',X258)))))) )),
  inference(superposition,[],[f281,f284])).
fof(f5731,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f341,f5575,f1209])).
fof(f5732,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f5575,f1259])).
fof(f5575,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f2746,f3206])).
fof(f5591,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f341,f2746,f1209])).
fof(f5592,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f2746,f1259])).
fof(f5544,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5514,f334])).
fof(f5514,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f5350,f1590])).
fof(f5498,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0)))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f5350,f356])).
fof(f5496,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0')))))),
  inference(unit_resulting_resolution,[],[f341,f5350,f1209])).
fof(f5497,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f5350,f1259])).
fof(f5350,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f1868,f3206])).
fof(f5366,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f341,f1868,f1209])).
fof(f5367,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1868,f1259])).
fof(f3291,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f230,f305])).
fof(f2746,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f230,f302])).
fof(f2965,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f334,f304])).
fof(f1868,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f230,f301])).
fof(f3367,plain,(
  ( ! [X14,X15] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(subsumption_resolution,[],[f3366,f230])).
fof(f3366,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(forward_demodulation,[],[f3358,f265])).
fof(f3358,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(superposition,[],[f3268,f1236])).
fof(f3365,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3354,f230])).
fof(f3354,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f3268,f356])).
fof(f3268,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),X49)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X48),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f3206,f306])).
fof(f3210,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X69) != s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68))) | s(bool,'T_0') = s(bool,X69) | s(X67,X70) = s(X67,sK7(X70,X68,X67))) )),
  inference(subsumption_resolution,[],[f3168,f230])).
fof(f3168,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,X69) != s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68))) | s(bool,'T_0') = s(bool,X69) | s(X67,X70) = s(X67,sK7(X70,X68,X67))) )),
  inference(superposition,[],[f359,f690])).
fof(f3143,plain,(
  ( ! [X167,X165,X166,X168] : (s('type.ImplicitQueue.digit'(X165),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X165),X166))) != s('type.ImplicitQueue.digit'(X165),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X165),X168))) | s('type.ImplicitQueue.times'(X165),X168) = s('type.ImplicitQueue.times'(X165),sK7(X166,X167,'type.ImplicitQueue.times'(X165))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ImplicitQueue.times'(X165),bool),X167)))) )),
  inference(superposition,[],[f281,f690])).
fof(f3138,plain,(
  ( ! [X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK7(X150,sK7(X148,X149,fun(X147,bool)),X147))))) | p(s(bool,'|dollar|exists_1'(s(fun(X147,bool),sK7(X148,X149,fun(X147,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X147,bool),bool),X149)))) )),
  inference(superposition,[],[f280,f690])).
fof(f3123,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,X102) != s(bool,'const.bool.DATATYPE_1'(s(X104,X105))) | s(bool,'T_0') = s(bool,sK7(X102,X103,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X103)))) )),
  inference(superposition,[],[f640,f690])).
fof(f3066,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK7(X8,X9,bool))) )),
  inference(superposition,[],[f690,f356])).
fof(f3065,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6))) | s(bool,X7) = s(bool,sK7(X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f690,f356])).
fof(f2903,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f2337,f356])).
fof(f2888,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2853,f341])).
fof(f2853,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f1697,f1590])).
fof(f2422,plain,(
  ( ! [X158,X156,X157] : (~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157)))) | s(bool,X158) = s(bool,sK7(X156,X157,bool)) | s(bool,'F_0') = s(bool,X158)) )),
  inference(superposition,[],[f280,f1590])).
fof(f2337,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f341,f1590])).
fof(f2274,plain,(
  ( ! [X167,X169,X168] : (~p(s(bool,happ(s(fun(bool,bool),X168),s(bool,X169)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) | s(bool,'F_0') = s(bool,X169) | s(bool,'F_0') = s(bool,sK7(X167,X168,bool))) )),
  inference(superposition,[],[f280,f1590])).
fof(f2151,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f1590,f1590])).
fof(f1705,plain,(
  ( ! [X182,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,'F_0') = s(bool,sK7(X182,X183,bool))) )),
  inference(forward_demodulation,[],[f1639,f265])).
fof(f1639,plain,(
  ( ! [X185,X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'const.bool.DATATYPE_1'(s(X184,X185)))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,'F_0') = s(bool,sK7(X182,X183,bool))) )),
  inference(superposition,[],[f280,f1236])).
fof(f1702,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK7(X90,X89,X88))))) )),
  inference(subsumption_resolution,[],[f1701,f230])).
fof(f1701,plain,(
  ( ! [X90,X88,X89] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK7(X90,X89,X88))))) )),
  inference(forward_demodulation,[],[f1617,f265])).
fof(f1617,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X91,X92)))) | p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK7(X90,X89,X88))))) )),
  inference(superposition,[],[f280,f1236])).
fof(f1544,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1513,f230])).
fof(f1513,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | s(bool,X20) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f359,f1447])).
fof(f1493,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1447,f356])).
fof(f1483,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.ImplicitQueue.digit'(X21),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X21),X24))) != s('type.ImplicitQueue.digit'(X21),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X21),X25))) | s('type.ImplicitQueue.times'(X21),X25) = s('type.ImplicitQueue.times'(X21),'const.pair.SND_1'(s('type.pair.prod'(X22,'type.ImplicitQueue.times'(X21)),'const.pair.,_2'(s(X22,X23),s('type.ImplicitQueue.times'(X21),X24)))))) )),
  inference(superposition,[],[f281,f308])).
fof(f1482,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.ImplicitQueue.digit'(X17),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X17),X18))) != s('type.ImplicitQueue.digit'(X17),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X17),X20))) | s('type.ImplicitQueue.times'(X17),X20) = s('type.ImplicitQueue.times'(X17),'const.ind_type.FCONS_3'(s('type.ImplicitQueue.times'(X17),X18),s(fun('type.num.num','type.ImplicitQueue.times'(X17)),X19),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f281,f276])).
fof(f1481,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.ImplicitQueue.digit'(X12),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X12),X14))) != s('type.ImplicitQueue.digit'(X12),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X12),X16))) | s('type.ImplicitQueue.times'(X12),X16) = s('type.ImplicitQueue.times'(X12),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(X12),X13),'const.pair.,_2'(s('type.ImplicitQueue.times'(X12),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f281,f310])).
fof(f1480,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.ImplicitQueue.digit'(X8),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X8),X11))) != s('type.ImplicitQueue.digit'(X8),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X8),'const.pair.FST_1'(s('type.pair.prod'('type.ImplicitQueue.times'(X8),X9),X10))))) | s('type.ImplicitQueue.times'(X8),X11) = s('type.ImplicitQueue.times'(X8),happ(s(fun('type.pair.prod'('type.ImplicitQueue.times'(X8),X9),'type.ImplicitQueue.times'(X8)),'const.pair.FST_0'),s('type.pair.prod'('type.ImplicitQueue.times'(X8),X9),X10)))) )),
  inference(superposition,[],[f281,f306])).
fof(f1479,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.ImplicitQueue.digit'(X3),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X3),X5))) != s('type.ImplicitQueue.digit'(X3),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X3),X7))) | s('type.ImplicitQueue.times'(X3),X7) = s('type.ImplicitQueue.times'(X3),happ(s(fun(X4,'type.ImplicitQueue.times'(X3)),'_dst_x0_1'(s('type.ImplicitQueue.times'(X3),X5))),s(X4,X6)))) )),
  inference(superposition,[],[f281,f309])).
fof(f1478,plain,(
  ( ! [X2,X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))) | s('type.ImplicitQueue.times'(X0),X2) = s('type.ImplicitQueue.times'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.times'(X0)),'_i_0_0'),s('type.ImplicitQueue.times'(X0),X1)))) )),
  inference(superposition,[],[f281,f266])).
fof(f1447,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f341,f1404,f1209])).
fof(f1448,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f1404,f1259])).
fof(f1404,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f230,f1259])).
fof(f1403,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,'|dollar|exists_1'(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97)))))) | ~p(s(bool,happ(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97))),s(X95,sK7(X98,happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97)),X95)))))) )),
  inference(forward_demodulation,[],[f1385,f306])).
fof(f1385,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97))),s(X95,sK7(X98,happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97)),X95))))) | p(s(bool,'|dollar|exists_1'(s(fun(X95,bool),happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97))))))) )),
  inference(superposition,[],[f280,f306])).
fof(f1366,plain,(
  ( ! [X45,X44] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),sK7(X45,'const.pair.FST_0','type.pair.prod'(bool,X44)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X44),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f280,f306])).
fof(f1269,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK7(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1253,f230])).
fof(f1253,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK7(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f280,f356])).
fof(f1251,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK7(X12,X13,bool))) )),
  inference(superposition,[],[f280,f237])).
fof(f1249,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK7(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f280,f356])).
fof(f1247,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK7(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f280,f356])).
fof(f1264,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | ~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f1246,f308])).
fof(f1246,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20))))))))) )),
  inference(superposition,[],[f280,f308])).
fof(f1263,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),X14)))) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13)))))) )),
  inference(forward_demodulation,[],[f1245,f276])).
fof(f1245,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))))) | p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f280,f276])).
fof(f1262,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X10)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8)))))) )),
  inference(forward_demodulation,[],[f1244,f310])).
fof(f1244,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11))))))))) )),
  inference(superposition,[],[f280,f310])).
fof(f1261,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | ~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3)))))) )),
  inference(forward_demodulation,[],[f1243,f309])).
fof(f1243,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6))))))) )),
  inference(superposition,[],[f280,f309])).
fof(f1260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1242,f266])).
fof(f1242,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f280,f266])).
fof(f1230,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1214,f230])).
fof(f1214,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1209,f356])).
fof(f1130,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f864,f230])).
fof(f864,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f359,f356])).
fof(f1129,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f863,f230])).
fof(f863,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f295,f356])).
fof(f846,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,'const.bool.DATATYPE_1'(s(bool,X28))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f265,f356])).
fof(f753,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X29,bool),'const.pair.,_2'(s(X29,X30),s(bool,X31))))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f356,f308])).
fof(f752,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,'const.ind_type.FCONS_3'(s(bool,X26),s(fun('type.num.num',bool),X27),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f356,f276])).
fof(f751,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X23),s(X22,X24))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f356,f310])).
fof(f750,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f356,f309])).
fof(f749,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f356,f266])).
fof(f739,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f356,f356])).
fof(f724,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f356,f308])).
fof(f723,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f356,f276])).
fof(f722,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X29),'const.pair.,_2'(s(bool,X30),s(X29,X31)))))) )),
  inference(superposition,[],[f356,f310])).
fof(f721,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26))),s(X25,X27)))) )),
  inference(superposition,[],[f356,f309])).
fof(f720,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X23)))) )),
  inference(superposition,[],[f356,f266])).
fof(f659,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X29) != s(bool,'const.bool.DATATYPE_1'(s(X30,X31))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29)))))) )),
  inference(superposition,[],[f640,f308])).
fof(f658,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X23) != s(bool,'const.bool.DATATYPE_1'(s(X25,X26))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f640,f276])).
fof(f657,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,X19) != s(bool,'const.bool.DATATYPE_1'(s(X21,X22))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X19),s(X18,X20)))))) )),
  inference(superposition,[],[f640,f310])).
fof(f656,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X14) != s(bool,'const.bool.DATATYPE_1'(s(X16,X17))) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),'_dst_x0_1'(s(bool,X14))),s(X13,X15)))) )),
  inference(superposition,[],[f640,f309])).
fof(f655,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,'const.bool.DATATYPE_1'(s(X11,X12))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10)))) )),
  inference(superposition,[],[f640,f266])).
fof(f640,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f627,f230])).
fof(f627,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f359,f265])).
fof(f330,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X12,X43,X14,X13,X41,X47] : ! [X15,X16,X42,X17] : s(fun(X13,fun(X41,X12)),'_dst_x00x00x11_3'(s(fun(X43,fun(X47,X12)),X15),s(X43,happ(s(fun(X14,X43),X16),s(X14,X17))),s(fun(X13,fun(X41,X47)),X42))) = s(fun(X13,fun(X41,X12)),happ(s(fun(X14,fun(X13,fun(X41,X12))),'_dst_x000x100x011_3'(s(fun(X43,fun(X47,X12)),X15),s(fun(X14,X43),X16),s(fun(X13,fun(X41,X47)),X42))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x000x100x011_ax')).
fof(f329,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X12,X41,X13,X43,X14] : ! [X15,X16,X42,X17] : s(fun(X13,X12),'_dst_x0x1x1_3'(s(fun(X41,fun(X43,X12)),X15),s(fun(X13,X41),X16),s(fun(X13,X43),happ(s(fun(X14,fun(X13,X43)),X42),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x01x11_3'(s(fun(X41,fun(X43,X12)),X15),s(fun(X13,X41),X16),s(fun(X14,fun(X13,X43)),X42))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x00x01x11_ax')).
fof(f328,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12,X13,X41,X14,X43] : ! [X15,X16,X42,X17] : s(fun(X13,X12),'_dst_x0x1x0_3'(s(fun(X41,fun(X43,X12)),X15),s(fun(X13,X41),happ(s(fun(X14,fun(X13,X41)),X16),s(X14,X17))),s(X43,X42))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x11x00_3'(s(fun(X41,fun(X43,X12)),X15),s(fun(X14,fun(X13,X41)),X16),s(X43,X42))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x00x11x00_ax')).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X12,X41,X13,X43,X14] : ! [X15,X16,X42,X17] : s(fun(X13,X12),'_dst_x0x0x1_3'(s(fun(X41,fun(X43,X12)),X15),s(X41,X16),s(fun(X13,X43),happ(s(fun(X14,fun(X13,X43)),X42),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x00x11_3'(s(fun(X41,fun(X43,X12)),X15),s(X41,X16),s(fun(X14,fun(X13,X43)),X42))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x00x00x11_ax')).
fof(f326,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12,X41,X14,X13,X43] : ! [X15,X16,X42,X17] : s(fun(X13,X12),'_dst_x0x0x1_3'(s(fun(X41,fun(X43,X12)),X15),s(X41,happ(s(fun(X14,X41),X16),s(X14,X17))),s(fun(X13,X43),X42))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x10x01_3'(s(fun(X41,fun(X43,X12)),X15),s(fun(X14,X41),X16),s(fun(X13,X43),X42))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x00x10x01_ax')).
fof(f325,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X12,X13,X41,X14] : ! [X15,X16,X42,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x1x1_3'(s(fun(X13,fun(X41,X12)),X15),s(fun(X14,X13),X16),s(fun(X14,X41),X42))),s(X14,X17))) = s(X12,happ(s(fun(X41,X12),happ(s(fun(X13,fun(X41,X12)),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))),s(X41,happ(s(fun(X14,X41),X42),s(X14,X17)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x0x1x1_ax')).
fof(f322,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X12,X13,X41,X14] : ! [X15,X16,X17] : s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X41,X12),X15),s(fun(X13,X41),happ(s(fun(X14,fun(X13,X41)),X16),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x11_2'(s(fun(X41,X12),X15),s(fun(X14,fun(X13,X41)),X16))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x00x11_ax')).
fof(f338,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),'_dst_x1x0_2'(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(fun('type.num.num',X0),X2))),s(X0,X1))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f321,f286])).
fof(f337,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(X1,fun(fun('type.num.num',X0),X0)),'_dst_x1x0_2'(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(X1,X5))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(backward_demodulation,[],[f321,f307])).
fof(f320,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13,X14] : ! [X15,X16,X17] : s(X12,happ(s(fun(X13,X12),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))) = s(X12,happ(s(fun(X14,X12),'_dst_x0x1_2'(s(fun(X13,X12),X15),s(fun(X14,X13),X16))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dst_x0x1_ax')).
fof(f319,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X12,X13,X14] : ! [X16,X17] : s(fun(fun(X13,X12),X12),'_dstsel_0_x0_1'(s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))) = s(fun(fun(X13,X12),X12),happ(s(fun(X14,fun(fun(X13,X12),X12)),'_dstsel_1_x10_1'(s(fun(X14,X13),X16))),s(X14,X17)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','_dstsel_1_x10_ax')).
fof(f314,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X2) = s(X1,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0'),s(X1,X2),s(fun('type.ImplicitQueue.times'(X0),X1),X3),s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),X1)),X4)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X24] : ! [X35,X26,X36] : s(X24,X35) = s(X24,'const.ImplicitQueue.digit_CASE_4'(s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0'),s(X24,X35),s(fun('type.ImplicitQueue.times'(X6),X24),X26),s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),X24)),X36)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_case_def|split|0')).
fof(f313,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.ImplicitQueue.digit_distinct|split|2',aHL_EXT)).
fof(f307,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK11(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X24,X6] : ! [X28] : ? [X26] : ! [X29,X30,X31] : s(X24,happ(s(fun('type.ind_type.recspace'(X6),X24),X26),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X29),s(X6,X30),s(fun('type.num.num','type.ind_type.recspace'(X6)),X31))))) = s(X24,happ(s(fun(fun('type.num.num',X24),X24),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X24),X24)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X24),X24))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X24),X24)))),X28),s('type.num.num',X29))),s(X6,X30))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X31))),s(fun('type.num.num',X24),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X24),X26),s(fun('type.num.num','type.ind_type.recspace'(X6)),X31)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ind_type.CONSTR_REC')).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X24] : ! [X21] : s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X24),X21))) = s(X6,happ(s(fun('type.pair.prod'(X6,X24),X6),'const.pair.FST_0'),s('type.pair.prod'(X6,X24),X21)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','ARITY_0_1_const.pair.FST')).
fof(f290,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ImplicitQueue.times'(X0),X1) != s('type.ImplicitQueue.times'(X0),X3) | s('type.ImplicitQueue.times'(X0),X2) != s('type.ImplicitQueue.times'(X0),X4) | s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X3),s('type.ImplicitQueue.times'(X0),X4)))) )),
  inference(cnf_transformation,[],[f214])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X3)))))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X3)))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X3)))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X26,X8,X7] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.times'(X6),X8))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.times'(X6),X7)))))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_size_def|split|2')).
fof(f286,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X21,X22,X23] : s(X6,'const.ind_type.FCONS_3'(s(X6,X21),s(fun('type.num.num',X6),X22),s('type.num.num',X23))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X21))),s(fun('type.num.num',X6),X22))),s('type.num.num',X23)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','ARITY_0_3_const.ind_type.FCONS')).
fof(f285,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X21,X22,X23] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X21),s(X6,X22),s(fun('type.num.num','type.ind_type.recspace'(X6)),X23))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X21))),s(X6,X22))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X23)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','ARITY_1_3_const.ind_type.CONSTR')).
fof(f284,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X9,X26,X27] : s(X6,happ(s(fun('type.num.num',X6),X26),s('type.num.num',X27))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ind_type.FCONS|split|1')).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X8] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X8),s('type.ImplicitQueue.times'(X6),X7)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_distinct|split|1')).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))) | s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))) | s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X2)) & (s('type.ImplicitQueue.times'(X0),X1) != s('type.ImplicitQueue.times'(X0),X2) | s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2)))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))) <=> s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X2))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))) <=> s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),X2))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X9,X34] : (s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9))) = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X34))) <=> s('type.ImplicitQueue.times'(X6),X9) = s('type.ImplicitQueue.times'(X6),X34))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_|49||49||split|0')).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (s('type.ImplicitQueue.times'(X0),X1) != s('type.ImplicitQueue.times'(X0),X2) | s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2)))) )),
  inference(cnf_transformation,[],[f212])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0))),'const.ImplicitQueue.Two_0'),s('type.ImplicitQueue.times'(X0),X1))),s('type.ImplicitQueue.times'(X0),X2)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0))),'const.ImplicitQueue.Two_0'),s('type.ImplicitQueue.times'(X0),X1))),s('type.ImplicitQueue.times'(X0),X2)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X0),X1),s('type.ImplicitQueue.times'(X0),X2))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),happ(s(fun('type.ImplicitQueue.times'(X0),fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0))),'const.ImplicitQueue.Two_0'),s('type.ImplicitQueue.times'(X0),X1))),s('type.ImplicitQueue.times'(X0),X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X21,X22] : s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Two_2'(s('type.ImplicitQueue.times'(X6),X21),s('type.ImplicitQueue.times'(X6),X22))) = s('type.ImplicitQueue.digit'(X6),happ(s(fun('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.digit'(X6)),happ(s(fun('type.ImplicitQueue.times'(X6),fun('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.digit'(X6))),'const.ImplicitQueue.Two_0'),s('type.ImplicitQueue.times'(X6),X21))),s('type.ImplicitQueue.times'(X6),X22)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','ARITY_0_2_const.ImplicitQueue.Two')).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.times'(X0),X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X26,X9] : s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.ImplicitQueue.times_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.times'(X6),X9)))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_size_def|split|1')).
fof(f275,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') != s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : ~s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X9] : ~s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0') = s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X9)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_distinct|split|0')).
fof(f270,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Twice_2'(s('type.ImplicitQueue.times'(X0),sK5(X1,X0)),s('type.ImplicitQueue.times'(X0),sK6(X1,X0)))) | s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,sK4(X1,X0))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : (s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,sK4(X1,X0)))) | s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Twice_2'(s('type.ImplicitQueue.times'(X0),sK5(X1,X0)),s('type.ImplicitQueue.times'(X0),sK6(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f101])).
fof(f101,plain,(
  ! [X0,X1] : (? [X4] : s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X4))) | ? [X2,X3] : s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Twice_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X3))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1] : (? [X4] : s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Once_1'(s(X0,X4))) | ? [X2,X3] : s('type.ImplicitQueue.times'(X0),X1) = s('type.ImplicitQueue.times'(X0),'const.ImplicitQueue.Twice_2'(s('type.ImplicitQueue.times'(X0),X2),s('type.ImplicitQueue.times'(X0),X3))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X39] : (? [X9] : s('type.ImplicitQueue.times'(X6),X39) = s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Once_1'(s(X6,X9))) | ? [X19,X40] : s('type.ImplicitQueue.times'(X6),X39) = s('type.ImplicitQueue.times'(X6),'const.ImplicitQueue.Twice_2'(s('type.ImplicitQueue.times'(X6),X19),s('type.ImplicitQueue.times'(X6),X40))))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.times_nchotomy')).
fof(f268,plain,(
  ( ! [X0,X1] : (s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X0),X1))) = s('type.ImplicitQueue.digit'(X0),happ(s(fun('type.ImplicitQueue.times'(X0),'type.ImplicitQueue.digit'(X0)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X0),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X21] : s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.One_1'(s('type.ImplicitQueue.times'(X6),X21))) = s('type.ImplicitQueue.digit'(X6),happ(s(fun('type.ImplicitQueue.times'(X6),'type.ImplicitQueue.digit'(X6)),'const.ImplicitQueue.One_0'),s('type.ImplicitQueue.times'(X6),X21)))),
  file('thm.ImplicitQueue.digit_distinct|split|2','ARITY_0_1_const.ImplicitQueue.One')).
fof(f267,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ImplicitQueue.digit'(X0),'const.ImplicitQueue.Zero_0')))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X26] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.ImplicitQueue.digit_size_2'(s(fun(X6,'type.num.num'),X26),s('type.ImplicitQueue.digit'(X6),'const.ImplicitQueue.Zero_0')))),
  file('thm.ImplicitQueue.digit_distinct|split|2','thm.ImplicitQueue.digit_size_def|split|0')).
fof(f11770,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760,f11762,f11315,f11271,f11265])).
fof(f11769,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11266,f230])).
fof(f11266,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f295])).
fof(f11768,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11268,f11256])).
fof(f11268,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f295])).
fof(f11767,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760,f11762,f11315,f11271])).
fof(f11766,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11272,f230])).
fof(f11272,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f295])).
fof(f11765,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11274,f11256])).
fof(f11274,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f295])).
fof(f11764,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760,f11762,f11315])).
fof(f11763,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760,f11762])).
fof(f11761,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753,f11760])).
fof(f11759,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11662,f230])).
fof(f11662,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f6994])).
fof(f11758,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11661,f230])).
fof(f11661,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f6995])).
fof(f11757,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11660,f230])).
fof(f11660,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f6891])).
fof(f11756,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11659,f230])).
fof(f11659,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f6890])).
fof(f11755,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11655,f230])).
fof(f11655,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f7210])).
fof(f11754,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751,f11753])).
fof(f11752,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749,f11751])).
fof(f11750,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744,f11749])).
fof(f11748,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11541,f11256])).
fof(f11541,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11319,f8680])).
fof(f11747,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11535,f11256])).
fof(f11535,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11319,f8076])).
fof(f11746,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11508,f11319])).
fof(f11508,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl20),
  inference(backward_demodulation,[],[f11319,f7218])).
fof(f11745,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740,f11744])).
fof(f11741,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f229,f230,f237,f241,f240,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f415,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1230,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1590,f1697,f1702,f1705,f1804,f1807,f2151,f2274,f2337,f2422,f2888,f2903,f3065,f3066,f3123,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f1868,f2965,f2746,f3291,f5367,f5366,f5350,f5497,f5496,f5498,f5544,f5592,f5591,f5575,f5732,f5731,f6188,f6348,f6331,f6330,f6538,f6539,f6584,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6741,f6889,f7052,f7074,f7073,f7400,f7434,f7528,f7561,f7560,f7553,f7542,f7615,f7614,f7707,f7692,f8730,f8733,f336,f1401,f1353,f1351,f3145,f3122,f3120,f3111,f3109,f3106,f3105,f3126,f1422,f1259,f3077,f1386,f6183,f6184,f6607,f6606,f3211,f3207,f3139,f3081,f3080,f3079,f3078,f3076,f3075,f688,f687,f686,f685,f684,f279,f7586,f7585,f7431,f7030,f7506,f7587,f7002,f7478,f7464,f8303,f8237,f6221,f6217,f6117,f6116,f7454,f7403,f6890,f6891,f6995,f6994,f7711,f7693,f7694,f7210,f3842,f3843,f7541,f6740,f6842,f6841,f5671,f5670,f5446,f5445,f3340,f2792,f8765,f6843,f5672,f5447,f1902,f3321,f3113,f3335,f3209,f6948,f6968,f7352,f7372,f8721,f8740,f3208,f3112,f7364,f7353,f7321,f7310,f6917,f6906,f8722,f3322,f3114,f3060,f3044,f305,f304,f303,f8306,f8243,f8177,f8160,f8118,f2775,f1891,f302,f301,f300,f6905,f6925,f7309,f7329,f8739,f7667,f7666,f7899,f3431,f7471,f7916,f7697,f7919,f7700,f7437,f7350,f7307,f6946,f6903,f318,f317,f7921,f7920,f7918,f7917,f7915,f7914,f298,f7924,f7664,f7702,f7701,f7699,f7698,f7696,f7695,f297,f7712,f3862,f3791,f3449,f3285,f3220,f3206,f299,f296,f8735,f8724,f8680,f8308,f8250,f8241,f8184,f8175,f8076,f7926,f7714,f7703,f7674,f7672,f7665,f7653,f7595,f7535,f7509,f7481,f7368,f7366,f7360,f7358,f7351,f7325,f7323,f7317,f7315,f7308,f7262,f7257,f7218,f7217,f7207,f7054,f7033,f7005,f6954,f6947,f6921,f6919,f6913,f6911,f6904,f6892,f6886,f6597,f6596,f6595,f6546,f6544,f6542,f6537,f6434,f6433,f6432,f6347,f6203,f6174,f6173,f6172,f6159,f6124,f6122,f6115,f5856,f5713,f5478,f3864,f3451,f3328,f3324,f3287,f3136,f3135,f3134,f3129,f3124,f3121,f3108,f3107,f3073,f3071,f3069,f3064,f3063,f3047,f2938,f2937,f2936,f2935,f2934,f2933,f2867,f2866,f2865,f2864,f2863,f2862,f2839,f2838,f2837,f2836,f2835,f2834,f2781,f2777,f2731,f2730,f2719,f2336,f2182,f2181,f2180,f2179,f2178,f2122,f2121,f2120,f2119,f2118,f2117,f1893,f1890,f1808,f1684,f1566,f1565,f1564,f1562,f1561,f1545,f1424,f1380,f1355,f1236,f1235,f1210,f1209,f1207,f1206,f1205,f1204,f585,f535,f489,f447,f389,f334,f245,f244,f9625,f9624,f9623,f9622,f9621,f9620,f9619,f9618,f9616,f9650,f11213,f11215,f11218,f11219,f11256,f11319,f11740])).
fof(f11737,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11494,f11256])).
fof(f11494,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11319,f6892])).
fof(f11734,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11373,f230])).
fof(f11373,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11319,f334])).
fof(f11733,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11732,f11256])).
fof(f11732,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f11371,f11319])).
fof(f11371,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f11319,f244])).
fof(f11320,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f341,f341,f11256,f1209])).
fof(f11316,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f266,f341,f11256,f359])).
fof(f11275,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f11256,f341,f11256,f295])).
fof(f11273,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f230,f341,f11256,f295])).
fof(f11269,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f11256,f341,f11256,f295])).
fof(f11267,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f230,f341,f11256,f295])).
fof(f11263,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f266,f11256,f244])).
fof(f11261,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f341,f11256,f241])).
fof(f11259,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f334,f11256])).
fof(f9229,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213,f9215,f9217,f9226,f9228])).
fof(f9228,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl24),
  inference(subsumption_resolution,[],[f9077,f230])).
fof(f9077,plain,(
  ( ! [X24,X23,X21,X25] : (~p(s(bool,'T_0')) | s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f7030])).
fof(f9226,plain,(
  ( ! [X24,X23,X21,X25] : (s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl24),
  inference(subsumption_resolution,[],[f9074,f230])).
fof(f9074,plain,(
  ( ! [X24,X23,X21,X25] : (~p(s(bool,'T_0')) | s(X21,X25) = s(X21,happ(s(fun('type.num.num',X21),X23),s('type.num.num',X24)))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f7002])).
fof(f9217,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9106,f230])).
fof(f9106,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(condensation,[],[f9038])).
fof(f9038,plain,(
  ( ! [X80,X81] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f8765])).
fof(f9215,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl24),
  inference(subsumption_resolution,[],[f9027,f230])).
fof(f9027,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f6968])).
fof(f9213,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X58)) | s(X56,X59) = s(X56,sK7(X59,X57,X56))) ) | $spl24),
  inference(subsumption_resolution,[],[f9023,f230])).
fof(f9023,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | s(X56,X59) = s(X56,sK7(X59,X57,X56))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f8740])).
fof(f9211,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9121,f230])).
fof(f9121,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(condensation,[],[f9008])).
fof(f9008,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f303])).
fof(f9209,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9128,f230])).
fof(f9128,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f8998])).
fof(f8998,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f300])).
fof(f9207,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl24),
  inference(subsumption_resolution,[],[f8996,f230])).
fof(f8996,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK7(X100,X98,X97))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f6925])).
fof(f9205,plain,(
  ( ! [X55] : (p(s(bool,X55))) ) | $spl24),
  inference(subsumption_resolution,[],[f8993,f230])).
fof(f8993,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f8739])).
fof(f9199,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f8976,f230])).
fof(f8976,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f317])).
fof(f9177,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9176,f230])).
fof(f9176,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(forward_demodulation,[],[f8945,f5591])).
fof(f8945,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f8241])).
fof(f9174,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9173,f230])).
fof(f9173,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(forward_demodulation,[],[f8943,f5366])).
fof(f8943,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f8175])).
fof(f9155,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f9154,f230])).
fof(f9154,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f9153])).
fof(f9153,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(condensation,[],[f9152])).
fof(f9152,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(forward_demodulation,[],[f8907,f8597])).
fof(f8907,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'T_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f6954])).
fof(f9150,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f9149,f230])).
fof(f9149,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f9148])).
fof(f9148,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(condensation,[],[f9147])).
fof(f9147,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(forward_demodulation,[],[f8902,f8597])).
fof(f8902,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f6911])).
fof(f9135,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f8773,f8597])).
fof(f8773,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8597,f244])).
fof(f9227,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213,f9215,f9217,f9226])).
fof(f9225,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9064,f230])).
fof(f9064,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f6890])).
fof(f9224,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9063,f230])).
fof(f9063,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f6891])).
fof(f9223,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9062,f230])).
fof(f9062,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f6995])).
fof(f9222,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9061,f230])).
fof(f9061,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f6994])).
fof(f9221,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9057,f230])).
fof(f9057,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f7210])).
fof(f9220,plain,(
  $false | ($spl16 | $spl24)),
  inference(global_subsumption,[],[f9219,f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213,f9215,f9217])).
fof(f9219,plain,(
  ( ! [X53] : (p(s(bool,X53))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9048,f230])).
fof(f9048,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | p(s(bool,X53))) ) | ($spl16 | $spl24)),
  inference(backward_demodulation,[],[f8597,f8737])).
fof(f8737,plain,(
  ( ! [X53] : (~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(bool,bool),bool),'?_1'),s(bool,X53))),s(fun(bool,bool),'_i_0_0')))) | p(s(bool,X53))) ) | $spl16),
  inference(forward_demodulation,[],[f8736,f324])).
fof(f8736,plain,(
  ( ! [X53] : (p(s(bool,X53)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X53))))) ) | $spl16),
  inference(subsumption_resolution,[],[f8700,f230])).
fof(f8700,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X53))))) ) | $spl16),
  inference(superposition,[],[f303,f4197])).
fof(f9218,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213,f9215,f9217])).
fof(f9216,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213,f9215])).
fof(f9214,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211,f9213])).
fof(f9212,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209,f9211])).
fof(f9210,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207,f9209])).
fof(f9208,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205,f9207])).
fof(f9206,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199,f9205])).
fof(f9200,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177,f9199])).
fof(f9180,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9179,f230])).
fof(f9179,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8948,f6758])).
fof(f8948,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8597,f8680])).
fof(f9178,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174,f9177])).
fof(f9175,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155,f9174])).
fof(f9172,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9171,f230])).
fof(f9171,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8942,f5366])).
fof(f8942,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8597,f8076])).
fof(f9170,plain,(
  $false | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f8938,f8597])).
fof(f8938,plain,(
  s(bool,'T_0') != s(bool,sK9('_i_0_0',bool)) | ($spl19 | $spl24)),
  inference(backward_demodulation,[],[f8597,f7708])).
fof(f7708,plain,(
  s(bool,'F_0') != s(bool,sK9('_i_0_0',bool)) | $spl19),
  inference(cnf_transformation,[],[f7708_D])).
fof(f7708_D,plain,(
  s(bool,'F_0') != s(bool,sK9('_i_0_0',bool)) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f9157,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f8914,f8597])).
fof(f8914,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl24),
  inference(backward_demodulation,[],[f8597,f7218])).
fof(f9156,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150,f9155])).
fof(f9151,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135,f9150])).
fof(f9146,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9145,f230])).
fof(f9145,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8900,f8597])).
fof(f8900,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8597,f6892])).
fof(f9137,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f8775,f230])).
fof(f8775,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8597,f334])).
fof(f9136,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f295,f292,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1366,f1403,f1404,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f2117,f2118,f2120,f2121,f2122,f2151,f2178,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2834,f2835,f2837,f2838,f2839,f2888,f2862,f2863,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f3064,f3065,f3066,f3069,f3071,f3073,f3121,f3123,f3134,f3135,f3136,f3138,f3143,f3210,f3268,f3365,f3367,f3364,f3858,f3841,f296,f299,f3206,f3220,f3285,f3287,f3449,f3451,f3791,f3862,f3864,f297,f298,f317,f318,f3431,f1868,f2965,f300,f301,f302,f1890,f1891,f2775,f2777,f303,f304,f305,f3044,f3047,f3060,f3063,f3114,f3322,f3324,f3112,f3208,f2746,f3291,f3209,f3335,f3113,f3321,f1902,f1893,f2792,f2781,f3340,f3328,f3842,f3843,f279,f684,f685,f686,f687,f688,f3075,f3076,f3078,f3079,f3080,f3081,f3139,f3207,f3211,f3077,f1386,f1259,f1422,f1424,f3129,f3126,f3105,f3106,f3107,f3108,f3109,f3111,f3120,f3122,f3124,f3145,f1351,f1353,f1355,f1380,f1401,f1684,f2119,f2179,f2836,f2864,f2935,f336,f5367,f5366,f5445,f5446,f5447,f5350,f5478,f5497,f5496,f5498,f5544,f5592,f5591,f5670,f5671,f5672,f5575,f5713,f5732,f5731,f5856,f6115,f6116,f6117,f6203,f6122,f6124,f6217,f6159,f6221,f6172,f6173,f6174,f6183,f6184,f6188,f6348,f6347,f6432,f6433,f6434,f6331,f6330,f6537,f6538,f6539,f6542,f6544,f6546,f6584,f6595,f6596,f6597,f6606,f6607,f6612,f6684,f6683,f6685,f6687,f6737,f6759,f6758,f6841,f6842,f6843,f6741,f6740,f6886,f6994,f6995,f6892,f6891,f6890,f6889,f7002,f6903,f6904,f6905,f6906,f7005,f6911,f6913,f6917,f6919,f6921,f6925,f7030,f6946,f6947,f6948,f7033,f6954,f6968,f7052,f7054,f7074,f7073,f7207,f7218,f7217,f7257,f7210,f7262,f7400,f7403,f7431,f7434,f7437,f7454,f7464,f7471,f7478,f7307,f7308,f7309,f7310,f7481,f7315,f7317,f7321,f7323,f7325,f7329,f7506,f7350,f7351,f7352,f7353,f7509,f7358,f7360,f7364,f7366,f7368,f7372,f7528,f7535,f7561,f7560,f7553,f7585,f7586,f7587,f7542,f7541,f7595,f7615,f7614,f7653,f7695,f7696,f7697,f7698,f7699,f7700,f7701,f7702,f7664,f7665,f7666,f7667,f7703,f7672,f7674,f7707,f7712,f7714,f7692,f7693,f7694,f7914,f7915,f7916,f7917,f7918,f7919,f7920,f7921,f7899,f7924,f7926,f8076,f8175,f8177,f8184,f8237,f8118,f8241,f8243,f8250,f8303,f8160,f8306,f8308,f8680,f8722,f8724,f8730,f8733,f8735,f8739,f8740,f8765,f8721,f8597,f9135])).
fof(f8731,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl28),
  inference(cnf_transformation,[],[f8731_D])).
fof(f8731_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8733,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8732])).
fof(f8732,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X41)))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f8695,f359])).
fof(f8695,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X41)))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f303,f356])).
fof(f8730,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8729])).
fof(f8729,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f8694,f359])).
fof(f8694,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f303,f356])).
fof(f8676,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8652,f7708])).
fof(f8652,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f7718,f245])).
fof(f8675,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8653,f7708])).
fof(f8653,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f334,f7718,f292])).
fof(f8674,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8655,f7708])).
fof(f8655,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f334,f7718,f292])).
fof(f8672,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8662,f7708])).
fof(f8662,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f341,f7718,f415])).
fof(f8598,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,sK9('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X61)) ) | $spl26),
  inference(cnf_transformation,[],[f8598_D])).
fof(f8598_D,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,sK9('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X61)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8526,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X61)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f8395])).
fof(f8395,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X60)) ) | $spl18),
  inference(superposition,[],[f1697,f7793])).
fof(f7793,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f7725,f341])).
fof(f7725,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK9('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f7709,f356])).
fof(f7711,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl20),
  inference(cnf_transformation,[],[f7711_D])).
fof(f7711_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7718,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | $spl23),
  inference(cnf_transformation,[],[f7718_D])).
fof(f7718_D,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f7692,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f297,f266])).
fof(f7709,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) | $spl18),
  inference(cnf_transformation,[],[f7709_D])).
fof(f7709_D,plain,(
  s(bool,'F_0') = s(bool,sK9('_i_0_0',bool)) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f7707,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK9('_i_0_0',bool))),
  inference(subsumption_resolution,[],[f7678,f230])).
fof(f7678,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK9('_i_0_0',bool))),
  inference(superposition,[],[f297,f389])).
fof(f4114,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f3861,f3371])).
fof(f4113,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100,f4104,f4106,f4108,f4110,f4112])).
fof(f4112,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(subsumption_resolution,[],[f3285,f3371])).
fof(f4110,plain,(
  ( ! [X6,X4,X3] : (s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(subsumption_resolution,[],[f3449,f3371])).
fof(f4108,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(subsumption_resolution,[],[f3862,f3371])).
fof(f4106,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK7(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f1269,f3371])).
fof(f4104,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl12),
  inference(subsumption_resolution,[],[f318,f3371])).
fof(f4100,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4099])).
fof(f4099,plain,(
  ( ! [X6,X4] : (s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(subsumption_resolution,[],[f1891,f3371])).
fof(f4097,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4096])).
fof(f4096,plain,(
  ( ! [X4,X2] : (s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f2775,f3371])).
fof(f4092,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(subsumption_resolution,[],[f3044,f3371])).
fof(f4090,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3060,f3371])).
fof(f4088,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4087])).
fof(f4087,plain,(
  ( ! [X4,X2] : (s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f3322,f3371])).
fof(f4085,plain,(
  ( ! [X39,X42,X40] : (s(X39,X42) = s(X39,sK7(X42,X40,X39))) ) | $spl12),
  inference(subsumption_resolution,[],[f3208,f3371])).
fof(f4081,plain,(
  ( ! [X43,X46,X44] : (s(X43,X46) = s(X43,sK7(X46,X44,X43))) ) | $spl12),
  inference(subsumption_resolution,[],[f3209,f3371])).
fof(f4079,plain,(
  ( ! [X54,X56,X53] : (s(X53,X56) = s(X53,sK7(X56,X54,X53))) ) | $spl12),
  inference(subsumption_resolution,[],[f3335,f3371])).
fof(f4077,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK7(X2,X1,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f279,f3371])).
fof(f4075,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f684,f3371])).
fof(f4073,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl12),
  inference(subsumption_resolution,[],[f685,f3371])).
fof(f4071,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X12) = s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))) ) | $spl12),
  inference(subsumption_resolution,[],[f686,f3371])).
fof(f4069,plain,(
  ( ! [X14,X15,X13,X16] : (s(X13,X16) = s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))) ) | $spl12),
  inference(subsumption_resolution,[],[f687,f3371])).
fof(f4067,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(X17,X21) = s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl12),
  inference(subsumption_resolution,[],[f688,f3371])).
fof(f4057,plain,(
  ( ! [X92,X95,X93] : (s(X92,X95) = s(X92,sK7(X95,X93,X92))) ) | $spl12),
  inference(subsumption_resolution,[],[f3211,f3371])).
fof(f4054,plain,(
  ( ! [X101,X99,X102,X100] : (s(X99,X102) = s(X99,sK7(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl12),
  inference(subsumption_resolution,[],[f1386,f3371])).
fof(f4052,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4051])).
fof(f4051,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(subsumption_resolution,[],[f1422,f3371])).
fof(f4050,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) ) | $spl12),
  inference(subsumption_resolution,[],[f4049,f3371])).
fof(f4049,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3105,f4048])).
fof(f4048,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100))) | s(bool,X101) = s(bool,sK7(X99,X100,bool))) ) | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047])).
fof(f4047,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) = s(bool,sK7(X99,X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3122,f3371])).
fof(f4042,plain,(
  ( ! [X68] : (s(bool,'F_0') = s(bool,X68)) ) | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935])).
fof(f4111,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100,f4104,f4106,f4108,f4110])).
fof(f4109,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100,f4104,f4106,f4108])).
fof(f4107,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100,f4104,f4106])).
fof(f4105,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100,f4104])).
fof(f4103,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1868,f3371])).
fof(f4102,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2965,f3371])).
fof(f4101,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097,f4100])).
fof(f4098,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092,f4097])).
fof(f4095,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4094,f3371])).
fof(f4094,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f304,f3371])).
fof(f4093,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090,f4092])).
fof(f4091,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088,f4090])).
fof(f4089,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085,f4088])).
fof(f4086,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081,f4085])).
fof(f4084,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2746,f3371])).
fof(f4083,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3291,f3371])).
fof(f4082,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079,f4081])).
fof(f4080,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077,f4079])).
fof(f4078,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075,f4077])).
fof(f4076,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073,f4075])).
fof(f4074,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071,f4073])).
fof(f4072,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069,f4071])).
fof(f4070,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067,f4069])).
fof(f4068,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057,f4067])).
fof(f4058,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054,f4057])).
fof(f4055,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052,f4054])).
fof(f4053,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f229,f230,f334,f237,f241,f240,f245,f244,f265,f266,f267,f268,f270,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f336,f295,f292,f297,f296,f299,f298,f302,f301,f300,f305,f304,f303,f306,f307,f308,f309,f310,f312,f313,f314,f315,f316,f318,f317,f319,f320,f321,f337,f338,f322,f323,f324,f325,f326,f327,f328,f329,f330,f341,f356,f359,f389,f415,f447,f489,f535,f585,f640,f655,f656,f657,f658,f659,f684,f685,f686,f687,f688,f690,f720,f721,f722,f723,f724,f739,f749,f750,f751,f752,f753,f846,f1129,f1130,f1204,f1205,f1206,f1207,f1209,f1210,f1230,f1235,f1236,f1260,f1261,f1262,f1263,f1264,f1247,f1249,f1251,f1265,f1269,f1258,f1259,f1351,f1353,f1355,f1366,f1380,f1401,f1403,f1386,f1404,f1422,f1424,f1448,f1447,f1478,f1479,f1480,f1481,f1482,f1483,f1493,f1544,f1545,f1561,f1562,f1684,f1564,f1565,f1566,f1590,f1697,f1702,f1705,f1804,f1807,f1808,f1868,f1890,f1891,f1893,f1902,f2117,f2118,f2119,f2120,f2121,f2122,f2151,f2178,f2179,f2180,f2181,f2182,f2274,f2336,f2337,f2719,f2730,f2731,f2422,f2746,f2775,f2777,f2781,f2792,f2834,f2835,f2836,f2837,f2838,f2839,f2888,f2862,f2863,f2864,f2865,f2866,f2867,f2903,f2933,f2934,f2936,f2937,f2938,f2965,f3044,f3047,f3060,f3063,f3064,f3065,f3066,f3069,f3071,f3073,f3075,f3076,f3077,f3078,f3079,f3080,f3081,f3206,f3105,f3106,f3107,f3108,f3109,f3111,f3112,f3113,f3114,f3120,f3121,f3122,f3123,f3124,f3126,f3129,f3134,f3135,f3136,f3138,f3139,f3143,f3145,f3207,f3208,f3209,f3210,f3211,f3220,f3285,f3287,f3268,f3291,f3322,f3324,f3328,f3335,f3340,f3321,f3365,f3367,f3364,f3431,f3449,f3451,f3791,f3858,f3862,f3864,f3841,f3842,f3843,f3371,f2935,f4042,f4047,f4048,f4050,f4052])).
fof(f4041,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f336,f3371])).
fof(f4040,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3972,f334])).
fof(f3972,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3859,f336])).
fof(f4039,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4038,f3992])).
fof(f3992,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f3887,f334])).
fof(f3887,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f318])).
fof(f4038,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3975,f334])).
fof(f3975,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3949])).
fof(f3949,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f1422])).
fof(f4037,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4036,f3992])).
fof(f4036,plain,(
  ( ! [X101,X99,X102,X100] : (s(X99,X102) = s(X99,sK7(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl14),
  inference(subsumption_resolution,[],[f3947,f334])).
fof(f3947,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'F_0')) | s(X99,X102) = s(X99,sK7(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f1386])).
fof(f4035,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4034,f3992])).
fof(f4034,plain,(
  ( ! [X92,X95,X93] : (s(X92,X95) = s(X92,sK7(X95,X93,X92))) ) | $spl14),
  inference(subsumption_resolution,[],[f3945,f3998])).
fof(f3998,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3899,f334])).
fof(f3899,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f301])).
fof(f3945,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'F_0')))))) | s(X92,X95) = s(X92,sK7(X95,X93,X92))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3211])).
fof(f4033,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4032,f3992])).
fof(f4032,plain,(
  ( ! [X154,X152,X151,X153] : (s(X151,X154) = s(X151,sK7(X154,sK7(X152,X153,fun(X151,bool)),X151))) ) | $spl14),
  inference(subsumption_resolution,[],[f4031,f3992])).
fof(f4031,plain,(
  ( ! [X154,X152,X151,X153] : (s(X151,X154) = s(X151,sK7(X154,sK7(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) ) | $spl14),
  inference(subsumption_resolution,[],[f3943,f334])).
fof(f3943,plain,(
  ( ! [X154,X152,X151,X153] : (p(s(bool,'F_0')) | s(X151,X154) = s(X151,sK7(X154,sK7(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3139])).
fof(f4030,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4029,f3992])).
fof(f4029,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(X17,X21) = s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl14),
  inference(subsumption_resolution,[],[f3936,f334])).
fof(f3936,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'F_0')) | s(X17,X21) = s(X17,sK7(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f688])).
fof(f4028,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4027,f3992])).
fof(f4027,plain,(
  ( ! [X14,X15,X13,X16] : (s(X13,X16) = s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))) ) | $spl14),
  inference(subsumption_resolution,[],[f3935,f334])).
fof(f3935,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'F_0')) | s(X13,X16) = s(X13,sK7(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f687])).
fof(f4026,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4025,f3992])).
fof(f4025,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X12) = s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))) ) | $spl14),
  inference(subsumption_resolution,[],[f3934,f334])).
fof(f3934,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'F_0')) | s(X8,X12) = s(X8,sK7(X12,'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),'const.pair.,_2'(s(fun(X8,bool),X10),s(X9,X11)))),X8))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f686])).
fof(f4024,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4023,f3992])).
fof(f4023,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl14),
  inference(subsumption_resolution,[],[f3933,f334])).
fof(f3933,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'F_0')) | s(X3,X7) = s(X3,sK7(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f685])).
fof(f4022,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4021,f3992])).
fof(f4021,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3932,f334])).
fof(f3932,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK7(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f684])).
fof(f4020,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4019,f3992])).
fof(f4019,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK7(X2,X1,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3931,f334])).
fof(f3931,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK7(X2,X1,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f279])).
fof(f4018,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4017,f3992])).
fof(f4017,plain,(
  ( ! [X54,X56,X53] : (s(X53,X56) = s(X53,sK7(X56,X54,X53))) ) | $spl14),
  inference(subsumption_resolution,[],[f3920,f334])).
fof(f3920,plain,(
  ( ! [X54,X56,X53] : (p(s(bool,'F_0')) | s(X53,X56) = s(X53,sK7(X56,X54,X53))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3335])).
fof(f4016,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4015,f3992])).
fof(f4015,plain,(
  ( ! [X43,X46,X44] : (s(X43,X46) = s(X43,sK7(X46,X44,X43))) ) | $spl14),
  inference(subsumption_resolution,[],[f3919,f334])).
fof(f3919,plain,(
  ( ! [X43,X46,X44] : (p(s(bool,'F_0')) | s(X43,X46) = s(X43,sK7(X46,X44,X43))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3209])).
fof(f4014,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3918,f334])).
fof(f3918,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3859,f3291])).
fof(f4013,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3917,f334])).
fof(f3917,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3859,f2746])).
fof(f4012,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4011,f3992])).
fof(f4011,plain,(
  ( ! [X39,X42,X40] : (s(X39,X42) = s(X39,sK7(X42,X40,X39))) ) | $spl14),
  inference(subsumption_resolution,[],[f3916,f334])).
fof(f3916,plain,(
  ( ! [X39,X42,X40] : (p(s(bool,'F_0')) | s(X39,X42) = s(X39,sK7(X42,X40,X39))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3208])).
fof(f4010,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4009,f3992])).
fof(f4009,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3976,f334])).
fof(f3976,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3913])).
fof(f3913,plain,(
  ( ! [X4,X2] : (p(s(bool,'F_0')) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3322])).
fof(f4008,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4007,f3992])).
fof(f4007,plain,(
  ( ! [X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(subsumption_resolution,[],[f4006,f3998])).
fof(f4006,plain,(
  ( ! [X76,X77] : (p(s(bool,sK7(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(subsumption_resolution,[],[f3912,f334])).
fof(f3912,plain,(
  ( ! [X76,X77] : (p(s(bool,'F_0')) | p(s(bool,sK7(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3114])).
fof(f4005,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4004,f3992])).
fof(f4004,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3908,f334])).
fof(f3908,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f3044])).
fof(f4003,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3979,f334])).
fof(f3979,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(condensation,[],[f3906])).
fof(f3906,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f304])).
fof(f4002,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4001,f3992])).
fof(f4001,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3981,f334])).
fof(f3981,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3903])).
fof(f3903,plain,(
  ( ! [X4,X2] : (p(s(bool,'F_0')) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f2775])).
fof(f4000,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3999,f3992])).
fof(f3999,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3982,f334])).
fof(f3982,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3902])).
fof(f3902,plain,(
  ( ! [X6,X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl14),
  inference(backward_demodulation,[],[f3859,f1891])).
fof(f3997,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3897,f334])).
fof(f3897,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3859,f2965])).
fof(f3996,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3896,f334])).
fof(f3896,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3859,f1868])).
fof(f3861,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl16),
  inference(cnf_transformation,[],[f3861_D])).
fof(f3861_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3841,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f299,f266])).
fof(f3858,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3828,f230])).
fof(f3828,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f299,f389])).
fof(f3370,plain,(
  ( ! [X32] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) ) | $spl10),
  inference(cnf_transformation,[],[f3370_D])).
fof(f3370_D,plain,(
  ( ! [X32] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3364,plain,(
  ( ! [X33,X32] : (~p(s(bool,X33)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f3268,f310])).
fof(f1805,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) ) | $spl6),
  inference(cnf_transformation,[],[f1805_D])).
fof(f1805_D,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1807,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15)))))) )),
  inference(subsumption_resolution,[],[f1783,f240])).
fof(f1783,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15)))))) )),
  inference(superposition,[],[f1424,f356])).
fof(f1804,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) )),
  inference(subsumption_resolution,[],[f1782,f240])).
fof(f1782,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1424,f356])).
fof(f1270,plain,(
  ( ! [X20] : (~p(s(bool,sK7(X20,'_i_0_0',bool)))) ) | $spl4),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  ( ! [X20] : (~p(s(bool,sK7(X20,'_i_0_0',bool)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1258,plain,(
  ( ! [X20] : (~p(s(bool,sK7(X20,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f280,f266])).
fof(f1266,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK7(X0,'_i_0_0',bool))) ) | $spl0),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK7(X0,'_i_0_0',bool))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1265,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1252,f230])).
fof(f1252,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK7(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f280,f389])).
% SZS output end Proof for thm.ImplicitQueue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 276
% Passive clauses: 3456
% Generated clauses: 45510
% Final active clauses: 153
% Final passive clauses: 2206
% Input formulas: 68
% Initial clauses: 102
% 
% Duplicate literals: 1097
% Trivial inequalities: 203
% Fw subsumption resolutions: 2102
% Fw demodulations: 8751
% Bw demodulations: 944
% Condensations: 2032
% Global subsumptions: 52
% 
% Simple tautologies: 258
% Equational tautologies: 282
% Forward subsumptions: 26092
% Fw demodulations to eq. taut.: 832
% Bw demodulations to eq. taut.: 43
% 
% Unit resulting resolution: 4130
% Factoring: 1
% Forward superposition: 13607
% Backward superposition: 11036
% Self superposition: 271
% Equality factoring: 96
% Equality resolution: 25
% 
% Split clauses: 77
% Split components: 154
% SAT solver clauses: 3940
% SAT solver unit clauses: 966
% SAT solver binary clauses: 1677
% 
% Sat splits: 77
% Sat splitting refutations: 276
% 
% Memory used [KB]: 48869
% Time elapsed: 2.225 s
% ------------------------------
% ------------------------------
% Success in time 3.58 s
